---
id: 646c47867800472a4ed5d2ea
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Počnite sa osnovnom strukturom HTML. Dodajte `DOCTYPE` referencu na `html` i polje `html` sa atributom polja `lang` postavljeno na `en`. Takođe, dodajte polje `head` i `body` unutar polja `html`.

# --hints--

Trebalo bi da imaš deklaraciju `DOCTYPE` od `html`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Tvoje saopštenje za `DOCTYPE` trebalo bi biti na početku tvog HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Trebalo bi da imate etiketu pakovanja `html` sa karakteristikom elementa `lang` od `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Trebalo bi da imaš oznaku za zatvaranje `html`.

```js
assert(code.match(/<\/html>/i));
```

Trebalo bi da imaš etiketu pakovanja za `head`.

```js
assert(code.match(/<head>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `head`.

```js
assert(code.match(/<\/head>/i));
```

Trebalo bi da imaš etiketu pakovanja za `body`.

```js
assert(code.match(/<body>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `body`.

```js
assert(code.match(/<\/body>/i));
```

Tvoj element `body` mora da dođe nakon elementa `head`.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

Vaše komponente za `head` i `body` moraju biti unutar elementa `html`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
