---
id: 646cf2249f02ca5233d9af7c
title: Korak 28
challengeType: 0
dashedName: step-28
---

# --description--

Sada možeš da radiš na desnom uvu mačke. Već imaš njegovo HTML.

Kada koristite selektor klase elemenata, dajte elementu `.cat-right-ear` levi i desni granice od `35px solid transparent` za oba. Takođe, postavite donju granicu na `70px solid #5e5e5e`.

# --hints--

Trebalo bi da imaš izabličac za `.cat-right-ear`.

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear'))
```

Vaš selektor za `.cat-right-ear` mora imati atribut `border-left` postavljen na `35px solid transparent`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid transparent')
```

Tvoj `.cat-right-ear` mora imati svojstvo `border-right` postavljeno na `35px solid transparent`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid transparent')
```

Vaš selektor za `.cat-right-ear` mora imati svojstvo `border-bottom` postavljeno na `70px solid #5e5e5e`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

--fcc-editable-region--

--fcc-editable-region--
```
