---
id: 646c5e727e56e743c9aed4a1
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Administrator za CSS dozvoljava da postavite kako želite da se element prikaže u pretraživaču. Ima svojstvo `position` koje možete postaviti na `static`, `absolute`, `relative`, `sticky` ili `fixed`.

Čim postavite svojstvo `position` elementa, možete pomeriti element postavljanjem vrednosti u pikselima ili procentima za jedno ili više svojstava `top`, `right`, `left`, ili `bottom`.

`static` je podrazumevana postavka za sve komponente. Ako ovo postavite za komponentu, nećete moći da pomerite tu komponentu koristeći `top`, `right`, `left` ili `bottom`.

Postavi `.cat-head` svojstvo `position` sa vrednošću `static`, zatim postavi svojstva za `top` i `left` na `100px` za svako.


# --hints--

Vaš selektor za `.cat-head` mora imati svojstvo `position` postavljeno na `static`. Uverite se da dodajete decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

Vaš selektor za `.cat-head` mora imati svojstvo `top` postavljeno na `100px`. Uverite se da dodajete decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Vaš selektor za `.cat-head` mora imati svojstvo `left` postavljeno na `100px`. Uverite se da dodate decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {


  
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
