---
id: 646c5ffef5598d449b52ec12
title: Korak 19
challengeType: 0
dashedName: step-19
---

# --description--

Sada treba da radiš na uši mačke. Biće desno uvo i levo uvo, a unutra svake će biti unutarnje uvo.

Unutar tvog elementa `.cat-head`, kreiraj element `div` koji ima klasu objekta `cat-ears`.

# --hints--

Ne smete da menjate postojeće polje `div` sa klasom objekta `cat-head`.

```js
assert(document.querySelectorAll('div.cat-head').length === 1);
```

Trebaće ti da napraviš jedan deo `div` unutar tvog dela `.cat-head`.

```js
assert(document.querySelectorAll('.cat-head div').length === 1);
```

Vaš element za `div` treba da ima klasu objekata od `cat-ears`.

```js
assert(document.querySelectorAll('.cat-head div')[0]?.getAttribute('class') === 'cat-ears');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      --fcc-editable-region--
      <div class="cat-head">
        
      </div>
      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
