---
id: 646c59e942f35541923104bf
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Idite u vašu datoteku HTML, kreirajte komponentu `main`. Unutar te komponente `main`, dodajte komponentu `div` koja ima klasu za `cat-head`.

# --hints--

Trebalo bi da imaš element `main`.

```js
assert(document.querySelectorAll('main').length === 1);
```

Trebalo bi da imaš element `div`.

```js
assert(document.querySelectorAll('div').length === 1);
```

Element za `div` mora imati klasu objekata `cat-head`.

```js
assert(document.querySelector('div')?.getAttribute('class') === 'cat-head');
```

Tvoja komponenta za `div` treba da bude unutar tvog označavanja/labela `main`.

```js
assert(document.querySelectorAll('main div').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
