---
id: 646c48df8674cf2b91020ecb
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Unutar tvog elementa sa `head`, dodaj oznaku za `meta` koja ima atribut elementa `charset` sa vrednošću `utf-8`. Takođe dodaj element za `title` koji ima tekst `fCC Cat Painting`.

# --hints--

Trebalo bi da dodate jedan element `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Vaše polje `meta` treba da ima svojstvo polja `charset` postavljeno na `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Treba da dodaš jedan element `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

Vaše polje za `title` treba da ima tekst `fCC Cat Painting`. Zapamtite da su pravopis i velika slova važni.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
