---
id: 646c5d7057c45f432fcdd46c
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Koristeći selektor klase elemenata, postavi elementu `.cat-head` širinu od `205px` i visinu od `180px`. Takođe, postavi granicu za `1px solid #000` i `border-radius` od `46%`.

# --hints--

Trebalo bi da imate birač za `.cat-head`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Vaš selektor za `.cat-head` treba da ima `width` postavljen na `205px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

Vaš selektor za `.cat-head` mora imati `height`, podešen na `180px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

Vaš selektor za `.cat-head` bi trebalo da ima `border` podešen na `1px solid #000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

Vaš selektor za `.cat-head` mora imati `border-radius` podešen na `46%`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
