---
id: 646c48df8674cf2b91020ecc
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Dodajte element iz `link` unutar vašeg elementa `head`. Za taj element `link`, postavite atribut elemenata `rel` na `stylesheet` i `href` na `./styles.css`.

# --hints--

Trebalo bi da imaš element `link`.

```js
assert.match(code, /<link/)
```

Vaš element za `link` mora da ima svojstva elemenata `rel` i `href`, postavljena na `stylesheet` i `./styles.css`.

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
