---
id: 6476f7a4827bcc61682f2347
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Sledeće svojstvo pozicije je `absolute`. Kada koristite vrednost `absolute` za vaše svojstvo `position`, element se uklanja iz normalnog toka dokumenta, a zatim mu pozicija određuju svojstva `top`, `right`, `bottom` i `left`.

Postavi atribut `position` tvog elementa `.cat-head` na `absolute`, zatim postavi atribute `top` i `left` na bilo koju pozitivnu vrednost piksela.

<!-- **Napomena**: Možete pokušati sa funkcijama `top`, `left`, `bottom`, na `right` ovde, ali test će proći samo sa `top` od `300px`, i levo od `400px`. -->

# --hints--

Tvoj selektor za `.cat-head` treba da ima atribut `position` postavljen na `absolute`. Proveri da dodaš decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Tvoj selektor za `.cat-head` mora imati svojstvo `top` postavljeno na bilo koju brojku u `px`. Proveri da dodaš decimalnu tačku.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

Vaš selektor za `.cat-head` mora imati svojstvo od `left` postavljeno za bilo koji broj u `px`. Obavezno dodajte decimalnu tačku.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
