---
id: 6476fc5cf14b276e6d04e82a
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Koristite selektor klase da mu date vaš atribut `.box` širine od `200px`, visine od `600px`, i pozadine boje od `#000`. Takođe, date mu `position` od `absolute`, `top` od `800px` i `left` od `650px`.

# --hints--

Trebalo bi da imaš birač za `.box`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box'))
```

Tvoj selektor za `.box` mora imati svojstvo od `width` postavljeno na `200px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.width === '200px')
```

Tvoj selektor za `.box` mora imati svojstvo `height` postavljeno na `600px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.height === '600px')
```

Tvoj selektor za `.box` mora da ima svojstvo `background-color` postavljeno na `#000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor === 'rgb(0, 0, 0)')
```

Vaš selektor za `.box` treba da ima svojstvo od `position` postavljeno na `absolute`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.position === 'absolute')
```

Vaš selektor za `.box` mora imati svojstvo od `left` postavljeno na `650px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.left === '650px')
```

Tvoj selektor za `.box` mora imati svojstvo `top` postavljeno na `800px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.top === '800px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
