---
id: 6476fd4213318f6ee211028a
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Sada promenite vrednost atributa ``position`` za ``.cat-head`` na ``fixed``. Ostavite ``top`` i ``left`` kako su.

Nakon toga, skrolujte ekran gore i dole da biste videli kako vrednost `fixed` funkcioniše.

# --hints--

Vaš selektor za `.cat-head` mora imati svojstvo `position` postavljeno na `fixed`. Proverite da dodate decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'fixed')
```

Vaš selektor za `.cat-head` mora imati svojstvo `top` postavljeno na `100px`. Uverite se da dodajete decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Vaš selektor za `.cat-head` mora imati svojstvo `left` postavljeno na `100px`. Uverite se da dodate decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
