---
id: 6476fb7cbaafa36d65e9cf35
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

`fixed` je vrednost svojstva `position` koja vam omogućava da postavite element fiksiran na stranici, bez obzira kako korisnik skroluje.

Trebaće vam dodati više znakova da biste videli kako funkcioniše prostor `fixed`. U vašem HTML, kreirajte element sa klasom `div` koji ima klasu `box`.

# --hints--

Trebalo bi da imaš element `div`.

```js
assert(document.querySelectorAll('div').length === 2);
```

Vaš element za `div` mora imati klasu `box`.

```js
assert(document.querySelectorAll('div')[1]?.getAttribute('class') === 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
