---
id: 64770351e8586671ec0911f0
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Vrednost atributa poslednje pozicije je `sticky`. Raspored `sticky` je kombinacija rasporeda `relative` i `fixed`. Omogućava **skupljanje** elementa na određenu poziciju unutar sadržajnog ili vidljivog područja, u zavisnosti od pregledača.

Promeni vrednost atributa `position` od `.cat-head` u `sticky`, promeni `top` u `0`, zatim obriši `left` i njegovu vrednost.

**Napomena**: Da biste videli kako `sticky` funkcioniše, morate staviti malo teksta pre i posle vašeg elementa za `.cat-head` `div`. Kada skrolujete dole posle toga, videćete da `.cat-head` ostaje pri vrhu i zadržava se tamo.

# --hints--

Tvoj selektor za `.cat-head` mora imati atribut `position` postavljen na `sticky`. Proveri da dodaš decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

Tvoj klijent ``.cat-head`` mora imati svojstvo ``top`` postavljeno na ``0``. Proveri da li si dodao tačku i zarez.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

Treba da ima atribut `left` i njegovu vrednost u vašem kodu.

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
