---
id: 647d855ac12cd436059acd39
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Koristite globalni preklopnik za dodavanje `box-sizing: border-box;` u vaš CSS. Ovo osigurava da elementi uključuju unutrašnji prostor i granice po definisanoj širini i visini.

# --hints--

Trebalo bi da imaš globalni preklopnik (`*`).

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*'))
```

Vaš globalni selektor mora imati svojstvo `box-sizing` postavljeno na `border-box`.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
