---
id: 6676a8a8d9861319ee901a90
title: Korak 24
challengeType: 0
dashedName: step-24
---

# --description--

Sada možeš početi da radiš na uši tvog mačka.

Očisti svoj kod za reviziju uklanjanjem selektora `.cat-right-ear` i svih njegovih atributa.

Kada birate koristeći selektor klase objekta, dodelite elementu za `.cat-left-ear` levu i desnu granicu od `35px solid transparent` za svaki. Takođe, postavi donju granicu na `70px solid #5e5e5e`.

# --hints--


Tvoj selektor za `.cat-right-ear` ne bi trebalo da ima svojstvo `height`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
```

Vaš selektor za `.cat-right-ear` ne sme imati svojstvo od `width`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
```

Tvoj selektor za `.cat-right-ear` ne bi trebalo da ima svojstvo od `border-left`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
```

Tvoj selektor za `.cat-right-ear` ne bi trebalo da ima svojstvo od `border-right`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
```

Vaš selektor za `.cat-right-ear` ne bi trebalo da ima svojstvo `border-top`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
```

Vaš selektor za `.cat-right-ear` ne bi trebalo da ima svojstvo `border-bottom`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
```

Ne bi trebalo da imaš selektor za `.cat-right-ear`.

```js
assert(!new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Trebalo bi da imaš birač za `.cat-left-ear`.

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear'))
```

Vaš selektor za `.cat-left-ear` mora imati svojstvo od `border-left` postavljeno na `35px solid transparent`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft === '35px solid transparent')
```

Vaš selektor za `.cat-left-ear` mora imati svojstvo `border-right` postavljeno na `35px solid transparent`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight === '35px solid transparent')
```

Tvoj selektor za `.cat-left-ear` mora imati svojstvo `border-bottom` postavljeno na `70px solid #5e5e5e`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
