---
id: 6676a8b01e56ec1a1e07c202
title: Korak 23
challengeType: 0
dashedName: step-23
---

# --description--

Primeti da sada imaš beli kvadrat sa debelim ivicama. Te ivice imaju podignute uglove koje se mogu koristiti za formiranje kruga.

Unutar istog selektora klase, izmeni `height` i `width` u `0`. Postavi levu, desnu i gornju granicu na `transparent`.

Ukloni atribut `background-color` sa ozbele/okvira.

# --hints--


Trebalo bi da imaš izabličac za `.cat-right-ear`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Tvoj selektor za `.cat-right-ear` treba da ima svojstvo `height` postavljeno na `0`. Ne zaboravi da dodaš decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height == '0px')
```

Vaš selektor za `.cat-right-ear` mora imati svojstvo `width` postavljeno na `0`. Nemojte zaboraviti da dodate decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '0px')
```

Vaš selektor za `.cat-right-ear` mora imati svojstvo `border-left` postavljeno na `35px solid transparent`. Ne zaboravite da dodate decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid transparent')
```

Vaš selektor za `.cat-right-ear` treba da ima svojstvo `border-right` postavljeno na `35px solid transparent`. Nemoj zaboraviti da dodaš decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid transparent')
```

Tvoj selektor za `.cat-right-ear` mora imati svojstvo od `border-top` postavljeno na `35px solid transparent`. Ne zaboravi da dodaš decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid transparent')
```

Tvoj selektor za `.cat-right-ear` mora imati atribut `border-bottom` postavljen na `35px solid red`. Ne zaboravi da dodaš decimalnu tačku.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

Tvoj selektor za `.cat-right-ear` ne bi trebalo da ima svojstvo od `background-color`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
