---
id: 66a9cb26ec6bd41cf6c82bc5
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Poslednji deo koji treba dodati tabeli biće element `<dfn>podnožje tabele</dfn>`, ``tfoot``. Element podnožja tabele će se koristiti za prikaz prosečne ocene svih učenika u tabeli.

Dodajte element ``tfoot`` u tabelu.

Unutar elementa `tfoot``, dodajte element ``tr`.

Unutar elementa `tr``, dodajte dva elementa ``td`.

Prvi element ``td`` treba da sadrži tekst ``Average Grade``. Drugi element ``td`` treba da sadrži tekst ``78.8``.

# --hints--

Trebalo bi da imate otvarajući tag `tfoot`.

```js
assert.match(code, /<tfoot>/);
```

Trebalo bi da imate zatvarajući `tfoot` tag.

```js
assert.match(code, /<\/tfoot>/);
```

Trebalo bi da imate element `tr` unutar vašeg elementa `tfoot`.

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

Vaš prvi `td` element treba da sadrži tekst `Average Grade`. Uverite se da je unutar elementa `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

Vaš drugi `td` element treba da sadrži tekst `78.8`. Uverite se da je unutar `tr` elementa.

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

Vaš `tr` element bi trebalo da ima dva `td` elementa unutar sebe.

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
