---
id: 66a9ca099e3ecb1be10a2696
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Sada je vreme da dodamo još dva učenika u tabelu.

Prateći isti obrazac kao prethodni koraci, dodajte četvrti red za učenike u tabelu. Koristite sledeće podatke za elemente podataka tabele:

- Prezime: `Thompson`
- Ime: `Jane`
- Razred: `77`

Za peti red za učenike u tabelu, koristite sledeće podatke za elemente podataka tabele:

- Prezime: `Williams`
- Ime: `Natalie`
- Razred: `83`

# --hints--

Trebalo bi da imate četvrti `tr` unutar vašeg `tbody` elementa.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

Trebalo bi da imate element `td` sa tekstom `Thompson` unutar četvrtog `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

Trebalo bi da imate element `td` sa tekstom `Jane` unutar četvrtog `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

Trebalo bi da imate element `td` sa tekstom `77` unutar četvrtog `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

Trebalo bi da imate peti `tr` unutar vašeg `tbody` elementa.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

Trebalo bi da imate element `td` sa tekstom `Williams` unutar petog `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

Trebalo bi da imate element `td` sa tekstom `Natalie` unutar petog `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

Trebalo bi da imate element `td` sa tekstom `83` unutar petog `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
