---
id: 66a9bcf00f13a418368a272e
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Element zaglavlja tabele se sastoji od elementa <dfn>red u tabeli</dfn>, `tr`, koji sadrži elemente <dfn>zaglavlje ćelije tabele</dfn>, `th`.

Evo primera korišćenjem elemenata `tr` i `th` za sportsku tabelu:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

Unutar vašeg `thead` elementa, dodajte element `tr`.

Unutar vašeg `tr` elementa, dodajte tri `th` elementa.

Prvi `th` element mora sadržati tekst `Last Name`. Drugi `th` element mora sadržati tekst `First Name`. Treći `th` element mora sadržati tekst `Grade`.

# --hints--

Trebalo bi da imate otvarajući tag `tr`.

```js
assert.match(code, /<tr>/i);
```

Trebalo bi da imate zatvarajući tag `tr`.

```js
assert.match(code, /<\/tr>/i);
```

Vaš `tr` element treba da bude unutar vašeg `thead` elementa.

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

Vaš prvi `th` element bi trebalo da ima tekst od `Last Name`.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

Vaš drugi `th` element treba da ima tekst od `First Name`.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

Vaš treći `th` element bi trebalo da ima tekst od `Grade`.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

Svi vaši elementi `th` bi trebalo da budu unutar vašeg elementa `tr`.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
