---
id: 66a9c006e4f58819396605b3
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Sada kada ste završili zaglavni deo, vreme je da dodate <dfn>telo tabele</dfn>, `tbody`. Telo tabele će predstavljati sva imena studenata i njihove ocene.

Dodajte element tela tabele u vašu tabelu.

# --hints--

Trebalo bi da imate otvarajući tag `tbody`.

```js
assert.match(code, /<tbody>/i);
```

Trebalo bi da imate zatvarajući tag `tbody`.

```js
assert.match(code, /<\/tbody>/i);
```

Vaš otvarajući `tbody` tag mora da stoji pre zatvarnog `table` taga.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
