---
id: 66a9c7a4521d2b1b1ec6dcf0
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Sada je vreme da dodamo još dva učenika u tabelu.

Prateći isti obrazac kao i prethodni korak, dodajte drugi red za studente u tabelu. Koristite sledeće podatke za elemente podataka tabele:

- Prezime: `Doe`
- Ime: `Samantha`
- Razred: `92`

Za treći red za studente u tabelu, koristite sledeće podatke za elemente podataka tabele:

- Prezime: `Rodriguez`
- Ime: `Marcus`
- Razred: `88`

# --hints--

Trebalo bi da imate drugi `tr` unutar vašeg `tbody` elementa.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

Trebalo bi da imate `td` element sa tekstom `Doe` unutar drugog `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

Trebalo bi da imate `td` element sa tekstom `Samantha` unutar drugog `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

Trebalo bi da imate element `td` sa tekstom `92` unutar drugog `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

Trebalo bi da imate treći `tr` unutar vašeg `tbody` elementa.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

Trebalo bi da imate element `td` sa tekstom `Rodriguez` unutar trećeg `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

Trebalo bi da imate element `td` sa tekstom `Marcus` unutar trećeg `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

Trebalo bi da imate element `td` sa tekstom `88` unutar trećeg `tr` elementa.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
