---
id: 66a9bb8578459a175432e7d0
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Za prvi deo tabele, želite da grupišete sadržaj zaglavlja koji predstavlja oznake kolona za ime učenika, prezime i ocenu završnog ispita.

Element `<dfn>table head</dfn>`, ``thead``, koristi se za grupisovanje sadržaja zaglavlja u tabeli.

Evo primera korišćenjem elementa ``thead`:`

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

Ispod vašeg elementa ``caption``, dodajte element za zaglavlje tabele.

# --hints--

Trebalo bi da imate otvarajući tag `thead`.

```js
assert.match(code, /<thead>/i);
```

Trebalo bi da imate zatvarajući tag `thead`.

```js
assert.match(code, /<\/thead>/i);
```

Vaš `thead` element treba da bude unutar vašeg `table` elementa.

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
