---
id: 68ed661e9f463c155bb677a1
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici ćete vežbati korišćenje CSS flexbox za dizajniranje seta kolor kutija.

Na početku, unutar elementa `body`, dodaj komponentu `header`. Postavi komponentu `h1` unutar elementa `header` koji ima tekst `Colored Boxes Layout`.

# --hints--

Trebalo bi da imaš element `header` unutar elementa `body`.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

Tvoja komponenta za `header` mora da ima komponentu `h1` unutar sebe.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

Vaš element za `h1` treba da sadrži tekst `Colored Boxes Layout` unutar sebe.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
