---
id: 68ef0b3237fef991c92d9c19
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Sada dodaj element za `div` sa klasom objekta `flex-container` ispod tvog `header`. Zatim kreiraj šest elemenata za `div` sa klasom objekta `box` unutar `div` koji si upravo kreirao.

Zapamtite da se atribut klase elementa često koristi za označavanje imena klase elementa u stilskom listu. U ovom slučaju, svi elementi `div` će biti jednako stilizovani prema definicijama stila od `.flex-container` i `.box`.

# --hints--

Trebalo bi da imaš element `div` ispod elementa `header`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

Element vašeg `div` ispod elementa `header` mora imati klasu od `flex-container`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

Trebalo bi da ima šest elemenata od `div` postavljeno unutar elementa `.flex-container`.

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

Sve šest elemenata od `div` unutar elementa `.flex-container` mora imati klasu `box`.

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
