---
id: 68f46b758eb8ea7043abd73c
title: Korak 33
challengeType: 0
dashedName: step-33
---

# --description--

Sada ćete kreirati kutije unutar vašeg flexbox kontejnera. Prvo, kreirajte selektor za `.box1` i dodajte svojstvo `background` sa vrednošću `#f16e79`.

Zatim, dodajte svojstvo `order` sa vrednošću `1`. Videćete da se prvi okvir pomera na kraj. To se dešava zato što klasa predmeta `box` već ima vrednost `order` od `0`, i predmeti sa višim rednim vrednostima pojavljuju se kasnije.

Dodaje atribut `flex-grow` sa vrednošću `1`. Ovo će učiniti da kutija raste kako bi popunila dodatni prostor na svojoj liniji.

# --hints--

Trebalo bi da imaš birač za `.box1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box1'));
```

Tvoj selektor za `.box1` mora imati svojstvo od `background` sa vrednošću `#f16e79`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('background'), 'rgb(241, 110, 121)');
```

Vaš selektor za `.box1` mora imati atribut `order` sa vrednošću `1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('order'), '1');
```

Tvoj selektor za `.box1` mora imati svojstvo `flex-grow` sa vrednošću `1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('flex-grow'), '1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

--fcc-editable-region--

--fcc-editable-region--
```
