---
id: 68f471d46fa193b002a5bfb8
title: Korak 43
challengeType: 0
dashedName: step-43
---

# --description--

Kreiraj selektor za `.box6` i dodaj svojstvo `background` sa vrednošću `#3f51b5` unutar njega.

Zatim dodajte atribut `order` sa vrednošću `5` za vaš novi selektor. Videćete da element `.box6` se pomera nakon elementa `.box5`.

Takođe, dodaj svojstvo `flex-grow` sa vrednošću `1`. Ovo će učiniti da kutija 6 raste i popuni bilo koji preostali prostor unutar kontejnera.

Dakle, radionica/projekat bojenih kutija je završen!

# --hints--

Trebalo bi da imaš birač za `.box6`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box6'));
```

Tvoj selektor za `.box6` mora imati svojstvo `background` sa vrednošću `#3f51b5`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.getPropVal('background'), 'rgb(63, 81, 181)');
```

Vaš selektor za `.box6` mora imati svojstvo od `order` sa vrednošću od `5`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.getPropVal('order'), '5');
```

Vaš selektor za `.box6` mora imati svojstvo `flex-grow` sa vrednošću `1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.getPropVal('flex-grow'), '1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box box2">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box box3">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box box4">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box box5">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box box6">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2;
  flex-shrink: 3; 
}

.box4 {
  background: #4caf50;
  order: 3; 
}

.box5 {
  background: #457b9d;
  order: 4;
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box box2">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box box3">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box box4">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box box5">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box box6">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
  align-content: space-evenly;
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1; 
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2; 
  flex-shrink: 3; 
}

.box4 {
  background: #4caf50;
  order: 3;
}

.box5 {
  background: #457b9d;
  order: 4; 
}

.box6 {
  background: #3f51b5;
  order: 5; 
  flex-grow: 1; 
}
```
