---
id: 68f932c803f4d3d9f511b436
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Sifa ya `flex-wrap` opisuje da li varijable moraju biti prisiljene na samo jednu liniju ili mogu da se prenose preko više linija.

Po podrazumevanju, prilagodljivi elementi imaju vrednost `nowrap`, što znači da ostaju na istoj liniji čak i ako pređu kontejner. Vrednost `wrap` umesto toga uzrokuje da se prilagodljivi elementi prenose na nove linije kada ne stanu u širinu kontejnera, stvarajući fleksibilniji raspored koji se prilagođava uređaju.

Da biste nastavili, dodajte svojstvo `flex-wrap` sa vrednošću `wrap` za selektor klase `.flex-container`.

# --hints--

Vaš selektor za `.flex-container` mora imati svojstvo `flex-wrap` sa vrednošću `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('flex-wrap'), 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
