---
id: 68f91df5776e693049fb8942
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sada je vreme da dodamo neke pravila za CSS u fajl `styles.css`. Počnite kreiranjem selektora za element `h1`.

Postavi sadržaj elementa za `h1` postavljanjem njegovog atributa od `text-align` na `center`. Takođe kreiraj atribut od `margin-bottom` sa vrednošću `10px` kako bi se dodala ivica između dna elementa `h1` i bilo kog elementa HTML koji sledi ispod.

# --hints--

Trebalo bi da imaš birač za `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Tvoj `h1` treba da postavi svojstvo `text-align` na `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

Твој селектор за `h1` треба да постави атрибут `margin-bottom` на `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
