---
id: 692f2d37f34ed866ed42fcd6
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Svojstvo `flex-direction` kontroliše kako su fleksibilni delovi raspoređeni unutar njihovog kontejnera. Podrazumevana vrednost je `row`, koja raspoređuje delove ravnomerno od leve ka desnoj strani. Ovo svojstvo radi samo za elemente koji imaju `display: flex` ili `display: inline-flex`.

Za vrednosti `column`, komadi se postavljaju vertikalno odozgo nadole. Ostale moguće vrednosti su `row`, `row-reverse` i `column-reverse`.

Sada dodaj atribut za `flex-direction` sa vrednošću `column` za selektor `.box`.

# --hints--

Vaš selektor za `.box` mora imati svojstvo `flex-direction` sa vrednošću `column`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex-direction'), 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
