---
id: 692f22a2ad44ae887a24c8f9
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Unutar svakog elementa od `.box`, dodaj element sa `h2` koji sadrži tekst `Box` praćen razmakom i sekvencijalnim brojevima počevši od `1`.

# --hints--

Treba da ima element `h2` koji sadrži tekst `Box 1` unutar prvog `div` sa klasom `box`.

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

Trebalo bi da ima element od `h2` sa tekstom `Box 2` unutar drugog koji ima klasu `div`, sa klasom `box`.

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

Trebalo bi da ima element sa `h2`, koji sadrži tekst `Box 3`, unutar trećeg `div` sa klasom `box`.

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

Trebalo bi da ima element `h2` koji sadrži tekst `Box 4` unutar četvrtog `div` sa klasom `box`.

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

Trebalo bi da ima element sa `h2` koji sadrži tekst `Box 5` unutar petog `div` sa klasom `box`.

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

Treba da imaš element `h2` sa tekstom `Box 6` unutar poslednjeg `div` elementa sa klasom `box`.

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">
			
		</div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
