---
id: 692f22933ba1cb86cd75a70c
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Sada je vreme za kreiranje imena boje za kontejnera. Da biste to uradili, dodajte element `p` ispod svakog elementa `h2` sa ovim imenima boja u sledećem redosledu: `Red`, `Orange`, `Yellow`, `Green`, `Blue`, i `Indigo`. Jedna boja za svako kutije.

# --hints--

Moraš dodati element `p` sa tekstom `Red` ispod elementa `h2`, unutar prvog elementa `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

Treba dodati element `p` sa tekstom `Orange` ispod elementa `h2` unutar drugog elementa `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

Treba da dodate element `p` sa tekstom `Yellow` ispod elementa `h2` unutar trećeg elementa `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

Morate dodati element `p` koji ima tekst `Green` ispod elementa `h2` unutar četvrtog elementa `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

Trebalo bi dodati element `p` sa tekstom `Blue` ispod elementa `h2`, unutar petog elementa `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

Treba da dodate element `p` sa tekstom `Indigo` ispod elementa `h2`, unutar šestog elementa `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
