---
id: 68e51d8f163d3f8b70f9c4e5
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, veštaćete se korišćenjem metoda `includes()` i `slice()`.

Počni, kreiraj parametar nazvan `fccSentence` i dodeli mu vrednost niza karaktera `"freeCodeCamp is a great place to learn web development."`.

# --hints--

Trebalo bi da objavite parametar nazvan `fccSentence`.

```js
assert.exists(fccSentence);
```

Vaše polje `fccSentence` mora biti niz slova.

```js
assert.isString(fccSentence);
```

Tvoj parametar za `fccSentence` mora dobiti vrednost stringa `"freeCodeCamp is a great place to learn web development."`.

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
