---
id: 68e6172975008fea742a63cb
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Sada je vreme da vidimo šta se dešava kada deo sekvence karaktera NIJE PRONAĐEN.

Deklariši parametar nazvan `hasJavaScript`. Zatim postavi njegovu vrednost na rezultat korišćenjem metode `includes()` da proveri da li `fccSentence` ima `"JavaScript"`.

Zatim koristi šablon niza ili spajanje nizova da ispišeš poruku `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.` u konzolu. Zameni `<hasJavaScript>` stvarnom vrednošću promenljive.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Trebalo bi da objavite parametar nazvan `hasJavaScript`.

```js
assert.exists(hasJavaScript);
```

Trebalo bi da koristiš `fccSentence.includes("JavaScript")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Treba da postaviš vrednost rezultata `fccSentence.includes("JavaScript")` u svoj parametar `hasJavaScript`.

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Rezultat od `fccSentence.includes("JavaScript")` bi trebalo da bude `false`.

```js
assert.isFalse(hasJavaScript);
```

Treba da ispišeš poruku `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`, gde `<hasJavaScript>` treba zameniti stvarnom vrednošću promenljive.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
