---
id: 68e6172975008fea742a63ce
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Funkcija `slice()` koristi deo niza karaktera. Prihvata dva argumenta: početni specifikator (uključeno) i krajnji specifikator (isključeno).

Evo primera:

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

Indeksi sekvence znakova počinju od 0. U sekvenci `message`, karakter na indeksu 11 je `f` (na početku "freeCodeCamp").

Kreiraj varijablu nazvanu `platform`. Zatim postavi njenu vrednost na rezultat korišćenja `message.slice(11, 23)` uzimanjem `"freeCodeCamp"`.

# --hints--

Trebalo bi da objavite parametar koji se zove `platform`.

```js
assert.exists(platform);
```

Tvoj parametar za `platform` mora biti sekvenca slova.

```js
assert.isString(platform);
```

Trebalo bi da koristiš `message.slice(11, 23)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Trebalo bi da postaviš vrednost rezultata iz `message.slice(11, 23)` u svoj parametar `platform`.

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Vaš parametar od `platform` mora biti jednak `"freeCodeCamp"`.

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
