---
id: 68e6172975008fea742a63cc
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Način od `includes()` proverava razliku između velikih i malih slova, što znači da razlikuje velika i mala slova.

Deklariši parametar nazvan `hasLowercaseFCC`. Zatim postavi njegovu vrednost za rezultat korišćenjem metode `includes()` da proveri da li `fccSentence` sadrži `"freecodecamp"` (sve mala slova).

Zatim koristite ili šablon niz znakova ili konkatenaciju niza znakova da napišete poruku `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` u konzolu. Zamijenite `<hasLowercaseFCC>` stvarnom vrijednošću parametra.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Trebalo bi da objavite parametar nazvan `hasLowercaseFCC`.

```js
assert.exists(hasLowercaseFCC);
```

Trebalo bi da koristiš `fccSentence.includes("freecodecamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Trebalo bi da postaviš vrednost rezultata `fccSentence.includes("freecodecamp")` u svoj parametar `hasLowercaseFCC`.

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Rezultat za `fccSentence.includes("freecodecamp")` bi trebalo da bude `false`.

```js
assert.isFalse(hasLowercaseFCC);
```

Treba da napišete poruku koja uključuje `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`, gde `<hasLowercaseFCC>` treba zameniti stvarnom vrednošću parametra.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
