---
id: 68e6172975008fea742a63cd
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Sada ćeš pregledati način `slice()` koja daje deo sekvence karaktera.

Definisiši parametar nazvan `message` i postavi vrednost na niz karaktera `"Welcome to freeCodeCamp!"`.

Zatim koristite `console.log()` da ukucate `"Here are some examples of the slice() method:"` u konzolu.

# --hints--

Trebalo bi da objaviš parametar koji se zove `message`.

```js
assert.exists(message);
```

Tvoj parametar za `message` mora biti niz znakova.

```js
assert.isString(message);
```

Vaš parametar za `message` treba biti `"Welcome to freeCodeCamp!"`.

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

Trebalo bi da ukucate `"Here are some examples of the slice() method:"` u konzolu.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
