---
id: 68e6172975008fea742a63ca
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada koristite niz znakova šablona ili spajanje niza znakova da biste napisali poruku `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` na konzoli. Zamenite `<hasFreeCodeCamp>` sa stvarnom vrednošću parametra.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Treba da napišete poruku `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.`, gde `<hasFreeCodeCamp>` treba zameniti stvarnom vrednošću parametra.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
