---
id: 68e6172975008fea742a63d1
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Putanja za `slice()` može koristiti negativne brojeve da broji od kraja sekvence karaktera. `-1` se odnosi na poslednji karakter.

Evo primera:

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

Kreiraj varijablu nazvanu `endPunctuation`. Zatim postavi njenu vrednost na rezultat korišćenja `message.slice(-1)` za preuzimanje poslednjeg karaktera.

Zatim koristi tekstualni niz šablona ili kombinaciju sekvenci znakova da napišeš poruku `The ending punctuation mark is a "<endPunctuation>"` u konzolu. Promeni `<endPunctuation>` sa stvarnom vrednošću parametra.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Trebalo bi da objaviš parametar koji se zove `endPunctuation`.

```js
assert.exists(endPunctuation);
```

Tvoj parametar za `endPunctuation` treba da bude niz slova.

```js
assert.isString(endPunctuation);
```

Treba ti da koristiš `message.slice(-1)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

Treba da postaviš vrednost rezultata `message.slice(-1)` u svoj parametar `endPunctuation`.

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

Tvoj parametar za `endPunctuation` mora biti jednak `"!"`.

```js
assert.strictEqual(endPunctuation, "!");
```

Treba da napišete poruku `The ending punctuation mark is a "<endPunctuation>"`, gde `<endPunctuation>` treba zameniti stvarnom vrednošću parametra.

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
