---
id: 68e6172975008fea742a63d0
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Možete koristiti `slice()` da izvučete od početka niza karaktera.

Kreirajte parametar nazvan `greetingWord`. Zatim mu dodelite vrednost koja je rezultat korišćenja `message.slice(0, 7)` da bi dobio prvu reč `"Welcome"`.

Zatim koristite ili niz teksta šablona ili spajanje niza karaktera da napišete poruku `The first word is "<greetingWord>".` u konzolu. Zamenite `<greetingWord>` sa stvarnom vrednošću parametra.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Trebalo bi da objavite parametar nazvan `greetingWord`.

```js
assert.exists(greetingWord);
```

Tvoj parametar `greetingWord` mora biti niz slova.

```js
assert.isString(greetingWord);
```

Trebalo bi da koristiš `message.slice(0, 7)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Trebalo bi da postaviš vrednost rezultata `message.slice(0, 7)` u svoj parametar `greetingWord`.

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Tvoj parametar za `greetingWord` bi trebalo da bude jednak `"Welcome"`.

```js
assert.strictEqual(greetingWord, "Welcome");
```

Treba da napišete poruku `The first word is "<greetingWord>".`, gde `<greetingWord>` treba da se zameni stvarnom vrednošću parametra.

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
