---
id: 68e6172875008fea742a63c9
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Kao što ste zapamtili iz ranijih lekcija, metoda `includes()` proverava da li niz znakova ima određenu podsekvencu i vraća `true` ili `false`.

Evo primera korišćenjem metode `includes()`:

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

Kreiraj parametar nazvan `hasFreeCodeCamp`. Zatim postavi njegovu vrednost na rezultat korišćenjem metode `includes()` na `fccSentence` kako bi proverio da li ima `"freeCodeCamp"`.

# --hints--

Trebalo bi da objaviš parametar nazvan `hasFreeCodeCamp`.

```js
assert.exists(hasFreeCodeCamp);
```

Trebalo bi da koristiš `fccSentence.includes("freeCodeCamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Trebalo bi da postaviš vrednost rezultata `fccSentence.includes("freeCodeCamp")` u svoj parametar `hasFreeCodeCamp`.

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Rezultati od `fccSentence.includes("freeCodeCamp")` treba da budu `true`.

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
