---
id: 68ef6e42fa6e8a9a18a401fa
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Koristite `console.log()` za pisanje `"Here are some examples of the includes() method:"` u konzoli.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Treba da napišete `"Here are some examples of the includes() method:"` u konzolu.

```js
assert.deepInclude(spy.calls, ["Here are some examples of the includes() method:"])
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
