---
id: 5a24c314108439a4d403616c
title: Ostali prosleđeni parametri
challengeType: 6
forumTopicId: 301399
dashedName: override-default-props
---

# --description--

Sposobnost definisanja podrazumevanih parametara je važna karakteristika u React. Način prepisivanja podrazumevanih parametara je eksplicitno postavljanje vrednosti parametara za određeni deo.

# --instructions--

Komponenta `ShoppingCart` sada prikazuje podređenu komponentu `Items`. Komponenta `Items` ima podrazumevani prop `quantity` postavljen na ceo broj `0`. Zameni podrazumevanu vrednost tako što ćeš proslediti vrednost `10` za `quantity`.

**Napomena:** Zapamtite da je sintaksa za dodavanje parametra komponenti ista kao i način na koji dodajete atribute `HTML`. Međutim, pošto je vrednost ``quantity`` ceo broj, neće biti unutar navodnika već treba da bude obuhvaćen uglastim zagradama. Na primer, ``{100}``. Ova sintaksa govori JSX-u da direktno tumači vrednost unutar zagrada kao `JavaScript`.

# --hints--

Dio `ShoppingCart` bi trebalo da prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Dio `Items` treba da prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

Deo `Items` treba da ima parametar `{ quantity: 10 }` koji je prosleđen iz dela `ShoppingCart`.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
      return (
        mockedComponent.find('Items').props().quantity == 10 &&
        code.replace(/ /g, '')
          .includes('<Itemsquantity={10}/>')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items />
    { /* Change code above this line */ }
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items quantity = {10} />
    { /* Change code above this line */ }
  }
};
```
