---
id: 5a24c314108439a4d4036163
title: Kreiraj komponentu za React
challengeType: 6
forumTopicId: 301386
dashedName: create-a-react-component
---

# --description--

Drugi način za definisanje dela React je korišćenjem ES6 sintakse `class`. U sledećem primeru, `Kitten` dodaje `React.Component`:

```jsx
class Kitten extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <h1>Hi</h1>
    );
  }
}
```

Ovo kreira klasu ES6 `Kitten` koja dodaje klasu `React.Component`. Dakle, klasa `Kitten` sada ima pristup mnogim važnim karakteristikama React, kao što su interni stanja i life cycle hooks. Ne brinite ako niste još upoznati sa ovim terminima, oni će biti detaljnije diskutovani u budućim izazovima. Takođe pogledajte klasu `Kitten` koja ima `constructor` definisanu unutra i koja poziva `super()`. Koristi `super()` za pozivanje komponente roditeljske klase, u ovom slučaju `React.Component`. Konstruktor je posebna metoda koja se koristi prilikom inicijalizacije objekata kreiranih pomoću ključne reči `class`. Najbolja praksa je pozivanje `constructor` komponente i `super`, i prosleđivanje `props` obe. Ovo osigurava da je komponenta pravilno inicijalizovana. Trenutno, znajte da je uobičajeno uključivanje ovog koda. Uskoro ćete videti i druge upotrebe komponente zajedno sa `props`.

# --instructions--

`MyComponent` je definisan u uređivaču koda koristeći klasnu sintaksu. Dovrši metodu `render` tako da vrati element `div` koji sadrži `h1` sa tekstom `Hello React!`.

# --hints--

Dio React treba da vrati element `div`.

```js
assert(Enzyme.shallow(React.createElement(MyComponent)).type() === 'div');
```

`div` treba da prikaže naslov sekcije od `h1` unutar sebe.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.shallow(React.createElement(MyComponent)).html()
  )
);
```

Naslov `h1` treba da sadrži tekst `Hello React!`.

```js
assert(
  Enzyme.shallow(React.createElement(MyComponent)).html() ===
    '<div><h1>Hello React!</h1></div>'
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line



    // Change code above this line
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line
    return (
      <div>
        <h1>Hello React!</h1>
      </div>
    );
    // Change code above this line
  }
};
```
