---
id: 5a24c314108439a4d4036179
title: Unda fomu inayodhibitiwa
challengeType: 6
forumTopicId: 301384
dashedName: create-a-controlled-form
---

# --description--

Poslednji primer je pokazao da React može kontrolisati interni status za određene elemente kao što su `input` i `textarea`, što ih čini kontrolisanim delovima. Ovo je takođe povezano sa drugim elementima forme, uključujući standardni element za HTML `form`.

# --instructions--

Komponenta `MyForm` je pripremljena sa praznim elementom `form` koji ima obrađivač za slanje obrasca. Taj obrađivač će biti pozvan kada se obrazac pošalje.

Dodato je dugme za slanje obrasca. Možeš videti da ima `type` postavljen na `submit`, što znači da upravlja slanjem obrasca. Dodaj element `input` unutar obrasca i postavi njegova svojstva `value` i `onChange()` kao u prethodnom izazovu. Zatim dovrši metodu `handleSubmit` tako da postavi svojstvo `submit` u stanju komponente na trenutnu vrednost unosa iz lokalnog `state`.

**Napomena:** Takođe moraš pozvati `event.preventDefault()` u obrađivaču slanja obrasca kako bi sprečio podrazumevano ponašanje koje osvežava stranicu. Ovde je to već predviđeno kako bi se sprečilo ponovno učitavanje koje bi resetovalo izazov.

Na kraju, napravi oznaku `h1` nakon elementa `form` koja prikazuje vrednost `submit` iz stanja komponente. Zatim možeš unositi tekst u obrazac i kliknuti na dugme (ili pritisnuti Enter) da vidiš svoj unos prikazan na stranici.

# --hints--

`MyForm` treba da vrati element `div` koji sadrži oznaku `form` i `h1`. Obrazac treba da sadrži `input` i `button`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyForm));
    return (
      mockedComponent.find('div').children().find('form').length === 1 &&
      mockedComponent.find('div').children().find('h1').length === 1 &&
      mockedComponent.find('form').children().find('input').length === 1 &&
      mockedComponent.find('form').children().find('button').length === 1
    );
  })()
);
```

Stanje komponente `MyForm` treba da bude inicijalizovano svojstvima `input` i `submit`, pri čemu oba treba da budu prazni tekstualni nizovi.

```js
assert(
  Enzyme.mount(React.createElement(MyForm)).state('input') === '' &&
    Enzyme.mount(React.createElement(MyForm)).state('submit') === ''
);
```

Unos u element `input` treba da ažurira svojstvo `input` u stanju komponente.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return mockedComponent.state('input');
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return {
      state: mockedComponent.state('input'),
      inputVal: mockedComponent.find('input').props().value
    };
  };
  const before = _1();
  const after = _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
})();
```

Slanje obrasca treba da pokrene `handleSubmit`, koji treba da postavi svojstvo `submit` u stanju na trenutnu vrednost unosa.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'SubmitInput' } });
    return mockedComponent.state('submit');
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.state('submit');
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'SubmitInput');
})();
```

`handleSubmit` inapaswa kuita `event.preventDefault`

```js
assert(
  __helpers.isCalledWithNoArgs(
    'event.preventDefault',
    MyForm.prototype.handleSubmit.toString()
  )
);
```

Naslov `h1` treba da prikazuje vrednost polja `submit` iz stanja komponente.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return mockedComponent.find('h1').text();
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.find('h1').text();
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'TestInput');
})();
```

# --seed--

## --seed-contents--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          {/* Change code below this line */}

          {/* Change code above this line */}
          <button type='submit'>Submit!</button>
        </form>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    event.preventDefault();
    this.setState(state => ({
      submit: state.input
    }));
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          <input value={this.state.input} onChange={this.handleChange} />
          <button type='submit'>Submit!</button>
        </form>
        <h1>{this.state.submit}</h1>
      </div>
    );
  }
}
```
