---
id: 5a24c314108439a4d4036189
title: Izmena CSS-a 'inline' na osnovu stanja elementa
challengeType: 6
forumTopicId: 301380
dashedName: change-inline-css-conditionally-based-on-component-state
---

# --description--

Do sada ste videli nekoliko upotreba uslovnog renderovanja i korišćenja stilova unutar linije. Evo još jednog primera koji kombinuje ove dve teme. Takođe možete prikazati CSS uslovno na osnovu stanja sekcije React. Da biste to uradili, proveravate određeno stanje, i ako je to stanje zadovoljeno, menjate stilsku svojstvo dato JSX sekciji u putanji renderovanja.

Ovaj sistem je važan za razumevanje jer predstavlja veliki prelazak od tradicionalnog načina korišćenja stilova menjanjem elemenata DOM direktno (što je veoma uobičajeno za jQuery, na primer). Na taj način, morate pratiti kada se elementi menjaju i takođe direktno rukovati tim promenama. Može biti teško pratiti promene, čime se vaš korisnički interfejs (UI) čini nepredvidivim. Kada postavite stil na osnovu stanja, definišete kako bi trebalo da izgleda UI kao funkcija stanja aplikacije. Postoji jasan tok informacija koji teče u samo jednom smeru. Ovo je preporučeni način prilikom razvoja aplikacija za React.

# --instructions--

Uređivač koda ima deo kontrolisanog polja za unos sa dekorisanim ivicom. Želite da dekorujete ovu ivicu crvenom bojom ako korisnik unese više od 15 slova u polje za unos. Dodajte uslov validacije koji, ukoliko je validan, postavi stil ivice polja za unos na `3px solid red`. Možete testirati tako što ćete uneti tekst u polje za unos.

# --hints--

Sekcija od `GateKeeper` treba da prikaže element za `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.find('div').length === 1;
  })()
);
```

Sekcija `GateKeeper` treba da se inicijalizuje pomoću ključa stanja `input`, postavljenog na prazan niz znakova.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.state().input === '';
  })()
);
```

Sekcija za `GateKeeper` bi trebalo da prikaže oznaku `h3` i oznaku `input`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('h3').length === 1 &&
      mockedComponent.find('input').length === 1
    );
  })()
);
```

Listovi `input` moraju imati obrazac `1px solid black` sa karakteristikom `border` na početku.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('input').props().style.border === '1px solid black'
    );
  })()
);
```

Polje `input` treba da ima granicu od `3px solid red` ako je vrednost unosa u stanju duža od 15 karaktera.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  let initialStyle = mockedComponent.find('input').props().style.border;
  const state_1 = () => {
    mockedComponent.setState({ input: 'this is 15 char' });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const state_2 = () => {
    mockedComponent.setState({
      input: 'A very long string longer than 15 characters.'
    });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const style_1 = await state_1();
  const style_2 = await state_2();
  assert(
    initialStyle === '1px solid black' &&
      style_1 === '1px solid black' &&
      style_2 === '3px solid red'
  );
```

# --seed--

## --seed-contents--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    // Change code below this line

    // Change code above this line
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    if (this.state.input.length > 15) {
      inputStyle.border = '3px solid red';
    };
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```
