---
id: 5a24c314108439a4d403616f
title: Proverite koristeći parametre koje su prosleđene sa sekcijama bez funkcionalnog statusa
challengeType: 6
forumTopicId: 301411
dashedName: review-using-props-with-stateless-functional-components
---

# --description--

Izuzimajući zadnji izazov, prosleđivali ste parametre bezstanja komponentama (stateless components). Ove komponente funkcionišu kao čiste funkcije (pure functions). Prihvataju parametre kao ulaz i vraćaju isti izgled svaki put kada im se proslede isti parametri. Možda vas zanima šta je stanje (state), a sledeći izazov će to detaljnije objasniti. Pre toga, evo pregleda terminologije komponenti.

***Funkcionalna komponenta bez stanja*** je bilo koja funkcija koju napišete koja prima prosleđene parametre i vraća JSX. ***Komponenta bez stanja***, s druge strane, je klasa stvari koja nasleđuje `React.Component`, ali ne koristi interno stanje (to će biti diskutovano u sledećem izazovu). Na kraju, ***komponenta sa stanjem*** je klasa komponente koja ima svoje interno stanje. Možda ćete videti komponente sa stanjem spomenute samo kao komponente ili komponente od React.

Standardni pristup je pokušaj smanjenja stanja i kreiranja funkcionalnih komponenti bez stanja gde god je to moguće. Ovo pomaže da se upravljanje vašim stanjem postavi na određeno mesto u vašoj aplikaciji. S druge strane, ovo povećava razvoj i održavanje vaše aplikacije jer olakšava praćenje kako promene stanja utiču na njeno ponašanje.

# --instructions--

Editor koda ima deo `CampSite` koji prikazuje deo `Camper` kao poddeo. Prevedite deo `Camper` i postavite podrazumevane parametre za `{ name: 'CamperBot' }`. Unutar dela `Camper`, prikažite bilo koji kod koji želite, ali osigurajte da postoji jedan element od `p` koji uključuje samo vrednost `name` prosleđenu kao `prop`. Na kraju, prevedite `propTypes` u deo `Camper` da zahteva `name` prosleđen kao parametar i proverite da li je tipa `string`.

# --hints--

Dio od `CampSite` treba da prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('CampSite').length === 1;
  })()
);
```

Dio od `Camper` treba da prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('Camper').length === 1;
  })()
);
```

Sekcija od `Camper` treba da sadrži osnovne prosleđene parametre koji postavljaju sekvencu karaktera `CamperBot` za ključ `name`.

```js
assert(
  /Camper.defaultProps={name:(['"`])CamperBot\1,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

Dio `Camper` treba da uključuje tipove prosleđenih parametara koji zahtevaju parametar od `name` da bude tipa `string`.

```js
assert(
  /Camper.propTypes={name:PropTypes.string.isRequired,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

Deo od `Camper` treba da ima element/karakteristiku `p` koji sadrži samo tekst iz parametra `name`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return (
      mockedComponent.find('p').text() ===
      mockedComponent.find('Camper').props().name
    );
  })()
);
```

# --before-all--

```js
var PropTypes = {
   string: { isRequired: true }
};
```

# --seed--

## --seed-contents--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line
```

# --solutions--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line

const Camper = (props) => {
   return (
     <div>
       <p>{props.name}</p>
     </div>
   );
};

Camper.propTypes = {
  name: PropTypes.string.isRequired
};

Camper.defaultProps = {
  name: 'CamperBot'
};
```
