---
id: 5a24c314108439a4d4036187
title: Koristi ternarni operator za uslovno prikazivanje
challengeType: 6
forumTopicId: 301414
dashedName: use-a-ternary-expression-for-conditional-rendering
---

# --description--

Pre nego što nastavimo sa metodama za prikaz promena, postoji jedan konačan način da koristite uslove ugrađene u JavaScript za prikaz onoga što želite: <dfn>ternarni operator</dfn>. Ternarni operator se često koristi kao kratak način za izjave od `if/else` u JavaScript. Nije tako moćan kao redovne izjave od `if/else`, ali je veoma popularan među razvojima za React. Jedan razlog je način na koji JSX interpretira; izjave od `if/else` ne mogu se direktno umetnuti u JSX kod. Možda ste ovo videli u nekim prethodnim izazovima — kada je bila potrebna izjava od `if/else`, uvek je bila *izvan* izjave od `return`. Ternarna ekspresija može biti dobra alternativa ako želite implementirati uslovnu logiku unutar vašeg JSX-a. Zapamtite da ternarni operator ima tri dela, ali možete kombinovati više ternarnih izraza zajedno. Evo osnovne sintakse:

```jsx
condition ? expressionIfTrue : expressionIfFalse;
```

# --instructions--

Editor koda ima tri konstante postavljene unutar putanje `CheckUserAge` sekcije `render()`. Zovu se `buttonOne`, `buttonTwo` i `buttonThree`. Svaka je dobila jednostavan JSX izraz koji predstavlja element dugmeta. Prvo, inicijalizujte stanje za `CheckUserAge`, `input` i `userAge` sa vrednošću praznog stringa.

Kada sekcija prikazuje informacije na stranici, korisnici bi trebalo da imaju način za interakciju sa njom. Unutar informacija sekcije `return`, postavite ternarni izraz koji izvršava sledeću logiku: kada se stranica prvi put učita, prikažite dugme za slanje (submit), `buttonOne`, na stranici. Zatim, kada korisnik unese svoj uzrast i klikne dugme, prikažite drugačije dugme u zavisnosti od uzrasta. Ako korisnik unese broj manji od `18`, prikažite `buttonThree`. Ako korisnik unese broj veći ili jednak `18`, prikažite `buttonTwo`.

# --hints--

Sekcija od `CheckUserAge` treba da prikaže jedan element od `input` i jedan element od `button`.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('input')
    .length === 1 &&
    Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('button')
      .length === 1
);
```

Stanje dela `CheckUserAge` treba da bude inicijalizovano sa atributom `userAge` i atributom `input`, oba sa vrednošću praznog niza karaktera.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).state().input === '' &&
    Enzyme.mount(React.createElement(CheckUserAge)).state().userAge === ''
);
```

Kada se komponenta `CheckUserAge` prvi put prikaže u DOM-u, tekst unutar dugmeta treba da bude `Submit`.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('button').text() ===
    'Submit'
);
```

Kada se u polje `input` unese broj manji od 18 i klikne na dugme, tekst na dugmetu treba da glasi `You Shall Not Pass`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter3AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '3' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter17AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '17' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge3 = enter3AndClickButton();
  const userAge17 = enter17AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge3 === 'You Shall Not Pass' &&
      userAge17 === 'You Shall Not Pass'
  );
})();
```

Kada se u polje `input` unese broj veći ili jednak 18 i klikne na dugme, tekst na dugmetu treba da glasi `You May Enter`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter35AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '35' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const userAge35 = enter35AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge18 === 'You May Enter' &&
      userAge35 === 'You May Enter'
  );
})();
```

Mara što se broj pošalje, i kada se vrednost `input` ponovo promeni, `button` treba da se ažurira (ili: vrati) čitajući `Submit`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const changeInputDontClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '5' } });
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter10AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '10' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const changeInput1 = changeInputDontClickButton();
  const userAge10 = enter10AndClickButton();
  const changeInput2 = changeInputDontClickButton();
  assert(
    userAge18 === 'You May Enter' &&
      changeInput1 === 'Submit' &&
      userAge10 === 'You Shall Not Pass' &&
      changeInput2 === 'Submit'
  );
})();
```

Твой код не сме да садржи никакве информације о `if/else`.

```js
assert(
  new RegExp(/(\s|;)if(\s|\()/).test(
    code
  ) === false
);
```

# --seed--

## --seed-contents--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      userAge: '',
      input: ''
    };
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {this.state.userAge === ''
          ? buttonOne
          : this.state.userAge >= 18
          ? buttonTwo
          : buttonThree}
      </div>
    );
  }
}
```
