---
id: 5a24c314108439a4d4036178
title: Unda sehemu ya ingizo inayodhibitiwa
challengeType: 6
forumTopicId: 301385
dashedName: create-a-controlled-input
---

# --description--

Vaša aplikacija može imati složeniju interakciju između `state` i prikazanog UI-a. Na primer, elementi kontrole formulara za unos teksta, kao što su `input` i `textarea`, čuvaju svoje lokalno stanje u DOM kada korisnik unosi podatke. Za React, možete prebaciti ovo dinamičko stanje unutar `state` dela React. Unos korisnika postaje deo `state` aplikacije, tako da React kontroliše vrednost tog polja za unos. Obično, ako imate elemente React sa poljima za unos u koje korisnik može da unese tekst, to će biti formular sa kontrolisanim unosom.

# --instructions--

Uređivač koda sadrži strukturu komponente `ControlledInput` za kreiranje kontrolisanog polja za unos (`input`). Stanje komponente već je inicijalizovano svojstvom `input` koje sadrži prazan tekstualni niz. Ova vrednost predstavlja tekst koji korisnik unosi u polje.

Najpre napravi metodu `handleChange()` koja prima argument `event`. Kada se metoda pozove, dobija objekat `event` koji sadrži tekst unet u polje `input`. Tom tekstu možeš pristupiti pomoću `event.target.value`. Ažuriraj svojstvo `input` u stanju komponente ovom novom vrednošću.

U metodi `render` napravi element `input` iznad oznake `h4`. Dodaj svojstvo `value` koje je jednako vrednosti `input` iz stanja komponente. Zatim dodaj svojstvo `onChange` koje poziva metodu `handleChange()`.

Kada kucaš u polje za unos, tekst obrađuje metoda `handleChange()`, čuva ga u svojstvu `input` lokalnog stanja i prikazuje ga kao vrednost u polju na stranici. Stanje komponente predstavlja jedini izvor istine za podatke unosa.

Na kraju, ne zaboravi da povežeš (`bind`) potrebne metode u konstruktoru.

# --hints--

`ControlledInput` treba da vrati element `div` koji sadrži `input` i oznaku `p`.

```js
assert(
  Enzyme.mount(React.createElement(ControlledInput))
    .find('div')
    .children()
    .find('input').length === 1 &&
    Enzyme.mount(React.createElement(ControlledInput))
      .find('div')
      .children()
      .find('p').length === 1
);
```

Stanje komponente `ControlledInput` treba da bude inicijalizovano svojstvom `input` postavljenim na prazan tekstualni niz.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(ControlledInput)).state('input'),
  ''
);
```

Kucanje u polje za unos treba da ažurira stanje i vrednost polja, a element `p` treba da prikazuje tu vrednost dok kucaš.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(ControlledInput));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return waitForIt(() => mockedComponent.state('input'));
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => ({
      state: mockedComponent.state('input'),
      text: mockedComponent.find('p').text(),
      inputVal: mockedComponent.find('input').props().value
    }));
  };
  const before = await _1();
  const after = await _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.text === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
```

# --seed--

## --seed-contents--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        { /* Change code below this line */}

        { /* Change code above this line */}
        <h4>Controlled Input:</h4>
        <p>{this.state.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  render() {
    return (
      <div>
        <input
          value={this.state.input}
          onChange={this.handleChange} />
        <h4>Controlled Input:</h4>

        <p>{this.state.input}</p>
      </div>
    );
  }
};
```
