---
id: 5a24c314108439a4d4036166
title: Napravi komponente za React
challengeType: 6
forumTopicId: 301381
dashedName: compose-react-components
---

# --description--

Kako izazovi nastavljaju da koriste sve složeniju kombinaciju delova React i JSX-a, postoji jedna važna stvar na koju treba obratiti pažnju. Prikazivanje delova objekta klase u stilu ES6 unutar drugog dela nije drugačije od prikazivanja jednostavnih delova koje ste koristili u nedavnim izazovima. Možete prikazivati JSX elemente, delove funkcije bez stanja (stateless function), i delove objekta klase ES6 unutar drugog dela.

# --instructions--

U uredniku koda, sekcija `TypesOfFood` već prikazuje sekciju nazvanu `Vegetables`. Takođe, postoji sekcija `Fruits` iz prethodnog izazova.

Ugnjezdi dve sekcije unutar `Fruits` — prvo `NonCitrus`, pa `Citrus`. Ove dve sekcije su vam pripremljene iza kulisa. Zatim, ugnjezdi sekciju klase objekata `Fruits` unutar sekcije `TypesOfFood`, ispod elementa zaglavlja sekcije `h1` i iznad `Vegetables`. Rezultat bi trebalo da bude niz ugnježdene sekcije, koje koriste dva različita tipa sekcija.

# --before-all--

```js
globalThis.NonCitrus = function NonCitrus() {
  return React.createElement('div', null);
};
globalThis.Citrus = function Citrus() {
  return React.createElement('div', null);
};
globalThis.Vegetables = function Vegetables() {
  return React.createElement('div', null);
};
```

# --hints--

Dio od `TypesOfFood` treba da vrati jednu funkciju od `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Sekcija ``TypesOfFood`` treba da vrati sekciju ``Fruits``.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

Dio `Fruits` mora vratiti deo `NonCitrus` i deo `Citrus`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return (
      mockedComponent.find('Fruits').children().find('NonCitrus').length ===
        1 &&
      mockedComponent.find('Fruits').children().find('Citrus').length === 1
    );
  })()
);
```

Dio `TypesOfFood` treba da premesti dio `Vegetables` ispod dijela `Fruits`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
        <Vegetables />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }
        <NonCitrus />
        <Citrus />
        { /* Change code above this line */ }
      </div>
    )
  }
}

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
    render() {
      return (
        <div>
        <h1>Types of Food:</h1>
          { /* Change code below this line */ }
          <Fruits />
          { /* Change code above this line */ }
          <Vegetables />
        </div>
      );
    }
};
```
