---
id: 5a24c314108439a4d4036181
title: Uvod šablona unutar linije
challengeType: 6
forumTopicId: 301395
dashedName: introducing-inline-styles
---

# --description--

Postoje druge složene koncepcije koje povećavaju značajnu sposobnost vašeg koda za React. Ali možda se pitate za jednostavniji problem kako stilizovati JSX komponente koje kreirate u React. Možda znate da neće biti potpuno isto kada radite sa HTML zbog <a href="/learn/front-end-development-libraries/react/define-an-html-class-in-jsx" target="_blank" rel="noopener noreferrer nofollow"> načina na koji koristite klase za JSX komponente</a>.

Ako uvozite stilove iz CSS fajla, nije bitno. Koristite klasu za svoj JSX element koristeći atribut klase `className`, a zatim koristite te stilove za tu klasu u vašem CSS fajlu. Druga opcija je korišćenje inline stilova, što je vrlo uobičajeno pri razvoju ReactJS-a.

Koristite ugrađene stilove za JSX elemente kao što radite u HTML, ali sa nekoliko razlika u JSX-u. Evo primera ugrađenog stila u HTML:

```jsx
<div style="color: yellow; font-size: 16px">Mellow Yellow</div>
```

Elementi JSX koriste atribut elementa `style`, ali zbog načina na koji je JSX interpretiran, ne možete postaviti vrednost na `string`. Umesto toga, postavljate ga jednako `object` od JavaScript. Evo primera:

```jsx
<div style={{color: "yellow", fontSize: 16}}>Mellow Yellow</div>
```

Pogledajte kako koristimo camelCase za svojstvo `fontSize`? Ovo je zato što React ne prihvata kebab-ključove case u stil objektu. React će koristiti ispravno ime svojstva za HTML u naše ime.

# --instructions--

Dodajte svojstvo elementu `style` za `div` u editor koda kako bi se podesila crvena boja teksta i veličina fonta za `72px`.

Zapamtite da možete takođe postaviti veličinu fonta na broj, ostavljajući jedinice od `px`, ili napisati kao `"72px"`.# --hints--

Sekcija treba da prikaže funkciju za `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Elemenat za `div` treba da ima boju `red`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().props().style.color === 'red';
  })()
);
```

Element za `div` treba da ima veličinu fonta od `72px`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return (
      mockedComponent.children().props().style.fontSize === 72 ||
      mockedComponent.children().props().style.fontSize === '72' ||
      mockedComponent.children().props().style.fontSize === '72px'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div>Big Red</div>
    );
  }
};
```

# --solutions--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div style={{color: "red", fontSize: 72}}>Big Red</div>
    );
  }
};
```
