---
id: 5a24c314108439a4d403617e
title: Povećajte slušaoce događaja
challengeType: 6
forumTopicId: 301377
dashedName: add-event-listeners
---

# --description--

Put putem `componentDidMount()` je takođe najbolje mesto za dodavanje bilo kog slušaoca događaja koji vam je potreban za određenu funkcionalnost. React pruža sintetički sistem događaja koji obuhvata izvorni sistem događaja prisutan u pregledačima. To znači da sintetički sistem događaja radi potpuno isto, bez obzira na pregledač korisnika – čak i ako se izvorni događaji mogu razlikovati između različitih pregledača.

Već ste počeli da koristite neke sintetičke rukovače događaja kao `onClick()`. Sintetički sistem događaja za React je dobar za korišćenje pri mnogim interakcijama koje se izvršavaju na elementima DOM. Međutim, ako želite da dodelite rukovača događaja na objekte dokumenta ili prozora, morate to uraditi direktno.

# --instructions--

Ambatisha slušalicu događaja na putanjom `componentDidMount()` za događaje `keydown` i učini da ovi događaji izazovu poziv za povrat `handleKeyPress()`. Možete koristiti `document.addEventListener()` koji uzima događaj (u navodniku) kao prvi argument i poziv za povrat kao drugi.

Zatim, u `componentWillUnmount()`, uklonite ovog istog slušaoca događaja. Možete preneti iste argumente na `document.removeEventListener()`. Dobra praksa je korišćenje ove metode životnog ciklusa za izvršavanje bilo kakvog čišćenja delova React pre nego što budu uklonjeni i uništeni. Uklanjanje slušaoca događaja je primer takvog čišćenja.

# --hints--

`MyComponent` treba da prikaže element `div` koji okružuje oznaku `h1`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').children().find('h1').length === 1;
  })()
);
```

Slušalac za `keydown` treba da bude priložen dokumentu u `componentDidMount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const didMountString = mockedComponent
      .instance()
      .componentDidMount.toString();
    return new RegExp(
      'document.addEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(didMountString);
  })()
);
```

Slušalac za `keydown` treba da se ukloni iz dokumenta u `componentWillUnmount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const willUnmountString = mockedComponent
      .instance()
      .componentWillUnmount.toString();
    return new RegExp(
      'document.removeEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(willUnmountString);
  })()
);
```

Nakon što je deo postavljen, klikatanje `enter` bi trebalo da ažurira njegov status sa prikazanom oznakom `h1`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const beforeState = mockedComponent.state('message');
  const beforeText = mockedComponent.find('h1').text();
  const pressEnterKey = () => {
    mockedComponent.instance().handleKeyPress({ keyCode: 13 });
    return waitForIt(() => {
      mockedComponent.update();
      return {
        state: mockedComponent.state('message'),
        text: mockedComponent.find('h1').text()
      };
    });
  };
  const afterKeyPress = await pressEnterKey();
  assert(
    beforeState !== afterKeyPress.state && beforeText !== afterKeyPress.text
  );
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleEnter = this.handleEnter.bind(this);
    this.handleKeyPress = this.handleKeyPress.bind(this);
  }
  // Change code below this line
  componentDidMount() {

  }
  componentWillUnmount() {

  }
  // Change code above this line
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleKeyPress = this.handleKeyPress.bind(this);
    this.handleEnter = this.handleEnter.bind(this);  }
  componentDidMount() {
    // Change code below this line
    document.addEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  componentWillUnmount() {
    // Change code below this line
    document.removeEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```
