---
id: 5a24c314108439a4d4036164
title: Generiš deo koristeći kompozitni obrazac
challengeType: 6
forumTopicId: 301383
dashedName: create-a-component-with-composition
---

# --description--

Sada ćemo pogledati kako možemo organizovati mnogo delova od React zajedno. Zamislite da gradite aplikaciju i kreirali ste tri dela: `Navbar`, `Dashboard`, i `Footer`.

Da biste grupisali ove komponente zajedno, možete kreirati *roditeljsku* komponentu `App` koja prikazuje svaku od ovih tri komponente kao *decu*. Prikazivanje komponente kao deteta u komponenti React, uključuje ime komponente napisano kao posebna oznaka (label) za HTML u JSX-u. Na primer, u putanji `render` možete napisati:

```jsx
return (
 <App>
  <Navbar />
  <Dashboard />
  <Footer />
 </App>
)
```

Kada React naiđe na specifičnu oznaku za HTML koja se odnosi na drugi deo (ime dela obavijeno sa `< />` kao u ovom primeru), on pokazuje indikatore tog dela na mestu te oznake. Ovo bi trebalo da pokaže roditeljsko/dete odnosite između dela `App` i `Navbar`, `Dashboard`, i `Footer`.

# --instructions--

U editoru koda, postoji jednostavna funkciona sekcija nazvana `ChildComponent` i deo klase objekata nazvan `ParentComponent`. Spojite ove dve tako što ćete prikazati `ChildComponent` unutar `ParentComponent`. Uverite se da zatvorite oznaku (tag) `ChildComponent` za početni red.

**Napomena:** `ChildComponent` je definisano pomoću ES6 strelične funkcije jer je ovo veoma uobičajena konvencija prilikom korišćenja React.

# --hints--

Deo React treba da vrati jedan element od `div`.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.type() === 'div';
  })()
);
```

Komponenta treba da vrati dva ugnježdena elementa.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.children().length === 2;
  })()
);
```

Deo mora vratiti `ChildComponent` kao svoj drugi deo.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ParentComponent));
    return (
      mockedComponent.find('ParentComponent').find('ChildComponent').length ===
      1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }


        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }
        <ChildComponent />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
