---
id: 5a24c314108439a4d4036170
title: Kreirajte sekciju sa statusom
challengeType: 6
forumTopicId: 301391
dashedName: create-a-stateful-component
---

# --description--

Jedna od najvažnijih tema u React-u je `state`. Stanje sadrži sve podatke koje aplikacija treba da zna, a koji se mogu menjati tokom vremena. Želiš da aplikacija reaguje na promene stanja i po potrebi prikaže ažuriran korisnički interfejs. React pruža elegantno rešenje za upravljanje stanjem u modernim veb aplikacijama.

Inicijalizuješ stanje u React komponenti deklarisanjem svojstva `state` u konstruktoru klase. Time se komponenta inicijalizuje sa svojim `state` kada se kreira. Svojstvo `state` mora imati vrednost JavaScript objekta. Deklaracija izgleda ovako:

```jsx
this.state = {

}
```

Objektu `state` imaš pristup tokom celog životnog ciklusa komponente. Možeš ga ažurirati, prikazati u korisničkom interfejsu i proslediti kao prop podređenim komponentama. Objekat `state` može biti jednostavan ili složen, u zavisnosti od potreba. Zapamti da moraš napraviti klasnu komponentu koja nasleđuje `React.Component` kako bi mogao da koristiš `state` na ovaj način.

# --instructions--

U uređivaču koda postoji komponenta koja pokušava da prikaže svojstvo `firstName` iz svog `state`, ali `state` nije definisan. Inicijalizuj komponentu sa `state` u konstruktoru i postavi svoje ime kao vrednost svojstva `firstName`.

# --hints--

`StatefulComponent` treba da postoji i da se prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return mockedComponent.find('StatefulComponent').length === 1;
  })()
);
```

`StatefulComponent` treba da prikaže elemente `div` i `h1`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

Stanje komponente `StatefulComponent` treba da bude inicijalizovano svojstvom `firstName` postavljenim na tekstualni niz.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return (
      typeof initialState === 'object' && typeof initialState.firstName === 'string'
    );
  })()
);
```

Svojstvo `firstName` iz stanja komponente `StatefulComponent` treba da bude prikazano u elementu `h1`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return mockedComponent.find('h1').text() === initialState.firstName;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    // Only change code below this line

    // Only change code above this line
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      firstName: 'freeCodeCamp!'
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```
