---
id: 5a24c314108439a4d4036184
title: Prikaži u slučaju ako-иначе
challengeType: 6
forumTopicId: 301410
dashedName: render-with-an-if-else-condition
---

# --description--

Druge upotrebe korišćenjem JavaScript za kontrolu prikazanog izgleda su kombinovanje prikazanih elemenata sa određenim uslovom. Kada je taj uslov tačan, prikazuje se jedan izgled. Kada nije tačan, prikazuje se drugi izgled. Ovo možete uraditi korišćenjem deklaracije `if/else` u metodi `render()` sekcije React.

# --instructions--

Komponenta MyComponent ima `boolean` u svom stanju koje prati da li želite da prikažete određenu funkciju (ili element) u korisničkom interfejsu (UI) ili ne. `button` menja stanje ove vrednosti. Trenutno, prikazuje isti UI uvek. Ponovo napiš metodu `render()` koristeći deklaraciju `if/else` tako da, ako je `display` jednako `true`, vrati trenutne oznake. Inače, vrati oznake bez funkcije `h1`.

**Napomena:** Morate napisati deklaraciju za `if/else` ili položiti ispite. Upotreba ternarnog operatora ovde nije dozvoljena.

# --hints--

`MyComponent` Mora da postoji da se prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length === 1;
  })()
);
```

Kada je `display` postavljeno na `true`, `div`, `button`, na `h1` treba da se prikaže.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 2 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 1
  );
```

Kada je `display` postavljen na `false`, treba da budu prikazani samo `div` i `button`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 0
  );
```

Način prikaza mora koristiti deklaraciju `if/else` za proveru stanja `this.state.display`.

```js
assert(code.includes('if') && code.includes('else'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line

    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
 }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    if (this.state.display) {
      return (
         <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
           <h1>Displayed!</h1>
         </div>
      );
    } else {
      return (
        <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
         </div>
      );
    }
  }
};
```
