---
id: 5a24c314108439a4d4036174
title: Poveži 'ovo' sa putanjom klase
challengeType: 6
forumTopicId: 301379
dashedName: bind-this-to-a-class-method
---

# --description--

Osim postavljanja i ispravljanja `state`, možete definisati metode klase vašeg dela. Metoda klase obično zahteva korišćenje glavne ključne reči `this` kako bi pristupila parametrima u klasi (kao što su `state` i `props`) unutar opsega te metode. Postoje malo načina da omogućite vašim metodama klase pristup `this`.

Jedan uobičajen način je da eksplicitno povežeš (`bind`) `this` u konstruktoru kako bi `this` bilo vezano za metode klase kada se komponenta inicijalizuje. Možda si primetio da je prethodni izazov koristio `this.handleClick = this.handleClick.bind(this)` za metodu `handleClick` u konstruktoru. Tada će se pri pozivu funkcije poput `this.setState()` unutar metode klase, `this` odnositi na instancu klase, a neće biti `undefined`.

**Napomena:** Glavno reči za `this` je jedna od najkonfuznijih stvari u JavaScript, ali ima važnu ulogu u React. Iako je njegovo ponašanje ovde potpuno normalno, ove lekcije nisu mesto za duboku verifikaciju `this`, pa molimo vas da se osvrnete na druge lekcije ako vam ovo objašnjenje zbunjuje!

# --instructions--

Kodni editor ima sekciju sa `state` koja prati tekst. Ima i metodu koja vam omogućava da postavite tekst na `You clicked!`. Međutim, ta metoda ne radi jer koristi glavnu ključnu reč `this` koja je nedefinisana. Ispravite eksplicitnim povezivanjem `this` i metodom `handleClick()` u komponenti te sekcije.

Nakon toga, dodajte handler za klik na element ``button`` u metodu prikaza. Trebalo bi da pokrene putanju ``handleClick()`` kada dugme primi događaj klikova. Imajte na umu da putanja koju prosleđujete handler-u za ``onClick`` zahteva kurzne zagrade jer treba da se interpretira direktno kao `JavaScript`.

Kada završite korake iznad, trebalo bi da možete kliknuti na dugme i videti `You clicked!`.

# --hints--

`MyComponent` treba da ponovi element `div` koji okružuje dva elementa, dugme i element `h1`, u tom redosledu.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).find('div').length === 1 &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(0)
      .type() === 'button' &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(1)
      .type() === 'h1'
);
```

Stanje `MyComponent` treba da se inicijalizuje par ključ-vrednost `{ text: "Hello" }`.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('text') === 'Hello'
);
```

Klik na element ``button`` treba da pokrene putanju ``handleClick`` i postavi stanje za ``text`` na ``You clicked!``.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ text: 'Hello' });
    return waitForIt(() => mockedComponent.state('text'));
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return waitForIt(() => mockedComponent.state('text'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Hello' && secondValue === 'You clicked!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    // Change code below this line

    // Change code above this line
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }
        <button>Click Me</button>
        { /* Change code above this line */ }
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```
