---
id: 5a24c314108439a4d4036177
title: Andika kielekezi rahisi
challengeType: 6
forumTopicId: 301425
dashedName: write-a-simple-counter
---

# --description--

Možete dizajnirati sekciju sa složenijom funkcionalnošću kombinujući koncepte koji su do sada diskutovani. Ovo uključuje inicijalizovanje `state`, pisanje ruta koje postavljaju `state`, i postavljanje rukovaoca za klik (click handlers) radi inicijalizovanja ovih ruta.

# --instructions--

Sekcija `Counter` čuva vrednost `count` u `state`. Postoje dva dugmeta koja pozivaju putem `increment()` i `decrement()`. Napišite ove funkcije tako da se vrednost brojača poveća ili smanji za 1 kada se klikne odgovarajuće dugme. Takođe, kreirajte funkciju za `reset()` tako da je brojač 0 kada se klikne dugme za resetovanje.

**Napomena:** Nemoj menjati `className` svojstva dugmadi. Takođe, ne zaboravi da povežeš (`bind`) nove metode u konstruktoru.

# --hints--

`Counter` treba da vrati element `div` koji sadrži tri dugmeta sa tekstom redom: `Increment!`, `Decrement!`, `Reset`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Counter));
    return (
      mockedComponent.find('.inc').text() === 'Increment!' &&
      mockedComponent.find('.dec').text() === 'Decrement!' &&
      mockedComponent.find('.reset').text() === 'Reset'
    );
  })()
);
```

Stanje `Counter` treba da se inicijalizuje sa atributom `count` postavljenim na `0`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
assert(mockedComponent.find('h1').text() === 'Current Count: 0');
```

Pritiskanje dugmeta za dodavanje bi trebalo da poveća brojač za `1`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.inc').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: 1');
```

Kubofya kifungo cha kupunguza kunapaswa kupunguza hesabu kwa `1`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.dec').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: -1');
```

Klik na dugme za resetovanje treba da postavi brojač na `0`.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.setState({ count: 5 });
const currentCountElement = mockedComponent.find('h1');
assert(currentCountElement.text() === 'Current Count: 5');
mockedComponent.find('.reset').simulate('click');
assert(currentCountElement.text() === 'Current Count: 0');
```

# --seed--

## --seed-contents--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
  this.increment = this.increment.bind(this);
 this.decrement = this.decrement.bind(this);
 this.reset = this.reset.bind(this);
 }
  reset() {
    this.setState({
      count: 0
    });
  }
  increment() {
    this.setState(state => ({
      count: state.count + 1
    }));
  }
  decrement() {
    this.setState(state => ({
      count: state.count - 1
    }));
  }
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```
