---
id: 5a24c314108439a4d4036180
title: Boresha upyaonyeshaji kwa shouldComponentUpdate
challengeType: 6
forumTopicId: 301398
dashedName: optimize-re-renders-with-shouldcomponentupdate
---

# --description--

Do sada, kada komponenta primi novi `state` ili nove `props`, ona ponovo renderuje sebe i svu svoju decu. To je uglavnom poželjno. Međutim, React pruža metodu životnog ciklusa `shouldComponentUpdate()` koja se poziva kada komponenta primi novi `state` ili `props` i omogućava da odrediš da li komponenta treba da se ažurira. Ova metoda prima `nextProps` i `nextState` kao argumente.

Ova metoda je važna za optimizaciju performansi. Na primer, podrazumevano ponašanje je da se komponenta ponovo renderuje kada primi nove `props`, čak i ako se njihove vrednosti nisu promenile. To možeš sprečiti korišćenjem metode `shouldComponentUpdate()` poređenjem trenutnih `props` (`this.props`) sa budućim (`nextProps`). Metoda mora vratiti `true` ili `false`, u zavisnosti od toga da li komponenta treba da se ažurira.

# --instructions--

Metoda `shouldComponentUpdate()` je dodata komponenti `OnlyEvens`. Trenutno uvek vraća `true`, pa se komponenta ponovo renderuje svaki put kada primi nove `props`. Izmeni metodu tako da se `OnlyEvens` ažurira samo kada je vrednost `value` iz novih propsa pozitivan paran ceo broj. Klikni na dugme `Add` i posmatraj redosled događaja u konzoli pregledača dok se metode životnog ciklusa pozivaju.

# --hints--

Komponenta `Controller` treba da prikazuje komponentu `OnlyEvens` kao dete.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Controller));
    return (
      mockedComponent.find('Controller').length === 1 &&
      mockedComponent.find('OnlyEvens').length === 1
    );
  })()
);
```

Metoda `shouldComponentUpdate` treba da bude definisana u komponenti `OnlyEvens`.

```js
assert(
  (() => {
    const child = React.createElement(OnlyEvens)
      .type.prototype.shouldComponentUpdate.toString()
      .replace(/s/g, '');
    return child !== 'undefined';
  })()
);
```

Komponenta `OnlyEvens` treba da vrati oznaku `h1` koja prikazuje vrednost `this.props.value`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 1000 });
    return mockedComponent.find('h1').html();
  };
  const second = () => {
    mockedComponent.setState({ value: 10 });
    return mockedComponent.find('h1').html();
  };
  const firstValue = first();
  const secondValue = second();
  assert(firstValue === '<h1>1000</h1>' && secondValue === '<h1>10</h1>');
})();
```

`OnlyEvens` treba da se ponovo prikaže samo kada `nextProps.value` je pun pozitivan broj greške.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 8 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ value: 7 });
    return mockedComponent.find('h1').text();
  };
  const third = () => {
    mockedComponent.setState({ value: 42 });
    return mockedComponent.find('h1').text();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(firstValue === '8' && secondValue === '8' && thirdValue === '42');
})();
```

# --seed--

## --seed-contents--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return true;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```

# --solutions--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return nextProps.value % 2 === 0;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```
