---
id: 5a24c314108439a4d4036160
title: Objasni HTML element u JSX-u
challengeType: 6
forumTopicId: 301393
dashedName: define-an-html-class-in-jsx
---

# --description--

Kada počnete da učite JSX, verovatno ćete se zapitati koje su mu razlike u odnosu na HTML.

Do sada izgleda kao da su HTML i JSX potpuno isti.

Glavna razlika u JSX-u je što više ne možete koristiti reč `class` za kreiranje klasa za HTML. Ovo je zato što `class` je rezervisana reč u JavaScript. Umesto toga, JSX koristi `className`.

Zapravo, konvencija imenovanja za sve svojstva elementa HTML i reference na događaje u JSX je camelCase. Na primer, klik događaj u JSX je `onClick`, umesto `onclick`. Slično tome, `onchange` se menja u `onChange`. Iako je ovo mala razlika, važno je imati na umu kada idete dalje.

# --instructions--

Primeni klasu `myDiv` na element `div` u JSX kodu.

# --hints--

Element `JSX` treba da vrati element `div`.

```js
assert.strictEqual(JSX.type, 'div');
```

`div` treba da ima klasu od `myDiv`.

```js
assert.strictEqual(JSX.props.className, 'myDiv');
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h1>Add a class to this div</h1>
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div className = 'myDiv'>
  <h1>Add a class to this div</h1>
</div>);
```
