---
id: 5a24c314108439a4d403616b
title: Koristite zadane parametre
challengeType: 6
forumTopicId: 301418
dashedName: use-default-props
---

# --description--

React takođe ima opciju za postavljanje podrazumevanih vrednosti za prosleđene propse. Možeš definisati podrazumevane propse kao svojstvo same komponente, a React će koristiti tu vrednost kada prop nije prosleđen. Na primer, ako napišeš `MyComponent.defaultProps = { location: 'San Francisco' }`, definisao si da će svojstvo `location` imati vrednost `San Francisco`, osim ako nije drugačije navedeno. React koristi podrazumevane propse samo kada prop nije prosleđen. Ako proslediš `null`, vrednost će ostati `null`.

# --instructions--

Uređivač koda prikazuje komponentu `ShoppingCart`. Deklariši podrazumevane propse za ovu komponentu tako da svojstvo `items` ima vrednost `0`.

# --hints--

Sekcija od `ShoppingCart` treba da prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Sekcija od `ShoppingCart` treba da ima svojstvo `{ items: 0 }`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    mockedComponent.setProps({ items: undefined });
    return mockedComponent.find('ShoppingCart').props().items === 0;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};
// Change code below this line
```

# --solutions--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};

// Change code below this line
ShoppingCart.defaultProps = {
  items: 0
}
```
