---
id: 5a24c314108439a4d403617c
title: Koristi životni put componentWillMount
challengeType: 6
forumTopicId: 301423
dashedName: use-the-lifecycle-method-componentwillmount
---

# --description--

React komponente imaju nekoliko posebnih metoda koje omogućavaju izvršavanje određenih radnji u različitim fazama životnog ciklusa komponente. One se nazivaju metode životnog ciklusa (lifecycle methods) ili lifecycle hooks i omogućavaju ti da reaguješ u određenim trenucima, kao što su pre renderovanja, pre ažuriranja, pre prijema propsa ili pre uklanjanja komponente. Neke od glavnih metoda su: `componentWillMount()`, `componentDidMount()`, `shouldComponentUpdate()`, `componentDidUpdate()` i `componentWillUnmount()`. Naredne lekcije će objasniti neke od njihovih osnovnih primena.

**Napomena:** Metoda životnog ciklusa `componentWillMount` biće zastarela u budućoj verziji 16.X i potpuno uklonjena u verziji 17. Više možeš saznati u <a href="https://www.freecodecamp.org/news/how-to-safely-use-reacts-life-cycles-with-fiber-s-async-rendering-fd4469ebbd8f/" target="_blank" rel="noopener noreferrer nofollow">ovom članku</a>.

# --instructions--

Metoda `componentWillMount()` poziva se pre metode `render()` kada se komponenta umeće u DOM. Ispiši nešto u konzolu unutar `componentWillMount()` — rezultat možeš videti u konzoli pregledača.

# --hints--

`MyComponent` treba da prikaže element `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').length === 1;
  })()
);
```

`console.log` inapaswa kuitwa ndani ya `componentWillMount`.

```js
assert(
  (function () {
    const lifecycle = React.createElement(MyComponent)
      .type.prototype.componentWillMount.toString()
      .replace(/ /g, '');
    return lifecycle.includes('console.log(');
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return <div />
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line
    console.log('Component is mounting...');
    // Change code above this line
  }
  render() {
    return <div />
  }
};
```
