---
id: 5a24c314108439a4d403617d
title: Koristi životni put componentDidMount
challengeType: 6
forumTopicId: 301422
dashedName: use-the-lifecycle-method-componentdidmount
---

# --description--

Mnogi web programeri, ponekad, moraju da pozovu krajnju tačku API kako bi dobili podatke. Ako radiš sa React, važno je znati gde da izvršiš ovu akciju.

Optimalna metoda sa React je postavljanje poziva za API ili bilo kog poziva ka vašem serveru u živom putu `componentDidMount()`. Ovaj metod se naziva nakon što segment bude postavljen na DOM. Bilo koji poziv od `setState()` ovde će uzrokovati ponovno prikazivanje vašeg segmenta. Kada pozovete API u ovom putu, i postavite svoj status i podatke koje API vraća, to će uzrokovati automatsko ažuriranje kada primite te podatke.

# --instructions--

Postoji poziv za API primer u `componentDidMount()`. Ažurira status nakon 2,5 sekunde simulirajući poziv na server radi dobijanja podataka. Ovaj primer traži ukupan broj trenutno aktivnih korisnika za sajt. U metodi prikaza, prikažite vrednost `activeUsers` u `h1` nakon teksta `Active Users:`. Pogledajte šta se dešava u prethodnom prikazu i slobodno promenite vreme čekanja da biste videli različite efekte.

# --hints--

`MyComponent` trebalo bi da prikaže element `div` koji okružuje oznaku `h1`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

Stanje dela treba ažurirati korišćenjem akcije čekanja u `componentDidMount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return new RegExp('setTimeout(.|\n)+setState(.|\n)+activeUsers').test(
      String(mockedComponent.instance().componentDidMount)
    );
  })()
);
```

Etiketa za `h1` treba da prikaže vrednost `activeUsers` iz stanja `MyComponent`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ activeUsers: 1237 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ activeUsers: 1000 });
    return mockedComponent.find('h1').text();
  };
  assert(new RegExp('1237').test(first()) && new RegExp('1000').test(second()));
})();
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        {/* Change code below this line */}
        <h1>Active Users: </h1>
        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        <h1>Active Users: {this.state.activeUsers}</h1>
      </div>
    );
  }
}
```
