---
id: 5a24c314108439a4d403617b
title: Prosleđi poziv za povratak kao parametre
challengeType: 6
forumTopicId: 301400
dashedName: pass-a-callback-as-props
---

# --description--

Možete proslediti `state` kao parametre za podkomponente deteta, ali nemate ograničenje samo na prosleđivanje podataka. Takođe možete proslediti rukovateljske rute ili bilo koju rutu definisanu u sekciji React za podkomponentu deteta. Ovo je način kojim omogućavate podkomponentama deteta da interaguju sa svojim roditeljskim komponentama. Prosleđujete rutu detetu kao običan prop. Daje se ime i imate pristup tom imenu rute ispod `this.props` u podkomponenti deteta.

# --instructions--

Postoji tri definisane sekcije u uređivaču koda. Sekcija `MyApp` je roditelj koji će prikazivati detinjaste sekcije `GetInput` i `RenderInput`. Dodajte sekciju `GetInput` na putanju prikaza u `MyApp`, zatim prosledite prop pod nazivom `input` koji ima vrednost `inputValue` iz `MyApp` u `state`. Takođe kreirajte prop pod nazivom `handleChange` i prosledite mu rukovatelja unosa (input handler) `handleChange`.

Sledeći, dodajte `RenderInput` u metodu za prikaz na putanji `MyApp`, zatim kreirajte prop nazvan(a) `input` i proslijedite `inputValue` iz `state` mu. Kada završite, možete pisati u polje `input` u dijelu `GetInput`, koji zatim poziva metodu rukovodioca u svom roditelju putem parametara. Ovo ažurira unos u `state` roditelja, što se proslijeđuje kao parametri oba potomka. Pogledajte kako podaci teku između komponenti i kako jedan izvor istine ostaje `state` roditeljske komponente. Zapravo, ovaj primjer je malo hipotetički, ali bi trebao pokazati kako se podaci i pozivi povratne funkcije mogu proslijediti između komponenti React.

# --hints--

Deo od `MyApp`, bi trebalo da prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('MyApp').length === 1;
  })()
);
```

Dio od `GetInput` treba da prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('GetInput').length === 1;
  })()
);
```

Dio `RenderInput` treba da prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('RenderInput').length === 1;
  })()
);
```

Dio `GetInput` treba da primi atribut stanja `MyApp` `inputValue` kao parametre i imati element/komponentu `input` koji menja stanje `MyApp`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: '' });
    return waitForIt(() => mockedComponent.state());
  };
  const state_2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => mockedComponent.state());
  };
  const updated_1 = await state_1();
  const updated_2 = await state_2();
  assert(updated_1.inputValue === '' && updated_2.inputValue === 'TestInput');
```

Deo `RenderInput` bi trebalo da primi atribut stanja `MyApp` `inputValue` kao parametre.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: 'TestName' });
    return waitForIt(() => mockedComponent);
  };
  const updated_1 = await state_1();
  assert(updated_1.find('p').text().includes('TestName'));
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
  this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
         <GetInput
           input={this.state.inputValue}
           handleChange={this.handleChange}/>
         <RenderInput
           input={this.state.inputValue}/>
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```
