---
id: 5a24c314108439a4d4036168
title: Napiši deo u React-u od početka
challengeType: 6
forumTopicId: 301424
dashedName: write-a-react-component-from-scratch
---

# --description--

Sada ste naučili osnove JSX-a i komponente za React, vreme je da napišete sopstvenu komponentu. Komponente za React su osnova projekata za React, stoga je važno naviknuti se na njihovo pravilno pisanje. Zapamtite, standardna komponenta React je ES6 `class` koja dodaje `React.Component`. Ima render metodu koja vraća HTML (iz JSX-a) ili `null`. Ovo je osnovna struktura komponente React. Ako ovo dobro razumete, bićete spremni da počnete sa izgradnjom još kompleksnijih projekata za React.

# --instructions--

Definišite klasu `MyComponent` koja proširuje `React.Component`. Njen metod za prikaz treba da vrati `div` sa oznakom `h1` sa tekstom: `My First React Component!` unutar njega. Koristite ovaj tekst tačno kako je, velika i mala slova, pa interpunkcija su važne. Takođe osigurajte da pozovete komponentu vašeg dela.

Pokažite ovaj deo na DOM koristeći `ReactDOM.render()`. Postoji `div` sa `id='challenge-node'` koji je dostupan za vas da koristite.

# --hints--

Treba da ima deo React koji se zove `MyComponent`.

```js
assert(__helpers.removeWhiteSpace(code).includes('classMyComponentextendsReact.Component{'));
```

`MyComponent` treba da ima oznaku `h1` sa tekstom `My First React Component!` Velika, mala slova i interpunkcija su važne.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('h1').text() === 'My First React Component!';
  })()
);
```

`MyComponent` treba da se prikaže na DOM.

```js
assert(document.getElementById('challenge-node').childNodes.length === 1);
```

`MyComponent` trebalo bi da ima komponentu nazvanu `super` i `props`.

```js
assert(
  MyComponent.toString().includes('MyComponent(props)') &&
    MyComponent.toString().includes('_super.call(this, props)')
);
```

# --seed--

## --seed-contents--

```jsx
// Change code below this line
```

# --solutions--

```jsx
// Change code below this line
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>My First React Component!</h1>
      </div>
    );
  }
};

ReactDOM.render(<MyComponent />, document.getElementById('challenge-node'));
```
