---
id: 5a24c314108439a4d4036165
title: Koristi React da prikažeš organizovane delove unutra
challengeType: 6
forumTopicId: 301420
dashedName: use-react-to-render-nested-components
---

# --description--

Posledji izazov pokazao je jednostavan način za povezivanje dve komponente, ali postoji mnogo različitih načina na koje možete spojiti deo pomoću React.

Struktura komponenti je jedna od najmoćnijih karakteristika React-a. Kada radiš sa React-om, važno je da počneš da razmišljaš o korisničkom interfejsu kao o skupu komponenti, kao što je komponenta App iz prethodnog izazova. Korisnički interfejs deliš na osnovne gradivne blokove, a ti blokovi postaju komponente. To pomaže da se kod za korisnički interfejs odvoji od logike aplikacije i znatno olakšava razvoj i održavanje složenih projekata.

# --instructions--

U uređivaču koda su deklarisane dve funkcionalne komponente, koje se zovu `TypesOfFruit` i `Fruits`. Uzmite deo `TypesOfFruit` i povežite ga, ili *ugnezdite*, unutar dela `Fruits`. Zatim uzmite deo `Fruits` i postavite ga unutar dela `TypesOfFood`. Rezultat bi trebalo da bude komponenta deteta, ugnježđena unutar komponente roditelja, koja je takođe ugnježđena unutar roditeljske komponente svoje vrste!

# --hints--

Dio od `TypesOfFood` treba da vrati jednu funkciju od `div`.

```js
assert(Enzyme.shallow(React.createElement(TypesOfFood)).type() === 'div');
```

Sekcija ``TypesOfFood`` treba da vrati sekciju ``Fruits``.

```js
assert(
  Enzyme.shallow(React.createElement(TypesOfFood)).props().children[1].type
    .name === 'Fruits'
);
```

Komponenta `Fruits` treba da vrati komponentu `TypesOfFruit`.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('h2').html() ===
    '<h2>Fruits:</h2>'
);
```

Komponenta `TypesOfFruit` treba da vrati elemente `h2` i `ul`.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('ul').text() ===
    'ApplesBlueberriesStrawberriesBananas'
);
```

# --seed--

## --seed-contents--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }

      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }
        <TypesOfFruit />
      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }
        <Fruits />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
