---
id: 5a24c314108439a4d403618a
title: Koristi Array.map() za dinamičko prikazivanje elemenata
challengeType: 6
forumTopicId: 301417
dashedName: use-array-map-to-dynamically-render-elements
---

# --description--

Uslovno prikazivanje je važno, ali možda će vam trebati da vaši komponenti prikažu nepoznat broj elemenata. Često u aplikacijama za promene, programer nema način da zna stanje aplikacije do vremena izvršavanja, jer mnoge zavise od interakcije korisnika sa tom aplikacijom. Programeri moraju da napišu svoj kod kako bi rešili to nepoznato stanje unapred. Korišćenje `Array.map()` u React pokazuje ovaj koncept.

Na primer, praviš jednostavnu aplikaciju „Lista obaveza“. Kao programer, ne možeš unapred znati koliko će stavki korisnik imati na svojoj listi. Zato komponenta treba da bude spremna da dinamički prikaže odgovarajući broj stavki bez obzira na to šta korisnik unese.

# --instructions--

Uređivač koda već sadrži najveći deo komponente `MyToDoList`. Deo koda će ti biti poznat ako si završio izazov sa kontrolisanim obrascem. Videćeš `textarea`, `button` i nekoliko metoda koje upravljaju njihovim stanjem, ali se još ništa ne prikazuje na stranici.

U konstruktoru napravi objekat `this.state` i definiši dva svojstva: `userInput` treba da bude inicijalizovan kao prazan tekstualni niz, a `toDoList` kao prazan niz. Zatim u metodi `render()` zameni vrednost `null` promenljive `items` tako što ćeš primeniti `map` na niz `toDoList` iz stanja komponente i dinamički prikazati po jedan element `li` za svaku stavku. Probaj da uneseš tekst `eat, code, sleep, repeat` u `textarea`, klikni na dugme i vidi šta će se desiti.

**Napomena:** Verovatno znaš da svi elementi na istom nivou koji nastaju korišćenjem metode `map` moraju imati jedinstveno svojstvo `key`. Ne brini, to je tema sledećeg izazova.

# --hints--

Sekcija MyToDoList mora da postoji i da se prikaže na stranici.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return mockedComponent.find('MyToDoList').length === 1;
  })()
);
```

Prvo dete komponente `MyToDoList` treba da bude element `textarea`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(0).type() ===
      'textarea'
    );
  })()
);
```

Drugo dete komponente `MyToDoList` treba da bude element `br`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(1).type() === 'br'
    );
  })()
);
```

Treće dete komponente `MyToDoList` treba da bude element `button`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(2).type() ===
      'button'
    );
  })()
);
```

Stanje `MyToDoList` treba da se inicijalizuje sa `toDoList` kao prazan niz podataka.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      Array.isArray(initialState.toDoList) === true &&
      initialState.toDoList.length === 0
    );
  })()
);
```

Stanje komponente `MyToDoList` treba da bude inicijalizovano tako da je `userInput` prazan tekstualni niz.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      typeof initialState.userInput === 'string' &&
      initialState.userInput.length === 0
    );
  })()
);
```

Kada se dugme `Create List` klikne, deo `MyToDoList` bi trebalo da prikaže dinamički neuređeni spisak koji sadrži stavku za svaki element liste predmeta razdvojene zarezima unetih u polje `textarea`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  const state_1 = () => {
    return mockedComponent.find('ul').find('li');
  };
  const setInput = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      'testA, testB, testC'
    );
  };
  const click = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_2 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const setInput_2 = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      't1, t2, t3, t4, t5, t6'
    );
  };
  const click_1 = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_3 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const awaited_state_1 = state_1();
  const awaited_setInput = setInput();
  const awaited_click = click();
  const awaited_state_2 = state_2();
  const awaited_setInput_2 = setInput_2();
  const awaited_click_1 = click_1();
  const awaited_state_3 = state_3();
  assert(
    awaited_state_1.length === 0 &&
      awaited_state_2.nodes.length === 3 &&
      awaited_state_3.nodes.length === 6 &&
      awaited_state_2.text === 'testAtestBtestC' &&
      awaited_state_3.text === 't1t2t3t4t5t6'
  );
})();
```

# --seed--

## --seed-contents--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = null; // Change this line
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```

# --solutions--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      toDoList: [],
      userInput: ''
    };
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = this.state.toDoList.map((item, i) => {
      return <li key={i}>{item}</li>;
    });
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```
