---
id: 5a24c314108439a4d403618b
title: Pružiti jedinstvenu karakteristiku oznake za srodne elemente
challengeType: 6
forumTopicId: 301394
dashedName: give-sibling-elements-a-unique-key-attribute
---

# --description--

Prethodni izazov pokazao je kako se metoda `map` koristi za dinamičko prikazivanje više elemenata na osnovu korisničkog unosa. Međutim, nedostajao je jedan važan deo. Kada kreiraš niz elemenata, svaki mora imati jedinstveno svojstvo `key`. React koristi ove ključeve da prati koji su elementi dodati, izmenjeni ili uklonjeni. To čini ponovno renderovanje efikasnijim kada se lista promeni.

**Napomena:** Ključevi bi trebalo da budu jedinstveni samo među srodnim elementima, ne moraju biti globalno jedinstveni u vašoj aplikaciji.

# --instructions--

Uređivač koda sadrži niz front-end okvira i funkcionalnu komponentu bez stanja pod nazivom `Frameworks()`. Komponenta `Frameworks()` treba da mapira niz u neuređenu listu, kao u prethodnom izazovu. Dovrši poziv metode `map` tako da vrati element `li` za svaki okvir iz niza `frontEndFrameworks`. Ovog puta obavezno svakom elementu `li` dodeli jedinstveno svojstvo `key`. Elementi `li` takođe treba da sadrže tekst iz niza `frontEndFrameworks`.

Obično želite da ključevi postanu nešto što jedinstveno identifikuje prikazani element. Kao poslednje sredstvo, može se koristiti model reda, ali obično biste trebali pokušati da koristite jedinstveni identifikator.

# --hints--

Sekcija od `Frameworks` mora da postoji kako bi se prikazala na stranici.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('Frameworks').length === 1
);
```

`Frameworks` treba da prikaže element `h1`.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('h1').length === 1);
```

`Frameworks` trebalo bi da prikaže funkciju od `ul`.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('ul').length === 1);
```

Element `ul` treba da sadrži 6 podređenih elemenata `li`.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('ul').children().length ===
    6 &&
    Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .childAt(0)
      .name() === 'li' &&
    Enzyme.mount(React.createElement(Frameworks)).find('li').length === 6
);
```

Svaki element liste mora imati jedinstveno svojstvo `key`.

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(Frameworks)).find('ul');
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key(),
      ul.childAt(4).key(),
      ul.childAt(5).key()
    ]);
    return keys.size === 6;
  })()
);
```

Svaki element liste mora imati tekst iz `frontEndFrameworks`.

```js
assert(
  (() => {
    const li = Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .children();
    return [...Array(5)].every((_, i) =>
      frontEndFrameworks.includes(li.at(i).text())
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = null; // Change this line
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```

# --solutions--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = frontEndFrameworks.map((fw, i) => <li key={i}>{fw}</li>);
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```
