---
id: 5a24c314108439a4d4036188
title: Prikaži po uslovima iz prosleđene parametre
challengeType: 6
forumTopicId: 301405
dashedName: render-conditionally-from-props
---

# --description--

Sada ste videli kako da koristite `if/else`, `&&` i ternarni operator (`condition ? expressionIfTrue : expressionIfFalse`) za donošenje uslovnih odluka o tome šta i kada prikazati. Međutim, postoji još jedna važna tema koju treba da pokrijemo, koja vam omogućava da kombinujete sve ili neke od ovih koncepata sa još jednom moćnom karakteristikom: propovima (vrednostima koje se prosleđuju) React. Korišćenje propova za uslovno prikazivanje koda je veoma uobičajeno kod programera React — to jest, koristite vrednost određenog prop-a da donosite direktne odluke o tome šta treba prikazati.

U ovom izazovu, treba da napraviš deo za dete da donosi odluke na osnovu prosleđenih kriterijuma. Takođe ćeš koristiti ternarni operator, ali ćeš videti kako neke druge koncepte diskutovane u prethodnim izazovima mogu biti korisne i u ovom kontekstu.

# --instructions--

Editor koda ima dva dela definisana za vas: roditelja nazvan `GameOfChance` i deteta nazvano `Results`. Oni se koriste za kreiranje jednostavne igre gde korisnik klikne na dugme da vidi da li je pobedio ili izgubio.

Prvo će ti biti potreban jednostavan matematički izraz koji pri svakom izvršavanju vraća nasumičnu vrednost. Možeš koristiti `Math.random()`. Ova metoda vraća broj između `0` (uključujući) i `1` (isključujući). Za verovatnoću 50/50 koristi izraz `Math.random() >= .5`. Ovaj izraz će vratiti `true` u približno 50% slučajeva, a `false` u preostalih 50%. U metodi `render` zameni `null` ovim izrazom da dovršiš deklaraciju promenljive.

Sada imate matematički izraz koji možete koristiti za donošenje nasumične odluke u kodu. Sledeće je implementacija toga. Prikazati sekciju `Results` kao dete od `GameOfChance`, i proslediti `expression` kao prop nazvanu `fiftyFifty`. U sekciji `Results`, napisati ternarni izraz da prikaže komponentu `h1` sa tekstom `You Win!` ili `You Lose!`, u zavisnosti od propa `fiftyFifty` koji se prosleđuje iz `GameOfChance`. Na kraju, osigurati da metoda `handleClick()` pravilno broji svaki potez kako bi korisnik znao koliko je puta igrao. Ovo takođe osigurava da korisnik zna zaista li se sekcija promenila ako osvoji ili izgubi dva puta zaredom.

# --hints--

Deo `GameOfChance` mora da postoji i bude prikazan na stranici.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('GameOfChance').length,
  1
);
```

`GameOfChance` trebalo bi da vrati jedan element od `button`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('button').length,
  1
);
```

`GameOfChance` treba da vrati jedan primer dela ``Results``, koji ima svojstvo (prop) nazvano ``fiftyFifty``.

```js
assert(
  Enzyme.mount(React.createElement(GameOfChance)).find('Results').length ===
    1 &&
    Enzyme.mount(React.createElement(GameOfChance))
      .find('Results')
      .props()
      .hasOwnProperty('fiftyFifty') === true
);
```

Stanje `GameOfChance` treba da se postavi atributom `counter` sa vrednošću `1`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).state().counter,
  1
);
```

Kada se deo `GameOfChance` prvi put prikazuje na DOM, element `p` treba da se vrati sa internim tekstom `Turn: 1`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('p').text(),
  'Turn: 1'
);
```

Svaki put kada se dugme klikne, status inkrementa treba da se poveća za vrednost od 1, i jedan element od `p` prikazuje se u DOM sa tekstom `Turn: N`, gde je `N` vrednost statusa inkrementa.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    count: comp.state('counter'),
    text: comp.find('p').text()
  });
  const _1 = () => {
    simulate();
    return result();
  };
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _1_val = _1();
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  assert(
    _1_val.count === 2 &&
      _1_val.text === 'Turn: 2' &&
      _2_val.count === 3 &&
      _2_val.text === 'Turn: 3' &&
      _3_val.count === 4 &&
      _3_val.text === 'Turn: 4' &&
      _4_val.count === 5 &&
      _4_val.text === 'Turn: 5' &&
      _5_val.count === 6 &&
      _5_val.text === 'Turn: 6'
  );
})();
```

Kada se deo `GameOfChance` inicijalizuje na DOM prvi put i svaki put kada se dugme klikne kasnije, treba da se vrati jedan element `h1` koji nasumično prikazuje jednu od `You Win!` ili `You Lose!`. Napomena: ovo može da ne uspe slučajno. Ako se to dogodi, molimo pokušajte ponovo.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    h1: comp.find('h1').length,
    text: comp.find('h1').text()
  });
  const _1 = result();
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _6 = () => {
    simulate();
    return result();
  };
  const _7 = () => {
    simulate();
    return result();
  };
  const _8 = () => {
    simulate();
    return result();
  };
  const _9 = () => {
    simulate();
    return result();
  };
  const _10 = () => {
    simulate();
    return result();
  };
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  const _6_val = _6();
  const _7_val = _7();
  const _8_val = _8();
  const _9_val = _9();
  const _10_val = _10();
  const __text = new Set([
    _1.text,
    _2_val.text,
    _3_val.text,
    _4_val.text,
    _5_val.text,
    _6_val.text,
    _7_val.text,
    _8_val.text,
    _9_val.text,
    _10_val.text
  ]);
  const __h1 = new Set([
    _1.h1,
    _2_val.h1,
    _3_val.h1,
    _4_val.h1,
    _5_val.h1,
    _6_val.h1,
    _7_val.h1,
    _8_val.h1,
    _9_val.h1,
    _10_val.h1
  ]);
  assert(__text.size === 2 && __h1.size === 1);
})();
```

# --seed--

## --seed-contents--

```jsx
class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    {/* Change code below this line */}
    return <h1></h1>;
    {/* Change code above this line */}
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      // Complete the return statement:
      return {
        counter: prevState
      }
    });
  }
  render() {
    const expression = null; // Change this line
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        {/* Change code below this line */}

        {/* Change code above this line */}
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```

# --solutions--

```jsx
// We want this to be deterministic for testing purposes.
const randomSequence = [true, false, false, true, true, false, false, true, true, false];
let index = 0;
const fiftyFifty = () => randomSequence[index++ % randomSequence.length];

class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h1>{this.props.fiftyFifty ? 'You Win!' : 'You Lose!'}</h1>;
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      return {
        counter: prevState.counter + 1
      }
    });
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        <Results fiftyFifty={fiftyFifty()} />
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```
