---
id: 5a24c314108439a4d4036169
title: Prolazni parametri za deo koji nema funkcionalni status
challengeType: 6
forumTopicId: 301402
dashedName: pass-props-to-a-stateless-functional-component
---

# --description--

Prethovni izazovi obuhvatili su mnogo o kreiranju i kombinovanju JSX elemenata, funkcionalskih komponenti i delova klasnih objekata u stilu ES6 u React. Na osnovu ovoga, vreme je da pogledamo još jednu funkciju koja se često ponavlja u React: **prosljeđeni parametri**. U React, možete prosleđivati parametre, ili *props*, za komponente dece. Recimo da imate komponentu `App` koja prikazuje komponentu dece nazvanu `Welcome`, a koja je nefunkcionalna komponenta. Možete proslediti `Welcome` *prop* sa `user` tako što ćete napisati:

```jsx
<App>
  <Welcome user='Mark' />
</App>
```

Koristiš **prilagođene HTML atribute koje React podržava** da proslediš podatke komponenti. U ovom slučaju, prilagođeni atribut `user` prosleđuje se komponenti `Welcome`. Pošto je `Welcome` funkcionalna komponenta bez stanja, ovoj vrednosti pristupa na sledeći način:

```jsx
const Welcome = (props) => <h1>Hello, {props.user}!</h1>
```

Često se naziva vrednost `props` i kada radiš sa komponentama koje nemaju funkcionalsko stanje, zapravo ih tretiraš kao propove za komponentu koja vraća JSX. Možeš dobiti vrednost tog propa unutar tela komponente. Za komponente klase objekata (Class Components), videćeš da je ovo malo drugačije.

# --instructions--

Postoje delovi za `Calendar` i `CurrentDate` u editoru koda. Kada prikazujete `CurrentDate` iz sekcije `Calendar`, prosledite svojstvo `date` koje je postavljeno na trenutni datum vrednost iz objekta `Date` od JavaScript. Zatim dohvatite ovu vrednost `prop` u sekciji `CurrentDate`, i prikažite njenu vrednost unutar tagova `p`. Zapamtite da bi vrednosti `prop` bile tretirane kao JavaScript, moraju biti zatvorene unutar zagrada, na primer `date={Date()}`.

# --hints--

Dio `Calendar` treba da vrati jedan element od `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Komponenta dva od sekcije `Calendar` treba da bude deo sekcije `CurrentDate`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).name() === 'CurrentDate';
  })()
);
```

Dio `CurrentDate` treba da ima svojstvo nazvano `date`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).props().date;
  })()
);
```

Prop za `date` za `CurrentDate` treba da ima niz slova.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    const prop = mockedComponent.children().childAt(1).props().date;
    return typeof prop === 'string' && prop.length > 0;
  })()
);
```

Prop za `date` mora biti kreiran pozivanjem `Date()`

```js
assert(/<CurrentDatedate={Date\(\)}\/>/.test(__helpers.removeWhiteSpace(code)));
```

Deo `CurrentDate` treba da prikaže vrednost iz svojstva `date` u labeli `p`.

```js
let date = 'dummy date';
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(CurrentDate, { date })
    );
    return mockedComponent.find('p').html().includes(date);
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: </p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: {props.date}</p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate date={Date()} />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
