---
id: 5a24c314108439a4d403617a
title: Prosledi stanje kao prop podređenim komponentama
challengeType: 6
forumTopicId: 301403
dashedName: pass-state-as-props-to-child-components
---

# --description--

U prethodnim izazovima video si mnogo primera prosleđivanja propsa podređenim JSX elementima i React komponentama. Možda se pitaš odakle ti propsi dolaze. Uobičajen obrazac je da komponenta sa stanjem (`state`) čuva podatke važne za aplikaciju, a zatim prikazuje podređene komponente. Tim komponentama prosleđuje deo svog `state` kao props.

Na primer, možda imaš komponentu `App` koja prikazuje `Navbar` zajedno sa drugim komponentama. U komponenti `App` postoji `state` koji sadrži razne informacije o korisniku, ali komponenti `Navbar` potrebno je samo korisničko ime kako bi ga prikazala. Zato taj deo `state` prosleđuješ komponenti `Navbar` kao prop.

Ovaj opis pokazuje neke ključne principe u React. Prvo je *jednosmerni tok podataka*. Stanje ulazi u jednom smeru kroz vaše komponentno stablo, od roditeljske komponente ka detinjoj. Detinjica prima samo podatke o stanju koje joj je potrebno. Drugo je da složene komponente sa stanjem mogu se razbiti na nekoliko malih delova, ili možda na jednu komponentu sa stanjem. Sve ostale komponente primaju stanje od roditelja kao prosleđene parametre i prikazuju UI iz tog stanja. Ovo počinje stvarati podelu gde upravljanje stanjem dešava u jednoj komponenti koda, a prikazivanje UI-a u drugoj. Ovaj princip razdvajanja logike stanja i logike UI je jedan od glavnih principa React. Kada se koristi pravilno, čini dizajn složenih komponenti sa stanjem lakšim za upravljanje.

# --instructions--

Komponenta `MyApp` ima stanje i prikazuje komponentu `Navbar` kao dete. Prosledi svojstvo `name` iz svog `state` komponenti `Navbar`, a zatim prikaži `name` u oznaci `h1` unutar metode `render` komponente `Navbar`. `name` treba da se pojavi nakon teksta `Hello, my name is:`.

# --hints--

Komponenta `MyApp` treba da prikazuje komponentu `Navbar` u sebi.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return (
      mockedComponent.find('MyApp').length === 1 &&
      mockedComponent.find('Navbar').length === 1
    );
  })()
);
```

Deo `Navbar` treba da primi atribut stanja `MyApp` `name` kao prosleđene parametre.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').props());
  };
  const navProps = await setState();
  assert(navProps.name === 'TestName');
```

Element [`h1`] u [`Navbar`] treba da prikaže parametar [`name`] .

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const navH1Before = mockedComponent.find('Navbar').find('h1').text();
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').find('h1').text());
  };
  const navH1After = await setState();
  assert(new RegExp('TestName').test(navH1After) && navH1After !== navH1Before);
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         {/* Change code below this line */}
         <Navbar />
         {/* Change code above this line */}
       </div>
    );
  }
};

class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      {/* Change code below this line */}
      <h1>Hello, my name is: </h1>
      {/* Change code above this line */}
    </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         <Navbar name={this.state.name}/>
       </div>
    );
  }
};
class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      <h1>Hello, my name is: {this.props.name}</h1>
    </div>
    );
  }
};
```
