---
id: 5a24c314108439a4d4036167
title: Prikaži deo elementa u DOM-u
challengeType: 6
forumTopicId: 301404
dashedName: render-a-class-component-to-the-dom
---

# --description--

Možda se sećaš da si u prethodnom izazovu koristio ReactDOM API za prikaz JSX elemenata u DOM-u. Proces prikazivanja React komponenti izgleda veoma slično. Nedavni izazovi bili su fokusirani na komponente i njihovu strukturu, pa je renderovanje bilo obavljeno umesto tebe. Međutim, nijedan React kod koji napišeš neće biti prikazan u DOM-u dok ne pozoveš ReactDOM API.

Ovo je sintaksička napomena: `ReactDOM.render(componentToRender, targetNode)`. Prvi argument je deo React koji želite prikazati. Drugi argument je čvor/node DOM u kojem želite prikazati taj deo.

Delovi za React se prosleđuju unutar `ReactDOM.render()` na malo drugačiji način od elemenata JSX. Za JSX elemente, prosleđujete ime elementa koji želite prikazati. Međutim, za delove React, morate koristiti istu sintaksu kao kada prikazujete ugnježđene delove, na primer `ReactDOM.render(<ComponentToRender />, targetNode)`. Koristite ovu sintaksu za delove klasa ES6 i funkcionalne delove.

# --instructions--

Sve sekcije za `Fruits` i `Vegetables` su kreirane u vaše ime iza kulisa. Prikažite sve sekcije kao potomke sekcije `TypesOfFood`, zatim prikažite `TypesOfFood` na DOM. Postoji `div` sa `id='challenge-node'` koje je dostupno za vas da koristite.

# --hints--

Dio od `TypesOfFood` treba da vrati jednu funkciju od `div`.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Deo ``TypesOfFood`` treba da prikaže deo ``Fruits`` nakon elementa ``h1``.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

Sekcija `TypesOfFood` treba da prikaže sekciju `Vegetables` nakon `Fruits`.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

Deo `TypesOfFood` treba da se prikaže na DOM unutar `div` sa identifikatorom `challenge-node`.

```js
assert(
  (function () {
    const hasRenderCall = /ReactDOM\.render\(<TypesOfFood\/>,document\.getElementById\((['"`])challenge-node\1\)\);?/.test(
      __helpers.removeWhiteSpace(code)
    );
    if (!hasRenderCall) {
      return false;
    }

    const e = React.createElement;
    globalThis.Fruits = function Fruits() {
      return e(
        'div',
        null,
        e('h2', null, 'Fruits:'),
        e('h4', null, 'Non-Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Apples'),
          e('li', null, 'Blueberries'),
          e('li', null, 'Strawberries'),
          e('li', null, 'Bananas')
        ),
        e('h4', null, 'Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Lemon'),
          e('li', null, 'Lime'),
          e('li', null, 'Orange'),
          e('li', null, 'Grapefruit')
        )
      );
    };
    globalThis.Vegetables = function Vegetables() {
      return e(
        'div',
        null,
        e('h2', null, 'Vegetables:'),
        e(
          'ul',
          null,
          e('li', null, 'Brussel Sprouts'),
          e('li', null, 'Broccoli'),
          e('li', null, 'Squash')
        )
      );
    };

    ReactDOM.render(
      React.createElement(TypesOfFood),
      document.getElementById('challenge-node')
    );

    const root = document.getElementById('challenge-node').childNodes[0];
    if (!root) {
      return false;
    }
    const html = root.innerHTML;
    return (
      html.includes(
        '<div><h2>Fruits:</h2><h4>Non-Citrus:</h4><ul><li>Apples</li><li>Blueberries</li><li>Strawberries</li><li>Bananas</li></ul><h4>Citrus:</h4><ul><li>Lemon</li><li>Lime</li><li>Orange</li><li>Grapefruit</li></ul></div>'
      ) &&
      html.includes(
        '<div><h2>Vegetables:</h2><ul><li>Brussel Sprouts</li><li>Broccoli</li><li>Squash</li></ul></div>'
      )
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
```

# --solutions--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}
          <Fruits />
           <Vegetables />
         {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
ReactDOM.render(<TypesOfFood />, document.getElementById('challenge-node'));
```
