---
id: 5a24c314108439a4d4036173
title: Postavi stanje koristeći this.setState
challengeType: 6
forumTopicId: 301412
dashedName: set-state-with-this-setstate
---

# --description--

Prethodni izazovi uključivali su deo `state` i kako da se inicijalizuje `state` u `constructor`. Postoji i način za menjanje/izmenu `state` dela. React pruža način za ažuriranje `state` dela nazvanog `setState`. Pozivate metodu `setState` unutar vaše klase delova kao što je sledeće: `this.setState()`, prosleđujući objekat sa parovima ključ-vrednost. Ključ su vaša svojstva stanja (state) a vrednosti su novi podaci stanja (state). Na primer, ako sačuvamo `username` u stanje (state) i želimo da ga promenimo, biće kao ovo:

```jsx
this.setState({
  username: 'Lewis'
});
```

React očekuje se da ne menjate `state` direktno, već koristite `this.setState()` svaki put kada se stanje promeni. Takođe, treba da razumete da React može da spoji više ažuriranja stanja radi poboljšanja performansi. To znači da ažuriranja stanja putem metode `setState` mogu biti asinhrona. Postoji alternativa sintaksa za metodu `setState` koja nudi način da se izbegne ovaj problem. Ovo se ne koristi često, ali je dobro ga zapamtiti! Molimo vas proverite <a href="https://www.freecodecamp.org/news/what-is-state-in-react-explained-with-examples/" target="_blank" rel="noopener noreferrer nofollow"> članak o React</a> za više detalja.

# --instructions--

Postoji element `button` u editoru koda koji ima rukovateljač za `onClick()`. Ovaj rukovateljač se inicijalizuje kada `button` primi događaj klika u pretraživaču, i izvršava metodu `handleClick` definisanu u `MyComponent`. Unutar metode `handleClick`, ažurirajte svojstvo `state` sekcije koristeći `this.setState()`. Postavite atribut `name` u `state` da bude jednak nizu karaktera `React Rocks!`.

Kliknite na dugme i posmatrajte ažuriranja stanja koja se prikazuju. Ne brinite ako još niste potpuno razumeli kako trenutno funkcioniše kod za obradu klikova. Ovo će biti diskutovano u budućim izazovima.

# --hints--

Stanje `MyComponent` treba da se inicijalizuje parom ključ-vrednost `{ name: Initial State }`.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'Initial State'
);
```

`MyComponent` trebalo bi da prikaže naslov sekcije za `h1`.

```js
assert(Enzyme.mount(React.createElement(MyComponent)).find('h1').length === 1);
```

Sekcijsko zaglavlje za `h1` koje je prikazano treba da ima tekst koji potiče iz stanja sekcije.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(/<h1>TestName<\/h1>/.test(firstValue));
```

Pozivanje metode `handleClick` na `MyComponent` treba da postavi atribut 'name' u stanje da bude jednak `React Rocks!`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'Before' });
    return waitForIt(() => mockedComponent.state('name'));
  };
  const second = () => {
    mockedComponent.instance().handleClick();
    return waitForIt(() => mockedComponent.state('name'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Before' && secondValue === 'React Rocks!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
     // Change code below this line
    this.setState({
      name: 'React Rocks!'
    });
    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```
