---
id: 5a24c314108439a4d403616d
title: Koristi PropTypes da definišeš propse koje očekuješ
challengeType: 6
forumTopicId: 301419
dashedName: use-proptypes-to-define-the-props-you-expect
---

# --description--

React pruža funkcije za proveru tipova pomoći kako bi se osiguralo da komponente primaju parametre odgovarajućeg tipa. Na primer, vaš program poziva API da dobije podatke koje očekujete u nizu podataka, koji se zatim prosleđuje komponenti kao argument. Možete postaviti `propTypes` u svoju komponentu da zahteva podatke tipa `array`. Ovo će nam prikazati upozorenje pomoći kada podaci nisu odgovarajućeg tipa.

Smatra se dobrom praksom da definišeš `propTypes` kada unapred znaš tip prop-a. Svojstvo `propTypes` možeš definisati na isti način kao i `defaultProps`. Time proveravaš da li određeni props postoji i da li je odgovarajućeg tipa. Evo primera gde se zahteva prop `handleClick` tipa `function`:

```js
MyComponent.propTypes = { handleClick: PropTypes.func.isRequired }
```

U gornjem primeru, deo `PropTypes.func` proverava da je `handleClick` funkcija. Dodavanje `isRequired` govori React da je `handleClick` potrebno svojstvo za taj deo. Videćete upozorenje ako kriterijum nije dostavljen. Takođe proverite da `func` predstavlja `function`. Među sedam osnovnih tipova od JavaScript, `function` i `boolean` (napisano kao `bool`) su jedini koji koriste neobično pravopisivanje. Pored osnovnih tipova, dostupni su i drugi tipovi. Na primer, možete proveriti da je kriterijum svojstvo od React. Molimo vas pogledajte u <a href="https://reactjs.org/docs/typechecking-with-proptypes.html#proptypes" target="_blank" rel="noopener noreferrer nofollow">dokument</a> za sve opcije.

**Napomena:** Od React v15.5.0, `PropTypes` se razlikuje od React, kao što sledi: `import PropTypes from 'prop-types';`

# --instructions--

Definišite `propTypes` za deo `Items`, zahtevajući `quantity` kao parametar i potvrđujući da je tip `number`.

# --hints--

Deo `ShoppingCart` bi trebalo da prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Deo `Items` treba da prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

Sekcija `Items` trebalo bi da uključuje proveru za `propTypes` koja zahteva vrednost za `quantity` i osigurava da je njena vrednost broj.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(noWhiteSpace.includes('quantity:PropTypes.number.isRequired') && noWhiteSpace.includes('Items.propTypes='));
```

# --before-all--

```js
var PropTypes = {
  number: { isRequired: true }
};
```

# --seed--

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line

// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line
Items.propTypes = {
  quantity: PropTypes.number.isRequired
};
// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```
