---
id: 66d6a7a3e1aa411e94bf2346
title: Dizajn kuće
challengeType: 25
dashedName: build-a-house-painting
demoType: onClick
---

# --description--

U ovoj laboratoriji, koristićeš HTML za kreiranje dizajna kuće. Zatim, koristićeš raspored od CSS da rasporediš komponente svoje kuće kao što su prozori i vrata.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba da imaš `div` element sa `id` vrednošću `house`.
1. Tvoj `#house` treba da ima podešen `position` na `relative` kako bi njegova deca mogla biti pozicionirana u odnosu na njega.
1. Tvoj `#house` treba da ima širinu od `500px` i visinu od `400px`. Nemoj postavljati manje ili veće vrednosti.
1. Tvoj `#house` treba da ima postavljenu boju pozadine i okvir.
1. Treba da imaš pet `div` elemenata unutar `#house` sa sledećim `id` vrednostima: `chimney`, `roof`, `window-1`, `window-2` i `door`.
1. Watoto wote wa moja kwa moja wa `#house` wanapaswa kuwa na `position` ya `absolute`.
1. `#roof`, `#chimney`, `#window-1`, `#window-2` i `#door` treba da imaju postavljene širinu, visinu, okvir i boju pozadine.
1. Tvoj `#roof` treba da ima vrednost `top` veću od `0`.
1. Tvoj `#door` treba da bude postavljen na dnu kuće.
1. Tvoji prozori treba da budu postavljeni ispod `#roof` i najmanje više od jedne trećine visine tvojih vrata.
1. Tvoji prozori i vrata treba da imaju podešen `left` ili `right` sa vrednošću koja ih postavlja unutar granica kuće.
1. Tvoj `#chimney` treba da ima vrednost `top` koja ga postavlja iznad `#house`.
1. Tvoj `#chimney` treba da ima `z-index` koji ga postavlja iza kuće.

**Napomena:** Uverite se da ste povezali vaš stilski list (stylesheet) u vaš HTML kako biste koristili vaš CSS.

# --hints--

Trebalo bi da imate `div` sa `id` od `house`.

```js
assert.exists(document.querySelector("div#house"));
```

Trebalo bi da ciljate na `#house` i postavite njegovu `position` na `relative`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("position"), "relative");
```

Trebalo bi da ciljate na `#house` i postavite mu atribut `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("border"));
```

Trebalo bi da ciljate na `#house` i postavite mu atribut `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("background-color"));
```

`#house` ne bi trebalo da ima svojstvo od `min-height`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("min-height"));
```

`#house` ne bi trebalo da ima svojstvo od `min-width`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("min-width"));
```

`#house` ne bi trebalo da ima karakteristiku `max-height`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("max-height"));
```

`#house` ne bi trebalo da ima osobinu od `max-width`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("max-width"));
```

Trebalo bi da ciljate `#house` i postavite mu širinu na `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("#house")?.getPropertyValue("width"), "500px")
```

Trebalo bi da ciljate `#house` i postavite mu visinu na `400px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#house')?.getPropertyValue("height"), "400px")
```

Trebalo bi da imaš `div` sa `id` u `chimney` unutar `#house`.

```js
assert.exists(document.querySelector("div#house div#chimney"));
```

`#chimney` mora imati atribut `position`, postavljen na `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("position"), "absolute");
```

Trebalo bi da ciljate na `#chimney` da postavite njegovu osobinu `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("width"));
```

Trebalo bi da ciljate na `#chimney` da biste postavili njegov atribut `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("height"));
```

Treba ti da ciljaš na `#chimney` da postaviš njegovu osobinu `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("border"));
```

Trebalo bi da ciljate na `#chimney` da biste postavili njegovu osobinu `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#chimney")?.getPropertyValue("background-color"));
```

Trebalo bi da imate `div` sa `id` u `roof` unutar `#house`.

```js
assert.exists(document.querySelector("div#house  div#roof"));
```

`#roof` mora imati atribut `position` postavljen na `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("position"), "absolute");
```

Trebalo bi da ciljate na `#roof` da biste postavili njegovu osobinu `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("width"));
```

Trebalo bi da ciljate na `#roof` da biste postavili njegovu osobinu `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("height"));
```

Treba ti da ciljaš na `#roof` da postaviš njegov atribut `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("border"));
```

Treba da ciljaš `#roof` kako bi mu dodelio atribut `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#roof")?.getPropertyValue("background-color"));
```

Trebalo bi da imaš `div` sa `id` unutar `window-1` u `#house`.

```js
assert.exists(document.querySelector("div#house  div#window-1"));
```

`#window-1` treba da ima svojstvo `position` postavljeno na `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#window-1"))?.getPropertyValue("position"), "absolute");
```

Trebalo bi da ciljate na `#window-1` da biste postavili njegov atribut `width`.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window1 || windows);
```

Trebalo bi da usmeriš `#window-1` i postavi njegovo svojstvo od `height`.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window1 || windows);
```

Treba da ciljate `#window-1` i postavite njegovo svojstvo od `border`.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window1 || windows);
```

Trebalo bi da ciljate na `#window-1` da biste postavili njegovu osobinu `background-color`.

```js
const window1 = new __helpers.CSSHelp(document).getStyle("#window-1")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window1 || windows);
```

Trebalo bi da imaš `div` sa `id` u `window-2` unutar `#house`.

```js
assert.exists(document.querySelector("div#house  div#window-2"));
```

`#window-2` mora imati svojstvo `position` postavljeno na `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#window-2"))?.getPropertyValue("position"), "absolute");
```

Trebalo bi da ciljate na `#window-2` da biste postavili njegovu osobinu `width`.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("width");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("width") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("width");
assert.isNotEmpty(window2 || windows);
```

Trebalo bi da ciljate na `#window-2` da biste postavili njegovu osobinu `height`.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("height");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("height") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("height");
assert.isNotEmpty(window2 || windows);
```

Trebalo bi da ciljate na `#window-2` da biste postavili njegovu osobinu `border`.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("border");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("border") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("border");
assert.isNotEmpty(window2 || windows);
```

Trebalo bi da ciljate na `#window-2` da biste postavili njegov atribut `background-color`.

```js
const window2 = new __helpers.CSSHelp(document).getStyle("#window-2")?.getPropertyValue("background-color");
const windows = new __helpers.CSSHelp(document).getStyle("#window-1, #window-2")?.getPropertyValue("background-color") || new __helpers.CSSHelp(document).getStyle("#window-2, #window-1")?.getPropertyValue("background-color");
assert.isNotEmpty(window2 || windows);
```

Trebalo bi da imaš `div` sa `id` u `door` unutar `#house`.

```js
assert.exists(document.querySelector("div#house  div#door"));
```

`#door` mora imati svojstvo `position` postavljeno na `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#door"))?.getPropertyValue("position"), "absolute");
```

Treba ti ciljati `#door` da bi mu dodelio njegovu atribut `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("width"));
```

Trebalo bi da ciljate na `#door` da biste postavili njegovo svojstvo `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("height"));
```

Treba da ciljaš `#door` kako bi mu dodelio atribut `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("border"));
```

Treba da ciljaš na `#door` kako bi mu dodelio njegovu osobinu `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("#door")?.getPropertyValue("background-color"));
```

`#roof` vaš treba da ima svojstvo od `top`, postavljeno na `0`.

```js
assert.equal(getComputedStyle(document.querySelector("#roof"))?.getPropertyValue("top"), "0px");
```

`#door` To treba postaviti ispod kuće.

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorBottom = Number(door.getPropertyValue("bottom").replace("px", ""));
assert.strictEqual(doorBottom, 0);
```

Treba da staviš `left` ili `right` u tvoje `#door` i da ga rasporediš unutar granica kućišta.

```js
const door = getComputedStyle(document.querySelector("#door"))
const doorLeft = Number(door.getPropertyValue("left").replace("px", ""));
const doorRight = Number(door.getPropertyValue("right").replace("px", ""));
assert.isTrue(doorLeft > 0 || doorRight > 0);
```

`#chimney` vaš treba da ima vrednost `top` koja ga postavlja iznad `#house` vašeg.

```js
const chimney = getComputedStyle(document.querySelector("#chimney"));
assert.approximately(parseFloat(chimney?.getPropertyValue("top")), - parseFloat(chimney?.getPropertyValue("height")), 2);
```

`#chimney` trebalo bi da ima `z-index` koja ga postavlja iza kuće.

```js
const houseZIndex = getComputedStyle(document.querySelector("#house"))?.getPropertyValue("z-index");
const chimneyZIndex = getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("z-index");
if (houseZIndex === "auto") {
    assert.isBelow(Number(chimneyZIndex), 0)
} else {
    assert.isBelow(Number(chimneyZIndex), Number(houseZIndex));
}
```

Vaši prozori bi trebalo postaviti ispod vašeg `#roof` i da budu na visini od najmanje jedne trećine visine vašeg `#door`.

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const door = getComputedStyle(document.querySelector("#door"))
const roof = getComputedStyle(document.querySelector("#roof"))
const w1Top = Number(window1.getPropertyValue("top").replace("px", ""));
const w1Bottom = Number(window1.getPropertyValue("bottom").replace("px", ""));
const w2Top = Number(window2.getPropertyValue("top").replace("px", ""));
const w2Bottom = Number(window2.getPropertyValue("bottom").replace("px", ""));
const doorHeight = Number(door.getPropertyValue("height").replace("px", ""));
const roofHeight = Number(roof.getPropertyValue("height").replace("px", ""));
assert.isAbove(w1Top, roofHeight);
assert.isAbove(w1Bottom, doorHeight/3);
assert.isAbove(w2Top, roofHeight);
assert.isAbove(w2Bottom, doorHeight/3);
```

Trebalo bi da postavite `left` ili `right` u vaše prozore i da ih rasporedite unutar granica kuće.

```js
const window1 = getComputedStyle(document.querySelector("#window-1"))
const window2 = getComputedStyle(document.querySelector("#window-2"))
const w1Left = Number(window1.getPropertyValue("left").replace("px", ""));
const w1Right = Number(window1.getPropertyValue("right").replace("px", ""));
const w2Left = Number(window2.getPropertyValue("left").replace("px", ""));
const w2Right = Number(window2.getPropertyValue("right").replace("px", ""));
assert.isTrue(w1Left > 0 || w1Right > 0);
assert.isTrue(w2Left > 0 || w2Right > 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div id="house">
        <div id="chimney">
          <div id="smoke"></div>
        </div>
        <div id="roof"></div>
        <div id="window-1"></div>
        <div id="window-2"></div>
        <div id="door">
          <div id="door-knob"></div>
        </div>
        <div id="welcome">WELCOME</div>
    </div>
</body>
</html>
```

```css
* {
    box-sizing: border-box;
  }

  body {
    margin: 0;
    padding: 0;
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    background: linear-gradient(
      #b3e6ff,
      #b3e6ff 60%,
      green 60%
    );
  }

  #house {
    flex: 0 0 auto;
    width: 500px;
    height: 400px;
    position: relative;
    background-color: #ff9980;
    border: 6px solid #b35900;
  }

  #chimney {
    position: absolute;
    width: 90px;
    height: 100px;
    top: -100px;
    left: 65%;
    background: repeating-linear-gradient(
      #e6e6e6,
      #e6e6e6 10%,
      black 10%,
      black 12%
    );
    border: 2px solid black;
    z-index: -1;
  }

  #smoke {
    width: 40px;
    height: 40px;
    background-color: rgba(0,0,0,0.3);
    position: absolute;
    top: -20px;
    left: 20px;
    border-radius: 50%;
    box-shadow: 0 -10px 10px 10px rgba(0,0,0,0.3);
    border: none;
  }

  #roof {
    border: 2px;
    width: 100%;
    height: 110px;
    background-color: green;
    position: absolute;
    top: 0;
    background: repeating-linear-gradient(
      45deg,
      #b35900,
      #b35900 2%,
      transparent 2%,
      transparent 5%
    ), repeating-linear-gradient(
      -45deg,
      #b35900,
      #b35900 2%,
      #ff9980 2%,
      #ff9980 5%
    )
  }

  #window-1, #window-2 {
    width: 100px;
    height: 100px;
    background: linear-gradient(
      to right,
      #ffffb3,
      #ffffb3 48%,
      #b35900 48%,
      #b35900 52%,
      #ffffb3 52%,
      #ffffb3 100%
    );
    position: absolute;
    top: 40%;
    border: 6px solid #b35900;
  }

  #window-1 {
    left: 7.5%;
  }

  #window-2 {
    right: 7.5%;
  }

  #door {
    width: 130px;
    height: 180px;
    position: absolute;
    bottom: 0%;
    left: 185px;
    background-color: #e6e6e6;
    border: 6px solid #b35900;
    transform: translateY(6px);
  }

  #door-knob {
    width: 20px;
    height: 20px;
    background-color: #b35900;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    right: 5%;
  }

  #welcome {
    font-family: sans-serif;
    font-weight: bold;
    width: 130px;
    height: 40px;
    position: absolute;
    bottom: -47px;
    left: 178px;
    background-color: #85e085;
    border: 2px solid black;
    transform: skewX(-20deg);
    display: flex;
    align-items: center;
    justify-content: center;
  }
```
