---
id: 6823f9df49cc206af5471a30
title: Uredi profilnu stranicu Camperbota
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Camperbot pokušava da kreira stranicu profila. Zamolio je svog prijatelja da proveri njegov kod, i njegov prijatelj je rekao da ima nekoliko grešaka.

Vaš zadatak je da popravite sve greške Camperbota kako bi mogao nastaviti kreiranje svoje stranice profila. Popunite stavke u pričama korisnika ispod i kliknite na "Check Your Code" da vidite da li ste uklonili sve greške.

Priče o korisniku:

1. Camperbot pokušava da koristi element `heading2`, ali taj element ne postoji. Ispravite te oznake tako da koriste odgovarajući naslovni podnaslov drugog nivoa.
2. Camperbot pokušava da doda dva odlomka koristeći `pp`, ali i oni ne postoje. Ispravite ga tako da koristi odgovarajuće oznake odlomaka.
3. Camperbot koristi element `h3` za podnaslov `Background and Interests`, ali ima sintaksičku grešku. Identifikujte problem i rešite ga.

# --hints--

Trebalo bi da imaš samo jedan element `h1` na stranici.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Ne smete da uklonite element `<h1>Hello from Camperbot!</h1>` sa stranice.

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

Trebalo bi da imate samo jedan element `h2` na stranici.

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

Trebalo bi da imaš element sa tekstom `h2` koji je `About`. Evo primera: `<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

Trebalo bi da imaš paragraf sa tekstom `My name is Camperbot and I love learning new things.` Evo primera: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

Trebalo bi da imate paragraf sa tekstom `I enjoy solving puzzles.` Ovde je primer: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

Trebalo bi da imaš dva elementa paragrafa na stranici.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Trebalo bi da imate samo jedan element `h3` na stranici.

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

Trebalo bi da imaš element sa tekstom `h3` koji sadrži tekst `Background and Interests`. Evo primera: `<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

Trebalo bi da imaš ukupno pet elemenata na stranici.

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
