---
id: 69a92f0b9cc04bb0d5327bb5
title: Implementiraj strukturu podataka stoga/za stack
challengeType: 26
dashedName: implement-a-stack
---

# --description--

U ovoj laboratoriji, implementiraÄ‡ete strukturu podataka 'stack' koristeÄ‡i funkcije. Stack je struktura podataka tipa Last-In-First-Out (LIFO), gde se elementi dodaju i uklanjaju sa vrha.

**Cilj:** Popunite korisniÄŤke priÄŤe navedene ispod i proÄ‘ite sve testove da biste zavrĹˇili laboratoriju.

**PriÄŤe korisnika:**

1. Potrebna je funkcija `initStack` koja vraÄ‡a objekat sa karakteristikom `collection`, postavljen u prazan niz podataka.
1. Potrebna je funkcija `push` koja dodaje element na vrh steka (stack).
1. Potrebna je funkcija `pop` koja uklanja i vraÄ‡a gornji element sa steka, ili `undefined` ako ne postoji.
1. Potrebna je funkcija `peek` koja vraÄ‡a gornji element sa steka bez uklanjanja, ili `undefined` ako ne postoji.
1. Potrebna je funkcija `isEmpty` koja vraÄ‡a `true` ako stek nema elemenata, i `false` inaÄŤe.
1. Potrebna je funkcija `clear` koja uklanja sve elemente sa steka.

**Napomena:** Mnoge mere zavise od funkcija `initStack` i `push`. IzvrĹˇite prvo, jer mere za `pop`, `peek`, `isEmpty`, i `clear` zahtevaju dodavanje elemenata na stak.

# --hints--

Trebalo bi da imaĹˇ funkciju `push`.

```js
assert.isFunction(push);
```

Funkcija `push` treba da doda element na vrh steka.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

Trebalo bi da imaĹˇ funkciju `pop`.

```js
assert.isFunction(pop);
```

Funkcija `pop` mora da ukloni i vrati najgornji element steka.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.strictEqual(pop(stack), 'second');
```

Funkcija `pop` treba da vrati `undefined` ako je stek prazan.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

Funkcija `pop` treba da vrati boolean vrednosti taÄŤno.

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

Trebalo bi da imaĹˇ funkciju `peek`.

```js
assert.isFunction(peek);
```

Funkcija `peek` bi trebalo da vrati najgornji element steka bez uklanjanja.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
```

Funkcija za `peek` treba da vrati `undefined` ako je stos prazan.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

Funkcija `peek` treba da vrati booleanske vrednosti taÄŤno.

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

Trebalo bi da imaĹˇ funkciju `isEmpty`.

```js
assert.isFunction(isEmpty);
```

Funkcija za `isEmpty` treba da vrati `true` na prazan stek.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

Funkcija `isEmpty` treba da vrati `false` na neprvu praznu stogu.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

Funkcija za `isEmpty` treba da vrati `false` kada gornji element ima vrednost laĹľno.

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

Trebalo bi da imate funkciju `clear`.

```js
assert.isFunction(clear);
```

Funkcija `clear` mora da ukloni sve elemente sa steka.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
