---
id: 61695d1fbc003856628bf561
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Svojstvo atributa `charset` opisuje kodiranje znakova koje koristi dokument. `utf-8` (Unicode Transformation Format – 8-bit) je standard kodiranja znakova koji se koristi za elektronsku komunikaciju.

Unutar elementa `head`, postavite unutar elementa `meta` sa atributom elementa `charset` koji je postavljen na `"utf-8"`.

# --hints--

Trebalo bi da imaš jedan element `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

Element sa vašim `meta` bi trebalo da bude element bez sadržaja, ne sme imati zatvarajuću oznaku `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Element vašeg `meta` mora imati svojstvo elementa `charset` postavljeno na vrednost `"utf-8"`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Vaš element od `meta` mora da bude unutar elementa `head`.

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
