---
id: 61696ef7ac756c829f9e4048
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

U ovom projektu ćete koristiti eksterni fajl CSS za stilizovanje stranice. Već smo vam pripremili fajl `styles.css`. Ali pre nego što ga koristite, morate ga povezati sa stranicom.

Postavite unutar elementa `link` u elementu `head`. Specifično svojstvo elementa `rel` mora imati vrednost `"stylesheet"`, a svojstvo elementa `href` mora imati vrednost `"styles.css"`.

# --hints--

Trebalo bi da imate jedan element `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Tvoj element za `link` treba da bude element bez zatvarajuće oznake `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Element za `link` mora imati atribut elementa `rel` sa vrednošću `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Element vašeg `link` treba da ima svojstvo polja `href` sa vrednošću `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Element vašeg `link` treba da bude unutar elementa `head`.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
