---
id: 616965351e74d4689eb6de30
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Na web stranici možete imati mnogo elemenata `meta`. Svaki element od `meta` dodaje informacije o stranici koje ne mogu biti prikazane drugim elementima od HTML.

Dodajte još jedan element za `meta` unutar `head`. Atribut `name` podešeno na `"viewport"` i atribut `content` podešeno na `"width=device-width, initial-scale=1.0"` kako bi vaša stranica izgledala dobro na svim uređajima.

# --hints--

Trebalo bi da ima dva elementa za `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

Tvoj novi element za `meta` mora biti void element, ne sme imati zatvarajući atribut `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Tvoja nova funkcija za `meta` mora imati atribut `name` postavljen na `"viewport"`, i atribut `content` postavljen na `"width=device-width, initial-scale=1.0"`.

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

Tva dva elementa za `meta` treba da budu unutar elementa `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
