---
id: 61695197ac34f0407e339882
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Kao što ste videli u prethodnim projektima, web stranice bi trebalo da počinju sa deklaracijom `DOCTYPE html`, praćeno elementom `html`.

Dodajte deklaraciju `DOCTYPE html` na sam vrh dokumenta, a zatim dodajte `html` element. Uklonite atribut `lang` iz `html` elementa sa vrednošću `en`.

# --hints--

Vaš kod treba da ima referencu na `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Treba da staviš razmak nakon reference `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Trebalo bi da definišete tip dokumenta kao `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Trebalo bi da zatvorite deklaraciju `DOCTYPE` sa etiketom za zatvaranje `>` nakon tipa.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Tvoj element `html` treba da ima labelu za otvaranje sa svojstvom elementa `lang` od `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Element vašeg `html` treba da ima zatvarajući tag.

```js
assert(code.match(/<\/html\s*>/gi));
```

Tvoje saopštenje za `DOCTYPE` trebalo bi biti na početku tvog HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
