---
id: 61695ab9f6ffe951c16d03dd
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Unesi element `head` unutar elementa `html`. Neposredno nakon elementa `head`, dodaj element `body`.

# --hints--

Trebalo bi da imate etiketu za otvaranje `head`.

```js
assert(code.match(/<head\s*>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Trebalo bi da imaš oznaku za otvaranje `body`.

```js
assert(code.match(/<body\s*>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Trebalo bi da postoji samo jedna etiketa za otvaranje `head`.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

Trebala bi postojati samo jedna etiketa za zatvaranje od `head`.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

Trebala bi da postoji samo jedna etiketa za otvaranje od `body`.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

Trebala bi da postoji samo jedna etiketa za zatvaranje `body`.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

Kipengele chako cha `head` treba da bude prazno.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

Vaše polje `body` treba da bude prazno.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

Tvoj element `body` mora biti smešten nakon tvog elementa `head`.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

Tvoj element `head` mora biti unutar elementa `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Tvoj element `body` treba da bude unutar elementa `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
