---
id: 61695c4aad56f95497c19583
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Zapamtite da element `title` pruža dodatne informacije za pretražne motore o stranici. Takođe prikazuje sadržaj elementa `title` na još dva načina:

* Na traku naslova kada je stranica otvorena
* Na kartici pretraživača za stranicu kada se mišem pređe preko nje. Čak i ako ta kartica ne radi, kada prevedete mišem preko kartice, prikazan je tekst `title`.

Unutar elementa `head`, postavi unutar elementa `title` koji ima tekst `Colored Markers`.

# --hints--

Trebalo bi da imaš etiketu pakovanja od `title`.

```js
assert(code.match(/<title\s*>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Tvoj projekat bi trebalo da ima zaglavlje od `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Zapamtite, velika i mala slova zajedno sa pravopisom su važni za naslove.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
