---
id: 616968c2c94c8071b349c146
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Sada ste spremni da počnete sa dodavanjem sadržaja na stranicu.

Unutar `body`, unesite u polje `h1` koje sadrži tekst `CSS Color Markers`.

# --hints--

Tvoj kod mora da ima element `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

Trebalo bi da imaš etiketu pakovanja za `h1`.

```js
assert(code.match(/<h1\s*>/i));
```

Trebalo bi da imaš etiketu za zatvaranje `h1`.

```js
assert(code.match(/<\/h1\s*>/i));
```

Tvoj element za `h1` treba da bude unutar elementa `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

Vaš element za `h1` treba da ima tekst od `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Zapamti, velika i mala slova zajedno sa pravopisom su važni za tekstualne sadržaje `h1`.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
