---
id: 616971b3cd990186b66c99a1
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Sada nakon što je vaš eksterni fajl CSS pripremljen, možete početi da dekorujete stranicu.

Kao podsjetnik, ovako ciljate na element paragrafa i poravnajte ga desno:

```css
p {
  text-align: right;
}
```

Napravi novi zakon za CSS koji cilja na element `h1`, i postavi mu atribut `text-align` da bude `center`.

# --hints--

Trebalo bi da koristiš alat `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

Vaše pravilo za CSS za `h1` trebalo bi da ima atribut `text-align` postavljen na `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
