---
id: 616d595270d902f0e7086e18
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Kada imate tri različita elementa oznaka `div`, oni izgledaju kao jedan veliki pravougaonik. Trebalo bi dodati razmak između njih kako bi bilo lakše videti svaki element.

Kada atribut `margin` sa kratkim putem ima dve vrednosti, postavlja `margin-top` i `margin-bottom` na prvu vrednost, a `margin-left` i `margin-right` na drugu vrednost.

U vašem pravilu od `.marker` za CSS, postavite atribut `margin` na `10px auto`.

# --hints--

Vaše pravilo od `.marker` za CSS trebalo bi da ima atribut `margin` postavljen na `10px auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```
