---
id: 616d3a67ccf800ad94ec89ae
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Sada ćete dodati elemente koje na kraju ćete organizovati bojama kao oznakama.

Prvo, unutar elementa sa `body`, dodaj element sa `div` i postavi atribut za `class` na `container`. Osigurajte da je element sa `div` ispod elementa sa `h1`.

# --hints--

Trebalo bi da imate etiketu pakovanja za `div`.

```js
assert(code.match(/<div.*>/i));
```

Trebalo bi da imate oznaku za pakovanje `div`.

```js
assert(code.match(/<\/div\s*>/i));
```

Vaš element za `div` mora biti unutar vašeg elementa za `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Vaš element `div` treba biti nakon elementa `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Vaš element za `div` treba da ima svojstvo `class` postavljeno na `container`.

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
