---
id: 616d4a84b756d9c4b8255093
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Vreme je da se doda boja oznaci. Zapamtite da je jedan od načina dodavanja boje elementu korišćenje <dfn>osnovno ime boje</dfn> kao `black`, `cyan`, ili `yellow`.

Kao podsjetnik, ovo je kako ciljate klasu objekata `freecodecamp`:

```css
.freecodecamp {
  
}
```

Kreiraj novo pravilo za CSS koje cilja klasu objekata `marker`, i postavi mu svojstvo od `background-color` da bude `red`.

**Napomena:** Nećete videti nikakve promene nakon dodavanja CSS.

# --hints--

Treba da kreirate selektor za klasu objekata ili da ciljate klasu objekata `marker`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Vaše pravilo za CSS za `.marker` trebalo bi da ima svojstvo `background-color` postavljeno na `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
