---
id: 616d50b93ba424d6282c99cf
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Pozadina je korišćena, ali pošto element oznake `div` nema sadržaj unutra, ne poseduje ni dužinu po defaultu.

U vašoj skemi za `.marker` od CSS, postavite atribut `height` na `25px` i atribut `width` na `200px`.

# --hints--

Vaše pravilo `.marker` za CSS trebalo bi da ima svojstvo `width` postavljeno na `200px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

Vaše pravilo od `.marker` za CSS mora imati svojstvo `height` postavljeno na `25px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
