---
id: 616d525007b8c5d8b3308b1c
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Vaš simbol će biti bolji ako je u centru stranice. Najlakši način za to je korišćenjem `margin` <dfn>funkcija skraćivanja</dfn>.

U prethodnom projektu, definisano je područje ivica elemenata pojedinačno sa svojstvima kao što su `margin-top` i `margin-left`. Svojstvo skraćivanja `margin` čini ga lakšim za postavljanje više područja ivice odjednom.

Da biste postavili svoj znak na sredini stranice, postavite mu atribut od `margin` na `auto`. Ovo će postaviti `margin-top`, `margin-right`, `margin-bottom`, i `margin-left` sve na `auto`.

# --hints--

Vaše pravilo za `.marker` CSS mora imati svojstvo od `margin` postavljeno na `auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
