---
id: 616d55bd160a95e22453a081
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Sada si postavio jednu oznaku u sredini sa bojom, vreme je da dodaš i druge oznake.

U `container` `div`, dodaj dva elementa od `div` i daj svakom od njih klasu `marker`.

# --hints--

Vaša prva nova funkcija od `div` mora imati oznaku za otvaranje.

```js
assert([...code.matchAll(/<div.*?>/gi)][2]);
```

Vaša prva nova komponenta za `div` mora imati zatvarajući tag.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][2]);
```

Vaša druga nova funkcija `div` treba da ima etiketu za otvaranje.

```js
assert([...code.matchAll(/<div.*?>/gi)][3]);
```

Vaš drugi novi element `div` treba da ima zatvarajuću oznaku.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][3]);
```

Tvoje nove komponente iz `div` bi trebalo da budu unutar `div` sa klasom `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.localName === 'div') && containerChildren.length === 3);
```

Svi tvoji novi elementi sa klasom `div` moraju imati klasu `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
