---
id: 61764fdda535587e1fefb3aa
title: Korak 19
challengeType: 0
dashedName: step-19
---

# --description--

Dodaj klasu stavki `two` za drugi znak `div`, i dodaj klasu stavki `three` za treći znak `div`.

# --hints--

Trebalo bi dodati klasu objekata `two` na atribut drugog elementa `div` u `container` `div`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(containerSecondChild?.classList?.contains('two'));
```

Tvoja druga oznaka `div` mora imati klase objekata `marker` i `two`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(containerSecondChild?.classList?.contains('marker') && containerSecondChild.classList?.contains('two'));
```

Treba da dodate klasu objekata `three` na element trećeg atributa `div` u `container` `div`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(containerThirdChild?.classList?.contains('three'));
```

Tvoj treći element `div` treba da ima klase objekata `marker` i `three`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(containerThirdChild?.classList?.contains('marker') && containerThirdChild?.classList?.contains('three'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

```
