---
id: 61764c602bee6974e7790f35
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Da biste dodelili različite boje elementima, moraćete dodati jedinstvenu klasu za svaki. Više klasa se mogu dodati elementu navođenjem ih u atribut ``class`` i razdvajanjem razmakom. Na primer, sledeće dodaje klase ``animal`` i ``dog`` elementu ``div``.

```html
<div class="animal dog">
```

Ako dodate mnogo klasama za element HTML, stilovi prvih navedenih klasa mogu biti nadjačani kasnijim klasama.

Počni, dodaj klasu za `one` za element prvog znaka `div`.


# --hints--

Trebalo bi da dodaš klasu `one` za prvi element oznake `div`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

Tvoja prva oznaka `div` mora imati klase `marker` i `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
