---
id: 6176519636a76b810ab1219a
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Napravi pravilo za CSS koje cilja klasu `two` i postavlja njeno svojstvo `background-color` na `green`.

Pia, kreiraj drugačiju pravilo za CSS koje cilja na klasu objekata `three` i postavi njegovo `background-color` da bude `blue`.

# --hints--

Trebalo bi da koristiš klasifikator objekata za ciljanje klase objekta `two`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two'));
```

Vaše pravilo za CSS od `.two` trebalo bi da ima svojstvo `background-color` postavljeno na `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'green');
```

Trebalo bi da koristiš selektor klase objekata da cilja na klasu objekata `three`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three'));
```

Vaše pravilo za CSS sa `.three` mora imati svojstvo od `background-color` postavljeno na `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
