---
id: 617bd6ec666b1da2587e4e37
title: Korak 34
challengeType: 0
dashedName: step-34
---

# --description--

Postoje tri boje za treći nivo: chartreuse zelena (zeleno + žuto), azurno (plavo + cijan), i ružičasta (crveno + magenta).

Da biste kreirali žalfrijev zeleni, ažurir funkciju za `rgb` u pravilu CSS od `.one`, da je crvena u `127`, i postavite zeleni na najvišu vrednost.

Za Azure, ažurirajte funkciju `rgb` u pravilu CSS od `.two` tako da je zelena na `127` i da plava ima najvišu vrednost.

I za ružičastu, koja se ponekad naziva svetlo roze, ažuriraj funkciju `rgb` u pravilo CSS od `.three` tako da je plava u `127` i da crvena ima najvišu vrednost.

# --hints--

Vaše pravilo za CSS za `.one` trebalo bi imati svojstvo od `background-color` postavljeno na `rgb(127, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(127, 255, 0)');
```

Tvoje pravilo za CSS za `.two` trebalo bi da ima svojstvo od `background-color` postavljeno na `rgb(0, 127, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 255)');
```

Vaše pravilo za CSS od `.three` mora imati svojstvo `background-color` postavljeno na `rgb(255, 0, 127)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
