---
id: 617b8b38f32bf2080a140675
title: Korak 22
challengeType: 0
dashedName: step-22
---

# --description--

Funkcija je deo koda koji može da primi ulaz i izvrši specifičnu akciju. Funkcija CSS `rgb` prihvata vrednosti, ili <dfn>hoja</dfn>, za crveno, zeleno i plavo, i stvara boje:

```css
rgb(red, green, blue);
```

Svaka vrednost crvene, zelene i plave je broj od `0` do `255`. `0` znači da nema te boje uopšte (0%) i da je crna. `255` znači da postoji 100 procenata te boje.

U pravilu CSS od `.one`, promeni osnovnu boju `red` u funkciji `rgb`. Za funkciju `rgb`, postavi crvenu vrednost na `255`, zelenu vrednost na `0` i plavu vrednost na `0`.

# --hints--

Vaše pravilo za CSS od `.one` ne bi trebalo da koristi osnovnu reč boje iz `red` za postavljanje atributa `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor !== 'red');
```

Vaše pravilo za CSS za `.one` trebalo bi da ima svojstvo `background-color` postavljeno na `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: red;
}
--fcc-editable-region--

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
