---
id: 617b954d9f4f6217a749380e
title: Korak 25
challengeType: 0
dashedName: step-25
---

# --description--

Sada povećajte vertikalni razmak malo više između vaših oznaka i ivica elementa `container` koji se nalaze unutra.

U sintaksi `.container` za CSS, koristite kratki atribut `padding` da povećate `10px` unutrašnji razmak za vrh i dno, i postavite unutrašnji razmak sa leve i desne strane na `0`. Ovo radi isto kao kratki atribut `margin` koji ste ranije koristili.

# --hints--

Ne možete ukloniti atribut `background-color` i njegovu vrednost iz pravila za `.container` odnosi se na CSS.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

Tvoja pravila za `.container` od CSS bi trebala imati svojstvo `padding` postavljeno na `10px 0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.padding === '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
