---
id: 617b65579ce424bf5f02ca73
title: Korak 21
challengeType: 0
dashedName: step-21
---

# --description--

Postoje dva glavna modela boja: model od <dfn>dodavanje</dfn> RGB (crveni, zeleni, plavi) koji se koristi u elektronskim uređajima, i model od <dfn>uklanjanje</dfn> CMYK (cijan, magenta, žuti, crni) koji se koristi u štampi.

U ovom projektu, koristićeš RGB model. To znači da boja počinje kao crna i menja se kako su uneti različiti nivoi crvene, zelene i plave. Jednostavan način da ovo vidiš je korišćenjem funkcije CSS `rgb`.

Kreiraj novo pravilo za CSS koje cilja na klasu objekata `container` i postavi njegov atribut `background-color` na crnu za `rgb(0, 0, 0)`.

# --hints--

Trebalo bi da koristiš klasifikator objekata za ciljanje klase objekta `container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

Tvoja pravila za CSS koja su u `.container` treba da imaju karakteristiku `background-color` postavljenu na `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
