---
id: 617b8e0d93a8d10d9a90e720
title: Korak 23
challengeType: 0
dashedName: step-23
---

# --description--

Molimo proverite da je vaš rezultat `background-color` i dalje crven. Ovo je zato što ste postavili crvenu vrednost parametra `rgb` na vrednost veću od `255`, ili crvenih 100 procenata, i postavili sve zelene i plave vrednosti na `0`.

Sada koristi funkciju `rgb` da postaviš druge boje.

U pravilima CSS od `.two`, koristi funkciju `rgb` da postavi `background-color` za najvecu zelenu vrednost, i `0` za ostale vrednosti. I u pravilima CSS od `.three`, koristi funkciju `rgb` da postavi `background-color` za najvecu plavu vrednost, i `0` za ostale vrednosti.

# --hints--

Vaše pravilo za CSS od `.two` ne bi trebalo da koristi osnovnu boju reč iz `green` za postavljanje svojstva/atributa `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

Tvoje pravilo za CSS od `.two` treba da ima svojstvo `background-color` postavljeno na `rgb(0, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

Vaše pravilo od CSS za `.three` ne bi trebalo da koristi osnovnu reč boje `blue` za postavljanje atributa/svojstva `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

Tvoje pravilo za CSS za `.three` treba da ima svojstvo `background-color` postavljeno na `rgb(0, 0, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
