---
id: 618a00ed1ca871a2b3aca0cb
title: Korak 35
challengeType: 0
dashedName: step-35
---

# --description--

Sada, nakon pregleda primarnih boja, sekundarnih boja i tercijarnih boja na kolornom krugu, biće lako razumeti ostale koncepte teorije boja i kako utiču na dizajn.

Prvo, u podešavanjima za CSS `.one`, `.two`, i na `.three`, prilagodite vrednost u funkciji `rgb` da svojstvo `background-color` svakog elementa bude postavljeno na čisti crni. Napomena: funkcija `rgb` koristi additivni model boja, gde boja počinje kao crna i menja se kako vrednosti crvene, zelene i plave rastu.

# --hints--

Vaše pravilo za CSS za `.one` trebalo bi da ima svojstvo od `background-color` postavljeno na `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

Vaše pravilo za CSS od `.two` trebalo bi da ima svojstvo `background-color` postavljeno na `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

Vaše pravilo za CSS od `.three` treba da ima svojstvo `background-color` postavljeno na `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
