---
id: 618a132676346ac9f7fd59dd
title: Korak 39
challengeType: 0
dashedName: step-39
---

# --description--

Pogledajte kako su oči prirodno privučene crvenoj boji u centru? Kada dizajnirate veb stranicu, možete koristiti ovaj efekat da privučete pažnju na važne naslove sekcija, dugmad ili linkove.

Postoji još jedna važna kombinacija boja izvan komplementarnih boja, ali to ćeš naučiti malo kasnije.

Trenutno, koristite funkciju `rgb` u pravilima `.two` od CSS da postavite `background-color` na crno.

# --hints--

Tvoje pravilo za `.two` za CSS trebalo bi da ima svojstvo od `background-color` postavljeno na `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
