---
id: 618a0927005553b74bfa05ae
title: Korak 36
challengeType: 0
dashedName: step-36
---

# --description--

Kolorni krug je krug gde su slične boje, ili <dfn>mionekano ya rangi</dfn>, blizu jedna drugoj, a različite su dalje. Na primer, čisti crveni je između nijansi ružičaste i narandžaste boje.

Dve boje koje su direktno nasuprot na točkići boja nazivaju se <dfn>komplementarne boje</dfn>. Ako se dve komplementarne boje pomešaju, biće sive. Ali kada se stave jedna pored druge, ove boje stvaraju veliku vidljivu razliku i izgledaju svetlije.

U funkciji `rgb` za pravilo CSS od `.one`, postavite crvenu vrednost na maksimalnu vrednost `255` da bi se dobio čist crveni. U funkciji `rgb` za pravilo CSS od `.two`, postavite zelene i plave vrednosti na maksimalnu vrednost `255` da bi se dobio cijan.

# --hints--

Vaše pravilo za CSS za `.one` trebalo bi da ima svojstvo `background-color` postavljeno na `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

Vaše pravilo od CSS za `.two` trebalo bi imati svojstvo `background-color` postavljeno na `rgb(0, 255, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
