---
id: 619b7fd56aa2253778bcf5f7
title: Korak 48
challengeType: 0
dashedName: step-48
---

# --description--

Primer boje <dfn>HSL</dfn>, ili nijansa, zasićenost i svetlina, je drugi način predstavljanja boja.

Funkcija CSS `hsl` prima vrednost 3: broj od 0 do 360 za nijansu, procenat od 0 do 100 za saturaciju i procenat od 0 do 100 za svetlinu.

Kada se pomisli na kolo boja, crvena nijansa je na 0 stepenima, zelena je na 120 stepenima, a plava je na 240 stepenima.

Saturacija je intenzitet boje od 0%, ili sive, do 100% za čistu boju. Morate dodati znak procenta `%` vrednostima saturacije i svetline.

Lightness je način na koji boja deluje u osvetljenosti, od 0%, što predstavlja potpuno crno, do 100%, što predstavlja potpuno belo, pri čemu je 50% srednja vrednost.

Prema pravilu CSS od `.blue`, koristi funkciju `hsl` da promeniš svojstvo `background-color` u čisti plavi. Podesi nijansu na `240`, saturaciju na `100%` i svetlinu na `50%`.

# --hints--

Vaše pravilo za CSS od `.blue` trebalo bi da ima svojstvo `background-color` postavljeno na `hsl(240, 100%, 50%)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

--fcc-editable-region--
.blue {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
