---
id: 619b761916dac02643940022
title: Korak 46
challengeType: 0
dashedName: step-46
---

# --description--

Najpoznatiji način korišćenja boja za element na CSS je korišćenje <dfn>hexadecimal</dfn> ili heksadecimalnih vrednosti. Iako heksadecimalne vrednosti izgledaju komplikovano, zapravo su samo drugi tip RGB vrednosti.

Hex vrednosti boja počinju sa karakterima `#` i uzimaju šest karaktera od 0-9 i A-F. Prvi deo karaktera predstavlja crvenu boju, drugi deo predstavlja zelenu, a treći deo predstavlja plavu. Na primer, `#4B5320`.

U selektoru klase `.green`, postavi atribut `background-color` za heksadecimalni kod boje sa vrednošću `00` za crvenu, `FF` za zelenu, i `00` za plavu.

# --hints--

Vaše pravilo za `.green` od CSS trebalo bi imati svojstvo `background-color` postavljeno na `#00FF00`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
