---
id: 619b7c3c83de403126b69c1e
title: Korak 47
challengeType: 0
dashedName: step-47
---

# --description--

Možda već znate decimalne brojeve, ili bazu 10, koji počinju od 0 do 9. Heksaadecimalni brojevi, ili baza 16, počinju od 0 do 9, zatim A do F:

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

Za heks boje, `00` je procenat 0 te boje, a `FF` je procenat 100. Dakle, `#00FF00` se prevodi u procenat 0 crvenosti, procenat 100 zelenosti i procenat 0 plavosti, i jednako je `rgb(0, 255, 0)`.

Smanjite zelenu snagu postavljanjem zelene vrednosti boje u heksadecimalnom formatu na `7F`.

# --hints--

Tvoje pravilo za `.green` za CSS trebalo bi da ima svojstvo od `background-color` postavljeno na `#007F00`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: #00FF00;
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
