---
id: 619b7424f43ec9215e538afe
title: Korak 44
challengeType: 0
dashedName: step-44
---

# --description--

Sledeće, promenite klasu objekata `two` u `green` na drugoj lokaciji `div`, i klasu objekata `three` u `blue` na trećoj lokaciji `div`.

# --hints--

Tvoja druga vrednost `div` ne sme imati klasu objekta `two`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(!containerSecondChild?.classList?.contains('two'));
```

Tvoja druga oznaka `div` mora imati klase objekata `marker` i `green`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(containerSecondChild?.classList?.contains('marker') && containerSecondChild?.classList?.contains('green'));
```

Tvoja treća oznaka `div` ne bi trebalo da ima klasu objekata `three`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(!containerThirdChild?.classList?.contains('three'));
```

Vaš treći element `div` mora imati klase objekata `marker` i `blue`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(containerThirdChild?.classList?.contains('marker') && containerThirdChild?.classList?.contains('blue'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      --fcc-editable-region--
      <div class="marker red">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}

```
