---
id: 619b74fa777a2b2473c94f82
title: Korak 45
challengeType: 0
dashedName: step-45
---

# --description--

Popravite konstruktor klase objekta za CSS `.two` ili da upočne na novu klasu `green`. Takođe, popravite konstruktor klase objekta za `.three` ili da upočne na novu klasu `blue`.

# --hints--

Vaš kod više ne sme imati selektor klase za `.two`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.two'));
```

Treba koristiti selektor klase objekata da cilja na klasu `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green'));
```

Vaše pravilo za CSS sa `.green` trebalo bi imati svojstvo `background-color` postavljeno na `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 0, 0)');
```

Vaš kod ne bi trebalo više da ima selektor klase objekata `.three`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.three'));
```

Treba da koristiš selektor klase elemenata za ciljanje klase `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue'));
```

Vaše pravilo za CSS od `.blue` trebalo bi da ima svojstvo `background-color` postavljeno na `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
