---
id: 61a493ead935c148d2b76312
title: Korak 50
challengeType: 0
dashedName: step-50
---

# --description--

Funkcija `linear-gradient` je veoma fleksibilna -- ovde je osnovna sintaksa koju ćete koristiti u ovom kursu:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` je direktija linije koja se koristi za transformacije. `color1` i `color2` su argumenti boja, koje su boje koje će se koristiti u samoj transformaciji. Ovo mogu biti bilo kakve vrste boja, uključujući ključne reči za boje, hex, `rgb`, ili `hsl`.

Sada ćete postaviti promenu boje od crvene do zelene na liniju od 90 stepeni za prvu oznaku.

Prvo, u pravilu za `.red` CSS, postavi svojstvo `background` da bude `linear-gradient()`, i prosledite vrednost `90deg` kao `gradientDirection`.

# --hints--

Vaše pravilo za `.red` CSS trebalo bi da ima svojstvo `background` sa vrednošću `linear-gradient(90deg)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
