---
id: 61a49d15bdbb5e57cc6fd280
title: Korak 54
challengeType: 0
dashedName: step-54
---

# --description--

Kontrole boja omogućavaju vam da ispravite gde su boje postavljene pored linije razlike boja. To je mera dužine kao što je `px` ili procenata koji prate boje u funkciji `linear-gradient`.

Na primer, u ovoj crveno-crnoj razlici boja, prelazak od crvene ka crnoj se dešava na približno 90% duž linije razlike boje, stoga crvena zauzima veliki deo dostupnog prostora:

```css
linear-gradient(90deg, red 90%, black);
```

Unutar funkcije `linear-gradient`, dodaj bojni marker `75%` nakon prve crvene instance. Ne dodaj bojne markere za druge boje.

# --hints--

Vaše pravilo za `.red` CSS mora imati karakteristiku od `background` postavljenu na `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
