---
id: 61a5bfe091060f1d6a629dd0
title: Korak 64
challengeType: 0
dashedName: step-64
---

# --description--

Čak i bez kontrolnih tačaka boja, možda ste primetili da se boje zelene oznake menjaju na istim mestima kao crvene oznake. Prva boja je na početku (0%), druga je u sredini (50%), a poslednja je na kraju (100%) linije gradijenta boje.

Funkcija `linear-gradient` izračunava ove vrednosti direktno za vas i postavlja boju jednako pored linije razlike boje po defaultu.

U pravilima CSS od `.red`, uklonite tri tačke boje iz funkcije `linear-gradient` kako biste malo očistili svoj kod.

# --hints--

Vaše pravilo za CSS od `.red` trebalo bi da ima svojstvo `background` postavljeno na `linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))`.

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}
--fcc-editable-region--

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
