---
id: 61b0a44a6b865738ba49b9fb
title: Korak 86
challengeType: 0
dashedName: step-86
---

# --description--

Poslednja stvar koju ćeš uraditi je dodati malu senku svakoj oznaci kako bi izgledale realnije.

Svojstvo `box-shadow` omogućava vam da koristite jednu ili više senki oko elementa. Evo osnovne sintakse:

```css
box-shadow: offsetX offsetY color;
```

Ovo je kako vrednosti za `offsetX` i `offsetY` funkcionišu:

* sve `offsetX` i `offsetY` prihvataju numeričke vrednosti u `px` i druge jedinice od CSS
* pozitivna vrednost `offsetX` pomera senku udesno, a negativna vrednost pomera ulevo
* pozitivna vrednost `offsetY` pomera senku dole, a negativna vrednost pomera gore
* ako želite nulu vrednost (`0`) za bilo koje ili svih `offsetX` i `offsetY`, ne morate dodavati jedinicu. Svaki pregledač razume da nula znači bez promena.

Visina i širina senke određuju visina i širina elementa koji se koristi. Takođe možete koristiti opcionalnu vrednost `spreadRadius` za proširenje raspona senke. Više o tome kasnije.

Počnite dodavanjem jednostavne senke sa crvenom oznakom.

U zakonu za CSS od `.red`, dodaj svojstvo `box-shadow` sa vrednošću `5px` za `offsetX`, `5px` za `offsetY`, i `red` za `color`.

# --hints--

Vaše pravilo za CSS od `.red` trebalo bi da ima svojstvo skraćivanja od `box-shadow` i vrednost `5px 5px red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow === 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}
--fcc-editable-region--

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
