---
id: 61b093179e7fc020b43b1bb3
title: Korak 73
challengeType: 0
dashedName: step-73
---

# --description--

<dfn>Prozirnost</dfn> opisuje kako je nešto neprozirno ili nije prozirno. Na primer, čvrsti zid je neproziran i ništa svetlo ne prolazi kroz njega. Ali čaša za piće je mnogo prozirnija, i možeš videti kroz čašu do druge strane.

Sa svojstvom CSS `opacity`, možete kontrolisati kako je element skriven ili vidljiv. Za vrednost `0`, ili 0%, element će biti potpuno vidljiv, a za `1.0`, ili 100%, element će biti potpuno skriven, kao podrazumevano.

U zakonu za CSS `.sleeve`, postavi atribut `opacity` da bude `0.5`.

# --hints--

Tvoje pravilo za CSS `.sleeve` trebalo bi da ima svojstvo od `opacity` postavljeno na `0.5`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity === '0.5');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
