---
id: acda2fb1324d9b0fa741e6b5
title: Napravi alat za validaciju kraja niza karaktera
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

U ovoj laboratoriji, ćete izvršiti funkciju koja proverava da li se sekvenca znakova završava sa zadatom ciljnom sekvencom znaka.

**Cilj:** Popunite korisničke priče ispod i prođite kroz sve testove kako biste završili ovaj laboratorijski rad.

Priče o korisniku:

1. Treba da kreirate funkciju sa imenom `confirmEnding` koja prima dva parametra: niz znakova za proveru i niz znakova za poređenje.
2. Funkcija bi trebalo da vrati `true` ako prvi niz znakova završi sa drugim nizom znakova, i `false` inače.
3. Ne smete da koristite metodu `.endsWith()`; umesto toga, koristite jednu od metoda za JavaScript podstring radi postizanja ovoga.

# --hints--

Treba da kreiraš funkciju sa imenom `confirmEnding`.

```js
assert.isFunction(confirmEnding);
```

`confirmEnding` Trebalo bi da primi dva parametra.

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian", "n")` treba da vrati `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` trebalo bi da vrati `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` treba da vrati `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` trebalo bi da vrati `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` treba da vrati `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` treba da vrati `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` treba da vrati `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` treba da vrati `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")` treba da vrati `false`.

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` treba da vrati `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Tvoj kod ne sme da koristi metodu koja je ugrađena unutar `.endsWith()` da reši ovaj laboratorijski zadatak.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
