---
id: 6143869bb45bd85e3b1926aa
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Unutar tvog elementa sa `.heading`, kreiraj element sa `header` koji ima svojstvo elementa `class` ako je `hero`.

Unutar tog elementa, kreiraj element `img` sa svojstvima elementa `src` postavljena na `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`, `alt` postavljeno na `freecodecamp logo`, i `class` postavljeno na `hero-img`.

Atribut elementa `loading` na elementu `img` može se postaviti na `lazy` kako bi se pretraživaču reklo da ne učitava resurs slike dok nije potreban (tj., kada korisnik skroluje sliku u vidokrug). Dodatno, elementi sa opuštenim učitavanjem neće početi sa učitavanjem sve dok se ne učitaju elementi bez opuštenog učitavanja – ovo znači da korisnici sa uspornom internet konekcijom mogu videti sadržaj vaše stranice bez čekanja da se slike učitaju.

Postavi svoj novi atribut za `img`, svojstvo elementa `loading` ako je `lazy`.

# --hints--

Moraš kreirati funkciju za `header`.

```js
assert.exists(document.querySelector('header'));
```

Element vašeg `header` mora biti unutar elementa vašeg `.heading`.

```js
assert(document.querySelector('header')?.parentElement?.className === 'heading');
```

Vaša funkcija od `header` mora imati `class` ako je `hero`.

```js
assert(document.querySelector('header')?.className === 'hero');
```

Treba da napraviš komponentu za `img`.

```js
assert.exists(document.querySelector('img'));
```

Tvoj element `img` treba da bude unutar tvog elementa `header`.

```js
assert(document.querySelector('img')?.parentElement?.localName === 'header');
```

Vaša komponenta za `img` mora da ima `src` ako je `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`.

```js
assert(document.querySelector('img')?.getAttribute('src') === 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

Element vašeg `img` bi trebalo da ima `alt` ako je `freecodecamp logo`.

```js
assert(document.querySelector('img')?.getAttribute('alt') === 'freecodecamp logo');
```

Vaš element za `img` mora imati svojstvo elementa `loading` ako je `lazy`.

```js
assert(document.querySelector('img')?.getAttribute('loading') === 'lazy');
```

Vaš element za `img` mora imati `class` ako je `hero-img`.

```js
assert(document.querySelector('img')?.className === 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
