---
id: 614385513d91ae5c251c2052
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Dodaj element `title` sa tekstom `Magazine`, element `link` za stilski list (stylesheet) `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`, `link` za FontAwesome stilski list (stylesheet) `https://use.fontawesome.com/releases/v5.8.2/css/all.css`, i `link` za tvoj stilski list (stylesheet) `./styles.css`.

Stilski list za fontove učitaće tri različita fonta: `Anton`, `Baskervville`, i `Raleway`.

# --hints--

Tvoj kod treba da ima tri elementa od `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

Vaši elementi za `link` moraju biti unutar vašeg elementa `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

Vaša tri elementa od `link` treba da imaju svojstvo elementa `rel` sa vrednošću `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Jedan od vaših komponentnih elemenata mora imati `href` podešen na `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Jedan od vaših komponentnih elemenata treba da ima `href` podešen na `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Jedan od vaših elemenata za `link` treba da ima `href` podešen na `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

Vaš kod mora imati karakteristiku `title`.

```js
const title = document.querySelector('title');
assert.isNotNull(title);
```

Vaš kod treba da ima jedan element od `title`.

```js
const title = document.querySelectorAll('title');
assert.strictEqual(title?.length, 1);
```

Element za `title` mora biti unutar elementa za `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

Tvoj projekat treba da ima naslov `Magazine`. Zapamti, velika i mala slova kao i pravopis su važni za naslov.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
