---
id: 614389f601bb4f611db98563
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Ispod vašeg elementa sa `.author`, kreiraj novi element sa `div` koji ima klasu objekata `social-icons`.

Dodajte pet elemenata iz `a` unutar novog `div`, i dodelite im sledeće karakteristike elementa za `href`.

- Elementat prvi sa `a` mora imati atribut elementa `href` postavljen na `https://www.facebook.com/freecodecamp`.
- Elementat drugi sa `a` mora imati atribut elementa `href` postavljen na `https://twitter.com/freecodecamp`.
- Elementat treći sa `a` mora imati atribut elementa `href` postavljen na `https://instagram.com/freecodecamp`.
- Elementat četvrti sa `a` mora imati atribut elementa `href` postavljen na `https://www.linkedin.com/school/free-code-camp`.
- Elementat peti sa `a` mora imati atribut elementa `href` postavljen na `https://www.youtube.com/freecodecamp`.

# --hints--

Trebalo bi da napraviš novu komponentu za `div`.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

Tvoj novi element za `div` mora da dođe nakon tvog elementa `.author`.

```js
assert(document.querySelector('.author')?.nextElementSibling?.localName === 'div');
```

Vaša nova funkcija za `div` mora imati klasu `social-icons`.

```js
assert(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

Vaš element `.social-icons` bi trebalo da ima pet komponenti `a`.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.length === 5);
```

Prvi atribut vašeg `a` treba da ima svojstvo atributa `href` postavljeno na `https://www.facebook.com/freecodecamp`.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href')?.includes('https://www.facebook.com/freecodecamp'));
```

Tvoj drugi element od `a` mora imati svojstvo elementa `href` postavljeno na `https://twitter.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

Tvoj treći element od `a` mora imati svojstvo polja `href` postavljeno na `https://instagram.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

Vaš četvrti element od `a` treba da ima svojstvo elementa `href` postavljeno na `https://www.linkedin.com/school/free-code-camp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

Petih vašeg elementa `a` treba da ima svojstvo elementa `href` postavljeno na `https://www.youtube.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
