---
id: 614387cbefeeba5f3654a291
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Nakon vašeg elementa `header`, kreirajte `div` koji se postavlja na `class` kao `author`.

Unutar `div`, kreiraj element za `p` i `class` postavi na `author-name`, a tekstualni sadržaj `By freeCodeCamp`. Pokrij deo `freeCodeCamp` unutar elementa `a` i `href` da bude postavljen na `https://freecodecamp.org`, a `target` da bude postavljen na `_blank`.

Ispod toga, dodaj drugi element za `p` koji ima klasu objekata `publish-date` i tekst `March 7, 2019`.

# --hints--

Treba da kreiraš novu komponentu za `div`.

```js
assert.exists(document.querySelector('div'));
```

Vaš element `div` treba da dođe posle vašeg elementa `header`.

```js
assert(document.querySelector('div')?.previousElementSibling?.localName === 'header');
```

Vaše polje za `div` mora imati `class` postavljeno na `author`.

```js
assert(document.querySelector('div')?.className === 'author');
```

Treba ti kreirati dve nove komponente za `p`.

```js
assert(document.querySelectorAll('p')?.length === 3);
```

Tvoja dva nova elementa ``p`` treba da budu unutar tvog elementa ``div``.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

Tvoj prvi novi element za `p` treba da ima `class` postavljen na `author-name`.

```js
assert(document.querySelector('div')?.querySelector('p')?.className === 'author-name');
```

Tvoj prvi novi element `p` treba da ima tekst `By freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.innerText === 'By freeCodeCamp');
```

Tvoj drugi novi element za `p` mora imati `class` postavljen na `publish-date`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className === 'publish-date');
```

Tvoj drugi novi element `p` treba da ima tekst `March 7, 2019`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText === 'March 7, 2019');
```

Trebalo bi da napraviš novu komponentu za `a`.

```js
assert.exists(document.querySelector('a'));
```

Vaš element `a` treba da bude unutar vašeg prvog novog elementa `p`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName === 'a');
```

Vaš element `a` mora imati `href` postavljen na `https://freecodecamp.org`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href') === 'https://freecodecamp.org');
```

Vaš atribut za `a` mora imati `target`, postavljen na `_blank`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target') === '_blank');
```

Element vašeg ``a`` mora da pokrije tekst od ``freeCodeCamp``.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText === 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
