---
id: 61437d575fb98f57fa8f7f36
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Počnite sa vašom osnovnom šablonom za HTML. Dodajte deklaraciju od `DOCTYPE`, funkciju `html` koja kaže da je ova stranica na engleskom jeziku, funkciju `head`, i funkciju `body`.

Dodajte oznaku `<meta>` sa odgovarajućim atributom elementa `charset` i oznaku `<meta>` za kompatibilnost sa telefonom unutar komponente `head`.

# --hints--

Vaš kod treba da ima referencu na `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Treba da staviš razmak nakon reference `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Trebalo bi da definišete tip dokumenta kao `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Trebalo bi da zatvorite deklaraciju `DOCTYPE` sa etiketom za zatvaranje `>` nakon tipa.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Mora imati etiketu za otvaranje `<html>`, i mora imati svojstvo elementa `lang` od `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Element vašeg `html` treba da ima zatvarajući tag.

```js
assert(code.match(/<\/html\s*>/));
```

Tvoje saopštenje za `DOCTYPE` trebalo bi biti na početku tvog HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Treba ti etiketa za otvaranje `<head>`.

```js
assert(code.match(/<head\s*>/i));
```

Trebalo bi da ima etiketu za zatvaranje `</head>`.

```js
assert(code.match(/<\/head\s*>/i));
```

Trebalo bi da ima etiketu za otvaranje `<body>`.

```js
assert(code.match(/<body\s*>/i));
```

Trebalo bi da imate etiketu za zatvaranje `</body>`.

```js
assert(code.match(/<\/body\s*>/i));
```

Elementi za `head` i `body` bi trebalo da budu sestrinski elementi.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Element `head` treba da bude unutar elementa `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Element za `body` treba da bude unutar elementa `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Trebalo bi da ima dva elementa za `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Jedan element od `meta` treba da ima atribut elementa `name` postavljen na `viewport`, a `content` je postavljen na `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Drugi element od `meta` mora imati svojstvo elementa `charset` postavljeno na `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
