---
id: 6165d3b702a5d92ad970b30c
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Nakon vašeg elementa sa `img`, dodajte element sa `h1` koji ima `class` postavljen na `hero-title` i tekst je `OUR NEW CURRICULUM`, praćen elementom sa `p` koji ima `class` postavljen na `hero-subtitle` i tekst je `Our efforts to restructure our curriculum with a more project-based focus`.

# --hints--

Trebalo bi da kreirate element `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Vaš element od `h1` mora da dođe nakon vašeg elementa od `img`.

```js
assert(document.querySelector('h1')?.previousElementSibling?.localName === 'img');
```

Element vašeg `h1` mora imati `class` podešen na `hero-title`.

```js
assert(document.querySelector('h1')?.className === 'hero-title');
```

Vaš element za `h1` mora imati tekst koji je `OUR NEW CURRICULUM`.

```js
assert(document.querySelector('h1')?.textContent === 'OUR NEW CURRICULUM');
```

Treba da kreiraš novi element za `p`.

```js
assert.exists(document.querySelector('p'));
```

Tvoj element `p` mora da sledi nakon tvog elementa `h1`.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

Vaša funkcija od `p` trebalo bi da ima `class` postavljeno na `hero-subtitle`.

```js
assert(document.querySelector('p')?.className === 'hero-subtitle');
```

Vaša sekcija za `p` treba da ima tekst koji je `Our efforts to restructure our curriculum with a more project-based focus`.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
