---
id: 696920c0c98a1ed58eb86293
title: Izložba proizvoda
challengeType: 25
dashedName: lab-product-showcase
demoType: onClick
---

# --description--

Za ovu vežbu, dobili ste sve HTML i CSS. Završite demonstraciju proizvoda koristeći TypeScript.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba da kreirate interfejs pod nazivom `Item` sa:
   - Svojstvo od `type` koje može biti samo `"book"`, `"electronics"` ili `"clothing"`.
   - Svojstvo od `id` tipa `string`.
   - Svojstvo od `price` tipa `number`.

1. Treba kreirati interfejs nazvan `Book` koji proširuje (ili nasleđuje) `Item`. Njegovo svojstvo (property) `type` treba postaviti na `"book"` tačno. Takođe treba da ima svojstva za `title` i `author` tipa `string`.

1. Treba definisati interfejs pod nazivom `Electronics` koji nasleđuje `Item`. Njegovo svojstvo `type` treba da bude postavljeno na stvarnu vrednost `"electronics"`. Takođe, trebalo bi mu da ima svojstva `item` i `model` tipa `string`, kao i opcionalno svojstvo `warranty` tipa `number`.

1. Trebalo bi da kreirate interfejs nazvan `Clothing` koji proširuje `Item`. Njegovo svojstvo `type` treba da se postavi na tačno `"clothing"`. Takođe bi trebalo da ima svojstva `item` i `brand` tipa `string`, kao i opcionalno svojstvo `size` koje može biti samo `"S"`, `"M"` ili `"L"`.

1. Morate kreirati tip konekcije nazvan `Product`, koji kombinuje interfejsi od `Book`, `Electronics`, na `Clothing`.

1. Treba kreirati opštu klasu nazvanu `Collection` sa parametrom tipa `T`. Konstruktor treba da primi niz elemenata tipa `T` i sačuva ih u svojstvu nazvanom `items`.

1. Klasa `Collection` treba da ima metodu `getAll()` koja vraća sve stavke u kolekciji.

1. Klasa `Collection` treba da ima metodu `filter()` koja prima callback funkciju i vraća samo elemente iz niza `items` koji zadovoljavaju navođeni uslov. Callback funkcija treba da primi element tipa `T` i vrati booleanski vrednost.

1. Treba da kreirate funkciju nazvanu `renderProduct` sa parametrom tipa `Product`. Moraće da vrati niz karaktera HTML koji predstavlja proizvod sa:
   - Elementom sa klasom `item` i identifikatorom sa vrednošću `id` proizvoda.
   - Elementom sa klasom `price` i tekstom/opisom `price` proizvoda.
   - Elementom sa dodatnim informacijama o proizvodu.

1. Funkcija `renderProduct` treba da koristi filtriranje tipa za prikaz različitih HTML u zavisnosti od vrste proizvoda:

- Za knjige, HTML koja se izdaje treba da ima format `Book: [title] by [author]`, gde se `[title]` i `[author]` moraju zameniti naslovom i autorom, respektivno.
   - Za elektronske uređaje, HTML koji se izdaju treba da imaju format `Electronics: [item] - [model]`. Ako proizvod ima garanciju, dodajte ` - Warranty: [warranty] year(s)` u format. Zamenite `[item]`, `[model]` i `[warranty]` stavkom, primerom i garancijom, respektivno.
   - Za odeću, HTML koja se izdaje treba da ima format `Clothing: [item] by [brand]`. Ako proizvod ima veličinu, dodajte ` - Size [size]` u format. Zamenite `[item]`, `[brand]` i `[size]` stavkom, markom i veličinom, respektivno.
   - Ako je tip proizvoda nevažeći, baciti grešku sa formatom: `Unknown product type: [product]`, gde se `[product]` mora zameniti nevažećim proizvodom konvertovanim u niz znakova korišćenjem `JSON.stringify`.

1. Treba da kreirate model od `Collection<Product>` koji uključuje najmanje jedan artikl za svaku vrstu proizvoda i postavite ga u polje nazvano `products`.

Trebalo bi da kreirate funkciju `showProducts` koja prikazuje proizvod u sekciji `#output`. Treba da prima opcionalni filter koji predstavlja vrstu proizvoda; ako je ostavljen prazan, prikažite sve proizvode; ukoliko se navede, prikažite samo proizvode te vrste.

1. Treba da prikažete proizvod konvertovanjem svakog proizvoda u niz znakova HTML koristeći funkciju `renderProduct` i unosom rezultata uključeno u polje `#output`.

1. Klik na dugmad `#all`, `#books`, `#electronics`, i `#clothing` bi trebalo da prikaže odgovarajući set proizvoda u sekciji `#output`.

1. Kada se stranica učitava, svi proizvodi bi trebalo da budu prikazani po defaultu.

# --hints--

Trebalo bi da ima interfejs koji se zove `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Item);
```

Vaš interfejs za `Item` mora imati svojstvo `type` koje može biti `"book"`, `"electronics"`, ili `"clothing"` samo.

```js
const explorer = await __helpers.Explorer(code);
const { Item } = explorer.interfaces;
const typeProp = Item.typeProps.type;
assert.isTrue(typeProp.isUnionOf(["'book'", "'electronics'", "'clothing'"]));
```

Vaše sučelje za `Item` trebalo bi imati svojstvo `id` tipa `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "id", type: "string"}));
```

Kiolesura chako cha `Item` mora imati funkcionalnost `price` tipa `number`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "price", type: "number"}));
```

Trebalo bi da ima interfejs koji se zove `Book`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Book);
```

Vaš interfejs za `Book` treba da proširi `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.doesExtend("Item"));
```

Kiolesura chako cha `Book` mora imati svojstvo `type` postavljeno na `"book"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "type", type: "'book'"}));
```

Vaš interfejs za `Book` mora imati svojstvo `title` tipa `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "title", type: "string"}));
```

Interfejs za `Book` mora imati svojstvo `author` tipa `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "author", type: "string"}));
```

Trebalo bi da ima interfejs koji se zove `Electronics`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Electronics);
```

Vaš interfejs za `Electronics` treba da proširi `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.doesExtend("Item"));
```

Tvoj interfejs za `Electronics` mora imati svojstvo `type` postavljeno na `"electronics"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "type", type: "'electronics'"}));
```

Struktura vašeg `Electronics` bi trebalo da ima svojstvo `item`, tipa `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "item", type: "string"}));
```

Vaš interfejs za `Electronics` treba da ima svojstvo `model` sa tipom `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "model", type: "string"}));
```

Vaš interfejs za `Electronics` mora imati opcionalno svojstvo `warranty` tipa `number`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "warranty", type: "number", isOptional: true}));
```

Trebalo bi da ima interfejs nazvan `Clothing`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Clothing);
```

Vaša struktura od `Clothing` trebalo bi da proširi `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.doesExtend("Item"));
```

Vaš interfejs za `Clothing` mora imati svojstvo od `type` postavljeno na `"clothing"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "type", type: "'clothing'"}));
```

Vaš interfejs za `Clothing` mora imati karakteristiku `item` tipa `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "item", type: "string"}));
```

Vaš interfejs za `Clothing` mora da ima karakteristiku `brand` tipa `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "brand", type: "string"}));
```

Interfejs vašeg `Clothing` treba da ima opcionalnu karakteristiku `size` koja može biti samo `"S"`, `"M"` ili `"L"`.

```js
const explorer = await __helpers.Explorer(code);
const { Clothing } = explorer.interfaces;
assert.isTrue(Clothing.hasTypeProps({name: "size", isOptional: true}));
assert.isTrue(Clothing.typeProps.size.isUnionOf(["'S'", "'M'", "'L'"]));
```

Trebalo bi da kreirate vrstu konekcije nazvanu `Product` koja kombinuje interfejse `Book`, `Electronics`, na `Clothing`.

```js
const explorer = await __helpers.Explorer(code);
const { Product } = explorer.types;
assert.isTrue(Product.isUnionOf(["Book", "Electronics", "Clothing"]));
```

Trebalo bi da imaš klasu `Collection` sa parametrom tipa `T`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.isTrue(Collection.typeParameters[0].matches("T"));
```

Komponenta vaše klase za `Collection` treba da primi niz stavki tipa `T` i sačuva ih u svojstvo nazvano `items`.

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.items, arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.items, arr);
```

Vaša klasa od `Collection` mora imati metod za `getAll`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
console.log(Collection.methods)
console.log(Collection.methods.getAll)
assert.exists(Collection.methods.getAll);
```

Tvoj put ``getAll`` bi trebalo da vrati sve stvari u kolekciju.

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.getAll(), arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.getAll(), arr);
```

Klasa za `Collection` treba da ima metodu `filter`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.exists(Collection.methods.filter);
```

Tvoja ruta ``filter`` bi trebalo da primi funkciju odgovora i vrati niz koji sadrži elemente atributa ``items`` koji zadovoljavaju zadato uslov.

```js
let arr = [1, 2, 3, 4, 5, 6];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.filter(n => n % 2 === 0), [2, 4, 6]);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.filter(s => s === "c"), ["c"]);
```

Trebalo bi da imaš funkciju nazvanu `renderProduct`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.renderProduct);
```

`renderProduct(product)` treba da vrati sekvencu karaktera koja uključuje element sa klasom `item` i identifikator jednak `product.id`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  assert.isNotNull(testEl.querySelector(`[class="item"][id="${p.id}"]`));
})
```

`renderProduct(product)` trebalo bi da ponovo stvori sekvencu karaktera koja uključuje element sa klasom `price` i tekst jednak `product.price`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const price = testEl.querySelector(`[class="price"]`);
  assert.include(price.innerText, p.price.toString());
})
```

Kada `renderProduct` bude pozvan od proizvoda `Book` `p`, treba da vrati sekvencu znakova koja uključuje element sa tekstom strukturiranim kao `Book: [title] by [author]`. Promeni `[title]` u `[author]` sa svojstvima koja odgovaraju `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const expectedText = `Book: ${p.title} by ${p.author}`;
  assert.include(testEl.innerText, expectedText);
})
```

Kada `renderProduct` je pozvan od proizvoda `Electronics` `p`, treba da vrati sekvencu znakova koja uključuje element sa tekstom strukturiranim kao `Electronics: [item] - [model]` ili `Electronics: [item] - [model] - Warranty: [warranty] year(s)` ako `p.warranty` nije vidljivo. Promeni `[item]`, `[model]`, i `[warranty]` sa atributima koji odgovaraju za `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Electronics: ${p.item} - ${p.model}`;
  const expectedText = !p.warranty ? baseFormat : baseFormat + ` - Warranty: ${p.warranty} year(s)`;
  assert.include(testEl.innerText, expectedText);
  if (!p.warranty) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

Kada je `renderProduct` pozvan od proizvoda `Clothing` `p`, trebalo bi da vrati niz znakova koji uključuje element sa tekstualnom strukturom `Clothing: [item] by [brand]` ili `Clothing: [item] by [brand] - Size [size]` ako `p.size` nije vidljiv. Promenite `[item]`, `[brand]` i `[size]` u svojstva koja odgovaraju `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Clothing: ${p.item} by ${p.brand}`;
  const expectedText = !p.size ? baseFormat : baseFormat + ` - Size ${p.size}`;
  assert.include(testEl.innerText, expectedText);
  if (!p.size) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

Kada je `renderProduct` pozvano sa nevažećim proizvodom, treba da baci grešku sa strukturom `Unknown product type: [p]`. Zamenite `[p]` sa nevažećim proizvodom konvertovanim u niz karaktera koristeći `JSON.stringify`.

```js
const invalidProduct = { item: "T-Shirt", price: 19.99 };
const expected = `Unknown product type: ${JSON.stringify(invalidProduct)}`;
try {
  renderProduct(invalidProduct);
  assert.fail();
} catch({name, message}) {
  if (name == "AssertionError") {
    assert.fail("renderProduct should throw with invalid product");
  } else {
    assert.equal(message, expected);
  }
}
```

Treba da kreiraš primer `Collection<Product>` i staviš ga u parametar nazvan `products`.

```js
assert.instanceOf(products, Collection);
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.products.value.typeArguments[0].matches("Product"));
```

Vaš kriterijum za `products` mora imati najmanje jedan element za svaku vrstu proizvoda.

```js
assert.isTrue(products.items.some(p => p.type === "book"));
assert.isTrue(products.items.some(p => p.type === "electronics"));
assert.isTrue(products.items.some(p => p.type === "clothing"));
```

Trebalo bi da imaš funkciju nazvanu `showProducts`.

```js
assert.isFunction(showProducts);
```

Klik na dugme `#all` bi trebalo da prikaže sve proizvode u sekciji `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
books.dispatchEvent(new Event("click", { bubbles: true }));
all.dispatchEvent(new Event("click", { bubbles: true }));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

Klik na dugme `#books` bi trebalo da prikaže relevantni skup proizvoda u sekciji `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "book") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click", { bubbles: true }));
books.dispatchEvent(new Event("click", { bubbles: true }));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Klik na dugme `#electronics` treba da prikaže odgovarajući skup proizvoda u sekciji `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "electronics") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click", { bubbles: true }));
electronics.dispatchEvent(new Event("click", { bubbles: true }));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Klik na dugme `#clothing` trebalo bi da prikaže relevantni set proizvoda u polju `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "clothing") {
    count++;
  }
});
books.dispatchEvent(new Event("click", { bubbles: true }));
clothing.dispatchEvent(new Event("click", { bubbles: true }));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

Svi proizvodi bi trebalo da budu vidljivi po podrazumevanju kada se stranica učita. Koristite događaj `DOMContentLoaded` da biste to osigurali.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
document.dispatchEvent(new Event("DOMContentLoaded"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts
interface Item {
  type: "book" | "electronics" | "clothing";
  id: string;
  price: number;
}

interface Book extends Item {
  type: "book";
  title: string;
  author: string;
}

interface Electronics extends Item {
  type: "electronics";
  item: string;
  model: string;
  warranty?: number;
}

interface Clothing extends Item {
  type: "clothing";
  item: string;
  brand: string;
  size?: "S" | "M" | "L";
}

type Product = Book | Electronics | Clothing;

class Collection<T> {
  items: T[];
  constructor(items: T[]) {
    this.items = items;
  }

  getAll(): T[] {
    return this.items;
  }

  filter(callback: (item: T) => boolean): T[] {
    return this.items.filter(callback);
  }
}

function renderProduct(p: Product): string {
  const createProductCard = (id: string, content: string, price: number): string => {
    return `
     <div class="item" id="${id}">
       ${content}
       <div class="price">$${price}</div>
     </div>
   `;
  };

  if (p.type === "book") {
    const content = `<strong>Book:</strong> ${p.title} by ${p.author}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "electronics") {
    const warranty = p.warranty ? ` - Warranty: ${p.warranty} year(s)` : "";
    const content = `<strong>Electronics:</strong> ${p.item} - ${p.model}${warranty}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "clothing") {
    const size = p.size ? ` - Size ${p.size}` : "";
    const content = `<strong>Clothing:</strong> ${p.item} by ${p.brand}${size}`;
    return createProductCard(p.id, content, p.price);
  }

  const _never: never = p;
  throw new Error(`Unknown product type: ${JSON.stringify(p)}`);
}


const products = new Collection<Product>([
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },

  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "b2", type: "book", title: "1984", author: "George Orwell", price: 9.99 },

  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 },

  { id: "b3", type: "book", title: "Brave New World", author: "Aldous Huxley", price: 11.99 },
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "e5", type: "electronics", item: "Smartwatch", model: "Chronex Pulse-Q", warranty: 1, price: 199.99 },

  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 },
]);

const output = document.getElementById("output");

// Track currently active filter category (or undefined for all)
let currentFilter: Product["type"] | undefined = undefined;

function showProducts(filter?: Product["type"]): void {
  if (!output) return;

  let itemsToShow: Product[];

  // Toggle logic
  if (currentFilter === filter || filter === undefined) {
    // Reset to all if same button clicked again or filter is undefined
    itemsToShow = products.getAll();
    currentFilter = undefined;
    setActiveButton(undefined);
  } else {
    itemsToShow = products.filter((p: Product): boolean => p.type === filter);
    currentFilter = filter;
    setActiveButton(filter);
  }

  output.innerHTML = itemsToShow.map(renderProduct).join("");
}

// Active Button Highlighting
function setActiveButton(filter: Product["type"] | undefined): void {
  const buttons: { [key: string]: HTMLElement | null } = {
    all: document.getElementById("all"),
    book: document.getElementById("books"),
    electronics: document.getElementById("electronics"),
    clothing: document.getElementById("clothing")
  };

  for (const key in buttons) {
    buttons[key]?.classList.remove("active");
  }

  if (filter === undefined) {
    buttons["all"]?.classList.add("active");
  } else {
    buttons[filter]?.classList.add("active");
  }
}

// Button Event Handlers
document.getElementById("all")?.addEventListener("click", (): void => showProducts());
document.getElementById("books")?.addEventListener("click", (): void => showProducts("book"));
document.getElementById("electronics")?.addEventListener("click", (): void => showProducts("electronics"));
document.getElementById("clothing")?.addEventListener("click", (): void => showProducts("clothing"));

// Initial Render
document.addEventListener("DOMContentLoaded", (): void => showProducts());
```
