---
id: 6565a536ba1f9f25bd30e88b
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovom projektu, implementiraćeš Luhnov algoritam. Ovaj algoritam je formula za verifikaciju različitih identifikacionih brojeva.

Počnite sa deklarisanjem funkcije nazvane `main`, ovo će biti kao deo inicijalizacije programa. Koristite ključnu reč `pass` da biste izbegli grešku.

# --hints--

Trebalo bi da imaš `def main():` u svom kodu.

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

Trebalo bi ti da imaš `pass` u svom kodu.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
