---
id: 65687db6a9fcf3c548a6c1b7
title: Korak 28
challengeType: 20
dashedName: step-28
---

# --description--

Deo algoritma je da se svaki drugi broj udvostru, počevši od desne strane. Ako je rezultat udvostručavanja broja veći ili jednak `10`, saberi te dva broja. Na primer, ako je broj `6`, udvostru ga da bi dobio `12`. Saberi `1` i `2` da bi dobio `3`. Ovo možete uraditi korišćenjem _integer division_-a da dobijete prvi broj i operatera modulo (`%`) da dobijete drugi broj:

```python
my_number = 12
first_digit = my_number // 10
second_digit = my_number % 10
```

Cjelobrojno deljenje (integer division) vraća količnik deljenja, zaokružen ka dole do najbližeg cjelog broja.

U deklaraciji `if`, postavi vrednost `number` kao rezultat `number // 10` (celobrojno deljenje) zajedno sa modulom `number` i `10`.

# --hints--

Trebalo bi da imaš `number = number // 10 + number % 10` unutar deklaracije `if`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const { block_body } = verify_card_number;
        const allowedMatches = [
            /number *= *\(?\s*number *\/\/ *10\s*\)? *\+ *\(?\s*number *% *10/,
            /number *= *\(?\s*number *% *10\s*\)? *\+ *\(?\s*number *\/\/ *10/
        ];

        const someMatch = allowedMatches.some(allowedMatch => allowedMatch.test(block_body));
        assert.isTrue(someMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
