---
id: 656878585631369a6b2d2191
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Definisati sam prevod ne prevodi sekvencu karaktera. Putanja za `translate` mora biti pozvana na sekvencu karaktera koja se želi prevesti koristeći tabelu prevoda kao argument:

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

Kreiraj parametar nazvan `translated_card_number` i postavi mu vrednost kao rezultat poziva metode `translate` na `card_number`, koristeći `card_translation` kao argument.

# --hints--

Treba da napraviš parametar `translated_card_number` unutar `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

Treba da postaviš vrednost `translated_card_number` na `card_number.translate(card_translation)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    
--fcc-editable-region--
```
