---
id: 656877c202f44995ed952608
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Promeni izraz `pass` u parametar nazvan `card_number` sa vrednošću `'4111-1111-4555-1142'`.

# --hints--

Trebalo bi da imaš `card_number = '4111-1111-4555-1142'` unutar funkcije `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_number\s*=\s*('|")4111-1111-4555-1142\1/);
    }
})
```

Moraš imati `pass` u svoj kod.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const no_comments = __helpers.python.removeComments(transformedCode);
        assert.notInclude(no_comments, "pass");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    pass

--fcc-editable-region--
```
