---
id: 656879c6f35a85a59c06b3a7
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Prethodno ste pristupili karakterima (slovima) iz niza koristeći regularni izraz `[]`. Takođe možete koristiti regularni izraz da biste pristupili opsegu karaktera u nizu za `string[start:stop:step]`:

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

Gde je `start` početni parametar (uključivo), `stop` je krajnji parametar (isključivo), a `step` je broj karaktera za preskočenje. Ako nije navedeno, `step` podrazumevana vrednost je 1.

Kreiraj parametar nazvan `card_number_reversed` i postavi vrednost prvih 4 slova iz `card_number`.

# --hints--

Trebalo bi da imaš `card_number_reversed = card_number[0:4]` unutar funkcije `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
