---
id: 65687b48f2201ab32e06c37c
title: Korak 21
challengeType: 20
dashedName: step-21
---

# --description--

Trenutno, vaš skript vraća `TypeError` jer pokušavate dodati niz slova na pun broj. Možete ovo popraviti promenom vrednosti `digit` u pun broj pre nego što ga dodate na `sum_of_odd_digits`, koristeći funkciju unutar `int`:

```python
my_string = '123'
my_int = int(my_string)
```

Izmeni vrednost `digit` u ceo broj pre nego što ga dodaš u `sum_of_odd_digits`. Zatim, premesti poziv `print` na kraj funkcije `verify_card_number` kako bi ispisao vrednost `sum_of_odd_digits`.

# --hints--

Trebalo bi da imaš `sum_of_odd_digits += int(digit)` unutar kruga `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*\+=\s*int\(\s*digit\s*\)/);
    }
})
```

Trebalo bi da imaš `print(sum_of_odd_digits)` na kraju funkcije `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*sum_of_odd_digits\s*\)/);
    }
})
```

Nemoj da imaš `print(digit)` više.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

--fcc-editable-region--
    for digit in odd_digits:
        print(digit)
        sum_of_odd_digits += digit
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
