---
id: 656879a66338b2a461d5d307
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

Algoritam Lhun je kao što sledi:

1. Počevši od desna ka levo, pomnožite sa dva vrednost svakog drugog broja; ako je proizvod veći od 9, saberi cifre tog proizvoda.
2. Uzmite zbir svih brojeva.
3. Ako je zbir svih brojeva višestrukost desetice, onda je taj broj validan; inače nije validan.

Uzmite primer broja računa "7992739871" koji će imati dodati proverni broj, čime postaje u formatu 7992739871x:

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

Promeni deklaraciju `pass` u parametar `sum_of_odd_digits` sa vrednošću `0`.

# --hints--

Trebalo bi da imaš `sum_of_odd_digits = 0` unutar funkcije `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

Ne smiješ imati `pass` unutar funkcije `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
