---
id: 656877f71bba2b97acedf9af
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Python dolazi sa klasama ugrađenih objekata koji nam mogu pomoći u obradi sekvenci karaktera. Jedan od njih je klasa `str`. Ima metodu nazvanu `maketrans` koja nam može pomoći da kreiramo tabelu prevoda. Ova tabela se može koristiti za zamenu karaktera u sekvenci karaktera:

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

Ovde iznad, kada se spomene u nizu karaktera, će promeniti sve karaktere za `t` u `c` i sve karaktere za `l` u `b`.

Kreiraj parametar nazvan `card_translation` i postavi mu vrednost u tabelu za translaciju da promeni sve karaktere od `-` na ` ` u prazan niz.

# --hints--

Treba da kreiraš parametar `card_translation` unutar `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

Trebalo bi da postaviš vrednost `card_translation` na `str.maketrans({'-': '', ' ': ''})`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    
--fcc-editable-region--
```
