---
id: a5deed1811a43193f9f1c841
title: Implementiraj funkciju za preskakanje elementa
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

U ovoj laboratoriji ćete kreirati funkciju koja iterira elemente u nizu podataka sve dok ne pronađe prihvatljiv element prema specifičnoj funkciji testa.

Na primer, za red podataka kao `[1, 1, 1, 2, 1, 1, 1]` i funkciju testiranja `function(n) {return n === 2}`, prvi prihvatljiv element je onaj iz figure 3, pa bi svi elementi pre njega trebalo da budu uklonjeni, a rezultat bi trebali biti preostali elementi `[2, 1, 1, 1]`.

**Cilj**: Završite korisničke priče navedene ispod i položite sve testove kako biste završili ovu laboratoriju.

**Priče korisnika:**

1. Treba da postoji funkcija ``dropElements`` koja prima niz podataka (``arr``) i funkciju (``func``) kao argumente.
1. Funkcija ``dropElements`` treba da iterira kroz niz podataka i ukloni elemente od prvog elementa do ``func``, vraćajući ``true`` za određeni element.
1. Funkcija ``dropElements`` treba da vrati preostale elemente u nizu podataka ako je uslov ispavljen.
1. Ako uslov nikada nije ispavljen, treba da vrati prazan niz podataka.

# --hints--

Trebalo bi da imaš funkciju `dropElements`.

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` trebalo bi vratiti `[3, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` treba da vrati `[1, 0, 1]`.

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` trebalo bi da vrati `[1, 2, 3]`.

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` treba da vrati `[]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` treba da vrati `[7, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` treba da vrati `[3, 9, 2]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
