---
id: 5f331e55dfab7a896e53c3a1
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` je jedan od nekoliko elemenata koji pružaju dodatne informacije koje nisu vidljive na web stranici, ali su važne za pretraživače ili način prikazivanja stranice.

Unutar elementa `head`, postavite unutar elementa `meta` koji ima atribut nazvan `charset` i postavljen na vrednost `utf-8` kako biste rekli pretraživaču kako da dekodira znakove stranice.

Imajte na umu da su elementi `meta` neograničeni elementi koji zahtevaju zatvarajući tag.

# --hints--

Trebalo bi da ima oznaku `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

Funkcija `meta` je funkcija koja nije konačna, i ne bi trebalo da ima oznaku za kraj `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Tvoj objekat/podatak za `meta` mora da ima svojstvo/atribut `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Atribut vašeg polja za `charset` mora imati vrednost `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
