---
id: 5f3326b143638ee1a09ff1e3
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Da biste se pripremili za kreiranje stvarnog sadržaja, dodajte polje `body` ispod polja `head`.

# --hints--

Trebalo bi da imaš oznaku za otvaranje `<body>`.

```js
assert(code.match(/<body>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `</body>`.

```js
assert(code.match(/<\/body>/i));
```

Ne smete da menjate svoj element `head`. Proverite da niste izbrisali vašu oznaku za zatvaranje.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

Element vašeg `body` treba da dođe nakon elementa vašeg `head`.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
