---
id: 5f3313e74582ad9d063e3a38
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Dodajte element `head` unutar elementa `html`, kako bi mogao dodati element `title`. Sadržaj elementa `title` treba da bude `Cafe Menu`.

# --hints--

Trebalo bi da imaš etiketu pakovanja za `<head>`.

```js
assert(code.match(/<head>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `</head>`.

```js
assert(code.match(/<head>/i));
```

Trebalo bi da imaš etiketu pakovanja za `<title>`.

```js
assert(code.match(/<title>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `</title>`.

```js
assert(code.match(/<\/title>/i));
```

Vaša komponenta od `title` mora biti postavljena unutar vaše komponente `head`.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

Vaše polje za `title` treba da ima tekst od `Cafe Menu`. Možda trebate proveriti pravopis.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
