---
id: 5f33071498eb2472b87ddee4
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovom projektu, ćeš naučiti osnove CSS (Cascading Style Sheets) za kreiranje menija kafe. CSS je jezik koji se koristi za stilizovanje dokumenta HTML. On definiše kako bi elementi HTML trebalo da budu prikazani na ekranu.

Kao što si naučio u prethodnim koracima Cat Photo App projekta, postoji osnovna struktura potrebna za početak izgradnje veb stranice. Svaki HTML dokument treba da ima deklaraciju `DOCTYPE` i `html` element. `DOCTYPE` objašnjava koju verziju HTML-a pregledač treba da koristi za taj dokument. `html` element predstavlja korenski element koji uključuje sve ostale elemente.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Dodaj deklaraciju `<!DOCTYPE html>` i `html` element koji ima atribut `lang` sa vrednošću `en`.

# --hints--

Trebalo bi da imate deklaraciju `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Trebalo bi da imate etiketu pakovanja za `<html lang="en">` na engleskom jeziku.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Trebalo bi da imaš oznaku za zatvaranje `</html>`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
