---
id: 5f33294a6af5e9188dbdb8f3
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Vreme je da se doda nešto sadržaja menija. Dodajte element `main` unutar postojećeg elementa `body`. Na kraju će imati informacije o cenama kafe i užina koje restoran nudi.

# --hints--

Vaš kod mora imati ključ za otključavanje `<main>`.

```js
assert(code.match(/<main>/i));
```

Vaš kod treba da ima zatvarajući tag od `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Nemojte menjati vaš element za `body`. Uverite se da niste slučajno izbrisali svoj zatvarajući tag.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Etiketa za `main` treba biti unutar etikete za `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
