---
id: 5f332a88dc25a0fd25c7687a
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Ime restorana je `CAMPER CAFE`. Dodajte element `h1` unutar vašeg elementa `main`. Dajte restoranu ime velikim slovima kako bi bilo dobro vidljivo.

# --hints--

Trebalo bi da imaš etiketu pakovanja za `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

Ne smete da menjate vašu sekciju `main` koja se nalazi.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Tvoja komponenta za `h1` treba da bude postavljena unutar tvoje komponente za `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Vaše polje za `h1` treba da ima tekst od `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```
