---
id: 5f332b23c2045fb843337579
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Da biste obavestili goste da je kafe osnovana 2020. godine, dodajte element ``p`` ispod elementa ``h1`` sa tekstom ``Est. 2020``.

# --hints--

Trebalo bi da imaš oznaku za otvaranje `<p>`.

```js
assert(code.match(/<p>/i));
```

Trebalo bi da ima oznaku za pakovanje `</p>`.

```js
assert(code.match(/<\/p>/i));
```

Ne smete da menjate vaš deo sa sadržajem `h1` koji se nalazi. Proverite da niste izbrisali oznaku za zatvaranje.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

Element vašeg `p` treba da bude ispod elementa vašeg `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

Vaš element za `p` mora imati tekst `Est. 2020`.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```
