---
id: 5f3477cb2e27333b1ab2b955
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Sada je potrebno povezati fajl `styles.css`, kako bi se stili ponovo koristili. Unutar elementa `head`, dodajte element `link`. Dodajte atribut elementa `rel` sa vrednošću `"stylesheet"` i atribut elementa `href` sa vrednošću `"styles.css"`.

Zapamtite da je element `link` element bez sadržaja.

# --hints--

Vaš kod mora imati element `link`.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

Element `link` je element bez sadržaja, ne bi trebalo da ima zatvarajući tag `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Ne smete da promenite svoj deo `head` koji se nalazi. Uverite se da niste izbrisali oznaku za zatvaranje.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

Trebalo bi da imate jedan element `link`.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

Tvoj element `link` treba da bude unutar tvog elementa `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Vaš element od `link` treba da ima svojstvo atributa `rel` sa vrednošću `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Vaš element za `link` mora imati atribut elementa `href` sa vrednošću `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
