---
id: 5f344f9c805cd193c33d829c
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Možete dodati stil elementu opisanog unutar elementa `style` i postaviti mu atribute kao što sledi:

```css
element {
 property: value;
}
```

Postavi sadržaj elementa `h1` postavljanjem atributa `text-align` na vrednost `center`.

# --hints--

Trebalo bi da imaš birač za `h1` unutar svog elementa `style`.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasSelector);
```

Tvoja vrednost za `text-align` bi trebalo da bude `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert(hasTextAlign);
```

Vaš selektor za `h1` treba da postavi svojstvo `text-align` na `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
