---
id: 5f344fbc22624a2976425065
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Kreiraj element sa `h2` unutar komponente `section` i popuni ga tekstom od `Coffee`.

# --hints--

Trebalo bi da imaš oznaku za otvaranje `<h2>`.

```js
assert(code.match(/<h2\s*>/i));
```

Trebalo bi da imaš etiketu za pakovanje `</h2>`.

```js
assert(code.match(/<\/h2\s*>/i));
```

Ne smete da menjate vaš deo `section` koji se nalazi. Uverite se da niste izbrisali oznaku za zatvaranje.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

Vaš element `h2` mora biti unutar vašeg elementa `section`.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

Vaš element za `h2` bi trebalo da ima tekst od `Coffee`.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
