---
id: 5f3477cbcb6ba47918c1da92
title: Korak 17
challengeType: 0
dashedName: step-17
---

# --description--

Da bi se dizajn stranice prikazao slično na mobilnom telefonu kako izgleda na računaru ili laptopu, morate dodati element `meta` koji ima svojstvo specifičnog elementa `content`.

Sledeći dodatak unutar elementa `head`:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

Vaš kod treba da ima dva elementa od `meta`.

```js
assert.lengthOf(code.match(/<meta.*?\/?>/g),2);
```

Tvoj element za `meta` bi trebalo da ima atribut polja `name` sa vrednošću `viewport`.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/name=('|")viewport\1/) || meta[1]?.outerHTML?.match(/name=('|")viewport\1/));
```

Vaš element za `meta` mora imati atribut elementa `content` sa vrednošću od `width=device-width, initial-scale=1.0`.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
