---
id: 5f344fc1520b6719f2e35605
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Biće dve sekcije na meniju, jedna za kafu i jedna za užine. Dodajte element `section` unutar elementa `main` kako biste imali mesto za sve dostupne kafe.

# --hints--

Trebalo bi da imaš oznaku za otvaranje `<section>`.

```js
assert(code.match(/<section\s*>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `</section>`.

```js
assert(code.match(/<\/section\s*>/i));
```

Ne smete da promenite svoj deo `main` koji se nalazi. Uverite se da niste izbrisali oznaku za zatvaranje.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

Vaš element `section` mora biti unutar vašeg elementa `main`.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
