---
id: 5f3477ae9675db8bb7655b30
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

U prethodnom koraku, koristili ste <dfn>selektor tipa</dfn> da kreirate stil elementa `h1`. Postavite sadržaj elemenata `h2` i `p` u sredinu dodajući novi selektor tipa za svaki od njih na postojeći element `style`.

# --hints--

Ne smete menjati postojeći selektor za `h1`.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

Trebaš dodati novu oznaku za `style`. Dodaj nove pravila za CSS na postojeću oznaku od `style`.

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

Treba da dodaš novi selektor za `h2`.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

Trebalo bi da dodate novi selektor za `p`.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

Vaša karakteristika za `h1` bi trebalo da ima `text-align` od `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

Vaš element `h2` bi trebalo da ima `text-align` od `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

Vaša karakteristika za `p` mora da ima `text-align` od `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
