---
id: 5f3477ae8466a9a3d2cc953c
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Sada, nakon što imate CSS u fajlu `styles.css`, nastavite i uklonite element `style` zajedno sa svim njegovim sadržajem. Nakon što se ukloni, tekst koji je bio u sredini ponovo će se premestiti na lijevu stranu.

# --hints--

Ne bi trebalo da imate tagove sa `style` u vašem kodu.

```js
assert(!code.match(/style/i));
```

Ne smete imati selektore za CSS u vašem fajlu HTML.

```js
assert(!editableContents.includes('style'));
assert(!editableContents.includes('text-align'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
