---
id: 5f344fad8bf01691e71a30eb
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Sada imate malo kontrole nad prezentacijom i izgledom vašeg sadržaja. Da biste to promenili, dodajte element `style` unutar elementa `head`.

# --hints--

Vaš kod bi trebalo da ima oznaku za otvaranje `<style>`.

```js
assert(code.match(/<style\s*>/i));
```

Vaš kod bi trebalo da ima zatvarajući label od `</style>`.

```js
assert(code.match(/<\/style\s*>/));
```

Vaš element `style` treba da bude postavljen unutar vašeg elementa `head`.

```js
assert(code.match(/<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
--fcc-editable-region--
</html>
```
