---
id: 5f356ed63c7807a4f1e6d054
title: Korak 21
challengeType: 0
dashedName: step-21
---

# --description--

Cilj je sada da `div` ne zauzme celu širinu stranice. Svojstvo CSS `width` je savršeno za ovo.

Možete koristiti selektor `id` da biste ciljali specifičan element sa atributom `id`. Selektor `id` je definisan postavljanjem haš znaka (hash #) `#` direktno ispred vrednosti `id` elementa. Na primer, ako element ima `id` sa `cat`, onda biste želeli da ciljate taj element ovako:

```css 
#cat {
  width: 250px;
}
```

Koristite alat `#menu` da dodate vašu širinu `300px`.

# --hints--

Trebalo bi da imaš selektor za `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

Treba da postaviš vrednost za `width` na `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

`div` Vaša širina mora biti 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```
