---
id: 5f356ed6199b0cdef1d2be8f
title: Korak 26
challengeType: 0
dashedName: step-26
---

# --description--

Do sada ste koristili tipove selektora i identifikator (id) za stilizovanje elemenata. Međutim, uobičajenije je koristiti različite selektore za stilizovanje vaših elemenata.

<dfn>Selektor klase elemenata</dfn> je definisan imenom sa jednom tačkom ispred njega, kao što je:

```css
.class-name {
  styles
}
```

Promijeni postojeći selektor za `#menu` u selektor klase objekata, zamjenom `#menu` sa klasom objekata pod nazivom `.menu`.

# --hints--

Treba da imaš selektor klase elemenata `.menu`.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

Morate imati selektor za `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```
