---
id: 5f356ed656a336993abd9f7c
title: Korak 25
challengeType: 0
dashedName: step-25
---

# --description--

Sledeći, želiš da postaviš `#menu` centrirano horizontalno. To možeš uraditi postavljanjem njegovih atributa za `margin-left` i `margin-right` na `auto`. Zamisli granicu kao nevidljiv kontejner koji okružuje element. Korišćenjem ovih dve graničke atribute, postavi element `#menu` centrirano unutar elementa `body`.

# --hints--

Trebalo bi da postaviš atribut `margin-left` na `auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-left'] === 'auto');
assert(hasMargin);
```

Trebalo bi da postaviš atribut `margin-right` na `auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-right'] === 'auto');
assert(hasMargin);
```

Trebalo bi da postavite atribute za `margin-left` i `margin-right` za vaš `#menu` da bude `auto`.

```js
const divMarginRight = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-right');
const divMarginLeft = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-left');
assert(divMarginRight === 'auto');
assert(divMarginLeft === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
}
--fcc-editable-region--
```
