---
id: 5f3c866d28d7ad0de6470505
title: Korak 32
challengeType: 0
dashedName: step-32
---

# --description--

Ukusi i cene su trenutno postavljeni jedan iznad drugog, a nalaze se u sredini zajedno sa komponentama od `p`. Biće bolje ako je ukus na levoj strani, a cena na desnoj strani.

Dodaj ime klase objekta `flavor` u polje `French Vanilla` `p`.

# --hints--

Treba da dodate klasu objekata `flavor` u vaš element `p`.

```js
const el = document.querySelector("p.flavor");
assert.exists(el);
```

Trebalo bi da imate samo jedan element sa klasom objekta `flavor`.

```js
const elements = document.querySelectorAll('.flavor');
assert.lengthOf(elements, 1);
```

Vaša klasa objekata za `flavor` treba da bude u elementu od `p` koji ima tekst `French Vanilla`.

```js
const el = document.querySelector(".flavor");
assert.equal(el.innerText.trim(), "French Vanilla");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
--fcc-editable-region--
            <p>French Vanilla</p>
            <p>3.00</p>
--fcc-editable-region--
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
