---
id: 5f3c866daec9a49519871816
title: Korak 30
challengeType: 0
dashedName: step-30
---

# --description--

Atributi za `article` obično imaju mnogo atributa povezanih sa informacijama. U ovom slučaju će imati ukus kafe i cenu tog ukusa.

Postavite dva elementa za `p` unutar vašeg elementa sa `article`. Prvi element treba da ima tekst `French Vanilla`, a drugi element treba da ima tekst `3.00`.

# --hints--

Ne možete promeniti postojeći element `article`.

```js
assert.lengthOf(document.querySelectorAll('article'),1);
```

Komponenta vašeg `article` bi trebalo da ima dva elementa od `p`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
assert.lengthOf(paragraphChildren,2);
```

Vaše prvo polje za `p` trebalo bi da sadrži tekst `French Vanilla`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const firstP = paragraphChildren?.[0];
assert.match(firstP?.innerText.trim(), /French Vanilla/i);
```

Vaš drugi element od `p`` treba da ima tekst ``3.00`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const secondP = paragraphChildren?.[1];
assert.match(secondP?.innerText.trim(), /3\.00/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
