---
id: 5f3c866de7a5b784048f94b1
title: Korak 36
challengeType: 0
dashedName: step-36
---

# --description--

To je blizu onoga što želiš, ali bi sada bilo dobro da ukus i cena budu na istoj liniji. Vipengele `p` su elementi <dfn>nivo bloka elementa</dfn>, tako da zauzimaju punu širinu roditeljskog elementa.

Da biste ih stavili na jednu liniju, morate koristiti određeni stil za elemente `p`` da bi delovali kao elemente `<dfn>` unutar linije `</dfn>.

Da biste to uradili, počnite dodavanjem atributa elementu `class` sa vrednošću `item` za prvi element `article` ispod naslova sekcije `Coffee`.

# --hints--

Trebalo bi da koristiš klasu `item` za svoj element `article`.

```js
const el = document.querySelector('article.item');
assert.exists(el);
```

Trebalo bi da ima samo jedan element sa klasom `item`.

```js
const elements = document.querySelectorAll('.item');
assert.lengthOf(elements, 1);
```

Tvoj prvi element za `article` treba da ima klasu od `item`.

```js
assert.equal(document.querySelector('article')?.className, 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
