---
id: 5f7b87422a560036fd03ccff
title: Korak 50
challengeType: 0
dashedName: step-50
---

# --description--

Za dva elementa iz `p` koje ste nedavno dodali, dodajte `dessert` kao vrednost svojstva elementa `p` prvog elementa iz `class` i `price` kao svojstvo elementa `p` drugog elementa iz `class`.

# --hints--

Treba da imaš jedan `p` element sa klasom `dessert`.

```js
assert.lengthOf(document.querySelectorAll('.dessert'), 1);
```

Vaš element `p` sa tekstom `Donut` treba da ima klasu objekata `dessert`.

```js
assert.match(document.querySelector('.dessert')?.innerText, /donut/i);
```

Vaš element za `p` koji sadrži tekst `1.50` treba da ima klasu objekata `price`.

```js
const lastElement = [...document.querySelectorAll('.price')].at(-1);
assert.match(lastElement?.textContent, /1\.50/);
```

Ne bi trebalo da postoji razmak između vaših elemenata za `p`.

```js
assert(!code.match(/<\/p>\s+<p/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
--fcc-editable-region--
            <p>Donut</p><p>1.50</p>
--fcc-editable-region--
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%;
}
```
