---
id: a8e512fbe388ac2f9198f0fa
title: Primeni filter za odgovarajući predmet
challengeType: 26
dashedName: implement-a-matching-object-filter
---

# --description--

U ovoj laboratoriji, kreiraćete funkciju koja filtrira niz elemenata i vraća samo elemente koji odgovaraju svim parovima ključ-vrednost u datom izvornom objektu.

**Cilj**: Popuni priče korisnika ispod i prođi sve testove kako bi se završio laboratorijum.

**Priče korisnika:**

1. Trebalo bi da postoji funkcija `whatIsInAName` koja prima dva argumenta: niz stavki i izvorni objekat.
1. Funkcija `whatIsInAName` treba da vrati novi niz koji sadrži samo stavke iz kolekcije koje imaju sve parove ključ-vrednost prisutne u izvornom objektu.
1. Ako ne postoji stavka koja odgovara svim parovima ključ-vrednost iz izvora, funkcija treba da vrati prazan niz. Primer:

```js
whatIsInAName(
  [
    { first: "Romeo", last: "Montague" },
    { first: "Mercutio", last: null },
    { first: "Tybalt", last: "Capulet" }
  ],
  { last: "Capulet" }
);
// should return [{ first: "Tybalt", last: "Capulet" }]
```

# --hints--

Trebalo bi da imaš funkciju `whatIsInAName`.

```js
assert.isFunction(whatIsInAName);
```

`whatIsInAName([{ first: "Romeo", last: "Montague" }, { first: "Mercutio", last: null }, { first: "Tybalt", last: "Capulet" }], { last: "Capulet" })` trebalo bi da vrati `[{ first: "Tybalt", last: "Capulet" }]`.

```js
assert.deepEqual(
  whatIsInAName(
    [{ first: "Romeo", last: "Montague" }, { first: "Mercutio", last: null }, { first: "Tybalt", last: "Capulet" }],
    { last: "Capulet" }
  ),
  [{ first: "Tybalt", last: "Capulet" }]
);
```

`whatIsInAName([{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }], { "apple": 1 })` treba da vrati `[{"apple": 1}, {"apple": 1}, {"apple": 1, "bat": 2}]`.

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }],
    { "apple": 1 }
  ),
  [{ "apple": 1 }, { "apple": 1 }, { "apple": 1, "bat": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }], { "apple": 1, "bat": 2 })` trebalo bi vratiti `[{"apple": 1, "bat": 2}, {"apple": 1, "bat": 2, "cookie": 2}]`.

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }],
    { "apple": 1, "bat": 2 }
  ),
  [{ "apple": 1, "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }], { "apple": 1, "cookie": 2 })` trebalo bi da vrati `[{"apple": 1, "bat": 2, "cookie": 2}]`.

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }],
    { "apple": 1, "cookie": 2 }
  ),
  [{ "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }, { "bat": 2 }], { "apple": 1, "bat": 2 })` treba da vrati `[{"apple": 1, "bat": 2}, {"apple": 1, "bat": 2, "cookie": 2}]`

```js
assert.deepEqual(
  whatIsInAName(
    [{ "apple": 1, "bat": 2 }, { "apple": 1 }, { "apple": 1, "bat": 2, "cookie": 2 }, { "bat": 2 }],
    { "apple": 1, "bat": 2 }
  ),
  [{ "apple": 1, "bat": 2 }, { "apple": 1, "bat": 2, "cookie": 2 }]
);
```

`whatIsInAName([{"a": 1, "b": 2, "c": 3}], {"a": 1, "b": 9999, "c": 3})` trebalo bi da vrati `[]`.

```js
assert.deepEqual(
  whatIsInAName([{ "a": 1, "b": 2, "c": 3 }], { "a": 1, "b": 9999, "c": 3 }),
  []
);
```

`whatIsInAName([{"a": 1, "b": 2, "c": 3, "d": 9999}], {"a": 1, "b": 9999, "c": 3})` trebalo bi vratiti `[]`.

```js
assert.deepEqual(
  whatIsInAName([{ "a": 1, "b": 2, "c": 3, "d": 9999 }], { "a": 1, "b": 9999, "c": 3 }),
  []
);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function whatIsInAName(collection, source) {
  const sourceKeys = Object.keys(source);

  return collection.filter(obj =>
    sourceKeys.every(key => obj.hasOwnProperty(key) && obj[key] === source[key])
  );
}
```
