---
id: 668f08ea07b99b1f4a91acab
title: Napravi stranicu sa receptima
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Mora da ima deklaraciju `!DOCTYPE html`.
1. Mora da ima element `html` sa `lang` podešen na `en`.
1. Mora da ima element `head` sa poljem `title` koje nosi naziv vašeg recepta, i element `meta` sa atributom polja `charset` podešen na `UTF-8`.
1. Mora da ima element `body`.
1. Mora da ima element `h1` sa nazivom vašeg recepta.
1. Mora da ima element `p` koji identifikuje recept pod `h1`.
1. Mora da ima jedan element `h2` sa tekstom `Ingredients` za sekciju sastojaka.
1. Mora da ima neuređeni spisak (element `ul`) sa najmanje četiri stavke spiska (elementi `li`) koje navode vaše sastojke ispod prvog elementa `h2`.
1. Mora da ima drugi element `h2` sa tekstom `Instructions` za sekciju uputstava.
1. Mora da ima uređeni spisak (element `ol`) sa najmanje četiri stavke spiska koje navode korake recepta po redosledu, ispod drugog elementa `h2`.
1. Mora da ima jedan element `img` sa atributom polja `src` podešen na validnu sliku, možete koristiti `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg` ako želite, i atribut polja `alt` koji opisuje tu sliku.

# --hints--

Vaša stranica sa receptima treba da ima deklaraciju `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Trebalo bi da imaš element sa `html` koji je postavljen na `lang` i `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

Trebalo bi da imaš element `head` unutar elementa `html`.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

Trebalo bi da imaš element `title` unutar svog elementa `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

Vaš element `title` treba da ima naslov vašeg recepta.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Trebalo bi da imaš element `meta` unutar svog elementa `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

Vaš element za `meta` treba da ima atribut elemenata `charset` postavljen na `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Trebalo bi da imaš element `body` unutar tvog elementa `html`.

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

Trebalo bi da imaš element `h1` sa nazivom svog recepta.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Trebalo bi da ima samo jedan element `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Trebalo bi da imaš element `p` ispod tvog elementa `h1`.

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

Tvoja prva sekcija za `p` bi trebalo da opisuje tvoje recepte.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Prvi element vašeg `h2` treba da ima tekst `Ingredients`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

Trebalo bi da imaš neuredan list element ispod svog prvog elementa `h2`.

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

Trebalo bi da ima najmanje četiri stavke liste u neuređenom spisku koje sadrži linkove.

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Tvoj drugi element od `h2` trebao bi da sadrži tekst od `Instructions`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

Trebalo bi da imaš element listovanog sadržaja ispod tvog drugog elementa `h2`.

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

Trebalo bi da imate najmanje četiri stavke u strukturiranoj listi sa uputstvima.

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Trebalo bi da ima barem jedan element `img`.

```js
assert.exists(document.querySelector('img'));
```

Svi elementi `img` moraju imati svojstvo validnog elementa `src` i vrednost.

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

Svi elementi sa `img` moraju imati atribut elementa `alt` za opisivanje slike.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
