---
id: 587d78b1367417b2b2512b0c
title: Neka tipografija se prilagođava uređaju
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/crzN7T8'
forumTopicId: 301141
dashedName: make-typography-responsive
---

# --description--

Umesto korišćenja `em` ili `px` za merenje veličine teksta, možete koristiti jedinice vidokruga (viewport units) za tipografiju koja se prilagođava uređaju. Jedinice vidokruga, kao što su procenati, su povezane jedinice, ali zavise od različitih stvari. Jedinice vidokruga se odnose na dimenzije vidokruga (širina ili visina) uređaja, dok procenati se odnose na veličinu elementa roditeljskog kontejnera.

Različite jedinice za oblast pregleda su:

<ul><li><code>vw</code> (širina vidljivog područja): <code>10vw</code> predstavljaće 10 procenata širine vidljivog područja.</li><li><code>vh</code> (visina vidljivog područja): <code>3vh</code> predstavljaće 3 procenata visine vidljivog područja.</li><li><code>vmin</code> (minimalna vrednost vidljivog područja): <code>70vmin</code> predstavljaće 70 procenata manje dimenzije vidljivog područja (visina ili širina).</li><li><code>vmax</code> (maksimalna vrednost vidljivog područja): <code>100vmax</code> predstavljaće 100 procenata veće dimenzije vidljivog područja (visina ili širina).</li></ul>

Evo primera koji postavlja oznaku `body` na 30 procenata širine vidljivog područja.

```css
body {
  width: 30vw;
}
```

# --instructions--

Postavi `width` elementa `h2` na 80% širine prikaza, a `width` pasusa na 75% manje dimenzije prikaza.
Postavi [Label] na 80% širine vidljivog područja i [Paragraph] na 75% manjeg dimenziona vidljivog područja.

# --hints--

Tvoja širina za `h2` treba da ima `width` od 80vw.

```js
assert.match(
  __helpers.removeCssComments(code),
  /h2\s*?{\s*?width:\s*?80vw;\s*?}/g
);
```

Vaš `p` treba da ima `width` od 75vmin.

```js
assert.match(
  __helpers.removeCssComments(code),
  /p\s*?{\s*?width:\s*?75vmin;\s*?}/g
);
```

# --seed--

## --seed-contents--

```html
<style>

</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  h2 {
    width: 80vw;
  }
  p {
    width: 75vmin;
  }
</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
