---
id: 587d78b1367417b2b2512b09
title: Napravite sliku i prilagodite je uređaju
challengeType: 0
forumTopicId: 301140
dashedName: make-an-image-responsive
---

# --description--

Izmeniti slike u zavisnosti od uređaja korišćenjem CSS je veoma lako. Samo treba da dodate ove parametre na sliku:

```css
img {
  max-width: 100%;
  height: auto;
}
```

`max-width` od `100%` će osigurati da slika ne prelazi širinu kontejnera unutar njega, i `height` od `auto` će naterati sliku da sačuva svoj originalni odnos stranica.

# --instructions--

Dodajte stilski pravilnike u klasu `responsive-img` kako bi se slika prilagođavala uređaju. Nikada ne sme biti šira od svog kontejnera (u ovom slučaju, to je prozor za prikaz) i mora očuvati svoje originalno aspektno ommerenje. Nakon dodavanja vašeg koda, promenite veličinu prozora za prikaz da biste videli kako vaše slike prilagođavaju.

# --hints--

Tvoja klasa za `responsive-img` mora imati `max-width` postavljeno na `100%`.

```js
assert.strictEqual(getComputedStyle(document.querySelector('.responsive-img')).maxWidth, '100%');
```

Vaša klasa za `responsive-img` mora imati `height` postavljeno na `auto`.

```js
assert.match(code, /height:\s*?auto;/g);
```

# --seed--

## --seed-contents--

```html
<style>
.responsive-img {


}

img {
  width: 600px;
}
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```

# --solutions--

```html
<style>
  .responsive-img {
    max-width: 100%;
    height: auto;
  }

  img {
    width: 600px;
  }
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```
