---
id: 587d78b0367417b2b2512b08
title: Unda media query
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cqwKrtm'
forumTopicId: 301139
dashedName: create-a-media-query
---

# --description--

Media Queries je nova tehnika uvedena u CSS3 koja menja izgled sadržaja u zavisnosti od različitih veličina vidokruga. Vidokrug je oblast koja je vidljiva korisniku web stranice, i menja se u zavisnosti od uređaja koji se koristi za pristup veb sajtu.

Media Queries uključuju tip medija, a ako se taj tip medija poklapa sa tipom uređaja na kojem je dokument prikazan, primenjuju se stilovi. Možete imati mnogo selektora i stilova unutar vaše media queries kako god želite.

Ovo je primer media query-a koji vraća sadržaj kada je širina uređaja manja ili jednaka `100px`:

```css
@media (max-width: 100px) {
  /* CSS Rules */
}
```

i sledeći medijski upit vraća sadržaj kada je širina uređaja veća ili jednaka `350px`:

```css
@media (min-height: 350px) {
  /* CSS Rules */
}
```

Zapamtite, CSS unutar `media query`-ja se koristi samo ako tip medija odgovara onom sa korišćenog uređaja.

# --instructions--

Dodajte medijski upit, da oznaka `p` ima `font-size` od `10px` kada je visina uređaja manja ili jednaka `800px`.

# --hints--

Trebalo bi da se navede pretraga po kriterijumu za `@media` za uređaje koji imaju `height` manji ili jednaki `800px`.

```js
const media = new __helpers.CSSHelp(document).getCSSRules('media');
assert.isTrue(
  media.some(x => x.media?.mediaText?.includes('(max-height: 800px)'))
);
```

Tvoj atribut od `p` treba da ima `font-size` od `10px` kada je `height` uređaja manje ili jednako od `800px`.

```js
const rules = new __helpers.CSSHelp(document).getRuleListsWithinMedia(
  '(max-height: 800px)'
);
assert.strictEqual(
  rules?.find(x => x.selectorText === 'p')?.style?.fontSize,
  '10px'
);
```

Vaša karakteristika za `p` treba da ima početnu vrednost od `font-size` sa `20px` kada je vrednost uređaja za `height` veća od `800px`.

```js
const ifPFirst = new __helpers.CSSHelp(document)
  .getCSSRules()
  ?.find(x => x?.selectorText === 'p' || x?.media);
assert.strictEqual(ifPFirst?.style?.fontSize, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  p {
    font-size: 20px;
  }

  /* Only change code below this line */

  /* Only change code above this line */
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  p {
    font-size: 20px;
  }

  @media (max-height: 800px) {
    p {
      font-size: 10px;
    }
  }
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
