---
id: 587d78b1367417b2b2512b0a
title: Koristite sliku retine sa najvišom rezolucijom
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cVZ4Rfp'
forumTopicId: 301142
dashedName: use-a-retina-image-for-higher-resolution-displays
---

# --description--

Zbog povećanja uređaja povezanih na internet, veličina i njihove karakteristike variraju, a ekrani koje koriste mogu biti različiti spolja i unutra. Pixela gustina je karakteristika koja može varirati između jednog uređaja i drugog, a ova gustina se naziva Pixel Per Inch (PPI) ili Dots Per Inch (DPI). Najpoznatiji tip ovih ekrana jeste onaj poznat kao "Retina Display" na modernim Apple MacBook Pro računarima, a nedavno i na iMac računarima. Usled razlike u pixela gustini između "Retina" i "Non-Retina" ekrana, neke slike koje nisu kreirane uzimajući u obzir visoku rezoluciju mogu izgledati "pikselasto" kada se prikazuju na ekranima sa visokom rezolucijom.

Najlakši način da vaše slike izgledaju adekvatno na ekranima visokog rezolucije, kao što su "Retina Display" ekrani MacBook Pro-a, jeste da se vrednosti `width` i `height` podele na pola originalne datoteke. Evo primera slike koja koristi samo polovinu originalne visine i širine:

```html
<style>
  img {
    height: 250px;
    width: 250px;
  }
</style>
<img src="coolPic500x500" alt="A most excellent picture" />
```

# --instructions--

Postavite `width` i `height` za etiketu/labelu `img` da budu pola njihovih originalnih vrednosti. U ovom slučaju, `height` i `width` su sve originalne vrednosti/za podešavanje `200px`.

# --hints--

Vaša oznaka za `img` treba da ima `width` od 100 piksela.

```js
assert.strictEqual(document.querySelector('img').width, 100);
```

Tvoja oznaka za `img` bi trebalo da ima `height` od 100 piksela.

```js
assert.strictEqual(document.querySelector('img').height, 100);
```

# --seed--

## --seed-contents--

```html
<style></style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```

# --solutions--

```html
<style>
  img {
    height: 100px;
    width: 100px;
  }
</style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```
