---
id: 681cb1afdab50c87ddb2e515
title: "Challenge 4: S  P  A  C  E  J  A  M"
challengeType: 29
dashedName: challenge-4
---

# --description--
Ako je dat niz slova, ukloni sve razmake iz tog niza, stavi dva razmaka između svakog slova, preтвори sva slova abecede u VELIKA SLOVA i vrati rezultat.

- Nealfabetski znakovi treba da ostanu nepromenjeni (osim razmaka).

# --hints--

`space_jam("freeCodeCamp")` should return `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("   free   Code   Camp   ")` should return `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("Hello World?!")` should return `"H  E  L  L  O  W  O  R  L  D  ?  !"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !")`)
}})
```

`space_jam("C@t$ & D0g$")` should return `"C  @  T  $  &  D  0  G  $"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $")`)
}})
```

`space_jam("allyourbase")` should return `"A  L  L  Y  O  U  R  B  A  S  E"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E")`)
}})
```

# --seed--

## --seed-contents--

```py
def space_jam(s):

    return s
```

# --solutions--

```py
def space_jam(s):
    s = s.replace(" ", "")
    s = "  ".join(s)
    return s.upper()
```
