---
id: 6821ec02237de8297eaee79a
title: "Challenge 24: Pangram"
challengeType: 29
dashedName: challenge-24
---

# --description--
Ako dobiješ reč ili rečenicu zajedno sa nizom malih slova, utvrdi da li ta reč ili rečenica koristi sva slova iz dato seta barem jednom i da nema nijedno drugo slovo.

- Ignoriši nealfabetičke znakove u reči ili rečenici.
- Ignoriši razliku između velikih i malih slova u reči ili rečenici.

# --hints--

`is_pangram("hello", "helo")` should return `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "helo"), True)`)
}})
```

`is_pangram("hello", "hel")` should return `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "hel"), False)`)
}})
```

`is_pangram("hello", "helow")` should return `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello", "helow"), False)`)
}})
```

`is_pangram("hello world", "helowrd")` should return `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("hello world", "helowrd"), True)`)
}})
```

`is_pangram("Hello World!", "helowrd")` should return `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("Hello World!", "helowrd"), True)`)
}})
```

`is_pangram("Hello World!", "heliowrd")` should return `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("Hello World!", "heliowrd"), False)`)
}})
```

`is_pangram("freeCodeCamp", "frcdmp")` should return `False`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("freeCodeCamp", "frcdmp"), False)`)
}})
```

`is_pangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz")` should return `True`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_pangram(sentence, letters):

    return sentence
```

# --solutions--

```py
import re
def is_pangram(sentence, letters):
    used_letters = []
    for char in sentence.lower():
        if re.match(r'[a-z]', char) and char not in used_letters:
            used_letters.append(char)

    sorted_letters = ''.join(sorted(letters.lower()))
    sorted_used_letters = ''.join(sorted(used_letters))

    return sorted_letters == sorted_used_letters
```
