---
id: 68adce01c0e1144d0a90295e
title: "Challenge 29: Acronym Builder"
challengeType: 29
dashedName: challenge-29
---

# --description--
Ako vam je dat niz karaktera koji sadrži jedno ili više reči, vratite početna slova tih reči korišćenjem sledećih uslova:

- Početna slova treba da budu velika (kapitala), osim ako nije drugačije navedeno.
- Početna slova sledećih reči ne bi trebalo da budu velika (kapitala), osim ako nisu prva reč u datoj sekvenci: `a`, `for`, `an`, `and`, `by`, i `of`.
- Početna slova treba da se vrate u redosleda kojim su data.
- Početna slova ne smeju imati razmak između sebe.

# --hints--

`build_acronym("Search Engine Optimization")` should return `"SEO"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Search Engine Optimization"), "SEO")`)
}})
```

`build_acronym("Frequently Asked Questions")` should return `"FAQ"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Frequently Asked Questions"), "FAQ")`)
}})
```

`build_acronym("National Aeronautics and Space Administration")` should return `"NASA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("National Aeronautics and Space Administration"), "NASA")`)
}})
```

`build_acronym("Federal Bureau of Investigation")` should return `"FBI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("Federal Bureau of Investigation"), "FBI")`)
}})
```

`build_acronym("For your information")` should return `"FYI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("For your information"), "FYI")`)
}})
```

`build_acronym("By the way")` should return `"BTW"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("By the way"), "BTW")`)
}})
```

`build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` should return `"AUHWPOTIMSH"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_acronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH")`)
}})
```

# --seed--

## --seed-contents--

```py
def build_acronym(s):

    return s
```

# --solutions--

```py
def build_acronym(s):
    small_words = {"a", "for", "an", "and", "by", "of"}
    words = s.split()
    acronym = ""

    for i, word in enumerate(words):
        if i == 0 or word.lower() not in small_words:
            acronym += word[0].upper()

    return acronym
```
