---
id: 68adce01c0e1144d0a902956
title: "Challenge 25: Vowel Repeater"
challengeType: 29
dashedName: challenge-25
---

# --description--
Ako uzmete sekvencu slova, vratite novu verziju te sekvence gde je svaki samoglasnik ponovljen za jedan put više nego prethodni susrećeni samoglasnik. Na primer, prvi samoglasnik u rečenici se ne menja. Drugi samoglasnik pojavljuje se dvaput uzastopno. Treći samoglasnik pojavljuje se trostruko uzastopno, i tako dalje.

- Slova `a`, `e`, `i`, `o`, na `u`, bilo da su velika ili mala slova, smatraju se vokalima.
- Originalni vokal treba da sačuva svoj tip slova.
- Ponovljeni vokali treba da budu mala slova.
- Sve slova koja nisu vokali treba da sačuvaju svoj tip slova kako su originalno.

# --hints--

`repeat_vowels("hello world")` should return `"helloo wooorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")` should return `"freeeCooodeeeeCaaaaamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")` should return `"AEeIiiOoooUuuuu"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")` should return `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
