---
id: 68af0687ef34c76c28ffa54b
title: "Challenge 32: Reverse Sentence"
challengeType: 29
dashedName: challenge-32
---

# --description--
Dati je niz slova iz reči, vrati novi niz slova sa rečima u obrnutom redosledu. Na primer, prva reč mora biti na kraju vraćenog niza, a poslednja reč mora biti na početku vraćenog niza.

- U datoj sekvenci, reči mogu biti razdvojene jednim ili više razmaka.
- Vraćena sekvenca mora imati samo jedan razmak između reči.

# --hints--

`reverse_sentence("world hello")` should return `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("world hello"), "hello world")`)
}})
```

`reverse_sentence("push commit git")` should return `"git commit push"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("push commit git"), "git commit push")`)
}})
```

`reverse_sentence("npm install apt sudo")` should return `"sudo apt install npm"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("npm  install   apt    sudo"), "sudo apt install npm")`)
}})
```

`reverse_sentence("import default function export")` should return `"export function default import"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("import    default   function  export"), "export function default import")`)
}})
```

# --seed--

## --seed-contents--

```py
def reverse_sentence(sentence):

    return sentence
```

# --solutions--

```py
def reverse_sentence(sentence):

   return " ".join(sentence.split()[::-1])
```
