---
id: 68b7687dded630607aceccab
title: "Challenge 45: Perfect Square"
challengeType: 29
dashedName: challenge-45
---

# --description--
Ako ti je dat pun broj, odluči da li je savršen kvadrat.

- Broj je savršen kvadrat ako možeš pomnožiti celobrojni broj sa njim da dobiješ taj broj. Na primer, 9 je savršen kvadrat jer možeš pomnožiti 3 sa sobom i dobiti 9.

# --hints--

`is_perfect_square(9)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(9), True)`)
}})
```

`is_perfect_square(49)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(49), True)`)
}})
```

`is_perfect_square(1)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(1), True)`)
}})
```

`is_perfect_square(2)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(2), False)`)
}})
```

`is_perfect_square(99)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(99), False)`)
}})
```

`is_perfect_square(-9)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(-9), False)`)
}})
```

`is_perfect_square(0)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(0), True)`)
}})
```

`is_perfect_square(25281)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_perfect_square(25281), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_perfect_square(n):

    return n
```

# --solutions--

```py
import math
def is_perfect_square(n):
    if n < 0:
        return False
    root = int(math.sqrt(n))
    return root * root == n
```
