---
id: 68c1a929005bf54d342aa8d3
title: "Challenge 56: Space Week Day 2: Exoplanet Search"
challengeType: 29
dashedName: challenge-56
---

# --description--
Drugog dana Svemirske nedelje, dat vam je niz karaktera gde svaki karakter predstavlja meru svetlosne jačine zvezde. Identifikujte da li mere otkrivaju egzoplanete koristeći metodu tranzita (transit method). Tranzitna metoda je kada planeta prođe ispred zvezde, smanjujući vidljivu svetlost.

- Mere osvetljenosti se sastoje samo od karaktera `0-9` i `A-Z`, gde svako merenje odgovara sledećim numeričkim vrednostima:
- Karakteri `0-9` predstavljaju nivoe osvetljenosti `0-9`.
- Karakteri `A-Z` predstavljaju nivoe osvetljenosti `10-35`.

Zvezda se smatra da ima eksoplanetu ako je bilo koja jedna mera manje ili jednaka 80 procenata proseka svih merenja. Na primer, ako je prosečna svetlina zvezde 10, smatraće se da ima eksoplanetu ako je bilo koja jedna mera 8 ili manja.

# --hints--

`has_exoplanet("665544554")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("665544554"), False)`)
}})
```

`has_exoplanet("FGFFCFFGG")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FGFFCFFGG"), True)`)
}})
```

`has_exoplanet("MONOPLONOMONPLNOMPNOMP")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("MONOPLONOMONPLNOMPNOMP"), False)`)
}})
```

`has_exoplanet("FREECODECAMP")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FREECODECAMP"), True)`)
}})
```

`has_exoplanet("9AB98AB9BC98A")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("9AB98AB9BC98A"), False)`)
}})
```

`has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_exoplanet(readings):

    return readings
```

# --solutions--

```py
def has_exoplanet(readings):
    values = [int(c, 36) for c in readings]
    average = sum(values) / len(values)
    threshold = average * 0.8
    for v in values:
        if v <= threshold:
            return True
    return False
```
