---
id: 68c1a929005bf54d342aa8d6
title: "Challenge 59: Space Week Day 5: Goldilocks Zone"
challengeType: 29
dashedName: challenge-59
---

# --description--
Na petom danu Nedelje svemira, izračunaćeš "Goldilocks zonu" zvezde – oblast koja okružuje zvezdu gde su uslovi "savršeni" za postojanje tečne vode.

Uzimajući u obzir masu zvezde, vrati red podataka sa početnom i krajnjom udaljenošću njegove Zone Goldiloks za svemirske mere.

Ili kuhesabu Eneo la Goldilocks:

1. Pronađite svetilnost zvezde podizanjem njene jačine na stepen 3,5.
2. Početak površine je 0,95 puta kvadratnog korena njene svetilnosti.
3. Kraj površine je 1,37 puta kvadratnog korena njene svetilnosti.

- Zaokruži smanjenu vrednost na dve decimale.

Na primer, ako dobijete `1` kao težinu, vratite `[0.95, 1.37]`.

# --hints--

`goldilocks_zone(1)` should return `[0.95, 1.37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` should return `[0.28, 0.41]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` should return `[21.85, 31.51]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` should return `[9.38, 13.52]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` should return `[179.69, 259.13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
