---
id: 68c1a929005bf54d342aa8d5
title: "Challenge 58: Space Week Day 4: Landing Spot"
challengeType: 29
dashedName: challenge-58
---

# --description--
U četvrtom danu Svemirske nedelje, dobijate matricu brojeva (mreža redova podataka), koja predstavlja zone gde vaš rover može da sleti. Pronađite najsigurnije mesto za slanje prema sledećim pravilima:

- Svaka ćelija u matrici će imati broj od `0-9`, uključujući.
- Svaki `0` predstavlja moguće mesto za sletanje.
- Svaki broj, osim `0`, je veoma opasan za sletanje. Kako se broj povećava, tako raste i opasnost.
- Najsigurnije mesto definisano je kao ćelija `0` čije okolne ćelije (do 4 suseda, ne uzimati uglove) imaju najniži ukupni nivo opasnosti.
- Ignorišite susede van granica (uglasi i ivice imaju samo malo suseda).
- Vratite indikatore najsigurnijeg mesta za sletanje. Uvek će postojati samo jedno najsigurnije mesto.

Na primer, ako dobijete:

```js
[
  [1, 0],
  [2, 0]
]
```

Vrati `[0, 1]`, indikatori za `0` u prvom redu.

# --hints--

`find_landing_spot([[1, 0], [2, 0]])` should return `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 0], [2, 0]]), [0, 1])`)
}})
```

`find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` should return `[1, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1])`)
}})
```

`find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` should return `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2])`)
}})
```

`find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` should return `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_landing_spot(matrix):

    return matrix
```

# --solutions--

```py
def find_landing_spot(matrix):
    best_spot = None
    lowest_neighbor_sum = float('inf')

    for i in range(len(matrix)):
        for j in range(len(matrix[i])):
            if matrix[i][j] == 0:
                current_neighbor_sum = 0

                if i > 0:
                    current_neighbor_sum += matrix[i - 1][j]
                if j < len(matrix[i]) - 1:
                    current_neighbor_sum += matrix[i][j + 1]
                if i < len(matrix) - 1:
                    current_neighbor_sum += matrix[i + 1][j]
                if j > 0:
                    current_neighbor_sum += matrix[i][j - 1]

                if current_neighbor_sum < lowest_neighbor_sum:
                    lowest_neighbor_sum = current_neighbor_sum
                    best_spot = [i, j]

    return best_spot
```
