---
id: 68c497f3aaefc9fd9f1b0e25
title: "Challenge 61: Space Week Day 7: Launch Fuel"
challengeType: 29
dashedName: challenge-61
---

# --description--
Za poslednji dan Nedelje svemira, biće vam dat težina u kilogramima (kg) tereta koji želite poslati u svemir. Odredite količinu goriva potrebnu za slanje vašeg tereta u svemir koristeći sledeća pravila:

- Raketni motori zahtevaju 1 kg goriva na svaki 5 kg podignutog tereta.
- Samo gorivo ima težinu. Stoga, kada dodate gorivo, povećava se ukupna težina dizanja, što zahteva više goriva, a to povišava težinu, i tako dalje.
- Da biste izračunali ukupno potrebno gorivo: počnite sa težinom tereta, izračunajte gorivo potrebno za tu težinu, dodajte to gorivo na ukupnu težinu i ponovo izračunajte. Ponovite ovaj proces dok ne smanji dodatno potrebno gorivo za 1 kg, zatim prestanite.
- Ignorišite težinu same rakete. Izračunajte samo gorivo potrebno za podizanje tereta i njegovog sopstvenog goriva.

Na primer, ako dobijete teret od 50 kg, biće vam potrebno 10 kg goriva da podignete taj teret (teret / 5), što povećava ukupnu težinu na 60 kg, koja zahteva 12 kg za podizanje (dodva 2 kg), što povećava ukupnu težinu na 62 kg, koja zahteva 12.4 kg za podizanje - dodatnih 0.4 kg - što je manje od dodatnog 1 kg, pa ovde prestajemo. Ukupna teža za podizanje je 62.4 kg, gde je 50 početni teret, a 12.4 je gorivo.

- Zaokruži potrebnu količinu goriva na jednu decimalnu zaostavku.

# --hints--

`launch_fuel(50)` should return `12.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(50), 12.4)`)
}})
```

`launch_fuel(500)` should return `124.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(500), 124.8)`)
}})
```

`launch_fuel(243)` should return `60.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(243), 60.7)`)
}})
```

`launch_fuel(11000)` should return `2749.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(11000), 2749.8)`)
}})
```

`launch_fuel(6214)` should return `1553.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(6214), 1553.4)`)
}})
```

# --seed--

## --seed-contents--

```py
def launch_fuel(payload):

    return payload
```

# --solutions--

```py
def launch_fuel(payload):
    total_mass = payload
    additional_fuel = total_mass / 5
    total_fuel = additional_fuel

    while additional_fuel >= 1:
        total_mass += additional_fuel
        additional_fuel = (total_mass / 5) - total_fuel
        total_fuel += additional_fuel

    return round(total_fuel, 1)
```
