---
id: 68cae5b538ff798bbd4da001
title: "Challenge 63: Battle of Words"
challengeType: 29
dashedName: challenge-63
---

# --description--
Uzimajući u obzir dve rečenice koje predstavljaju vaš tim i protivnički tim, gde svaka reč iz vašeg tima se takmiči sa odgovarajućom reči protivničkog tima, utvrdite koji tim pobeđuje koristeći sledeća pravila:

- Pružene rečenice će uvek imati isti broj reči.
- Reči su razdvojene samo jednim razmakom i sadržiće samo slova.
- Vrednost svake reči je zbir njenih slova.
- Slova `a` do `z` odgovaraju vrednosti `1` do `26`. Na primer, `a` je `1`, a `z` je `26`.
- Velika slova udvostručavaju vrednost tog slova. Na primer, `A` je `2`, a `Z` je `52`.
- Reči se takmiče po redu: prva reč vašeg tima takmiči se sa prvom rečju protivničkog tima, i tako dalje.
- Reč pobeđuje ako je njena vrednost veća od vrednosti reči protivnika.
- Tim sa najviše pobedničkih reči je pobjednik.

Vrati `"We win"` ako je tima vaš pobjednik, `"We lose"` ako je tima izgubila, i `"Draw"` ako obe ekipe imaju jednak broj pobeda.

# --hints--

`battle("hello world", "hello word")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "hello word"), "We win")`)
}})
```

`battle("Hello world", "hello world")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello world", "hello world"), "We win")`)
}})
```

`battle("lorem ipsum", "kitty ipsum")` should return `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("lorem ipsum", "kitty ipsum"), "We lose")`)
}})
```

`battle("hello world", "world hello")` should return `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "world hello"), "Draw")`)
}})
```

`battle("git checkout", "git switch")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("git checkout", "git switch"), "We win")`)
}})
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` should return `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose")`)
}})
```

`battle("We must never surrender", "Our team must win")` should return `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("We must never surrender", "Our team must win"), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(our_team, opponent):

    return our_team
```

# --solutions--

```py
def get_word_value(word):
    value = 0
    for char in word:
        val = ord(char.lower()) - ord('a') + 1
        if char.isupper():
            val *= 2
        value += val
    return value

def battle(our_team, opponent):
    my_wins = 0
    opp_wins = 0

    for my_word, opp_word in zip(our_team.split(), opponent.split()):
        my_val = get_word_value(my_word)
        opp_val = get_word_value(opp_word)

        if my_val > opp_val:
            my_wins += 1
        elif opp_val > my_val:
            opp_wins += 1

    if my_wins > opp_wins:
        return "We win"
    elif opp_wins > my_wins:
        return "We lose"
    else:
        return "Draw"
```
