---
id: 68cae5b538ff798bbd4da007
title: "Challenge 69: Missing Socks"
challengeType: 29
dashedName: challenge-69
---

# --description--
Uzimajući u obzir ukupan broj koji predstavlja početni broj parova čarapa, i drugi broj koji predstavlja koliko ste vešeračkih ciklusa uradili, vratite trenutni broj kompletnih parova čarapa koristeći sledeće uslove:

- Svakih 2 pranja (cikla pranja), gubite jednu čarapeu.
- Svakih 3 pranja (cikla pranja), pronađete jednu izgubljenu čarapeu.
- Svakih 5 pranja (cikla pranja), jedna čarapa se ošteti i mora biti bacena.
- Svakih 10 pranja (cikla pranja), kupujete par novih čarapa.
- Ne možete imati manje od nule čarapa.
- Pravila mogu da se kombinuju. Na primer, u ciklusu pranja od 10, gubite jednu čarapeu, bacate jednu čarapeu i kupujete par novih čarapa.
- Vratite broj kompletnih para čarapa.

# --hints--

`sock_pairs(2, 5)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(2, 5), 1)`)
}})
```

`sock_pairs(1, 2)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 2), 0)`)
}})
```

`sock_pairs(5, 11)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(5, 11), 4)`)
}})
```

`sock_pairs(6, 25)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(6, 25), 3)`)
}})
```

`sock_pairs(1, 8)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 8), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sock_pairs(pairs, cycles):

    return pairs
```

# --solutions--

```py
def sock_pairs(pairs, cycles):
    socks = pairs * 2

    for i in range(1, cycles + 1):
        if i % 2 == 0:
            socks -= 1
        if i % 3 == 0:
            socks += 1
        if i % 5 == 0:
            socks -= 1
        if i % 10 == 0:
            socks += 2

        if socks < 0:
            socks = 0

    return socks // 2
```
