---
id: 68cae5b538ff798bbd4da004
title: "Vežba 66: Ukloni HTML oznake"
challengeType: 29
dashedName: challenge-66
---

# --description--
Uklanjanje sekvence znakova koda za HTML, uklonite oznaku i vratite sadržaj običnog teksta.

- Sekvenca znakova unosa će imati samo važeći HTML.
- Etikete za HTML mogu biti interno konfigurisane.
- Uklonite sve oznake i atribute polja.

Na primer, `'<a href="#">Click here</a>'` trebalo bi da vrati `"Click here"`.

# --hints--

`strip_tags('<a href="#">Click here</a>')` inapaswa kurudisha `"Click here"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<a href="#">Click here</a>'), "Click here")`)
}})
```

`strip_tags('<p class="center">Hello <b>World</b>!</p>')` inapaswa kurudisha `"Hello World!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!")`)
}})
```

`strip_tags('<img src="cat.jpg" alt="Cat">')` should return `""`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<img src="cat.jpg" alt="Cat">'), "")`)
}})
```

`strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` inapaswa kurudisha `sectionsection`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection")`)
}})
```

# --seed--

## --seed-contents--

```py
def strip_tags(html):

    return html
```

# --solutions--

```py
import re
def strip_tags(html):
    return re.sub(r'<[^>]*>', '', html)
```
