---
id: 68d2ba1468508398389487ce
title: "Challenge 73: Speak Wisely, You Must"
challengeType: 29
dashedName: challenge-73
---

# --description--
Ako napišeš rečenicu, vrati joj verziju koja zvuči kao savet od mudrog učitelja koristeći sledeća pravila:

- Reči su razdvojene samo jednim razmakom.
- Pronađite prvi pojavljivanje bilo koje od sledećih reči u rečenici: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Premestite sve reči pre, zajedno sa tom reči, na kraj rečenice i:
  - Sačuvajte redosled reči dok ih premještate.
  - Pretvorite sve u mala slova.
  - I dodajte zarez i razmak ispred njih.
- Napravite veliko slovo prve reči nove rečenice.
- Sve izvađene rečenice će se završiti jednim interpunkcijskim znakom. Sačuvajte originalni interpunkcijski znak rečenice i premjestite ga na kraj nove rečenice.
- Vratite novu rečenicu, osiguravši da postoji jedan razmak između svake riječi i da nema razmaka na početku ili kraju rečenice.

Na primer, ako dobijete `"You must speak wisely."` vratite `"Speak wisely, you must."`

# --hints--

`wise_speak("You must speak wisely.")` should return `"Speak wisely, you must."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You must speak wisely."), "Speak wisely, you must.")`)
}})
```

`wise_speak("You can do it!")` should return `"Do it, you can!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You can do it!"), "Do it, you can!")`)
}})
```

`wise_speak("Do you think you will complete this?")` should return `"Complete this, do you think you will?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("Do you think you will complete this?"), "Complete this, do you think you will?")`)
}})
```

`wise_speak("All your base are belong to us.")` should return `"Belong to us, all your base are."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("All your base are belong to us."), "Belong to us, all your base are.")`)
}})
```

`wise_speak("You have much to learn.")` should return `"Much to learn, you have."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You have much to learn."), "Much to learn, you have.")`)
}})
```

# --seed--

## --seed-contents--

```py
def wise_speak(sentence):

    return sentence
```

# --solutions--

```py
def wise_speak(sentence):
    triggers = ["have", "must", "are", "will", "can"]
    punctuation = sentence[-1]
    words = sentence[:-1].split(" ")

    index = next(i for i, w in enumerate(words) if w in triggers)

    to_move = [w.lower() for w in words[:index + 1]]
    remaining = words[index + 1:]

    if remaining:
        remaining[0] = remaining[0][0].upper() + remaining[0][1:]

    return " ".join(remaining) + ", " + " ".join(to_move) + punctuation
```
