---
id: 68d30fc57588d97fd3027b30
title: "Challenge 79: Navigator"
challengeType: 29
dashedName: challenge-79
---

# --description--

Mnamo Oktoba 28, 1994, Netscape Navigator ilizinduliwa, ikisaidia mamilioni ya watu kuchunguza mtandao wa mwanzo.

Uzimajući u obzir niz podataka o komandama pregledača koje ste izvršili na Netscape Navigator-u, vratite trenutnu stranicu koju imate nakon izvršenja svih komandi koristeći sledeće pravilo:

- Uvijek počinjete na stranici `"Home"`, koja nije uključena u listu komandi.
- Važeće komande su:
  - `"Visit Page"`: Gdje je `"Page"` naziv stranice koju posjećujete. Na primjer, `"Visit About"` vodi vas na stranicu `"About"`. Kada posjetite novu stranicu, pobrinite se da obrišete bilo koju istoriju napred (forward history) koju imate.
  - `"Back"`: Vodi vas na prethodno posjećenu stranicu u vašoj istoriji ili ostaje na trenutnoj stranici ako nema takve.
  - `"Forward"`: Vodi vas napred u istoriju do stranice sa koje ste otišli ili ostaje na trenutnoj stranici ako nema takve.

Na primer, ako dobiješ `["Visit About Us", "Back", "Forward"]`, vrati `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` should return `"About Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Back", "Forward"]), "About Us")`)
}})
```

`navigate(["Forward"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Forward"]), "Home")`)
}})
```

`navigate(["Back"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Back"]), "Home")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery"])` should return `"Gallery"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery"]), "Gallery")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home")`)
}})
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` should return `"Contact"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact")`)
}})
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` should return `"Visit Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate(commands):

    return commands
```

# --solutions--

```py
def navigate(commands):
    history = ["Home"]
    current_index = 0

    for command in commands:
        if command.startswith("Visit "):
            history = history[:current_index + 1]
            history.append(command[6:])
            current_index += 1
        elif command == "Back" and current_index > 0:
            current_index -= 1
        elif command == "Forward" and current_index < len(history) - 1:
            current_index += 1

    return history[current_index]
```
