---
id: 68d30845cc08266018fc46bd
title: "Challenge 77: Duration Formatter"
challengeType: 29
dashedName: challenge-77
---

# --description--
Ako dobijete pun broj sekundi, vratite niz znakova koji predstavlja isto vreme u formatu `"H:MM:SS"`, gde je `"H"` broj sati, `"MM"` broj minuta, i `"SS"` broj sekundi. Vratite vreme koristeći sledeće pravilo:

- Sekunde: Moraju imati dve cifre uvek.
- Minute: Moraju ukloniti vodeće nule kada nisu potrebne. Koristite `"0"` ako je vreme manje od jedne minute.
- Sati: Moraju biti uključeni samo ako su veći od nule.

# --hints--

`format(500)` should return `"8:20"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(500), "8:20")`)
}})
```

`format(4000)` should return `"1:06:40"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(4000), "1:06:40")`)
}})
```

`format(1)` should return `"0:01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(1), "0:01")`)
}})
```

`format(5555)` should return `"1:32:35"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(5555), "1:32:35")`)
}})
```

`format(99999)` should return `"27:46:39"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(99999), "27:46:39")`)
}})
```

# --seed--

## --seed-contents--

```py
def format(seconds):

    return seconds
```

# --solutions--

```py
def format(seconds):
    h = seconds // 3600
    m = (seconds % 3600) // 60
    s = seconds % 60

    seconds_str = f"{s:02d}"
    minutes_str = str(m)

    if h > 0:
        hours_str = str(h)
        return f"{hours_str}:{minutes_str.zfill(2)}:{seconds_str}"
    else:
        return f"{minutes_str}:{seconds_str}"
```
