---
id: 68e39ed6106dac2f0a98fd66
title: "Challenge 84: Infected"
challengeType: 29
dashedName: challenge-84
---

# --description--
Dana 2. novembra 1988., pušten je prvi veliki internet crv, ostavljajući za sobom približno 10 procenata povezanih računara zaraženih nakon samo jednog dana.

U ovom izazovu, dobijate broj dana koji su prošli od oslobađanja internetskog crva i morate utvrditi koliko računara je zaražano koristeći sledeća pravila:

- Dan 0, prvi računar je zaražen.
- Svaki sledeći dan, broj zaraženih računara se udvostručuje.
- Svaki treći dan, koristi se zaštita nakon što virus proširi i smanji broj zaraženih računara za 20 procenata. Zaokružite broj zaštićenih računara na sledeći ceo broj.

Na primer, dan:

- Dan 0: računar 1 ukupno je zaražen.
- Dan 1: računar 2 ukupno su zaražene.
- Dan 2: računar 4 ukupno su zaražene.
- Dan 3: računar 8 ukupno su zaražene. Zatim, koristi vakcinu: računar 8 zaraženih * %20 = 1.6 dobili je zaštitu. Zaokruži 1.6 na 2. Računar 8 zaraženih - 2 dobili je zaštitu = računar 6 ukupno je zaraženo nakon Dana 3.

Rudisha idadi ya kompyuta jumla zilizambukizwa baada ya idadi ya siku zilizotolewa kupita.

# --hints--

`infected(1)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(1), 2)`)
}})
```

`infected(3)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(3), 6)`)
}})
```

`infected(8)` should return `152`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(8), 152)`)
}})
```

`infected(17)` should return `39808`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(17), 39808)`)
}})
```

`infected(25)` should return `5217638`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(25), 5217638)`)
}})
```

# --seed--

## --seed-contents--

```py
def infected(days):

    return days
```

# --solutions--

```py
import math
def infected(days):
    infected = 1
    for day in range(1, days + 1):
        infected *= 2
        if day % 3 == 0:
            patched = math.ceil(infected * 0.2)
            infected -= patched

    return infected
```
