---
id: 68ee9e3066cfd4eb2328e8a8
title: "Challenge 89: Counting Cards"
challengeType: 29
dashedName: challenge-89
---

# --description--
Standardni komplet karata ima 13 jedinstvenih karata u svakoj boji. Ako dobijete kompletan broj koji predstavlja broj karata koje treba izabrati iz pakovanja, vratite broj jedinstvenih kombinacija karata koje možete izabrati.

- Redosled nije važan. Izbor karte A, a zatim karte B je jednak izboru karte B, a zatim karte A.

Na primer, ako dobiješ `52`, vrati `1`. Postoji samo jedna kombinacija od 52 karte iz paketa od 52 karte. I ako dobiješ `2`, vrati `1326`. Postoji kombinacija od 1326 karata koje možeš dobiti kada izabereš 2 karte iz paketa.

# --hints--

`combinations(52)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(52), 1)`)
}})
```

`combinations(1)` should return `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(1), 52)`)
}})
```

`combinations(2)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(2), 1326)`)
}})
```

`combinations(5)` should return `2598960`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(5), 2598960)`)
}})
```

`combinations(10)` should return `15820024220`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(10), 15820024220)`)
}})
```

`combinations(50)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(50), 1326)`)
}})
```

# --seed--

## --seed-contents--

```py
def combinations(cards):

    return cards
```

# --solutions--

```py
from math import factorial
def combinations(cards):
    n = 52
    
    return factorial(n) // (factorial(cards) * factorial(n - cards))
```
