---
id: 68ee9e3066cfd4eb2328e8a5
title: "Challenge 86: Image Search"
challengeType: 29
dashedName: challenge-86
---

# --description--
Na 4. novembra 2001., Google je lansirao svoju pretragu slika, omogućavajući ljudima da pronađu slike koristeći reči za pretraživanje. U ovom izazovu, ćete simulirati pretragu slika.

Ako pružite red podataka sa imenima slika i terminom pretrage, vratite red podataka imena slika koja sadrže termin pretrage.

- Ignorišite razlike između velikih i malih slova prilikom poređenja traženih reči.
- Vratiti slike u istom redosledu kako su prikazane u unosnim podacima.

# --hints--

`image_search(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` should return `["dog.png"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"])`)
}})
```

`image_search(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` should return `["Sunset.jpg", "sunflower.jpeg"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"])`)
}})
```

`image_search(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` should return `["Moon.png", "stars.png"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"])`)
}})
```

`image_search(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` should return `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(image_search(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"])`)
}})
```

# --seed--

## --seed-contents--

```py
def image_search(images, term):

    return images
```

# --solutions--

```py
def image_search(images, term):
    lower_term = term.lower()
    return [img for img in images if lower_term in img.lower()]
```
