---
id: 68f6587287ad1f4ad39b0c84
title: "Challenge 98: Rectangle Count"
challengeType: 29
dashedName: challenge-98
---

# --description--
Ako date dva pozitivna celobrojka koja predstavljaju širinu i dužinu pravougaonika, utvrdite koliko linija/segmenta tog pravougaonika se može smestiti unutar datog.

- Broj samo pravougaonike čija su širina i visina celi brojevi.

Na primer, ako uklonite `1` i `3`, vratite `6`. Tri linije od 1x1, dve linije od 1x2, i jedna linija od 1x3.

# --hints--

`count_rectangles(1, 3)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 3), 6)`)
}})
```

`count_rectangles(3, 2)` should return `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(3, 2), 18)`)
}})
```

`count_rectangles(1, 2)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 2), 3)`)
}})
```

`count_rectangles(5, 4)` should return `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(5, 4), 150)`)
}})
```

`count_rectangles(11, 19)` should return `12540`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(11, 19), 12540)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_rectangles(width, height):

    return width
```

# --solutions--

```py
def count_rectangles(width, height):
    total = 0
    for w in range(1, width + 1):
        for h in range(1, height + 1):
            total += (width - w + 1) * (height - h + 1)
    return total
```
