---
id: 68f6587287ad1f4ad39b0c83
title: "Challenge 97: GCD"
challengeType: 29
dashedName: challenge-97
---

# --description--

Uzimajući dva pozitivna cela broja, vrati njihov najveći zajednički delilac (GCD).

- Najveći zajednički delilitelj (NZD) dve celobrojke je najveći broj koji deli oba broja po jednakosti, bez ostanka.

Na primer, delitelji od `4` su `1`, `2`, i `4`. Delitelji od `6` su `1`, `2`, `3`, i `6`. Stoga, uzimajući u obzir `4` i `6`, vratite `2`, najveći broj koji se pojavljuje u oba skupa delitelja.

# --hints--

`gcd(4, 6)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(4, 6), 2)`)
}})
```

`gcd(20, 15)` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(20, 15), 5)`)
}})
```

`gcd(13, 17)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(13, 17), 1)`)
}})
```

`gcd(654, 456)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(654, 456), 6)`)
}})
```

`gcd(3456, 4320)` should return `864`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(3456, 4320), 864)`)
}})
```

# --seed--

## --seed-contents--

```py
def gcd(x, y):

    return x
```

# --solutions--

```py
def gcd(x, y):
    while y != 0:
        x, y = y, x % y
    return x
```
