---
id: 68f6587287ad1f4ad39b0c7c
title: "Challenge 90: Character Limit"
challengeType: 29
dashedName: challenge-90
---

# --description--
U ovom izazovu, dobijate niz slova i morate utvrditi da li je usklađen sa objavom na društvenim mrežama. Vratite sledeći niz slova u skladu sa datiim pravilima:

- `"short post"` ako odgovara dužini od 40 karaktera.
- `"long post"` ako je više od 40 karaktera i odgovara dužini od 80 karaktera.
- `"invalid post"` ako je preduga da bi odgovarala bilo kojim od dva standarda.

# --hints--

`can_post("Hello world")` should return `"short post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("Hello world"), "short post")`)
}})
```

`can_post("This is a longer message but still under eighty characters.")` should return `"long post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This is a longer message but still under eighty characters."), "long post")`)
}})
```

`can_post("This message is too long to fit into either of the character limits for a social media post.")` should return `"invalid post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This message is too long to fit into either of the character limits for a social media post."), "invalid post")`)
}})
```

# --seed--

## --seed-contents--

```py
def can_post(message):

    return message
```

# --solutions--

```py
def can_post(message):
    if len(message) <= 40:
        return "short post"
    elif len(message) <= 80:
        return "long post"
    else:
        return "invalid post"
```
