---
id: 68ffb91507a5b645769328c4
title: "Challenge 101: Markdown Heading Converter"
challengeType: 29
dashedName: challenge-101
---

# --description--
Ako vam je dat niz slova koji predstavlja zaglavlje sekcije u Markdown-u, vratite odgovarajuće zaglavlje sekcije za HTML.

Naslov važeće sekcije u Markdown-u mora:

- Počinje sa nula ili više razmaka, praćenih
- Slova od 1 do 6 karaktera hasha (`#`) uzastopno, zatim
- Najmanje jedan razmak. I na kraju,
- Tekst naslova sekcije.

Broj znaka haša određuje nivo zaglavlja sekcije. Na primer, jedan znak haša odgovara oznaci `h1`, a šest znakova haša odgovara oznaci `h6`.

Ako dati string ne odgovara tačno navedenom formatu, vrati `"Invalid format"`.

Na primer, ukoliko primiš `"# My level 1 heading"`, vrati `"<h1>My level 1 heading</h1>"`.

Napomena: Konzola možda neće prikazati HTML oznake unutar stringa prilikom ispisivanja. Pogledajte konzolu pregledača da biste videli ispis sa umetnutim oznakama.

# --hints--

`convert("# My level 1 heading")` should return `"<h1>My level 1 heading</h1>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>")`)
}})
```

`convert("My heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("My heading"), "Invalid format")`)
}})
```

`convert("##### My level 5 heading")` should return `"<h5>My level 5 heading</h5>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>")`)
}})
```

`convert("#My heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("#My heading"), "Invalid format")`)
}})
```

`convert("  ###  My level 3 heading")` should return `"<h3>My level 3 heading</h3>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>")`)
}})
```

`convert("####### My level 7 heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("####### My level 7 heading"), "Invalid format")`)
}})
```

`convert("## My #2 heading")` should return `"<h2>My #2 heading</h2>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("## My #2 heading"), "<h2>My #2 heading</h2>")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert(heading):

    return heading
```

# --solutions--

```py
import re
def convert(heading):
    trimmed = heading.strip()

    if not re.match(r'^#{1,6}\s+\S+', trimmed):
        return "Invalid format"

    hash_count = trimmed.find(' ')
    text = trimmed[hash_count:].lstrip()

    return f"<h{hash_count}>{text}</h{hash_count}>"
```
