---
id: 68ffb91507a5b645769328c6
title: "Challenge 103: LCM"
challengeType: 29
dashedName: challenge-103
---

# --description--
Ako razmatrate dva celobrojna broja, vratite najmanji zajednički sadržalj (LCM) tih dva broja.

Najmanji zajednički višekratnik (LCM) za dva broja je najmanji pozitivan ceo broj koji je višestrukost oba broja. Na primer, uzimajući u obzir `4` i `6`, vrati `12` zato što:

- Izvor za `4` je `4`, `8`, `12` i tako dalje.
- Izvor za `6` je `6`, `12`, `18` i tako dalje.
- `12` je najmanji broj koji je izvor za oba.

# --hints--

`lcm(4, 6)` should return `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(4, 6), 12)`)
}})
```

`lcm(9, 6)` should return `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(9, 6), 18)`)
}})
```

`lcm(10, 100)` should return `100`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(10, 100), 100)`)
}})
```

`lcm(13, 17)` should return `221`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(13, 17), 221)`)
}})
```

`lcm(45, 70)` should return `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(45, 70), 630)`)
}})
```

# --seed--

## --seed-contents--

```py
def lcm(a, b):

    return a
```

# --solutions--

```py
def lcm(a, b):
    def gcd(x, y):
        return x if y == 0 else gcd(y, x % y)

    return abs(a * b) // gcd(a, b)
```
