---
id: 69162d64f96574d9bb629efb
title: "Challenge 110: Word Guesser"
challengeType: 29
dashedName: challenge-110
---

# --description--
Kada uzmete u obzir sekvencu od dva slova iste dužine, tajnu reč i pogađanu reč, uporedite pogađanu reč sa tajnom rečju koristeći sledeća pravila:

- Tajna i pogačka reč će sadržati samo velika slova (`"A"` do `"Z"`);
- Za svako slovo u pogačkoj reči, zamenite brojem u zavisnosti od toga kako odgovara tajnoj reči:
  - `"2"` ako je slovo deo tajne reči i na pravom mestu.
  - `"1"` ako je slovo deo tajne reči, ali na pogrešnom mestu.
  - `"0"` ako slovo nije u tajnoj reči.
- Svako slovo u tajnoj reči može se koristiti samo jednom.
- Potpuno odgovaranje (`"2"`) se prvo unosi, a zatim delimično odgovaranje (`"1"`) se unosi od leve ka desnoj za preostala slova.
- Ako slovo wystupa često u pogačkoj reči, može se uporediti samo onoliko puta koliko wystupa u tajnoj reči.

Na primer, uzimajući u obzir lozinku `"APPLE"` i nasumičnu reč `"POPPA"`, vrati `"10201"`:

Prvi karakter `"P"` nije na pravom mestu (`"1"`), `"O"` nije u tajnoj reči (`"0"`), drugi karakter `"P"` je na pravom mestu (`"2"`), treći karakter `"P"` je nula (`"0"`) jer su dva karaktera iz `"P"` u upotrebi, i `"A"` nije na pravom mestu (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` should return `"10201"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("APPLE", "POPPA"), "10201")`)
}})
```

`compare("REACT", "TRACE")` should return `"11221"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("REACT", "TRACE"), "11221")`)
}})
```

`compare("DEBUGS", "PYTHON")` should return `"000000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("DEBUGS", "PYTHON"), "000000")`)
}})
```

`compare("JAVASCRIPT", "TYPESCRIPT")` should return `"0000222222"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222")`)
}})
```

`compare("ORANGE", "ROUNDS")` should return `"110200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("ORANGE", "ROUNDS"), "110200")`)
}})
```

`compare("WIRELESS", "ETHERNET")` should return `"10021000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("WIRELESS", "ETHERNET"), "10021000")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare(word, guess):

    return word
```

# --solutions--

```py
def compare(word, guess):
    n = len(word)
    result = ["0"] * n
    used_secret = [False] * n

    for i in range(n):
        if guess[i] == word[i]:
            result[i] = "2"
            used_secret[i] = True

    for i in range(n):
        if result[i] == "0":
            for j in range(n):
                if not used_secret[j] and guess[i] == word[j]:
                    result[i] = "1"
                    used_secret[j] = True
                    break

    return "".join(result)
```
