---
id: 69162d64f96574d9bb629f00
title: "Challenge 115: Markdown Ordered List Item Converter"
challengeType: 29
dashedName: challenge-115
---

# --description--
Ukoliko vam je dat niz znakova koji predstavlja stavku u listi formatiranoj u Markdownu, vratite odgovarajući niz znakova od HTML.

Validni element spiska formatiran u Markdown-u mora:

- Počinje sa nula ili više razmaka, praćenim
- Broj (jedan ili više) i tačka (`.`), praćenim
- Najmanje jedan razmak, zatim
- Tekst stavke liste.

Ako sekvenca karaktera nije u potpunosti usklađena sa gore navedenom strukturom, vratite `"Invalid format"`. Inače, obavijte tekst stavke liste unutar oznaka za `li` i vratite sekvencu karaktera.

Na primer, ako primiš `"1. My item"`, vrati `"<li>My item</li>"`.

Napomena: Konzola možda neće prikazati HTML oznake unutar stringa prilikom ispisivanja. Pogledajte konzolu pregledača da biste videli ispis sa umetnutim oznakama.

# --hints--

`convert_list_item("1. My item")` should return `"<li>My item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1. My item"), "<li>My item</li>")`)
}})
```

`convert_list_item(" 1.  Another item")` should return `"<li>Another item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(" 1.  Another item"), "<li>Another item</li>")`)
}})
```

`convert_list_item("1 . invalid item")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1 . invalid item"), "Invalid format")`)
}})
```

`convert_list_item("2. list item text")` should return `"<li>list item text</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("2. list item text"), "<li>list item text</li>")`)
}})
```

`convert_list_item(". invalid again")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(". invalid again"), "Invalid format")`)
}})
```

`convert_list_item("A. last invalid")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("A. last invalid"), "Invalid format")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_list_item(markdown):

    return markdown
```

# --solutions--

```py
import re
def convert_list_item(markdown):
    match = re.match(r'^\s*(\d+)\.\s+(.+)$', markdown)
    if not match:
        return "Invalid format"
    text = match.group(2)
    return f"<li>{text}</li>"
```
