---
id: 69162d64f96574d9bb629efd
title: "Challenge 112: AI Detector"
challengeType: 29
dashedName: challenge-112
---

# --description--
Danasov izazov nastao je usled lansiranja ChatGPT-a dana 30. novembra 2022.

Uzevši u obzir sekvencu slova jedne ili više rečenica, utvrdite da li je možda kreiran od strane AI-ja koristeći sledeća pravila:

- Imaš dve ili više crtica (`-`).
- Imaš dve ili više skupova zagrada (`()`). Tekst može biti unutar zagrada.
- Imaš tri ili više reči sa sedam ili više slova.

- Reči su razdvojene samo jednim razmakom i sastoje se isključivo od slova (`A-Z`). Ne uključujte interpunkcijske znakove ili karaktere koji nisu slova kao deo reči.

Ako data rečenica zadovoljava bilo koje od ovih pravila, vrati `"AI"`, inače, vrati `"Human"`.

# --hints--

`detect_ai("The quick brown fox jumped over the lazy dog.")` should return `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The quick brown fox jumped over the lazy dog."), "Human")`)
}})
```

`detect_ai("The hypersonic brown fox - jumped (over) the lazy dog.")` should return `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The hypersonic brown fox - jumped (over) the lazy dog."), "Human")`)
}})
```

`detect_ai("Yes - you're right! I made a mistake there - let me try again.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("Yes - you're right! I made a mistake there - let me try again."), "AI")`)
}})
```

`detect_ai("The extraordinary students were studying vivaciously.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The extraordinary students were studying vivaciously."), "AI")`)
}})
```

`detect_ai("The (excited) student was (coding) in the library.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The (excited) student was (coding) in the library."), "AI")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_ai(text):

    return text
```

# --solutions--

```py
import re
def detect_ai(text):
    if len(re.findall(r'-', text)) >= 2:
        return "AI"

    if len(re.findall(r'\([^)]*\)', text)) >= 2:
        return "AI"

    words = text.split()
    long_word_count = sum(1 for w in words if len(re.sub(r'[^A-Za-z]', '', w)) >= 7)
    if long_word_count >= 3:
        return "AI"

    return "Human"
```
