---
id: 69162d64f96574d9bb629eff
title: "Challenge 114: Camel to Snake"
challengeType: 29
dashedName: challenge-114
---

# --description--
Ako vam je dat string slova u camel case, konvertujte snake case verziju tog niza koristeći sledeća pravila:

- Ulazni niz će sadržati samo slova (`A-Z` na `a-z`) i uvek mora početi malim slovom.
- Svako veliko slovo u nizu camel case započinje novu reč.
- Prebacite sva slova na mala.
- Razdvojite reči donjom crtom (underscore) (`_`).

# --hints--

`to_snake("helloWorld")` should return `"hello_world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("helloWorld"), "hello_world")`)
}})
```

`to_snake("myVariableName")` should return `"my_variable_name"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("myVariableName"), "my_variable_name")`)
}})
```

`to_snake("freecodecampDailyChallenges")` should return `"freecodecamp_daily_challenges"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_snake(camel):

    return camel
```

# --solutions--

```py
import re
def to_snake(camel):
    return re.sub(r'([A-Z])', r'_\1', camel).lower()
```
