---
id: 69162d64f96574d9bb629f01
title: "Challenge 116: Permutation Count"
challengeType: 29
dashedName: challenge-116
---

# --description--
Ako vam je dat niz slova, vratite broj različitih permutacija koje se mogu formirati iz tih slova.

- Permutacija je bilo kakvo ponovno poređanje slova u sekvenci.
- Ne računajte ponavljajuće permutacije.
- Ako sekvenca sadrži ponavljajuća slova, ponovljeno poređanje treba da se izračuna samo jednom.
- Sekvenca će sadržati samo slova (`A-Z`, `a-z`).

Na primer, ako dobijete ``"abb"``, vratite ``3`` jer postoji tri jedinstvena načina za aranžiranje slova: ``"abb"``, ``"bab"``, i ``"bba"``.

# --hints--

`count_permutations("abb")` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abb"), 3)`)
}})
```

`count_permutations("abc")` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abc"), 6)`)
}})
```

`count_permutations("racecar")` should return `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("racecar"), 630)`)
}})
```

`count_permutations("freecodecamp")` should return `39916800`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("freecodecamp"), 39916800)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_permutations(s):

    return s
```

# --solutions--

```py
from math import factorial
from collections import Counter
def count_permutations(s):
    freq = Counter(s)
    n = len(s)

    result = factorial(n)
    for count in freq.values():
        result //= factorial(count)

    return result
```
