---
id: 691b559495c5cb5a37b9b483
title: "Challenge 123: Roman Numeral Builder"
challengeType: 29
dashedName: challenge-123
---

# --description--
Ukoliko nađeš ceo broj, vrati mu ekvivalentnu vrednost u rimskim brojevima.

Rimski brojevi koriste sledeće simbole:

| Symbol | Value |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

Rimski brojevi se pišu od najvećeg ka najmanjem, sa leve na desno, koristeći sledeća pravila:

- Dodavanje se koristi kada je cifra praćena drugom cifrom iste ili manje vrednosti. Na primer, `18` piše se kao `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- Oduzimanje se koristi kada je manja cifra vidljiva pre veće, predstavljajući 4 ili 9 u bilo kojoj poziciji vrednosti. Na primer, 19 piše se kao `XIX` (10 + (10 - 1)).
- Nije moguće ponoviti cifru više od tri puta zaredom. Oduzimanje se koristi kada je potrebno napisati cifru više od tri puta zaredom. Stoga, najveći broj koji možete napisati je 3999.

Ovo je još jedan primer: ako dobiješ `1464`, vrati `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`to_roman(18)` should return `"XVIII"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(18), "XVIII")`)
}})
```

`to_roman(19)` should return `"XIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(19), "XIX")`)
}})
```

`to_roman(1464)` should return `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(1464), "MCDLXIV")`)
}})
```

`to_roman(2025)` should return `"MMXXV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(2025), "MMXXV")`)
}})
```

`to_roman(3999)` should return `"MMMCMXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(3999), "MMMCMXCIX")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_roman(num):

    return num
```

# --solutions--

```py
def to_roman(num):
    symbols = [
        ("M", 1000),
        ("CM", 900),
        ("D", 500),
        ("CD", 400),
        ("C", 100),
        ("XC", 90),
        ("L", 50),
        ("XL", 40),
        ("X", 10),
        ("IX", 9),
        ("V", 5),
        ("IV", 4),
        ("I", 1)
    ]

    result = ""
    for sym, val in symbols:
        while num >= val:
            result += sym
            num -= val
    return result
```
