---
id: 691b559495c5cb5a37b9b485
title: "Challenge 125: Game of Life"
challengeType: 29
dashedName: challenge-125
---

# --description--
Ako vam je data matrica (redovi redova) koja predstavlja trenutno stanje u Konvejevoj igri života, vratite sledeće stanje matrice koristeći ova pravila:

- Svaka ćelija je `1` (živ) ili `0` (smrt).
- Komšije ćelije su do osam ćelija koje okružuju (vertikalno, horizontalno i po dubini).
- Ćelije blizu granica imaju manje komšija od osam.

Pravila za ažuriranje svake ćelije:

- Svaka živa ćelija sa manje od dva živa suseda umire (usled usamljenosti).
- Svaka živa ćelija sa dva ili tri živa suseda nastavlja da živi.
- Svaka živa ćelija sa više od tri živa suseda umire (zbog prenaseljenosti).
- Svaka mrtva ćelija sa tačno tri živa suseda oživljava (rađanje).

Na primer, ukoliko dobijete:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

return:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Svaka ćelija se ažurira u zavisnosti od broja živih suseda. Na primer, `[0][0]` ostaje mrtva (2 živih suseda), `[0][1]` ostaje živa (2 živih suseda), `[0][2]` umire (3 živih suseda), i tako dalje.

# --hints--

`game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` should return `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]])`)
}})
```

`game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` should return `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]])`)
}})
```

`game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` should return `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]])`)
}})
```

`game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` should return `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def game_of_life(grid):

    return grid
```

# --solutions--

```py
def game_of_life(grid):
    rows, cols = len(grid), len(grid[0])

    def count_live_neighbors(r, c):
        count = 0
        for i in range(r-1, r+2):
            for j in range(c-1, c+2):
                if 0 <= i < rows and 0 <= j < cols and (i != r or j != c):
                    count += grid[i][j]
        return count

    next_state = [row[:] for row in grid]

    for r in range(rows):
        for c in range(cols):
            live_neighbors = count_live_neighbors(r, c)

            if grid[r][c] == 1:
                if live_neighbors < 2 or live_neighbors > 3:
                    next_state[r][c] = 0
            else:
                if live_neighbors == 3:
                    next_state[r][c] = 1

    return next_state
```
