---
id: 691b559495c5cb5a37b9b489
title: "Challenge 129: Markdown Blockquote Parser"
challengeType: 29
dashedName: challenge-129
---

# --description--
Ukoliko vam je dat niz znakova koji sadrži citat bloka u Markdown formatu, vratite odgovarajuću sekvencu znakova HTML.

> Blockquote u Markdown-u je bilo koja linija koja ima:

- Počinje sa nula ili više razmaka
- Praćen znakiem veće od (>)(`>`)
- Zatim, jedan ili više razmaka
- I na kraju, tekst u bloku citata (blockquote).

Vrati tekst blockquote-a okružen otvarajućom i zatvarajućom HTML oznakom `blockquote`.

Na primer, ako primiš `"> This is a quote"`, vrati `<blockquote>This is a quote</blockquote>`.

Napomena: Konzola možda neće prikazati HTML oznake unutar stringa prilikom ispisivanja. Pogledajte konzolu pregledača da biste videli ispis sa umetnutim oznakama.

# --hints--

`parse_blockquote("> This is a quote")` should return `"<blockquote>This is a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>")`)
}})
```

`parse_blockquote(" > This is also a quote")` should return `"<blockquote>This is also a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>")`)
}})
```

`parse_blockquote("  >    So  Is  This")` should return `"<blockquote>So  Is  This</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_blockquote(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_blockquote(markdown):
    pattern = r"^\s*>\s+(.+)$"
    match = re.match(pattern, markdown)

    text = match.group(1)
    return f"<blockquote>{text}</blockquote>"
```
