---
id: 6925e2068081f40f549ced1a
title: "Challenge 136: Markdown Image Parser"
challengeType: 29
dashedName: challenge-136
---

# --description--
Ako dobijete niz znakova sa slike u Markdownu, vratite isti niz znakova kao HTML.

Markdown slika ima sledeći format: `"![alt text](image_url)"`. Gde:

- `alt text` je opis slike (vrednost atributa elementa od `alt`).
- `image_url` je URL izvor slike (vrednost atributa elementa od `src`).

Vrati HTML oznaku `img` sa atributom `src` postavljenim na URL slike, a atributom `alt` na alternativni tekst.

Na primer, ako dobijete `"![Cute cat](cat.png)"` vratite `'<img src="cat.png" alt="Cute cat">'`;

- Osigurajte da su etiketa, raspored atributa elementa, razmak i korišćenje navodnika u skladu sa onim što je gore.

Napomena: Konsola možda neće prikazati oznake HTML u sekvenci karaktera prilikom snimanja poruka — proverite konzolu pregledača da vidite snimke sa priloženim oznakama.

# --hints--

`parse_image("![Cute cat](cat.png)")` should return `'<img src="cat.png" alt="Cute cat">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">')`)
}})
```

`parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` should return `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">')`)
}})
```

`parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` should return `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_image(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_image(markdown):
    match = re.search(r'!\[(.*?)\]\((.*?)\)', markdown)
    if not match:
        return ""
    alt, src = match.groups()
    return f'<img src="{src}" alt="{alt}">'
```
