---
id: 69306364df283fcaff2e1ad7
title: "Challenge 146: Left-Handed Seat at the Table"
challengeType: 29
dashedName: challenge-146
---

# --description--
Prikazana je tabela od 4x2 (red po red) koja predstavlja raspored stolica za obroke; utvrdi koliko sedišta osoba koja piše levom rukom može da zauzme.

- Čovek koji piše levom rukom ne može da sedi na mestu gde bi čovek koji piše desnom rukom sedeo na stolicu koja je sa leve strane te osobe.

U matrici koja je data:

- `"R"` je stolica zauzeta od strane osobe koja piše desnom rukom.
- `"L"` je stolica zauzeta od strane osobe koja piše levom rukom.
- `"U"` je prazna stolica.
- Ovo su samo prazne stolice koje se mogu koristiti za sedenje.
- Stolice u gornjem redu okrenute su "nadole", dok stolice u donjem redu gledaju "gore" (kao sto), pa zavisi da li je levo ili desno od orijentacije stolice.
- U ćošnim mestima je samo jedna stolica pored njih.

Primer, u priloženoj tabeli:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

Sedište gore-levo ne može se zauzeti jer neko piše desnom rukom na levoj strani. Ostala dva slobodna sedišta mogu se zauzeti jer niko ne piše desnom rukom na levoj strani.

# --hints--

`find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2)`)
}})
```

`find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` should return `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0)`)
}})
```

`find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_left_handed_seats(table):

    return table
```

# --solutions--

```py
def find_left_handed_seats(table):
    available_seats = 0
    top_row, bottom_row = table

    for i in range(len(top_row)):
        if top_row[i] == "U":
            if i == len(top_row) - 1:
                available_seats += 1
            elif top_row[i + 1] != "R":
                available_seats += 1

    for i in range(len(bottom_row)):
        if bottom_row[i] == "U":
            if i == 0:
                available_seats += 1
            elif bottom_row[i - 1] != "R":
                available_seats += 1

    return available_seats
```
