---
id: 69373793f5a867f769cde139
title: "Challenge 154: Par for the Hole"
challengeType: 29
dashedName: challenge-154
---

# --description--
Uzimajući u obzir dve pune cifre, par za golf rupu i broj udaraca koje je golfer napravio na toj rupi, vratite rezultat igrača koristeći golf terminologiju.

Return:

- `"Hole in one!"` ikiwa ilichukua mchoro mmoja.
- `"Eagle"` ako su bila potrebna dva udarca manje od para.
- `"Birdie"` ako je bio potreban jedan udarac manje od para.
- `"Par"` ikiwa ilichukua idadi sawa ya michoro na par.
- `"Bogey"` ako je bio potreban jedan udarac više od para.
- `"Double bogey"` ako su bila potrebna dva udarca više od para.

# --hints--

`golf_score(3, 3)` should return `"Par"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 3), "Par")`)
}})
```

`golf_score(4, 3)` should return `"Birdie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 3), "Birdie")`)
}})
```

`golf_score(3, 1)` should return `"Hole in one!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 1), "Hole in one!")`)
}})
```

`golf_score(5, 7)` should return `"Double bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 7), "Double bogey")`)
}})
```

`golf_score(4, 5)` should return `"Bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 5), "Bogey")`)
}})
```

`golf_score(5, 3)` should return `"Eagle"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 3), "Eagle")`)
}})
```

# --seed--

## --seed-contents--

```py
def golf_score(par, strokes):

    return par
```

# --solutions--

```py
def golf_score(par, strokes):
    if strokes == 1:
        return "Hole in one!"

    diff = strokes - par

    if diff == -2:
        return "Eagle"
    if diff == -1:
        return "Birdie"
    if diff == 0:
        return "Par"
    if diff == 1:
        return "Bogey"
    if diff == 2:
        return "Double bogey"
```
