---
id: 69373793f5a867f769cde13a
title: "Challenge 155: Plant the Crop"
challengeType: 29
dashedName: challenge-155
---

# --description--
Uzimajući u obzir kompletan broj koji predstavlja veličinu vaše poljoprivredne parcele, i `"acres"` ili `"hectares"` koji predstavljaju meru veličine vašeg polja, zajedno sa vrstom usjeva, utvrdite koliko sadnica tog tipa možete zasaditi na svom polju.

- 1 ekar je jednak 4046,86 kvadratnih metara.
- 1 hektar je jednak 10.000 kvadratnih metara.

Ovde je lista usova koji pružaju prinos zajedno sa površinom koju jedna biljka zauzima:

|Usivojstvo|Prostor po biljci|
|-|-|
|`"corn"`|Jedan kvadratni metar|
|`"wheat"`|0.1 kvadratnog metra|
|`"soybeans"`|0.5 kvadratnog metra|
|`"tomatoes"`|0.25 kvadratnog metra|
|`"lettuce"`|0.2 kvadratnog metra|

Vrati broj biljaka koje mogu da stanu na parcelu, zaokružujući nadole na najbliži ceo broj biljaka.

# --hints--

`get_number_of_plants(1, "acres", "corn")` should return `4046`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(1, "acres", "corn"), 4046)`)
}})
```

`get_number_of_plants(2, "hectares", "lettuce")` should return `100000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(2, "hectares", "lettuce"), 100000)`)
}})
```

`get_number_of_plants(20, "acres", "soybeans")` should return `161874`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(20, "acres", "soybeans"), 161874)`)
}})
```

`get_number_of_plants(3.75, "hectares", "tomatoes")` should return `150000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(3.75, "hectares", "tomatoes"), 150000)`)
}})
```

`get_number_of_plants(16.75, "acres", "tomatoes")` should return `271139`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(16.75, "acres", "tomatoes"), 271139)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_of_plants(field_size, unit, crop):

    return field_size
```

# --solutions--

```py
import math
def get_number_of_plants(field_size, unit, crop):
    crop_space = {
        "corn": 1,
        "wheat": 0.1,
        "soybeans": 0.5,
        "tomatoes": 0.25,
        "lettuce": 0.2
    }

    field_in_sqm = field_size * 4046.86 if unit == "acres" else field_size * 10000

    return math.floor(field_in_sqm / crop_space[crop.lower()])
```
