---
id: 694596b0585c11170ac7c7fc
title: "Challenge 164: Markdown Inline Code Parser"
challengeType: 29
dashedName: challenge-164
---

# --description--
Ako vam je dat Markdown niz znakova koji uključuje jednu ili više kodnih sekcija unutar linije, vratite odgovarajući niz znakova HTML.

Sekcije koda unutar linije u Markdown-u koriste jednu backtick oznaku (```` ` ```) Na početku i na kraju teksta dela koda.

Vratite izbačenu sekvencu gde su svi kodni delovi zamenjeni oznakama HTML `code`.

Na primer, ako dobiješ string `"Use \`let\` to declare the variable."`, vrati `"Tumia <code>let</code> kutangaza kigezo."`.

Napomena: Konzola možda neće prikazati HTML oznake unutar stringa prilikom ispisivanja. Pogledajte konzolu pregledača da biste videli ispis sa umetnutim oznakama.

# --hints--

```parse_inline_code("Use `let` to declare the variable.")``` should return `"Use <code>let</code> to declare the variable."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Use \`let\` to declare the variable."), "Use <code>let</code> to declare the variable.")`)
}})
```

```parse_inline_code("Use `let` or `const` to declare a variable.")``` should return `"Use <code>let</code> or <code>const</code> to declare a variable."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Use \`let\` or \`const\` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.")`)
}})
```

```parse_inline_code("Run `npm install` then `npm start`.")``` should return `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Run \`npm install\` then \`npm start\`."), "Run <code>npm install</code> then <code>npm start</code>.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_inline_code(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_inline_code(markdown):
    return re.sub(r"`([^`]+)`", r"<code>\1</code>", markdown)
```
