---
id: 696655d24b614176d4c9b78d
title: "Challenge 170: Odd or Even Day"
challengeType: 29
dashedName: challenge-170
---

# --description--

Ukitoa timestamp (idadi ya millisekunde tangu mwanzo wa Unix epoch), rudisha:

- `"odd"` ako je broj dana meseca za taj vremenski pečat neispravan.
- `"even"` ako je broj dana meseca za taj vremenski pečat ispravan.

Na primer, kada date `1769472000000`, vremenski pečat za datum 27. januara 2026., vratite `"odd"` jer je broj dana (`27`) pogrešan broj.

Napomena: Timestamp je u milisekundama i treba da koristite datum po UTC vremenu, a ne vaše lokalno vreme.

# --hints--

`odd_or_even_day(1769472000000)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769472000000), "odd")`)
}})
```

`odd_or_even_day(1769444440000)` should return `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769444440000), "even")`)
}})
```

`odd_or_even_day(6739456780000)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(6739456780000), "odd")`)
}})
```

`odd_or_even_day(1)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1), "odd")`)
}})
```

`odd_or_even_day(86400000)` should return `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(86400000), "even")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even_day(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime

def odd_or_even_day(timestamp):
    dt = datetime.utcfromtimestamp(timestamp / 1000)
    day = dt.day
    return "even" if day % 2 == 0 else "odd"
```
