---
id: 696655d24b614176d4c9b78b
title: "Challenge 168: Scaled Image"
challengeType: 29
dashedName: challenge-168
---

# --description--

Ako je dat string koji predstavlja širinu i visinu slike, zajedno sa faktorom skaliranja, vrati novu širinu i visinu slike nakon skaliranja.

- Sekvenca karaktera unetog teksta je u formatu `"WxH"`. Na primer, `"800x600"`.
- Merenje je broj koji množi širinu i dužinu.

Vratite ponovno izmerene dimenzije za istu strukturu od `"WxH"`.

# --hints--

`scale_image("800x600", 2)` should return `"1600x1200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("800x600", 2), "1600x1200")`)
}})
```

`scale_image("100x100", 10)` should return `"1000x1000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("100x100", 10), "1000x1000")`)
}})
```

`scale_image("1024x768", 0.5)` should return `"512x384"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("1024x768", 0.5), "512x384")`)
}})
```

`scale_image("300x200", 1.5)` should return `"450x300"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("300x200", 1.5), "450x300")`)
}})
```

# --seed--

## --seed-contents--

```py
def scale_image(size, scale):

    return size
```

# --solutions--

```py
def scale_image(size, scale):
    width, height = map(int, size.split("x"))

    new_width = round(width * scale)
    new_height = round(height * scale)

    return f"{new_width}x{new_height}"
```
