---
id: 697a49e6ff50d756c9b6935e
title: "Challenge 181: 2026 Winter Games Day 2: Snowboarding"
challengeType: 29
dashedName: challenge-181
---

# --description--
Otpuštanje položaja snowboardiste kada započne sa rotacijom u stepenima, određuje mu položaj prilikom pada.

- Položaj snoubordista je `"Regular"` ili `"Goofy"`.
- Cikli tehnike su često 90 stepeni. Pozitiv pokazuje rotaciju u smeru kazaljke na satu, a negativ pokazuje rotaciju suprotno od kazaljke na satu.
- Položaj prilikom pada menja se svakog ciklusa od 180 stepeni.

Na primer, ako se ukloni `"Regular"` i `90`, vrati `"Regular"`. Ako se ukloni `"Regular"` i stepen `180`, vrati `"Goofy"`.

# --hints--

`get_landing_stance("Regular", 90)` should return `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 90), "Regular")`)
}})
```

`get_landing_stance("Regular", 180)` should return `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 180), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -270)` should return `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -270), "Regular")`)
}})
```

`get_landing_stance("Regular", 2340)` should return `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 2340), "Goofy")`)
}})
```

`get_landing_stance("Goofy", 2160)` should return `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 2160), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -540)` should return `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -540), "Regular")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_landing_stance(start_stance, rotation):

    return start_stance
```

# --solutions--

```py
def get_landing_stance(start_stance, rotation):
    flips = abs(rotation) // 180
    if flips % 2 == 0:
        return start_stance
    else:
        return "Goofy" if start_stance == "Regular" else "Regular"
```
