---
id: 698a1a73ade5ac0e19180fa9
title: "Challenge 206: Playing Card Values"
challengeType: 29
dashedName: challenge-206
---

# --description--
Kada izvučete red podataka igračkih karata, vratite novi red podataka sa numeričkom vrednošću svake karte.

Card Values:

- As (`"A"`) ima vrednost od 1.
- Kartice sa brojevima (`"2"` - `"10"`) imaju vrednost jednaku brojevima koji stoje ispred njih: 2 - 10, po redosledu.
- Figurativne kartice: Džek (`"J"`), Kraljica (`"Q"`), i Kralj (`"K"`) svaka ima vrednost od 10.

Suits:

- Kila kadi ina aina: Spades (`"S"`), Clubs (`"C"`), Diamonds (`"D"`), au Hearts (`"H"`).

Card Format:

- Svaka kartica je predstavljena kao niz znakova: `"valueSuit"`. Na primer: `"AS"` je As Pik, `"10H"` je Deset Srca, i `"QC"` je Kraljica Špila.

# --hints--

`card_values(["3H", "4D", "5S"])` should return `[3, 4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["3H", "4D", "5S"]), [3, 4, 5])`)
}})
```

`card_values(["AS", "10S", "10H", "6D", "7D"])` should return `[1, 10, 10, 6, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7])`)
}})
```

`card_values(["8D", "QS", "2H", "JC", "9C"])` should return `[8, 10, 2, 10, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9])`)
}})
```

`card_values(["AS", "KS"])` should return `[1, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "KS"]), [1, 10])`)
}})
```

`card_values(["10H", "JH", "QH", "KH", "AH"])` should return `[10, 10, 10, 10, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def card_values(cards):

    return cards
```

# --solutions--

```py
def card_values(cards):
    values = []
    for card in cards:
        value_str = card[:-1]

        if value_str == "A":
            value = 1
        elif value_str in ["J", "Q", "K"]:
            value = 10
        else:
            value = int(value_str)

        values.append(value)
    return values
```
