---
id: 699c8e045ee7cb94ed2322dd
title: "Challenge 222: Equinox Shadows"
challengeType: 29
dashedName: challenge-222
---

# --description--
Danas je dan ekvinokcija, kada je sunce direktno iznad centralne linije Zemlje i najviše tačno podne. U nekom trenutku, izmerite senku koju baca vertikalna šipka od 4 stopa.

- Vreme će biti sekvenca znakova u formatu od 24 sata `"HH:MM"` (na primer, `"15:00"` je 3 ujutru).
- Biće vam dato vreme samo u intervalima od 30 minuta.

Rules:

- Sunce izlazi tačno u 6 ujutro `"east"`, i zalazi tačno u 6 uveče `"west"`.
- Senka se uvek kreće suprotno od sunca.
- Dužina senke (u stopama) je broj sati od podne, povišen za treći stepen.
- Nema senke pre nego što sunce izađe (pre 6 ujutro), posle zalaska sunca (nakon 6 uveče ili kasnije), niti podne.

Return:

- Ako postoji senka, vrati `"(length)ft (direction)"`. Na primer, `"8ft west"`.
- Inače, vrati `"No shadow"`.

Na primer, ako dobiješ `"10:00"`, vrati `"8ft west"` jer je deset ujutro dva od šest popodne, pa 2<sup>3</sup> = 8 stopa, i senka usmerena je na zapad jer je sunce na istoku u 10 ujutro.

# --hints--

`get_shadow("10:00")` should return `"8ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` should return `"27ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` should return `"166.375ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` should return `"216ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` should return `"91.125ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```
