---
id: 69a890af247de743333bd4cd
title: "Challenge 224: Coffee Roast Detector"
challengeType: 29
dashedName: challenge-224
---

# --description--
Niz slova koji predstavlja zrna korišćena za pravljenje šolje kafe, određuje stepen prženja te kafe.

- Niz znakova koji je dat imaće sledeća slova, pri čemu svako slovo predstavlja vrstu krokodila:

- Znak apostrofa (`'`) je znak za sagorevanje na nivou svetlosti sa vrednošću od 1 poen svaki.
  - Znak crte (`-`) je znak za sagorevanje srednje jačine sa vrednošću od 2 poena svaki.
  - Znak tačke (`.`) je znak za sagorevanje u mraku sa vrednošću od 3 poena svaki.

- Nivo pečenja se određuje po proseku svih grahova.

Return:

- `"Light"` ako je prosečna vrednost ispod 1.75.
- `"Medium"` ako je prosečna vrednost između 1.75 i 2.5.
- `"Dark"` ako je prosečna vrednost veća od 2.5.

# --hints--

`detect_roast("''-''''''-'-''--''''")` should return `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("''-''''''-'-''--''''"), "Light")`)
}})
```

`detect_roast(".'-''-''..'''.-.-''-")` should return `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".'-''-''..'''.-.-''-"), "Medium")`)
}})
```

`detect_roast("--.''--'-''.--..-.--")` should return `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("--.''--'-''.--..-.--"), "Medium")`)
}})
```

`detect_roast("-...'-......-..-...-")` should return `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-...'-......-..-...-"), "Dark")`)
}})
```

`detect_roast(".--.-..-......----.'")` should return `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".--.-..-......----.'"), "Medium")`)
}})
```

`detect_roast("..-..-..-..-....-.-.")` should return `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("..-..-..-..-....-.-."), "Dark")`)
}})
```

`detect_roast("-'-''''''..-'.''-'.'")` should return `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-'-''''''..-'.''-'.'"), "Light")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_roast(beans):

    return beans
```

# --solutions--

```py
def detect_roast(beans):
    total = 0

    for bean in beans:
        if bean == "'":
            total += 1
        elif bean == "-":
            total += 2
        elif bean == ".":
            total += 3

    avg = total / len(beans)

    if avg < 1.75:
        return "Light"
    elif avg <= 2.5:
        return "Medium"
    else:
        return "Dark"
```
