---
id: 69b1028d6e265413d0198a2a
title: "Challenge 231: ISBN-10 Validator"
challengeType: 29
dashedName: challenge-231
---

# --description--
Given niz karaktera, utvrđuje da li je validan ISBN-10.

ISBN-10 sadrži crtice (`"-"`) i još 10 znakova. Nakon uklanjanja crtica (`"-"`):

- Prvih 9 karaktera mora biti brojevi, a
- Poslednji karakter može biti broj ili slovo `"X"`, koji predstavlja broj 10.

Ili kuikagua:

- Pomnoži svaki broj (ili vrednost) sa svojom pozicijom (pomnoži prvi broj sa 1, drugi sa 2, i tako dalje).
- Saberi sve rezultate.
- Ako je zbir deljiv sa 11, onda je validan.

# --hints--

`is_valid_isbn10("0-306-40615-2")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-306-40615-2"), True)`)
}})
```

`is_valid_isbn10("0-306-40615-1")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-306-40615-1"), False)`)
}})
```

`is_valid_isbn10("0-8044-2957-X")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-8044-2957-X"), True)`)
}})
```

`is_valid_isbn10("X-306-40615-2")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("X-306-40615-2"), False)`)
}})
```

`is_valid_isbn10("0-6822-2589-4")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-6822-2589-4"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_isbn10(s):

    return s
```

# --solutions--

```py
def is_valid_isbn10(s):
    isbn = s.replace("-", "")

    if len(isbn) != 10:
        return False

    total = 0

    for i in range(9):
        if not isbn[i].isdigit():
            return False
        total += int(isbn[i]) * (i + 1)

    last = isbn[9]

    if last == "X":
        total += 10 * 10
    elif last.isdigit():
        total += int(last) * 10
    else:
        return False

    return total % 11 == 0
```
