---
id: 69bc6cb30c1d112a2e110a07
title: "Challenge 250: Hidden Key"
challengeType: 29
dashedName: challenge-250
---

# --description--
Dobrodošli na dnevnu vežbu od 250!

Ako imate kodiran niz znakova, dešifrujte koristeći ključ za dekodiranje i vratite rezultat.

Ili kupata ufunguo:

- Pogledajte sve dnevne izazove do danas čiji je broj izazova deo 25 (uključujući i ovaj).
- Uzmite prvo slovo iz naslova svih tih izazova i spojite ih u niz slova. Ako naslov počinje sa karakterom koji nije slovo, pronađite njegovo prvo slovo.

Da biste dešifrovali poruku, proverite svako slovo u kodiranoj poruci sa:

1. Pronađite odgovarajuće slovo na tastaturi (ponovite tastaturu ako je poruka duža od kapaciteta tastature).
2. Pretvorite slovo sa tastature u odgovarajući broj: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
3. Pritisnite slovo koje odgovara tom broju u alfabetu.
4. Ako pritiskanje dođe pre `"A"`, vratite se na `"Z"`.

Na primer, ako kodirana poruka počinje sa `"Y"` i prvi karakter ključa je `"V"` (22), pomerite `"Y"` unazad za poziciju 22 da biste dobili `"C"`. Ponovite ovaj proces za svaki karakter kako biste dešifrovali celu poruku.

- Samo se miču slova, razmaci se vraćaju kako su bili.
- Sva slova koja su uklonjena i vraćena su velika slova.

# --hints--

`decode("YAVJYNXE")` should return `"CONGRATS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YAVJYNXE"), "CONGRATS")`)
}})
```

`decode("YALLUT PQUMJP")` should return `"CODING LEGEND"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YALLUT PQUMJP"), "CODING LEGEND")`)
}})
```

`decode("UAC DYR EISAKYM")` should return `"YOU ARE AWESOME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME")`)
}})
```

`decode("GQMS NBMZU")` should return `"KEEP GOING"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("GQMS NBMZU"), "KEEP GOING")`)
}})
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` should return `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message):

    return message
```

# --solutions--

```py
def decode(message):
    key = "VLHCGMDLNH"
    result = []
    key_index = 0
    for char in message:
        if char == ' ':
            result.append(char)
        else:
            shift = ord(key[key_index % len(key)]) - 64
            key_index += 1
            result.append(chr((ord(char) - 65 - shift) % 26 + 65))
    return ''.join(result)
```
