---
id: 69e2383af7832c8032603b93
title: "Challenge 277: Mirror Image"
challengeType: 29
dashedName: challenge-277
---

# --description--
S obziма na sekvencu od dva slova, utvrđuje da li je druga sekvenca ogledsko slika prve sekvence.

Ogledalni odraz se stvara menjanjem redosleda slova i uzimanjem svakog slova i zamenjivanjem ga njegovim ekvivalentom u ogledalu.

- Odgovarajuća slova izgledaju kao što su u ogledalu:

`W`, `T`, `Y`, `U`, `I`, `O`, `H`, `A`, `X`, `V`, `M`, `w`, `o`, `x`, `v`, `0`, `8`, `=`, `+`, `:`, `|`, `-`, `_`, `*`, `^`, `!`, `.`, na nafasi (` `).

- Parovi koji se razmenjuju u ogledalu, razmenjuju među sobom:

| Herufi | Hubadilishana na |
| - | - |
| `[` | `]` |
| `{` | `}` |
| `<` | `>` |
| `b` | `d` |
| `p` | `q` |
| `(` | `)` |

Ako bilo koja sekvenca sadrži znak koji nije na navedenoj listi, ne postoji ogledska slika koja se može kreirati iz tih znakova.

Na primer, sliku ogledala od `"[HOW]"` je `"[WOH]"`.

# --hints--

`is_mirror_image("[HOW]", "[WOH]")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("[HOW]", "[WOH]"), True)`)
}})
```

`is_mirror_image("MOM", "MOM")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("MOM", "MOM"), True)`)
}})
```

`is_mirror_image("vow", "wov")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("vow", "wov"), True)`)
}})
```

`is_mirror_image("TIM", "TIM")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("TIM", "TIM"), False)`)
}})
```

`is_mirror_image("{WOW}", "}WOW{")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("{WOW}", "}WOW{"), False)`)
}})
```

`is_mirror_image("XXVII", "IIV%X")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("XXVII", "IIV%X"), False)`)
}})
```

`is_mirror_image("><(((*>", "<*)))><")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("><(((*>", "<*)))><"), True)`)
}})
```

`is_mirror_image("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_mirror_image(s1, s2):

    return s1
```

# --solutions--

```py
def is_mirror_image(s1, s2):
    symmetric = set('WTYUIOHA XVMwoxv08=+:|-_*^!.')
    pairs = {'[': ']', ']': '[', '{': '}', '}': '{', '<': '>', '>': '<',
             'b': 'd', 'd': 'b', 'p': 'q', 'q': 'p', '(': ')', ')': '('}
    mirrored = []
    for c in reversed(s1):
        if c in symmetric:
            mirrored.append(c)
        elif c in pairs:
            mirrored.append(pairs[c])
        else:
            return False
    return ''.join(mirrored) == s2
```
