---
id: 69f8c998d78ad3171a0713ba
title: "Vežba 287: Ispravi rimske brojeve"
challengeType: 29
dashedName: challenge-287
---

# --description--
Ako izlazite sekvencu Rimskih brojeva koja nije pravilno formatirana, vratite njenu vrednost koristeći standardnu notaciju Rimskih brojeva.

Ulaz će koristiti samo operaciju sabiranja, stoga svaka oznaka dodaje svoju vrednost ukupnom iznosu. Kao podsetnik, ovo su oznake i njihove vrednosti:

| Simbol | Vrednost |
| ------ | ------- |
| `"I"` | 1 |
| `"V"` | 5 |
| `"X"` | 10 |
| `"L"` | 50 |
| `"C"` | 100 |
| `"D"` | 500 |
| `"M"` | 1000 |

Pri ponovnom kodiranju, koristite najveći mogući znak u svakom koraku, koristeći subtractive pairs (`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`) ako je to potrebno.

# --hints--

`fix_numerals("XIIIII")` inapaswa kurudisha `"XV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("XIIIII"), "XV")`)
}})
```

`fix_numerals("IIIILX")` inapaswa kurudisha `"LXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("IIIILX"), "LXIV")`)
}})
```

`fix_numerals("XXVVVIIIII")` inapaswa kurudisha `"XL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("XXVVVIIIII"), "XL")`)
}})
```

`fix_numerals("MDCCLXXXXVIIII")` inapaswa kurudisha `"MDCCXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("MDCCLXXXXVIIII"), "MDCCXCIX")`)
}})
```

`fix_numerals("IIIIVVVVXXXXLLLLCCDD")` inapaswa kurudisha `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV")`)
}})
```

`fix_numerals("ILCDMIVDIIXLCVCXDL")` inapaswa kurudisha `"MMCMLXXXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV")`)
}})
```

# --seed--

## --seed-contents--

```py
def fix_numerals(s):

    return s
```

# --solutions--

```py
def fix_numerals(s):
    values = {"I": 1, "V": 5, "X": 10, "L": 50, "C": 100, "D": 500, "M": 1000}
    encoding = [
        (1000, "M"), (900, "CM"), (500, "D"), (400, "CD"),
        (100, "C"), (90, "XC"), (50, "L"), (40, "XL"),
        (10, "X"), (9, "IX"), (5, "V"), (4, "IV"), (1, "I")
    ]

    total = sum(values[c] for c in s)

    result = ""
    for value, symbol in encoding:
        while total >= value:
            result += symbol
            total -= value
    return result
```
