---
id: 69f8c998d78ad3171a0713bf
title: "Zoezi la 292: Uwiano mpana wa sura"
challengeType: 29
dashedName: challenge-292
---

# --description--
Ako dobijete sekvencu od dve slike sa različitim dimenzijama, vratite aspekt razmera slike koja ima najveći odnos širine na visinu.

- Generisana sekvenca karaktera biće u formatu `"WxH"`, na primer, `"1920x1080"`.
- Aspektno ommerje je odnos širine na visinu, svedeno na najmanji ceo broj. Na primer, `"1920x1080"` se smanjuje na `"16:9"`.
- Konvertujte sekvencu karaktera u format `"W:H"`, na primer, `"16:9"`.

# --hints--

`get_wider_aspect_ratio("1920x1080", "800x600")` inapaswa kurudisha `"16:9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1920x1080", "800x600"), "16:9")`)
}})
```

`get_wider_aspect_ratio("1080x1350", "2048x1536")` inapaswa kurudisha `"4:3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1080x1350", "2048x1536"), "4:3")`)
}})
```

`get_wider_aspect_ratio("640x480", "2440x1220")` inapaswa kurudisha `"2:1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("640x480", "2440x1220"), "2:1")`)
}})
```

`get_wider_aspect_ratio("360x640", "1080x1920")` inapaswa kurudisha `"9:16"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("360x640", "1080x1920"), "9:16")`)
}})
```

`get_wider_aspect_ratio("3440x1440", "2048x858")` inapaswa kurudisha `"43:18"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("3440x1440", "2048x858"), "43:18")`)
}})
```

`get_wider_aspect_ratio("12345x61234", "12534x51234")` inapaswa kurudisha `"2089:8539"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("12345x61234", "12534x51234"), "2089:8539")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_wider_aspect_ratio(a, b):

    return a
```

# --solutions--

```py
from math import gcd

def get_wider_aspect_ratio(a, b):
    def parse(s):
        w, h = map(int, s.split("x"))
        g = gcd(w, h)
        return w / h, f"{w // g}:{h // g}"

    ratio_a, simplified_a = parse(a)
    ratio_b, simplified_b = parse(b)

    return simplified_a if ratio_a >= ratio_b else simplified_b
```
