---
id: 69f8c998d78ad3171a0713bb
title: "Zoezi la 288: Nambari ya siri"
challengeType: 29
dashedName: challenge-288
---

# --description--
Uzmite u obzir tajni broj i procenu, utvrdite da li je procena tačna.

Rudisha:

- `"higher"` ako je taj broj veći od procene.
- `"lower"` ako je taj broj manji od procene.
- `"you got it!"` ako je procena tačna.

# --hints--

`guess_number(50, 30)` inapaswa kurudisha `"higher"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(50, 30), "higher")`)
}})
```

`guess_number(85, 99)` inapaswa kurudisha `"lower"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(85, 99), "lower")`)
}})
```

`guess_number(2026, 2026)` inapaswa kurudisha `"you got it!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(2026, 2026), "you got it!")`)
}})
```

`guess_number(92904, 11283)` inapaswa kurudisha `"higher"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(92904, 11283), "higher")`)
}})
```

`guess_number(230495, 423920)` inapaswa kurudisha `"lower"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(230495, 423920), "lower")`)
}})
```

`guess_number(120349, 120349)` inapaswa kurudisha `"you got it!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(120349, 120349), "you got it!")`)
}})
```

# --seed--

## --seed-contents--

```py
def guess_number(secret, guess):

    return secret
```

# --solutions--

```py
def guess_number(secret, guess):
    if guess == secret:
        return "you got it!"
    return "higher" if guess < secret else "lower"
```
