---
id: 69f90c1329a94b37e2a2086d
title: "Vežba 294: Kombinacija standardnih zagradi
challengeType: 29
dashedName: challenge-294
---

# --description--
Ako je dat kompletan broj, `n`, vrati broj validnih kombinacionih parova za standardne zagrade `n`.

- Važeća kombinacija je niz znakova gde svaka otvarajuća vitičasta zagrada ima odgovarajuću zatvarajuću vitičastu zagradu, i nijedna zatvarajuća vitičasta zagrada ne nastaje pre odgovarajuće otvarajuće vitičaste zagrade.

Na primer, ako je dato `2`, postoje dve važeće kombinacije:

```json
(())
()()
```

# --hints--

`get_combinations(2)` inapaswa kurudisha `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(2), 2)`)
}})
```

`get_combinations(3)` inapaswa kurudisha `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(3), 5)`)
}})
```

`get_combinations(5)` inapaswa kurudisha `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(5), 42)`)
}})
```

`get_combinations(8)` inapaswa kurudisha `1430`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(8), 1430)`)
}})
```

`get_combinations(13)` inapaswa kurudisha `742900`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(13), 742900)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_combinations(n):

    return n
```

# --solutions--

```py
def get_combinations(n):
    count = 0

    def generate(op, cl):
        nonlocal count
        if op == 0 and cl == 0:
            count += 1
            return
        if op > 0:
            generate(op - 1, cl)
        if cl > op:
            generate(op, cl - 1)

    generate(n, n)
    return count
```
