---
id: 69f90c1329a94b37e2a2086c
title: "Zoezi la 293: Mkono bora"
challengeType: 29
dashedName: challenge-293
---

# --description--
Ako vam je dato red podataka sa sekvencijom od pet karaktera koji predstavljaju igrane karte, vratite naziv najbolje ruke.

- Svaka kartica je predstavljena kao sekvenca od dvocifrenih znakova: rang praćen bojom, `"2h"` na primer.
  - Rang, od najnižeg do najvišeg, je: `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"`, i `"A"`.
  - Boja je: `"h"`, `"d"`, `"c"`, i `"s"`.
- Asovi (`"A"`) mogu se koristiti kao najviši ili najniži rang u sekvenci.

Koraci, u redosledu od lošeg do boljeg, su:

| Naziv | Opis |
| - | - |
| `"High Card"` | Nema para ili bolje |
| `"Pair"` | Jedna para istog ranga |
| `"Two Pair"` | Dve pare različitih ranga |
| `"Three of a Kind"` | Tri iste karte (ili: Tri iste rangove) |
| `"Straight"` | Pet uzastopnih ranga |
| `"Flush"` | Pet istog ranga/boje |
| `"Full House"` | Tri iste karte i jedna para drugog ranga |
| `"Four of a Kind"` | Četiri iste karte (ili: Četiri iste rangove) |
| `"Straight Flush"` | Pet uzastopnih ranga istog ranga/boje |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` istog ranga/boje |

Vrati ime najbolje ruke.

# --hints--

`get_best_hand(["7s", "7h", "7d", "2c", "5h"])` inapaswa kurudisha `"Three of a Kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind")`)
}})
```

`get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"])` inapaswa kurudisha `"Full House"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House")`)
}})
```

`get_best_hand(["2h", "5h", "7h", "9h", "Jh"])` inapaswa kurudisha `"Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["2h", "5h", "7h", "9h", "Jh"]), "Flush")`)
}})
```

`get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"])` inapaswa kurudisha `"Four of a Kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind")`)
}})
```

`get_best_hand(["Ts", "Th", "9d", "9c", "8h"])` inapaswa kurudisha `"Two Pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair")`)
}})
```

`get_best_hand(["9c", "8c", "7c", "6c", "5c"])` inapaswa kurudisha `"Straight Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush")`)
}})
```

`get_best_hand(["As", "Kh", "Jd", "8c", "5h"])` inapaswa kurudisha `"High Card"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Kh", "Jd", "8c", "5h"]), "High Card")`)
}})
```

`get_best_hand(["As", "2h", "3d", "4c", "5h"])` inapaswa kurudisha `"Straight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "2h", "3d", "4c", "5h"]), "Straight")`)
}})
```

`get_best_hand(["Ts", "Th", "7c", "6d", "5h"])` inapaswa kurudisha `"Pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "7c", "6d", "5h"]), "Pair")`)
}})
```

`get_best_hand(["As", "Ks", "Qs", "Js", "Ts"])` inapaswa kurudisha `"Royal Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_best_hand(cards):

    return cards
```

# --solutions--

```py
def get_best_hand(cards):
    rank_order = "23456789TJQKA"

    ranks = [c[0] for c in cards]
    suits = [c[1] for c in cards]

    is_flush = len(set(suits)) == 1
    rank_indices = sorted(rank_order.index(r) for r in ranks)

    def is_straight(indices):
        is_regular = all(indices[i] == indices[i - 1] + 1 for i in range(1, len(indices)))
        is_low_ace = indices == [0, 1, 2, 3, 12]
        return is_regular or is_low_ace

    straight = is_straight(rank_indices)

    counts = {}
    for r in ranks:
        counts[r] = counts.get(r, 0) + 1
    freq = sorted(counts.values(), reverse=True)

    if is_flush and rank_indices == [8, 9, 10, 11, 12]:
        return "Royal Flush"
    if is_flush and straight:
        return "Straight Flush"
    if freq[0] == 4:
        return "Four of a Kind"
    if freq[0] == 3 and freq[1] == 2:
        return "Full House"
    if is_flush:
        return "Flush"
    if straight:
        return "Straight"
    if freq[0] == 3:
        return "Three of a Kind"
    if freq[0] == 2 and freq[1] == 2:
        return "Two Pair"
    if freq[0] == 2:
        return "Pair"
    return "High Card"
```
