---
id: 6a0dcd03ee4e68698080ef68
title: "Zoezi la 304: Mpangilio wa Ratiba ya Safari"
challengeType: 29
dashedName: challenge-304
---

# --description--
Izvući red podataka sa najmanje dve opcione stanice za jednodnevno putovanje, vratiti broj validnih podešavanja rasporeda.

Raspored uvijek uključuje `"breakfast"`, `"lunch"` i `"dinner"`; oni se neće prihvatiti kao argumenti. Opcioni objekti mogu se postaviti bilo gdje u rasporedu, uz poštovanje sljedećih pravila:

- `"breakfast"` je uvek prvi, i najmanje jedna stanica pre `"lunch"`.
- `"lunch"` mora da se pojavi pre `"dinner"`, sa najmanje jednom stanicom između njih.
- Više od svega, jedna opcionalna stanica može biti viđena nakon `"dinner"`.

Rudisha idadi ya mipangilio halali.

# --hints--

`get_itinerary_count(["library", "park"])` inapaswa kurudisha `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])` inapaswa kurudisha `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])` inapaswa kurudisha `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` inapaswa kurudisha `840`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` inapaswa kurudisha `55440`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
