---
id: 6a0dcd03ee4e68698080ef6b
title: "Zoezi la 307: Kanuni za Mipangilio ya Miji"
challengeType: 29
dashedName: challenge-307
---

# --description--
Ako vam je data 2D mreža (matrica redova podataka) koja predstavlja raspored gradskih zgrada, vratite koordinata svih zgrada koje krše pravila urbanističkog planiranja.

Svaka ćelija u mrežici ima jednu od oznaka iz tabele ispod. Zgrada krši pravilo ako je bilo koji od njenih suseda (do četiri), vertikalno ili horizontalno, tip koji ne može biti njen sused.

| Oznaka | Tip | Ne može biti pored |
| - | - | - |
| `"i"` | viwanda | `"R"`, `"I"` |
| `"A"` | kilimo | `"C"` |
| `"R"` | makazi | `"i"`, `"C"` |
| `"I"` | taasisi | `"i"` |
| `"C"` | biashara | `"R"`, `"A"` |
| `""` (prazan string) | nije ugnježden | nema ograničenja |

Vrati sve koordinate ćelija koje ne važe kao niz parova od `[row, col]`, u bilo kom redosledu. Ako nema nikakvih kršenja, vrati prazan niz.

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])` inapaswa kurudisha `[[0, 0], [0, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])` inapaswa kurudisha `[[0, 1], [1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` inapaswa kurudisha `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` inapaswa kurudisha `[[0, 1], [0, 2], [0, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` inapaswa kurudisha `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
