---
id: 6a0dcc730cb92a616f86f0c5
title: "Vežba 301: Konačni prtljag
challengeType: 29
dashedName: challenge-301
---

# --description--
Uzimajući u obzir broj kašika deterdženta za veš koje imate i red podataka o kašikama koje ste koristili u prethodnim danima, vratite broj punih dana sapuna/deterdženta koji vam ostaje.

Izračunaj prosečnu potrošnju za svaki dan iz istorije potrošnje i pretpostavi da ta potrošnja traje svakog dana.

# --hints--

`last_load_date(10, [2, 2, 2, 2, 2, 2, 2])` inapaswa kurudisha `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(10, [2, 2, 2, 2, 2, 2, 2]), 5)`)
}})
```

`last_load_date(16, [2, 3, 0, 3, 4, 2, 1])` inapaswa kurudisha `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(16, [2, 3, 0, 3, 4, 2, 1]), 7)`)
}})
```

`last_load_date(33, [5, 0, 4, 3, 3, 2])` inapaswa kurudisha `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(33, [5, 0, 4, 3, 3, 2]), 11)`)
}})
```

`last_load_date(50, [2, 0, 2, 9, 12, 0, 2])` inapaswa kurudisha `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(50, [2, 0, 2, 9, 12, 0, 2]), 12)`)
}})
```

`last_load_date(20, [13, 9, 12, 10, 8])` inapaswa kurudisha `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(20, [13, 9, 12, 10, 8]), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def last_load_date(scoops, usage):

    return scoops
```

# --solutions--

```py
def last_load_date(scoops, usage):
    avg = sum(usage) / len(usage)
    return int(scoops / avg)
```
