---
id: 6a151d32d772271dd2448c2c
title: "Vežba 310: Britanski engleski u američki engleski"
challengeType: 29
dashedName: challenge-310
---

# --description--
Kada date rečenicu, promenite sve pravopisne razlike iz britanskog engleskog da odgovaraju američkom engleskom koristeći tabelu za pretraživanje koja se nalazi ispod i vratite ažuriranu rečenicu:

| British | American |
| - | - |
| `"colour"` | `"color"` |
| `"flavour"` | `"flavor"` |
| `"honour"` | `"honor"` |
| `"neighbour"` | `"neighbor"` |
| `"labour"` | `"labor"` |
| `"humour"` | `"humor"` |
| `"centre"` | `"center"` |
| `"fibre"` | `"fiber"` |
| `"defence"` | `"defense"` |
| `"offence"` | `"offense"` |
| `"organise"` | `"organize"` |
| `"recognise"` | `"recognize"` |
| `"analyse"` | `"analyze"` |

- Ove promene treba da budu nezavisne od velikog ili malog slova. Na primer, `"Colour"` bi trebalo da se promeni u `"Color"`.
- Ulazni podaci mogu sadržati reči koje su sastavljene na osnovu stvarnog pravopisa korena u tabeli, a koje takođe zahtevaju promenu. Na primer, `"colouring"` bi trebalo da se promeni u `"coloring"`, i `"disorganised"` bi trebalo da se promeni u `"disorganized"`.

# --hints--

`british_to_american("I love the colour blue.")` inapaswa kurudisha `"I love the color blue."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(british_to_american("I love the colour blue."), "I love the color blue.")`)
}})
```

`british_to_american("The fibre optic cable is new.")` inapaswa kurudisha `"The fiber optic cable is new."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(british_to_american("The fibre optic cable is new."), "The fiber optic cable is new.")`)
}})
```

`british_to_american("It's an honour to meet someone with such humour.")` inapaswa kurudisha `"It's an honor to meet someone with such humor."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(british_to_american("It's an honour to meet someone with such humour."), "It's an honor to meet someone with such humor.")`)
}})
```

`british_to_american("The unrecognised artist analysed his colour palette at the centre.")` inapaswa kurudisha `"The unrecognized artist analyzed his color palette at the center."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(british_to_american("The unrecognised artist analysed his colour palette at the centre."), "The unrecognized artist analyzed his color palette at the center.")`)
}})
```

`british_to_american("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful.")` inapaswa kurudisha `The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(british_to_american("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful."), "The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.")`)
}})
```

# --seed--

## --seed-contents--

```py
def british_to_american(sentence):

    return sentence
```

# --solutions--

```py
import re

def british_to_american(sentence):
    lookup = {
        "colour": "color",
        "flavour": "flavor",
        "honour": "honor",
        "neighbour": "neighbor",
        "labour": "labor",
        "humour": "humor",
        "centre": "center",
        "fibre": "fiber",
        "defence": "defense",
        "offence": "offense",
        "organise": "organize",
        "recognise": "recognize",
        "analyse": "analyze"
    }

    result = sentence
    for british, american in lookup.items():
        def replace(match, american=american):
            return american[0].upper() + american[1:] if match.group()[0].isupper() else american
        result = re.sub(british, replace, result, flags=re.IGNORECASE)

    return result
```
