---
id: 6a151d32d772271dd2448c2d
title: "Zoezi la 311: Mchawi wa maneno"
challengeType: 29
dashedName: challenge-311
---

# --description--

Ako je dat string kodova reči koje koristiš, izračunaj ukupan broj poena.

Svaka slovo u sekvenci predstavlja reč:

| Kod | Reč | Kategorija | Osnovni broj poena |
| - | - | - | - |
| `"f"` | Moto | Uharibifu | 3 |
| `"l"` | Umeme | Uharibifu | 3 |
| `"i"` | Barafu | Udhibiti | 2 |
| `"w"` | Upepo | Udhibiti | 2 |
| `"h"` | Tiba | Urejeshaji | 1 |
| `"s"` | Mlinzi | Urejeshaji | 1 |

Kombinaciona inkrementacija je postavljena u zavisnosti od broja uzastopnih reči korišćenih iz različitih kategorija:

- Prva reč uvek dobija osnovnu poen vrednost.
- Svaka sledeća reč iz druge kategorije od prethodne povećava množilac za 1.
- Upotreba reči iz iste kategorije kao i prethodna vraća množilac na 1.
- Poen za svaku reč je njena osnovna vrednost pomnožena trenutnim množilacem.

Restauriraj interpunkcijske znake iz niza reči.

# --hints--

`cast("fihwl")` inapaswa kurudisha `33`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` inapaswa kurudisha `45`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` inapaswa kurudisha `37`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` inapaswa kurudisha `50`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` inapaswa kurudisha `101`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
