---
id: 685523c05fb41d0848c4b52f
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Unutar prvog elementa od `div`, kreiraj element `span` sa tekstom `Soundflow`, element `h1` sa tekstom `Discover New Music`, i element `p` sa tekstom `Stream your favorite tracks and discover new artists.`

# --hints--

Trebalo bi da kreirate element `span` sa tekstom `Soundflow` unutar prvog elementa `div`.

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

Treba da napravite `h1` element sa tekstom `Discover New Music` unutar prvog `div` elementa.

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

Treba da napravite `p` element sa tekstom `Stream your favorite tracks and discover new artists.` unutar prvog `div` elementa.

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
