---
id: 68552464ef378e1051a55c28
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Unutar `div` drugog, kreiraj dva anker elementa. Prvi treba da ima tekst `Learn more`, a drugi treba da ima tekst `Start listening`.

S obzirom na `href` dva anker elementa, postavi sve na `#`.

# --hints--

Moraš da kreiraš dva elementa od `a` unutar drugog elementa `div`.

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

Tvoj prvi element ``a`` treba da ima tekstualni sadržaj ``Learn more`` i atribut od ``href`` za ``#`` unutar drugog elementa ``div``.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

Tvoj drugi element od `a` treba da ima tekst `Start listening` i svojstvo `href` od `#` unutar drugog elementa `div`.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
