---
id: 68551b677a5f6898af364b85
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Na ovom radionici, naučićete više o Tailwind CSS koristeći njegove flexbox utility klase da napravite sekciju koja se prilagođava uređaju za poziv na akciju (CTA).

Osnovni dizajn za HTML je već spreman za vas, pa prvo morate izgraditi komponente za HTML.

Počnite kreiranjem elementa za `div`. Unutar `div`, napravite dva različita elementa od `div`.

# --hints--

Trebalo bi da kreirate element za `div`.

```js
assert.exists(document.querySelector("div"));
```

Trebalo bi da imaš još dva elementa `div` unutar elementa `div`.

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
