---
id: 6858ee73ed8106ccb214942a
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Sada je sve poravnato na lijevoj strani, moraš centrirati sav sadržaj dodavanjem automatske margine na X-osu.

Dodajte klasu Tailwind koja odgovara tome na eksterni element ``div``. Da biste videli taj efekat, podesite karticu za prikaz tako da je sav sadržaj centriran.

# --hints--

Tvoja spoljna komponenta za `div` treba da ima servisnu klasu od `mx-auto`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--    
  </body>
</html>
```
