---
id: 6858f3a67aa8580ff37016c7
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Ako podesite širinu ekranskog taba da bude jednaka širini mobilnog uređaja, videćete sadržaj uvek jedan pored drugog. Ovo nije ono što želite na mobilnim uređajima.

Da biste to popravili, dodajte klasu elemenata `flex-col` u `div`. Ovo će promeniti fleksibilni pravac (flex direction) u red, tako da će elementi biti postavljeni jedan preko drugog.

Evo kako da se koristi klasa servisa za `flex-col`:

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

Vaš element otvaranja za `div` treba da ima klasu usluge `flex-col`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
