---
id: 6858f90c54647a48ddb8822d
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Trenutno, sadržaj, to jest, dva elementa `div` unutar spoljnog `div` biće postavljena jedan preko drugog bez obzira na veličinu ekrana.

Ovo je zato što imate klasu elemenata za `flex-col` u atributu zatvaranja `div`. Da biste to popravili, koristite prefiks za `lg:` da dodate klasu elemenata `flex-row` na `div`.

Ovo znači na uređajima sa velikim ekranima, fleksibilna orijentacija će biti `row`. Da biste videli da to ponovo radi, prilagodite pregledač širem opsegu/širini.

# --hints--

Vaša eksterni komponenta `div` mora imati klasu usluge `lg:flex-row`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
