---
id: 6858ebf2552046ab3de9f17c
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Trenutno, videćete da postoji previše prostora sa desne strane sadržaja u `div`. Možete ovo smanjiti povećanjem širine.

Dodajte klasu elemenata `md:w-1/2` u `div` da smanjite širinu do 50% na srednjim i većim ekranima. To jest, ekrani sa 768 piksela i više.

Da biste videli taj efekat, proširite predložak prikaza.

# --hints--

Tvoj spoljašnji element `div` treba da ima klasu usluge `md:w-1/2`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
