---
id: 6858ef974c8ad5dcd7abb9f1
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Sadržaj unutar `div` može se organizovati i stilizovati na najbolji mogući način. Za raspored i pozicioniranje elemenata, potrebno je koristiti flexbox.

Dodaj klasu elemenata `flex` za `div` ili dva detinjstva elementa `div` su postavljena jedno pored drugog.

Takođe, postavite ``justify-around`` ili da prostor oko sadržaja bude ravnomerno distribuiran, i ``items-center`` ili da sav sadržaj bude potpuno centriran unutar ``div``.

# --hints--

Vaš spoljni element `div` mora imati servisnu klasu `flex`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

Tvoj spoljni element za `div` bi trebalo da ima servisnu klasu od `justify-around`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

Vaš eksterni komponent `div` mora imati klasu usluge `items-center`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
