---
id: 68597114730d7c6b8838d204
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Sekcija ispod `h1` izgleda dobro, pa ne zahteva više stilizovanja. To znači da možete nastaviti sa anchor elementima unutar drugog elementa od `div`.

Dodaj klasu elemenata `bg-white`, `hover:bg-gray-100`, na `text-indigo-600` za prvi anker element. Ovo mu daje belu pozadinsku boju, indigo boju teksta i nešto svetliji pozadinski deo kada klikneš na njega.

Za drugi jangkir element, dodeli klasu objekata `bg-pink-500`, `hover:bg-pink-600`, i `text-white`.

# --hints--

Tvoj prvi anker element mora imati servisnu klasu `bg-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

Tvoj prvi jangkorni element treba da ima servisnu klasu `hover:bg-gray-100`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

Vaš prvi anker element treba da ima servisnu klasu `text-indigo-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

Tvoj drugi anker element bi trebalo da ima servisnu klasu od `bg-pink-500`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

Vaš drugi ankorni element treba da ima servisnu klasu `hover:bg-pink-600`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

Tvoj drugi anker element mora imati servisnu klasu od `text-white`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
