---
id: 685972d14c6d8182fbae71f1
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Da bi sidarne oznake izgledale kao dugme, dodajte klase elemenata `px-4` i `py-2` svakom od njih. `px-4` će kreirati unutrašnji prostor desno i levo za `1rem`, a `py-2` će kreirati unutrašnji prostor gore i dole za `0.5rem`.

# --hints--

Tvoj prvi anker element treba da ima servisnu klasu objekata `px-4`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("px-4"));
```

Vaš prvi anker element bi trebalo da ima klasu objekta usluge `py-2`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("py-2"));
```

Vaš drugi element ankora bi trebalo da ima klasu objekta usluge `px-4`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("px-4"));
```

Tvoja druga karakteristika sidra bi trebalo da ima klasu objekata usluge `py-2`.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("py-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a
          href="#"
          class="bg-white hover:bg-gray-100 text-indigo-600"
        >
          Learn more
        </a>
        <a
          href="#"
          class="bg-pink-500 hover:bg-pink-600 text-white"
         >
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
