---
id: 673b567e3ba535dda140d278
title: Izgradi sistem glasanja
challengeType: 26
dashedName: build-a-voting-system
saveSubmissionToDB: true
---

# --description--

U ovoj laboratoriji, gradiš sistem glasanja koji koristi `Map` za kreiranje glasova i `Set` sprečava dvostrano glasanje.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Morate inicijalizovati parametar `poll` na novu vrednost `Map`.

2. Trebalo bi da imate funkciju `addOption` koja prima parametar `option`.

3. U funkciji `addOption`:

- Ako `option` još nije u glasanju, trebalo je ga dodati u glasanje sa `Set` praznim kao vrednošću za praćenje birača. Takođe je potrebno vratiti poruku `Option "<option>" added to the poll.`

- Ako `option` već postoji, treba da vrati poruku `Option "<option>" already exists.`

- Ako pokušate dodati praznu opciju, funkcija bi trebalo da vrati poruku `Option cannot be empty.`

4. Treba da ima funkciju `vote` koja prima dva parametra, `option` (izbor za glasanje) i `voterId` (jedinstveni identifikator birača).

5. U funkciji `vote`:

- Ako `option` nije u glasanju, funkcija bi trebalo da vrati poruku `Option "<option>" does not exist.`

- Ako `option` postoji, funkcija bi trebalo da proveri da li `voterId` je već glasao za ovaj `option`.

- Ako je birač već glasovao, funkcija bi trebalo da vrati poruku `Voter <voterId> has already voted for "<option>".`

- Ako birač nije glasovao, `voterId` treba ga dodati u `Set` od `voters` za ovu opciju. Funkcija bi trebalo da vrati poruku `Voter <voterId> voted for "<option>".`

6. Treba vam najmanje tri opcije u vašem `poll`.

7. `poll` Tvoje bi trebalo da ima najmanje tri glasa.

8. Trebalo bi da ima funkciju `displayResults` koja vraća rezultate glasanja u sledećem formatu:

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

Trebalo bi da imaš kriterijum `poll` inicijalizovan kao nova vrednost `Map`.

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

Morate definisati funkciju `addOption` koja prima parametar `option`.

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

Morate definisati funkciju `vote` koja prima dva parametra, `option` i `voterId`.

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

Treba da definišete funkciju `displayResults` koja prikazuje rezultate glasanja.

```js
assert.isFunction(displayResults);
```

Trebalo bi da imaš najmanje tri opcije u vašem `poll`.

```js
assert.isAtLeast(poll.size, 3);
```

`poll` bi trebalo da ima najmanje tri glasa.

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

Morate osigurati da je svaka opcija glasanja povezana sa `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")` treba vratiti `Option "Egypt" added to the poll.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Dodavanje prazne opcije trebalo bi da vrati `"Option cannot be empty."`

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

Kada `Turkey` je već dodato, `addOption("Turkey")` bi trebalo da vrati `Option "Turkey" already exists.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Kada `Malaysia` bude u opcijama glasanja, `vote("Malaysia", "traveler1")` treba da vrati `Voter traveler1 voted for "Malaysia".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` treba da ažurira `Set` podatke birača za tu opciju.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Kada `traveler1` pokušava da glasa za `Algeria` ponovo, `vote("Algeria", "traveler1")` treba da vrati `Voter traveler1 has already voted for "Algeria".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Dva puta ne bi trebalo da poveća veličinu `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Kada `Nigeria` nije u opcijama glasanja, `vote("Nigeria", "traveler2")` treba da vrati `Option "Nigeria" does not exist.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Jedinstveni izbor može da primi mnogo glasova.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Kada opcija `Turkey` dobije dve glasa, opcija `Morocco` dobija jedan glas, a opcija `Spain` ne dobija nijedan glas, `displayResults()` treba da vrati/pošalje `"Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes"`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');
  addOption('Spain');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` Trebalo bi da vrati rezultate u ispravnom formatu.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
