---
id: 63b607af6fcdb119aae9b16a
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Na vašem formularu, korisnici bi trebalo da imaju mogućnost dodavanja različitih vrsta obroka zajedno sa brojem njihovih kalorija.

Kreiraj element `fieldset` sa atributom `id` postavljenim na `breakfast`.

Unutar tog elementa, kreiraj `legend` sa tekstom `Breakfast`, i u praznom `div` podešenim na `class` koji je postavljen na `input-container`.

# --hints--

Trebalo bi da kreiraš element `fieldset` unutar svog `form`.

```js
assert.exists(document.querySelector('form fieldset'));
```

Tvoj element `fieldset` treba da se nalazi nakon elementa `input`.

```js
assert.equal(document.querySelector('form fieldset')?.previousElementSibling?.tagName, "INPUT");
```

Vaš element `fieldset` treba da ima `id` koji je postavljen na `breakfast`.

```js
assert.equal(document.querySelector('form fieldset')?.id, "breakfast");
```

Treba da kreiraš element `legend` unutar svog `fieldset`.

```js
assert.exists(document.querySelector('form fieldset legend'));
```

Vaš element `legend` mora imati tekst od `Breakfast`.

```js
assert.equal(document.querySelector('form fieldset legend')?.innerText, "Breakfast");
```

Treba da kreiraš element `div` unutar svog elementa `fieldset`.

```js
assert.exists(document.querySelector('form fieldset div'));
```

Element `div` mora da sledi nakon elementa `legend`.

```js
assert.equal(document.querySelector('form fieldset div')?.previousElementSibling?.tagName, "LEGEND");
```

Tvoj element `div` treba da ima `class` postavljen na `input-container`.

```js
assert.equal(document.querySelector('form fieldset div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
