---
id: 684aaf9ec670c68d20efd0d0
title: Konstruktor administratora konfiguracije korisnika
challengeType: 27
dashedName: build-a-user-configuration-manager
saveSubmissionToDB: true
---

# --description--

U ovoj laboratoriji, kreiraćeš Upravljač podešavanjima korisnika koji omogućava korisniku da upravlja svojim podešavanjima kao što su tema, jezik i obaveštenja. Implementiraćeš funkcije za dodavanje, uređivanje, brisanje i prikazivanje podešavanja korisnika.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba da definiše funkciju pod nazivom `add_setting`, koja ima dva parametra koji predstavljaju rečnik podešavanja i tuple koja sadrži par ključ-vrednost.

1. Funkcija `add_setting` treba da:
    
   - Promeni ključ i vrednost u mala slova.
   - Ako postoji podešavanje ključa, vrati `Setting '[key]' already exists! Cannot add a new setting with this name.`
   - Ako ne postoji podešavanje ključa, dodaj par ključ-vrednost u dostavljeni rečnik podešavanja i vrati `Setting '[key]' added with value '[value]' successfully!`.
   - Vraćena poruka treba da ima ključ i vrednost u malim slovima.

1. Treba da definišete funkciju nazvanu `update_setting` sa dva parametra koja predstavljaju rečnik podešavanja i tuple koji sadrži par ključ-vrednost.

1. Funkcija `update_setting` mora da:
       
   - Konvertuje ključ i vrednost u mala slova.
   - Ako postoji podešavanje ključa, ažurira njegovu vrednost u dato rečnik podešavanja i vraća: `Setting '[key]' updated to '[value]' successfully!`
   - Ako ne postoji podešavanje ključa, vraća `Setting '[key]' does not exist! Cannot update a non-existing setting.`
   - Vraćena poruka mora imati ključ i vrednost u malim slovima.

1. Morate definisati funkciju nazvanu `delete_setting` sa dva parametra koja predstavljaju rečnik podešavanja i ključ.

1. Funkcija za `delete_setting` treba da:
    
   - Promeni poslati ključ u mala slova.
   - Ako postoji podešavanje ključa, ukloni par ključ-vrednost iz dato rečnika podešavanja i vrati `Setting '[key]' deleted successfully!`
   - Ako ne postoji podešavanje ključa, vrati `Setting not found!`
   - Vraćena poruka mora imati ključ u malim slovima.

Treba da definišete funkciju nazvanu `view_settings` koja ima jedan parametar koji predstavlja rječnik podešavanja.

1. Funkcija za `view_settings` treba da:
   
   - Vrati `No settings available.` ako je dati rečnik podešavanja prazan.
   - Ako rečnik ima bilo koje podešavanje, vrati niz znakova koji prikazuje podešavanja. Taj niz treba da počinje sa `Current User Settings:` i praćen parovima ključ-vrednost, svaki na novoj liniji, a ključ mora započeti velikim slovom. Na primer, `view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})` treba da vrati:

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. Za testiranje koda, trebate kreirati rečnik pod nazivom `test_settings` za čuvanje nekih korisničkih podešavanja konfiguracije.

# --hints--

Trebalo bi da kreirate rječnik pod nazivom `test_settings` dodajući mu neke vrijednosti unutra.

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

Treba da definišete funkciju nazvanu `add_setting`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

Funkcija za `add_setting` treba da ima dva parametra.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting` Trebalo bi da promeni ključ u mala slova.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting` Trebalo bi da promeni vrednost u mala slova.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))` treba da vrati poruku o grešci `Setting 'theme' already exists! Cannot add a new setting with this name.`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))` treba da doda novi par ključ-vrednost i vrati poruku o uspehu `Setting 'volume' added with value 'high' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting` Trebalo bi dodati par ključ-vrednost koji je dostavljen u rečnik.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

Trebalo bi da definišete funkciju nazvanu `update_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

Funkcija `update_setting` treba da ima dva parametra.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

Funkcija za `update_setting` treba da konvertuje ključ u mala slova.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

Funkcija ``update_setting`` treba da konvertuje vrednost u mala slova.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))` trebalo bi da ažurira postojeći ključ i vrati poruku o uspehu `Setting 'theme' updated to 'dark' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))` treba da vrati poruku o grešci `Setting 'volume' does not exist! Cannot update a non-existing setting.` kada ključ ne postoji.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting` Trebalo bi da ažurira par ključ-vrednost koji je dobijen iz rečnika.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

Trebalo bi da definišete funkciju nazvanu `delete_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

Funkcija `delete_setting` treba da ima dva parametra.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting` Trebalo bi da promeni ključ u mala slova.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')` treba da obriše postojeći ključ i vrati poruku o uspehu `Setting 'theme' deleted successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

`delete_setting` mora da vrati poruku o grešci `Setting not found!` kada ključ nedostaje.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting` Trebalo bi da precizno ukloni izvučeni ključ iz rečnika.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

Trebalo bi da definišete funkciju pod nazivom `view_settings`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

Funkcija `view_settings` mora imati jedan parametar.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings` trebalo bi vratiti poruku `No settings available.` ako je dostavljeni rečnik prazan.

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings` Treba vratiti podešavanja konfigurisana za nepristojan rečnik.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings` Trebalo bi da se stavi veliko slovo na početak svakog imena polja.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings` Trebalo bi da prikaže ispravne rezultate i završi ponašanjem novog reda.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "english", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: english\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "english",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
