---
id: 673339d96c56d36e963e888e
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, kreiraćete katalog biljke bašte dok vežbate JavaScript objekte Map i Set.

Da biste predstavili svaku biljku u vašem katalogu, koristićete objekte, svaki sa tri svojstva:

- `commonName`, uobičajeno ime biljke
- `scientificName`, naučno ime (ili latinsko ime) biljke
- `cultivar`, varijanta biljke

Počnite kreiranjem objekta sa gore navedenim tri svojstva za `Spanish lavender` (naučno ime `Lavandula stoechas`) kultivara pod nazivom `Ballerina` i dodelite ga promenljivoj nazvanoj `ballerina`.

# --hints--

You should have a variable named `ballerina`.

```js
assert.exists(ballerina);
```

You should assign an object to your `ballerina` variable.

```js
assert.isObject(ballerina);
```

Your `ballerina` object should have a `commonName` property set to `"Spanish lavender"`.

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

Your `ballerina` object should have a `scientificName` property set to `"Lavandula stoechas"`.

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

Your `ballerina` object should have a `cultivar` property set to `"Ballerina"`.

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

Your `ballerina` object should have three keys.

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
