---
id: 6733b5e6f210c1bd1afdf8b3
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Preostali objekti biljaka dodati su za vas. Sada je vreme da kreirate svoj katalog.

Kao i obični JavaScript objekti, objekti `Map` sadrže kolekciju parova ključ-vrednost. Da biste kreirali objekat `Map`, možete koristiti konstruktor `Map`:

```js
const myMap = new Map();
```

Deklarišite varijablu `catalog` i dodelite joj objekat `Map`.

# --hints--

You should have a variable named `catalog`.

```js
assert.exists(catalog);
```

You should assign a `Map` object to your `catalog` variable.

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

--fcc-editable-region--

--fcc-editable-region--
```
