---
id: 67349ba230cff237c0f2bdf2
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Objekti mapa pamte redosled kojim su ključevi umetnuti u mapu. Da biste postavili par ključ-vrednost u vaš objekat mape, možete koristiti metod `set`. Sledeće postavlja ključ `"name"` na vrednost `"Jerry"` u objektu mape `myMap`.

```js
myMap.set("name", "Jerry");
```

Postavite ključ `"ballerina"` na vrednost vašeg objekta `ballerina` u vašoj mapi `catalog`. Zatim, logujte svoj katalog u konzolu.

Napomena: Da biste pravilno videli svoj objekat mape, moraćete da otvorite konzolu pretraživača.

# --hints--

You should set the `"ballerina"` key to the value of your `ballerina` object in your `catalog` map.

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

Your `catalog` map should contain a single key-value pair.

```js
assert.lengthOf(catalog, 1);
```

You should log your `catalog` map to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
--fcc-editable-region--

--fcc-editable-region--
```
