---
id: 587d825a367417b2b2512c89
title: Implementiraj algoritam Quicksort
challengeType: 26
dashedName: implement-the-quicksort-algorithm
---

# --description--

U ovom laboratoriju ćeš implementirati sortirujući algoritam: quicksort.

Quicksort je efikasan metod za rekurentne podelu i osvajanje (recursive divide-and-conquer) za sortiranje niza podataka. U ovom metodu, vrednost pivota se bira iz početnog niza podataka. Niz se deli na dva manja niza sa vrednostima koje su manje i veće od vrednosti pivota. Zatim se rezultati pozivanja Quicksort algoritma rekurentno na oba manja niza spajaju. Ovo se nastavlja sve dok se ne dosegne osnovni slučaj praznog ili jednoelementnog niza, koji se vraća. Smanjenje rekurentnih poziva vraća sortirani niz.

Iako je izbor vrednosti pivota važan, bilo koji pivot biće dovoljan za naše svrhe ovde. Za jednostavnost, može se koristiti prvi ili poslednji element.

Quicksort je vrlo efikasna metoda sortiranja, pružajući performanse od *O(nlog(n))* u proseku. Takođe je laka za implementaciju. Ove karakteristike ga čine popularnom i korisnom metodom sortiranja.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da postoji funkcija za `quicksort` koja prima niz celobrojnih brojeva kao ulaz.
1. Funkcija za `quicksort` treba da vrati niz podataka sa celobrojnim brojevima jednakim ulazu, ali u redosledu od dna ka vrhu.
1. Funkcija za `quicksort` treba da se pozove rekurzivno radi sortiranja niza podataka.

# --hints--

Trebalo bi da imaš funkciju `quicksort`.

```js
assert.isFunction(quicksort);
```

`quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` Trebalo bi da vrati nepromenjiv red podataka, osim po redu.

```js
assert.sameMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` trebalo bi da vrati sortiran red podataka (od najmanjeg do najvećeg).

```js
assert.sameOrderedMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`quicksort` ne bi trebalo da koristi metod ugrađen u `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quicksort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function quicksort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
  }
}
```
