---
id: 655cd899f8de09431eabb40c
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovom projektu, naučićete o strukturama podataka kako biste izgradili algoritam za merge sort.

Ovo je algoritam za sortiranje koji koristi princip razdvajanja i osvajanja kako bi sortirao skupove podataka. To jest, deli skup na manje delove, sortira te delove nezavisno, a zatim spaja sortirane delove.

Kreiraj funkciju nazvanu `merge_sort`. Ova funkcija će obraditi vežbu sortiranja liste brojeva.

Koristite glavnu reč `pass` unutar tela funkcije.

# --hints--

Treba da objaviš funkciju koja se zove `merge_sort`. Ne zaboravi da koristiš ključnu reč `pass` unutar tela funkcije.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
