---
id: 6564a9fe51964c229d5b7f4c
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Biće vam potreban parametar koji prikazuje kolekciju podataka koja treba da se sortira. Kreirajte parametar pod nazivom `array` u funkciji `merge_sort`.

# --hints--

Tvoja funkcija za `merge_sort` mora da primi samo jedan parametar: `array`.

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
