---
id: 6564aee9c077774ea49c3faf
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Algoritam za merge sort obavlja tri glavna koraka:

- Razdvojite neorganizovan niz predmeta u manje delove
- Rasporedite predmete u manje delove
- Spojite organizovane manje delove

Sve ovo se dešava u sekvenci ponavljanja dok manji delovi ne formiraju jednu kompletnu, organizovanu sekvencu. Počnite sa podešavanjem sekvence.

Prvo, promeni token `pass` u varijablu `middle_point`. Koristi operator za celobrojno deljenje (`//`) da podeli dužinu liste/niza `array` sa 2 i dodeli tu vrednost novoj varijabli `middle_point`. Zapamti da unosiš svoj kod u ispravnom redosledu.

# --hints--

Ne smete imati `pass` unutar tela funkcije.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.notMatch(function_body, /^\s{4}pass/m);
  }
})
```

Trebalo bi da deklariš varijablu `middle_point` i dodeli joj vrednost `len(array) // 2`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}middle_point\s*=\s*len\(\s*array\s*\)\s*\/\/\s*2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
