---
id: 656680b0fc79f2c38a34d90e
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Sada je vreme za organizovanje i spajanje liste (`left_part` i `right_part`) u prethodni `array`.

Možete ovo učiniti poređenjem elemenata u obe liste i spajanjem najmanjeg elementa u glavnu listu. Izvršićete ovo poređenje za sve primere u `left_part` i `right_part`.

Kreiraj tri kriterijuma: `left_array_index`, `right_array_index`, na `sorted_index` i postavi njihove vrednosti kao `0`. Ovi kriterijumi će vam pomoći da pratite svaki element tokom procesa sortiranja.

# --hints--

Trebalo bi da imaš parametar nazvan `left_array_index` unutar svoje funkcije `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index/);
  }
})
```

Vrednost `left_array_index` bi trebalo da bude 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_array_index\s*=\s*0/);
  }
})
```

Trebalo bi da imaš parametar nazvan `right_array_index` unutar svoje funkcije `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index/);
  }
})
```

Vrednost `right_array_index` treba da bude 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_array_index\s*=\s*0/);
  }
})

```

Trebalo bi da imaš parametar nazvan `sorted_index` unutar svoje funkcije `merge_sort`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index/);
  }
})
```

Vrednost `sorted_index` treba da bude 0

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /sorted_index\s*=\s*0/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)

--fcc-editable-region--

--fcc-editable-region--
```
