---
id: 656639771fed09b5c0e8fe71
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Sada, nakon deljenja spiska `array` na dva različita spiska, nastavićete da delite svaki spisak sve dok svaki element nije u svom pojedinačnom spisku. Spisak sa samo jednim stavkom je već podešen.

Da biste to uradili, itajiita `merge_sort` za petlju unutar vaše funkcije. Trenutno ne prosleđujte nikakve argumente.

# --hints--

Trebalo bi da nazoveš `merge_sort` na najdonjem delu vagine.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
