---
id: 656627b47bd2d2a4afbc945d
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

U prethodnom koraku, dobili ste srednju vrednost. Možete je koristiti za podelu `array` na dva dela i dodeljivanje vrednosti svakog dela u nove varijable.

Koristi sintaksu za slice da izvuče levi deo od `array` i postavi tu vrednost u parametar nazvan `left_part`.


# --hints--

Trebalo bi da imaš parametar nazvan `left_part` u svojoj funkciji `merge_sort`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

Trebalo bi da postaviš vrednost `array[:middle_point]` u parametar `left_part`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
