---
id: 656707e11e671ed54c96f7ec
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Kada je stanje `if` procenjeno kao `True`, to znači da je stavka u listi `left_part` manja od stavke sa kojom se poredi u listi `right_part`.

U tom slučaju, možete postaviti polje `left_part` u red strukturiranih podataka.

Unutar bloka `if`, ukloni `pass` i postavi vrednost `left_part[left_array_index]` za `array[sorted_index]`.

# --hints--

Treba da zameniš `pass` sa `array[sorted_index] = left_part[left_array_index]` unutar tela deklaracije `if`.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
