---
id: 656710d1e0ec62253426db24
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Ciklus koji je kreirao `while` upoređuje jedan element iz `left_part` sa drugim elementom iz `right_part`, a zatim dodaje još jedan element u glavni spisak `array`.

Nastavite sa ovom operacijom dok više ne postoji poređivački elementi. Međutim, `left_part` i dalje može imati preostale elemente kada `right_part` nema, i obrnuto.

Kreiraj drugi ciklus od `while` koji kopira preostale elemente iz `left_part` unutar liste `array`. Koristi `left_array_index < len(left_part)` za stanje `while`.

# --hints--

Tvoj ciklus od `while` treba da ima ovaj status: `left_array_index < len(left_part)`. Ne zaboravi da koristiš ključnu reč `pass` unutar tela ciklusa.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):/m);
  }
})
```

Trebalo bi da imaš ključnu reč `pass` unutar tela svog ciklusa `while`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
