---
id: 656709e50ed928da35260276
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

U prethodnom koraku, postavili ste vrednost elementa u `left_part` u listi `array` jer je bila mala. Ali ovo neće biti slučaj uvek. U nekim slučajevima poređenja, element sa desne strane može biti mali.

Generišite deo za `else` koji će biti izvršen kada šablon `left_part` nije ispod šablona `right_part`.

Unutar bloka od `else`, postavi vrednost `right_part[right_array_index]` za `array[sorted_index]`.

# --hints--

Trebalo bi da imaš `array[sorted_index] = right_part[right_array_index]` unutar tela svog `else`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /else:\s*[^}]*array\s*\[\s*sorted_index\s*\]\s*=\s*right_part\s*\[\s*right_array_index\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
