---
id: 656702f8b4cbd8cbf0a433c6
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Unutar vaše funkcije, kreirajte ciklus `while` koji povezuje element u `left_part` i elementu u `right_part`, i spaja najmanji element sa glavnom listom za `array`.

Napravi dva uslova za petlju: jedan koji proverava da li `left_array_index` je manji od dužine `left_part` i drugi uslov koji proverava da li `right_array_index` je manji od dužine `right_part`.

# --hints--

Vaš ciklus za `while` mora imati ove uslove: `left_array_index < len(left_part) and right_array_index < len(right_part)`. Ne zaboravite da koristite ključnu reč `pass` unutar tela ciklusa.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/);
  }
})
```

Koristi glavnu reč `pass` unutar tela ciklusa

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--

--fcc-editable-region--
```
