---
id: 6569b743630ee592a65a7e2f
title: Korak 26
challengeType: 20
dashedName: step-26
---

# --description--

Pre testiranja funkcije `merge_sort()`, neophodno je kreirati osnovni slučaj koji zaustavlja izvršavanje funkcije kada je dužina `array` jednaka ili manja od `1`.

Ovaj osnovni slučaj će zaustaviti poziv petlje. Bez njega, operacija sortiranja pomoću merge sort-a bi se nastavila izvršavati čak i ako je lista već sortirana ili nema nijedan element unutra.

Nakon deklaracije funkcije, kreiraj deklaraciju `if` sa ovim uslovom: `len(array) <= 1`. Koristi glavnu reč/ključnu reč `pass` unutar tela funkcije.

# --hints--

Treba da kreirate deklaraciju za `if` ili proverite da li `len(array) <= 1`. Ne zaboravite da koristite ključnu reč za `pass` unutar tela deklaracije za `if`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):

--fcc-editable-region--
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
