---
id: 6577320da0d4c2e594d418e2
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Pregledajte `right_part` kao komentar za poziv `merge_sort()` koji ste dodali u poslednji korak.

# --hints--

Trebalo bi da prosledite `right_part` kao argument vašem poslednjem pozivu `merge_sort()`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*right_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort()
--fcc-editable-region--
```
