---
id: 657f3a661730891aa64f3568
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Koristi sintaksu za *slice* da uzmeš desnu polovinu od `array` i postaviš tu vrednost u parametar nazvan `right_part`.

# --hints--

Trebalo bi da imaš kriterijum nazvan `right_part`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_part\s*(?!=)/m);
  }
})
```

Trebalo bi da postavite vrednost `array[middle_point:]` u parametar `right_part`

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_part\s*=\s*array\s*\[\s*middle_point\s*:\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
