---
id: 657f59751d5a5c9b069ae0f3
title: Korak 28
challengeType: 20
dashedName: step-28
---

# --description--

Možete koristiti varijablu `__name__` da utvrdite da li skripta Python radi kao glavni program ili je učitana kao modul (kod napisan u drugoj datoteci od Python).

Ako je vrednost `__name__` postavljena na `'__main__'`, ukazuje da je trenutni skript glavni program, a ne modul.

U ovom projektu, koristićeš trenutni skript kao glavni program.

Generiši izjavu za `if` koja proverava da li je vrednost `__name__` jednaka `'__main__'`.

Koristite glavnu reč `pass` unutar tela deklaracije `if`.

# --hints--

Vaša izjava za `if` treba da proveri da li `__name__` ima vrednost od `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
