---
id: 5a24c314108439a4d403614c
title: Dohvati `state` iz Redux skladišta
challengeType: 6
forumTopicId: 301443
dashedName: get-state-from-the-redux-store
---

# --description--

Komponenta Redux skladišta nudi nekoliko načina kojima možete da sa njim interagujete. Na primer, možete dobiti trenutnu vrednost `state` koja se nalazi u komponenti Redux skladišta koristeći metodu/način `getState()`.

# --instructions--

Kod iz prethodnog izazova je ponovo napisan kraće u uređivaču koda. Koristite `store.getState()` da biste dobili `state` iz `store`, i postavite ovu vrednost u novi parametar `currentState`.

# --hints--

Redux store treba da ima vrednost od 5 za inicijalno `state`.

```js
assert(store.getState() === 5);
```

Parametar `currentState` mora da ima i bude definisan trenutnom vrednošću iz Redux skladišta `state`.

```js
assert(currentState === 5 && code.includes('store.getState()'));
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
```

# --solutions--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
const currentState = store.getState();
```
