---
id: 5a24c314108439a4d403615a
title: Uklonite element iz niza podataka
challengeType: 6
forumTopicId: 301447
dashedName: remove-an-item-from-an-array
---

# --description--

Doslo je vreme za vežbanje uklanjanja elemenata iz niza podataka. Spread operator se takođe može koristiti ovde. Drugi važni načini za JavaScript su `slice()` i `concat()`.

# --instructions--

Reducer-i i kreatori akcija su modifikovani da uklone element iz niza podataka na osnovu predikata tog elementa. Završi pisanje reducer-a kako bi se vratila novi niz stanja sa uklonjenim elementom koji odgovara određenom kriterijumu.

# --hints--

Redux skladište treba da postoji i da se inicijalizuje sa stanjem jednakim `[0,1,2,3,4,5]`

```js
const initialState = store.getState();
assert.isArray(initialState) 
assert.deepEqual(initialState, [0, 1, 2, 3, 4, 5])
```

`removeItem` i `immutableReducer` sve treba da bude funkcionalno.

```js
assert(
  typeof removeItem === 'function' && typeof immutableReducer === 'function'
);
```

Akcija kreatora od `removeItem` treba da ukloni elemente iz stanja i NE ŽELI da direktno promeni to stanje.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(removeItem(3));
const state_1 = store.getState();
store.dispatch(removeItem(2));
const state_2 = store.getState();
store.dispatch(removeItem(0));
store.dispatch(removeItem(0));
store.dispatch(removeItem(0));
const state_3 = store.getState();

assert(isFrozen)
assert.deepEqual(state_1, [0, 1, 2, 4, 5]) 
assert.deepEqual(state_2, [0, 1, 4, 5]) 
assert.deepEqual(state_3, [5])

```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = [0,1,2,3,4,5], action) => {
  switch(action.type) {
    case 'REMOVE_ITEM':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const removeItem = (index) => {
  return {
    type: 'REMOVE_ITEM',
    index
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = [0,1,2,3,4,5], action) => {
  switch(action.type) {
    case 'REMOVE_ITEM':
      return [
        ...state.slice(0, action.index),
        ...state.slice(action.index + 1)
      ];
    default:
      return state;
  }
};

const removeItem = (index) => {
  return {
    type: 'REMOVE_ITEM',
    index
  }
}

const store = Redux.createStore(immutableReducer);
```
