---
id: 5a24c314108439a4d4036158
title: Nemoj menjati stanje direktno
challengeType: 6
forumTopicId: 301445
dashedName: never-mutate-state
---

# --description--

Ovi poslednji izazovi opisuju nekoliko načina za implementaciju glavnog principa neizmenjavanja stanja u Redux-u. Neizmenljivo stanje znači da ne možete direktno menjati stanje, već vraćate novu kopiju tog stanja.

Ako uzmete snimak stanja Redux aplikacije na trenutak, videćete nešto poput `state 1`, `state 2`, `state 3`,`state 4`, `...` i slično, gde može svaki stani biti isti kao prethodni, ali je svaki drugačiji komad podataka. Nepromenljivost ovog stanja zapravo omogućava funkcije kao što je debagovanje korišćenjem putovanja kroz vreme o kojima ste možda čuli.

Redux ne nameće imutabilnost stanja u svom store-u ili reducerima; to je odgovornost programa. Srećom, JavaScript (posebno ES6) pruža nekoliko korisnih alata koje možete koristiti za implementaciju imutabilnosti vašeg stanja, bilo da je to `string`, `number`, `array`, ili `object`. Zapamtite da su nizovi slova i brojeva primitive vrednosti koje su po prirodi nemeneškim. Drugim rečima, 3 je uvek 3. Ne možete promeniti vrednost broja 3. Međutim, `array` ili `object` su mutabilni. U praksi, vaše stanje može uključivati `array` ili `object`, jer su ovo strukture podataka važne za predstavljanje različitih vrsta informacija.

# --instructions--

Postoji `store` i `reducer` u editoru koda za upravljanje akcijama. Završite pisanje slučaja `ADD_TO_DO` u reduceru kako biste dodali novu akcionu karakteristiku u stanje. Postoji nekoliko načina za ovo korišćenjem standardnog JavaScript ili ES6-a. Proverite da li možete pronaći način za vraćanje novog niza podataka koji ima karakteristiku iz `action.todo` dodatu na kraj.

# --hints--

Redux skladište treba da postoji i bude inicijalizovano sa stanjem sličnim liniji `todos` u editoru koda.

```js
assert(
  (function () {
    const todos = [
      'Go to the store',
      'Clean the house',
      'Cook dinner',
      'Learn to code'
    ];
    const initialState = store.getState();
    return (
      Array.isArray(initialState) && initialState.join(',') === todos.join(',')
    );
  })()
);
```

`addToDo` na `immutableReducer` sve bi trebalo da budu putevi.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Slanje akcije tipa `ADD_TO_DO` u Redux store bi trebalo da doda funkciju `todo` i ne treba da menja stanje po komadu.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
  '__TEST__TO__DO__'
];
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

# --seed--

## --seed-contents--

```js
const ADD_TO_DO = 'ADD_TO_DO';

// A list of strings representing tasks to do:
const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const ADD_TO_DO = 'ADD_TO_DO';

const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      return state.concat(action.todo);
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
