---
id: 5a24c314108439a4d4036157
title: Napiši indikator u Reduxu
challengeType: 6
forumTopicId: 301453
dashedName: write-a-counter-with-redux
---

# --description--

Sada si naučio sve osnovne principe Redux-a! Video kako da kreiraš akcije i *action creators*, kreiraš Redux store, dispatchuje svoje akcije ka store-u, i dizajnira ažuriranja stanja korišćenjem čistih reducera. Video si čak ni kako da upravljaš kompleksnim stanjem sa strukturom reducera i obrađuješ asinhrone akcije. Ovi primeri su jednostavni, ali ovi koncepti su osnovni principi Redux-a. Ako ih dobro razumeš, spreman si da počneš da gradiš svoju Redux aplikaciju. Sledeći izazovi se tiču nekih detalja o nemenjivosti (immutability) za `state`, ali prvo, evo pregleda svega što si naučio do sada.

# --instructions--

U ovoj lekciji/kursu, kreiraćete jednostavan store koristeći Redux od nule. Osnove su date u editoru koda, ali ćete morati da popunite detalje! Koristite zadate imena i definišite *action creator*-e `incAction` i `decAction`, tipove akcija `counterReducer()`, `INCREMENT` i `DECREMENT`, a na kraju Redux store `store`. Kada završite, trebalo bi da budete u stanju da dispatch-ujete akcije `INCREMENT` ili `DECREMENT` kako biste povećali ili smanjili stanje sačuvano u `store`. Srećno sa izgradnjom vaše prve Redux aplikacije!

# --hints--

Kreator akcije `incAction` mora vratiti objekat akcije sa `type`, jednak vrednosti `INCREMENT`.

```js
assert(incAction().type === INCREMENT);
```

Generator akcije `decAction` mora da vrati objekat akcije koji ima `type`, jednak vrednosti `DECREMENT`

```js
assert(decAction().type === DECREMENT);
```

Pokretanje `store.getState()` treba da vrati broj

```js
assert(typeof store.getState() === 'number');
```

Redux skladište treba da se inicijalizuje sa `state` od 0.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    return _store.getState() === 0;
  })()
);
```

Slanje `incAction` u Redux store bi trebalo da poveća `state` za 1.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(incAction());
    const incState = _store.getState();
    return initialState + 1 === incState;
  })()
);
```

Slanje `decAction` u Redux store bi trebalo da smanji `state` za 1.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(decAction());
    const decState = _store.getState();
    return initialState - 1 === decState;
  })()
);
```

`counterReducer` Trebalo bi da bude funkcija

```js
assert(typeof counterReducer === 'function');
```

# --seed--

## --seed-contents--

```js
const INCREMENT = null; // Define a constant for increment action types
const DECREMENT = null; // Define a constant for decrement action types

const counterReducer = null; // Define the counter reducer which will increment or decrement the state based on the action it receives

const incAction = null; // Define an action creator for incrementing

const decAction = null; // Define an action creator for decrementing

const store = null; // Define the Redux store here, passing in your reducers
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const incAction = () => {
  return {
    type: INCREMENT
  }
};

const decAction = () => {
  return {
    type: DECREMENT
  }
};

const store = Redux.createStore(counterReducer);
```
