---
id: 5a24c314108439a4d4036156
title: Koristite centralnu aplikaciju za obradu više zadataka istovremeno
challengeType: 6
forumTopicId: 301451
dashedName: use-middleware-to-handle-asynchronous-actions
---

# --description--

Do sada su ove izazove izbegli diskutovati asinhrone akcije, ali je to neizbežan deo razvoja mreže. Ponekad ćete morati da pozovete asinhrone akcije u vašoj Redux aplikaciji, pa kako rešavate ovakve zahteve? Redux pruža *middleware* dizajniran posebno za ove svrhe, nazvan Redux Thunk. Evo kratkog objašnjenja kako koristiti ovo sa Redux-om.

Da biste uključili Redux Thunk middleware, ga šaljete kao argument za `Redux.applyMiddleware()`. Ova izjava se zatim pruža kao drugi opcioni parametar funkciji `createStore()`. Pogledajte kod ispod u editoru da biste ovo videli. Zatim, da biste kreirali asinhroni akciju, umotavate funkciju u kreator akcija koji uzima `dispatch` kao argument. Unutar ove funkcije, možete poslati akcije i izvršiti asinhrone zahteve.

U ovom primeru, zahtev koji nije trenutni je kreiran pozivom `setTimeout()`. Uobičajeno je poslati akciju pre pokretanja bilo kakvog ne-trenutnog ponašanja kako bi stanje vaše aplikacije znalo da su podaci potrebni (ovo stanje može prikazivati ikonu učitavanja, na primer). Zatim, kada povratite podatke, koristite drugu akciju koja nosi podatke kao opterećenje (payload) zajedno sa informacijom da je akcija završena.

Zapamtite da prosljeđujete `dispatch` kao parametar ovom specifičnom kreatoru akcije. Ovo ćete koristiti za slanje svojih akcija, prosleđujući akciju direktno na dispatch, a middleware obrađuje sve ostalo.

# --instructions--

Napišite oba dispača u kreator akcije `handleAsync()`. Pošaljite `requestingData()` pre `setTimeout()` (kreirani poziv za API). Zatim, nakon prijema podataka (dummy), pošaljite akciju `receivedData()`, prosleđujući ove podatke. Sada znate kako da rukujete asinhronim akcijama u Reduxu. Sve ostalo i dalje radi kao pre.

# --hints--

Kreator akcije `requestingData` treba da vrati nešto istog tipa kao vrednost `REQUESTING_DATA`.

```js
assert(requestingData().type === REQUESTING_DATA);
```

Kreator akcije `receivedData` treba da vrati objekat istog tipa kao vrednost od `RECEIVED_DATA`.

```js
assert(receivedData('data').type === RECEIVED_DATA);
```

`asyncDataReducer` Trebalo bi da bude funkcija.

```js
assert(typeof asyncDataReducer === 'function');
```

Radnja slanja `requestingData` treba da ažurira svojstvo fetchovanja u skladištu `state` na `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(requestingData());
    const reqState = store.getState();
    return initialState.fetching === false && reqState.fetching === true;
  })()
);
```

Slanje `handleAsync` mora prvo poslati akciju zahteva podataka, a zatim poslati akciju sa primljenim podacima nakon kašnjenja.

```js
assert(
  (function () {
    const noWhiteSpace = __helpers.removeWhiteSpace(handleAsync.toString());
    return (
      noWhiteSpace.includes('dispatch(requestingData())') === true &&
      noWhiteSpace.includes('dispatch(receivedData(data))') === true
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    // Dispatch request action here

    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      // Dispatch received data action here

    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```

# --solutions--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    dispatch(requestingData());
    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      dispatch(receivedData(data));
    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```
