---
id: 5a24c314108439a4d403614f
title: Pošalji događaj akcije
challengeType: 6
forumTopicId: 301442
dashedName: dispatch-an-action-event
---

# --description--

Put za `dispatch` je ono što koristite za dispatch-ovanje akcija u Redux store. Pozivanjem `store.dispatch()` i prosleđivanjem vrednosti koja se vraća od kreatora akcija, dispatch-uje se akcija u store.

Zapamtite da kreatori akcija vraćaju objekat sa svojstvom tip koje određuje vrstu izvršene akcije. Zatim ta funkcija šalje objekat akcije u Redux store. Prema primeru prethodnog izazova, sledeće linije su ispravne, i sve šalju akciju tipa `LOGIN`:

```js
store.dispatch(actionCreator());
store.dispatch({ type: 'LOGIN' });
```

# --instructions--

Redux store u editoru koda ima definisano stanje koje je objekat sa svojstvom `login`, a trenutno je postavljeno na `false`. Takođe postoji kreator akcija (action creator) pod nazivom `loginAction()` koji vraća akciju tipa `LOGIN`. Pošalj (Dispatch) akciju `LOGIN` u Redux store pozivom metode `dispatch`, i prosledite akciju kreiranu od strane `loginAction()`.

# --hints--

Kuita kitendakazi `loginAction` kunapaswa kurudisha kitu chenye sifa ya `type` iliyowekwa kuwa mfuatano wa herufi `LOGIN`.

```js
assert(loginAction().type === 'LOGIN');
```

Pohrana mora biti inicijalizovana objektom sa svojstvom `login` postavljeno na `false`.

```js
assert(store.getState().login === false);
```

Njia ya `store.dispatch()` inapaswa kutumika kutuma kitendo cha aina `LOGIN`.

```js
assert(
    (function () {
      let noWhiteSpace = code.replace(/\s/g, '');
      return (
        noWhiteSpace.includes('store.dispatch(loginAction())') ||
        noWhiteSpace.includes("store.dispatch({type: 'LOGIN'})") === true
      );
    })()
);
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = {login: false}) => state
);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};

// Dispatch the action here:
```

# --solutions--

```js
const store = Redux.createStore(
  (state = {login: false}) => state
);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};

store.dispatch(loginAction());
```
