---
id: 5a24c314108439a4d403614e
title: Prevod funkcionalne akcije
challengeType: 6
forumTopicId: 301441
dashedName: define-an-action-creator
---

# --description--

Nakon kreiranja akcije, sledeći korak je slanje te akcije u Redux store kako bi mogao ažurirati svoje stanje (*state*). U Redux-u, definišete *action creator*-e da biste ovo postigli. *Action creator* je funkcija od JavaScript koja vraća akciju. Drugim rečima, *action creator*-i kreiraju objekte koji predstavljaju akcione događaje.

# --instructions--

Funkcija nazvana `actionCreator()` koja vraća objekat za `action` kada se pozove.

# --hints--

Funkcija `actionCreator` mora postojati.

```js
assert(typeof actionCreator === 'function');
```

Funkcija `actionCreator` treba da vrati objekat od `action`.

```js
assert(typeof action === 'object');
```

Vraćeni objekat za `action` mora imati atribut ključa `type` sa vrednošću `LOGIN`.

```js
assert(action.type === 'LOGIN');
```

# --seed--

## --seed-contents--

```js
const action = {
  type: 'LOGIN'
}
// Define an action creator here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
const actionCreator = () => {
  return action;
};
```
