---
id: 5a24c314108439a4d4036155
title: Pošalji podatke o akciji u skladište
challengeType: 6
forumTopicId: 301448
dashedName: send-action-data-to-the-store
---

# --description--

Sada ste naučili kako da šaljete akcije u Redux store, ali do sada ove akcije nisu uključivale nikakve informacije osim `type`. Takođe možete poslati specifične podatke zajedno sa svojim akcijama. Zapravo, ovo je vrlo uobičajena stvar jer akcije obično nastaju usled interakcije korisnika i dolaze sa određenim podacima. Redux store često mora da zna za ove podatke.

# --instructions--

Postoji osnovni konstruktor za `notesReducer()` i `addNoteText()` definisan u editoru koda. Dopunite telo funkcije/metode za `addNoteText()` da bi vratila objekat `action`. Objekat mora da sadrži svojstvo `type` sa vrednošću `ADD_NOTE`, kao i svojstvo `text` postavljeno za podatke `note` koji se prosleđuju konstruktoru. Kada pozovete konstruktor, prosleđujete specifične informacije o napomeni koje možete dobiti za taj objekat.

Sledeći korak je da završiš pisanje deklaracije za `switch` u `notesReducer()`. Moraš dodati slučaj koji obrađuje akcije/operacije za `addNoteText()`. Ovaj slučaj treba da se inicijalizuje svaki put kada postoji akcija tipa `ADD_NOTE` i mora vratiti atribut/svojstvo `text` na `action` koji dolazi kao novi `state`.

Akcija se šalje ispod koda. Kada završite, pokrenite kod i pogledajte konzolu. Ovo je potrebno za slanje specifičnih podataka o akciji u skladište i korišćenje prilikom ažuriranja `state` skladišta.

# --hints--

Funkcija za kreiranje/izvršavanje `addNoteText` treba da vrati objekt sa ključevima `type` i `text`.

```js
assert(
  (function () {
    const addNoteFn = addNoteText('__TEST__NOTE');
    return addNoteFn.type === ADD_NOTE && addNoteFn.text === '__TEST__NOTE';
  })()
);
```

Slanje akcije tipa `ADD_NOTE` kreatoru akcije `addNoteText` mora ažurirati `state` u niz znakova prosleđen kreatoru akcije.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(addNoteText('__TEST__NOTE'));
    const newState = store.getState();
    return initialState !== newState && newState === '__TEST__NOTE';
  })()
);
```

# --seed--

## --seed-contents--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line

    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello!'));
console.log(store.getState());
```

# --solutions--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line
    case ADD_NOTE:
      return action.text;
    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line
  return {
    type: ADD_NOTE,
    text: note
  }
  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello Redux!'));
console.log(store.getState());
```
