---
id: 5a24c314108439a4d4036159
title: Koristite operaciju raspršivanja na redove podataka
challengeType: 6
forumTopicId: 301452
dashedName: use-the-spread-operator-on-arrays
---

# --description--

Jedno rešenje iz ES6 za pomoć u implementaciji nemenjivosti stanja u Reduxu je spread operater: `...`. Spread operater ima različite upotrebe, jedna od njih je pogodna za prethodni izazov kreiranja novog niza podataka iz postojećeg niza. Ovo je relativno nova sintaksa, ali veoma korišćena. Na primer, ako imate niz podataka `myArray` i pišete:

```js
let newArray = [...myArray];
```

`newArray` sada je kopija `myArray`. Oba niza i dalje postoje zasebno u memoriji. Ako napravite promene kao što su `newArray.push(5)`, `myArray` se ne menja. `...` efikasno *raspršuje* vrednosti koje se nalaze u `myArray` u novi niz. Da biste kopirali niz podataka, ali dodali dodatne vrednosti u novi niz, možete napisati `[...myArray, 'new value']`. Ovo će vratiti novi niz koji uključuje vrednosti iz `myArray` i sekvencu karaktera `new value` kao poslednju vrednost. Spread sintaksa se često može koristiti prilikom kreiranja niza poput ovog, ali je važno razumeti da radi samo površinsku kopiju niza. To znači da pruža samo nepromenljive operacije na nizovima podataka dimenzije jedne.

# --instructions--

Koristi operaciju širenja (spread) za vraćanje novog stanja kada se dodaje to-do.

# --hints--

Redux store treba da postoji i da se inicijalizuje sa stanjem jednakim `["Do not mutate state!"]`.

```js
assert(
  (function () {
    const initialState = store.getState();
    return (
      Array.isArray(initialState) === true &&
      initialState[0] === 'Do not mutate state!'
    );
  })()
);
```

`addToDo` na `immutableReducer` sve bi trebalo da budu funkcionalne.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Slanje akcije tipa `ADD_TO_DO` u Redux skladište bi trebalo da doda element sa `todo` i NE MENJA stanje.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = ['Do not mutate state!', '__TEST__TO__DO__'];
assert(isFrozen)
assert.deepEqual(finalState, expectedState);
```

Spread operacija treba da se koristi za vraćanje novog stanja.

```js
assert(code.includes('...state'));
```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      return [
        ...state,
        action.todo
      ];
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
