---
id: 5a24c314108439a4d4036152
title: Koristi `const` za tipove akcija
challengeType: 6
forumTopicId: 301450
dashedName: use-const-for-action-types
---

# --description--

Uobičajena praksa kada koristite Redux je definisanje tipova akcija kao nepromenljivih konstanti, a zatim referenciranje tih konstanti svuda gde se koriste. Možete poboljšati kod koji sa njima radi tako što ćete pisati tipove akcija kao deklaracije `const`.

# --instructions--

Deklarišite `LOGIN` i `LOGOUT` kao vrednosti za `const` i postavite ih u sekvencu karaktera `'LOGIN'` i `'LOGOUT'`, respektivno. Zatim, izmenite `authReducer()` i naterajte programere da se osvrću na ove konstante umesto vrednosti sekvence karaktera.

**Napomena:** Obično je konvencija pisati konstante u velikim slovima, i ovo je takođe uobičajena praksa u Redux-u.

# --hints--

Pozivanje funkcije `loginUser` trebalo bi da vrati objekat sa svojstvom `type` koji ima vrednost stringa `LOGIN`.

```js
assert(loginUser().type === 'LOGIN');
```

Kuitisha kitendakazi `logoutUser` inapaswa kurudisha kitu chenye sifa ya `type` ikiwa na thamani ya mfuatano wa herufi `LOGOUT`.

```js
assert(logoutUser().type === 'LOGOUT');
```

Skladište mora biti inicijalizovano objektom sa svojstvom `login`, ako ima vrednost `false`.

```js
assert(store.getState().authenticated === false);
```

Kutuma `loginUser` treba da ažurira atribut `login` u stanju skladištenja na `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

Kutuma `logoutUser` kunapaswa kusasisha sifa ya `login` katika hali ya hifadhi kuwa `false`.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

Kitendakazi cha `authReducer` kinapaswa kushughulikia aina nyingi za vitendo kwa kutumia taarifa ya kubadilisha.

```js
assert(
    (function () {
      return (
        typeof authReducer === 'function' &&
        code.toString().includes('switch') &&
        code.toString().includes('case') &&
        code.toString().includes('default')
      );
    })()
  );
```

`LOGIN` na `LOGOUT` zinapaswa kutangazwa kama thamani za `const` na kuwekwa mfuatano wa herufi wa `LOGIN` na `LOGOUT`.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(LOGIN === 'LOGIN' && LOGOUT === 'LOGOUT')
assert(noWhiteSpace.includes('const'))
```

Watengeneza vitendo na reducer wanapaswa kurejea constants za `LOGIN` na `LOGOUT`.

```js
  assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(
        code.toString()
      );
      return (
        noWhiteSpace.includes('caseLOGIN:') &&
        noWhiteSpace.includes('caseLOGOUT:') &&
        noWhiteSpace.includes('type:LOGIN') &&
        noWhiteSpace.includes('type:LOGOUT')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js


const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {
    case 'LOGIN': 
      return {
        authenticated: true
      }
    case 'LOGOUT': 
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case LOGIN:
      return {
        authenticated: true
      }

    case LOGOUT:
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: LOGIN
  }
};

const logoutUser = () => {
  return {
    type: LOGOUT
  }
};
```
