---
id: 5a24c314108439a4d4036153
title: Registrujte listener za repozitorijum
challengeType: 6
forumTopicId: 301446
dashedName: register-a-store-listener
---

# --description--

Drugi način na koji možete pristupiti objektu Redux `store` je `store.subscribe()`. Ovo vam omogućava da se pretplatite na magacine, koji se pozivaju svaki put kada je akcija poslata protiv magacina. Jednostavna upotreba ovog načina je pretplata funkcije na vaš magacin koja snima poruku svaki put kada se akcija primi i magacin ažurira.

# --instructions--

Napišite trigger funkciju koja povećava vrednost globalnog parametra `count` svaki put kada baza podataka primi akciju, i prođite ovu funkciju kroz rutu `store.subscribe()`. Videćete da `store.dispatch()` je pozvana tri puta zaredom, svaki put kada direktno prolazi objekat akcije. Pogledajte izlaz konzole između slanja akcija da biste videli kako se ažuriranja dešavaju.

# --hints--

Kutuma kitendo cha `ADD` kwenye hifadhi kunapaswa kuongeza hali kwa `1`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch({ type: 'ADD' });
    const newState = store.getState();
    return newState === initialState + 1;
  })()
);
```

Kufaa kuwe na kitendakazi cha msikilizaji kilichojiandikisha kwenye hifadhi kwa kutumia `store.subscribe`.

```js
assert.match(code, /store\s*\.\s*subscribe\(/gm);
```

`store.subscribe` inapaswa kupokea kitendakazi.

```js
assert.match(code, /(\s*function\s*)|(\s*\(\s*\)\s*=>)/gm);
```

Funkcija koja je prosleđena za `store.subscribe` ne bi trebalo da se poziva.

```js
assert.notMatch(code, /store\.subscribe\(.+\(\)\)/);
```

Wito wa kurudisha kwa `store.subscribe` pia unapaswa kuongeza kigezo cha kimataifa `count` wakati hifadhi inaposasishwa.

```js
assert.strictEqual(store.getState(), count);
```

# --seed--

## --seed-contents--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);

// Global count variable:
let count = 0;

// Change code below this line

// Change code above this line

store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
```

# --solutions--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);
 let count = 0;
// Change code below this line

store.subscribe( () =>
 {
 count++;
 }
);

// Change code above this line

store.dispatch({type: ADD});
store.dispatch({type: ADD});
store.dispatch({type: ADD});
```
