---
id: 5a24c314108439a4d4036151
title: Koristi informaciju za preusmeravanje da obrađuje više akcija
challengeType: 6
forumTopicId: 301449
dashedName: use-a-switch-statement-to-handle-multiple-actions
---

# --description--

Možete objasniti Redux skladište kako da rukuje više tipova akcija. Recimo da upravljate autentifikacijom korisnika u vašem Redux skladištu. Želite da imate predstavljanje stanja kada se korisnici prijave i kada se korisnici odjave. Predstavljate ovo jednim objektom stanja sa svojstvima `authenticated`. Takođe vam trebaju *action creators* koji kreiraju akcije koje odgovaraju prilogovanju i odjavljivanju korisnika, zajedno sa samim action objekatima.

# --instructions--

Mhariri wa msimbo una hifadhi, vitendo, na watengeneza vitendo vimeandaliwa kwa ajili yako. Jaza kitendakazi cha `reducer` kushughulikia vitendo vingi vya uthibitisho wa utambulisho. Tumia taarifa ya kubadilisha ya JavaScript `switch` katika `reducer` kujibu matukio tofauti ya vitendo. Hii ni mtindo wa kawaida katika kuandika reducers za Redux. Taarifa ya kubadilisha inapaswa kubadilisha juu ya `action.type` na kurudisha hali inayofaa ya uthibitisho wa utambulisho.

**Kumbuka:** Kwa sasa, usijali kuhusu kutokubadilika kwa hali, kwa sababu ni ndogo na rahisi katika mfano huu. Kwa kila kitendo, unaweza kurudisha kitu kipya — kwa mfano, `{authenticated: true}`. Pia, usisahau kuandika kesi ya `default` katika taarifa yako ya kubadilisha ambayo inarudisha `state` ya sasa. Hii ni muhimu kwa sababu mara programu yako itakapokuwa na reducers nyingi, zote zinaendeshwa kila wakati kitendo kinapotumwa, hata wakati kitendo hakihusiani na reducer hiyo. Katika hali kama hiyo, unataka kuhakikisha unarudisha hali ya sasa `state`.

# --hints--

Pozivanje funkcije `loginUser` treba da vrati objekat sa svojstvom tipa ako ima niz znakova `LOGIN`.

```js
assert(loginUser().type === 'LOGIN');
```

Kupiga wito kitendakazi `logoutUser` kunapaswa kurudisha kitu chenye sifa ya type ikiwa na mfuatano wa herufi `LOGOUT`.

```js
assert(logoutUser().type === 'LOGOUT');
```

Skladište mora biti inicijalizovano sa objektom koji ima atribut `authenticated` sa vrednošću `false`.

```js
assert(store.getState().authenticated === false);
```

Pošiljanje `loginUser` mora ažurirati atribut `authenticated` u stanju pohrane na `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

Kutuma `logoutUser` kunapaswa kusasisha sifa `authenticated` katika hali ya hifadhi kuwa `false`.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

Funkcija `authReducer` trebalo bi da obrađuje različite vrste operacija za ažuriranje informacija `switch`.

```js
assert(
    code.toString().includes('switch') &&
      code.toString().includes('case') &&
      code.toString().includes('default')
  );
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case 'LOGIN':
      return {
        authenticated: true
      }

    case 'LOGOUT':
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```
