---
id: 5a24c314108439a4d403614d
title: Tangaza tukio la Redux
challengeType: 6
forumTopicId: 301440
dashedName: define-a-redux-action
---

# --description--

Pošto je Redux sistem za upravljanje stanjem, ažuriranje stanja je jedna od njegovih glavnih funkcija. U Reduxu, sve ažuracije stanja započinju slanjem događaja. Događaj je samo objekat tipa JavaScript koji sadrži informacije o dogodnutom događaju. Redux store prima ove objekte događaja, a zatim ažurira svoje stanje odgovarajuće. Ponekad Redux događaj nosi i neke podatke. Na primer, događaj nosi korisničko ime nakon što se korisnik prijavi. Iako su podaci opcionali, događaji moraju imati atribut `type` koji određuje vrstu dogodka.

Zamislite Redux akcije kao kurire koji prenose informacije o događajima koji se dešavaju u vašoj aplikaciji do Redux store-a. Taj store zatim izvršava posao ažuriranja stanja na osnovu dogodljenog događaja.

# --instructions--

Kuandika tukio la Redux ni rahisi kama kutangaza kitu chenye sifa ya aina. Tangaza kitu `action` na kipa sifa `type` iliyowekwa kwa mfuatano wa herufi `'LOGIN'`.

# --hints--

Element za `action` mora da postoji.

```js
assert(
  (function () {
    return typeof action === 'object';
  })()
);
```

Predmet `action` mora imati svojstvo ključeva `type` sa vrednošću `LOGIN`.

```js
assert(
  (function () {
    return action.type === 'LOGIN';
  })()
);
```

# --seed--

## --seed-contents--

```js
// Define an action here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
```
