---
id: 5a24c314108439a4d403615b
title: Kopiraj objekat sa Object.assign
challengeType: 6
forumTopicId: 301437
dashedName: copy-an-object-with-object-assign
---

# --description--

Nedavni izazovi su bili povezani sa redovima podataka, ali postoje načini za pomoć u primeni nemene promenljive čak i kada je stanje `object`. Važan alat za rad sa objektima je alat `Object.assign()`. `Object.assign()` uzima ciljni objekat i izvorne objekte, pa spaja svojstva iz izvornih objekata na ciljni objekat. Bilo koja poklapanja svojstva prepisuju se svojstvima koji postoje u izvornim objektima. Ova funkcija se često koristi za kreiranje dubokih kopija objekata, prosleđivanjem praznog objekta kao prvog argumenta, praćenog objektom (izvorom) koji želite da kopirate. Evo primera:

```js
const newObject = Object.assign({}, obj1, obj2);
```

Ovo kreira `newObject` kao novi `object`, koji ima funkcije trenutno dostupne u `obj1` i `obj2`.

# --instructions--

Redux stanje i akcije su modifikovane da obrađuju `object` za `state`. Izmijenite kod kako biste vratili novu vrijednost od `state` za akcije tipa `ONLINE`, koje postavljaju svojstvo `status` na niz znakova `online`. Pokušajte da koristite `Object.assign()` da završite izazov.

# --hints--

Redux skladište treba da postoji i bude inicijalizovano stanjem koje odgovara objektu `defaultState` deklarisanom na liniji 1.

```js
const expectedState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};
const initialState = store.getState();
assert.deepEqual(expectedState, initialState);
```

`wakeUp` na `immutableReducer` sve mora biti funkcionalno.

```js
assert(typeof wakeUp === 'function' && typeof immutableReducer === 'function');
```

Slanje akcije tipa `ONLINE` treba da ažurira svojstvo `status` u stanje koje je `online` i NE SME menjati stanje direktno.

```js

const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch({ type: 'ONLINE' });
const finalState = store.getState();
const expectedState = {
  user: 'CamperBot',
  status: 'online',
  friends: '732,982',
  community: 'freeCodeCamp'
};
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

`Object.assign` Trebalo bi da se koristi za vraćanje novog stanja.

```js
assert(code.includes('Object.assign'));
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      return Object.assign({}, state, {
        status: 'online'
      });
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```
