---
id: 5a24c314108439a4d4036154
title: Povežite više reduktora
challengeType: 6
forumTopicId: 301436
dashedName: combine-multiple-reducers
---

# --description--

Kada počinje stanje vaše aplikacije da postaje sve složenije, može biti privlačno razbiti to stanje na mnogo delova. Umesto toga, zapamtite prvi princip Reduxa: celokupno stanje aplikacije se čuva u jednom objektu stanja u *store*-u (skladištu). Zbog toga nudi strukturu reducersa kao rešenje za složeno stanje. Vi definišete više reducersa koji obrađuju različite delove stanja vaše aplikacije, a zatim ih spajate zajedno u jedan glavni reducer. Taj glavni reducer se potom prosleđuje Redux putanji `createStore()`.

Da bismo vam omogućili da kombinujete više reduktora zajedno, Redux pruža metodu `combineReducers()`. Ova metoda prihvata objekat gde definišete ključeve koji će se mapirati na odgovarajuće reducer funkcije. Ime koje date ključu će Redux koristiti kao ime dela stanja.

Obično je dobra praksa da kreirate reducer za svaki deo stanja aplikacije kada je to logički odvojen ili specifičan. Na primer, u aplikaciji za beleške sa autentifikacijom, jedan reducer može da upravlja autentifikacijom dok drugi upravlja tekstom i detaljima koje korisnik unosi. Za takvu aplikaciju, možemo napisati poziv `combineReducers()` kao što sledi:

```js
const rootReducer = Redux.combineReducers({
  auth: authenticationReducer,
  notes: notesReducer
});
```

Sasa, funguo `notes` itakuwa na hali yote inayohusiana na maelezo yetu na inayoshughulikiwa na `notesReducer`. Hivi ndivyo reducers nyingi zinavyoweza kuunganishwa kusimamia hali tata ya programu. Katika mfano huu, hali iliyohifadhiwa katika hifadhi ya Redux itakuwa kitu kimoja chenye sifa za `auth` na `notes`.

# --instructions--

Postoje funkcije za `counterReducer()` i `authReducer()` koje su dostupne u editoru koda, zajedno sa Redux skladištem. Završite pisanje funkcije `rootReducer()` koristeći metod `Redux.combineReducers()`. Postavite `counterReducer` u funkciju nazvanu `count`, i `authReducer` u funkciju nazvanu `auth`.

# --hints--

`counterReducer` bi trebalo da se poveća i smanji `state`.

```js
assert(
  (function () {
    const initialState = store.getState().count;
    store.dispatch({ type: INCREMENT });
    store.dispatch({ type: INCREMENT });
    const firstState = store.getState().count;
    store.dispatch({ type: DECREMENT });
    const secondState = store.getState().count;
    return firstState === initialState + 2 && secondState === firstState - 1;
  })()
);
```

`authReducer` inapaswa kubadilisha `state` ya `authenticated` kati ya `true` na `false`.

```js
assert(
  (function () {
    store.dispatch({ type: LOGIN });
    const loggedIn = store.getState().auth.authenticated;
    store.dispatch({ type: LOGOUT });
    const loggedOut = store.getState().auth.authenticated;
    return loggedIn === true && loggedOut === false;
  })()
);
```

Skladište za `state` bi trebalo da ima dva ključa: `count`, koji čuva broj, i `auth`, koji čuva objekat. Objekat za `auth` bi trebalo da ima svojstvo `authenticated`, koje čuva vrednost istina ili laž.

```js
assert(
  (function () {
    const state = store.getState();
    return (
      typeof state.auth === 'object' &&
      typeof state.auth.authenticated === 'boolean' &&
      typeof state.count === 'number'
    );
  })()
);
```

`rootReducer` trebalo bi da bude funkcija koja povezuje `counterReducer` i `authReducer`.

```js
assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(code);
      return (
        typeof rootReducer === 'function' &&
        noWhiteSpace.includes('Redux.combineReducers')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = // Define the root reducer here

const store = Redux.createStore(rootReducer);
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = Redux.combineReducers({
  count: counterReducer,
  auth: authReducer
});

const store = Redux.createStore(rootReducer);
```
