---
id: 6823d9ac8bdc3853df65a1ff
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Postoji šest elemenata naslova odeljka u HTML: `h1` do `h6`. Koriste se da pokažu važnost odeljka na vašoj web stranici, gde je `h1` najvažnije i `h6` malo manje važno.

Ispod vašeg elementa ``p``, dodajte zaglavlje sekcije ``h3`` sa tekstom:

```md
Introduction to HTML
```

# --hints--

Tvoj element `h3` bi trebalo da ima otvarajući tag od `<h3>`.

```js
assert.exists(document.querySelector("h3"));
```

Vaš element za `h3` mora imati zatvarajući tag od `</h3>`.

```js
assert.match(code, /<\/h3\s*\>/);
```

Tekst za element `h3` treba da bude `Introduction to HTML`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
