---
id: 6823ac607bfdbc46331b2559
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML znači HyperText Markup Language. To je kod koji opisuje strukturu i sadržaj web stranice. Ovo je vaš editor koda, gde ćete pisati HTML.

Pronađite liniju 1 u editoru i napišite ovaj tekst:

```md
Welcome to freeCodeCamp
```

U slučaju problema, kliknite na dugme "Proverite vaš kod" kako biste videli da li je tačan.

# --hints--

Trebalo bi ti da imaš tekst `Welcome to freeCodeCamp` u svom editoru. Proveri pažljivo pravopis.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
