---
id: 6823d6244511f252c8300eed
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Kada trebate dodati paragraf na veb stranicu, možete koristiti element `p`, kao što sledi:

```html
<p>This is a paragraph element.</p>
```

Zameni `Learn the skills to become a full-stack developer` u paragrafski element.

# --hints--

Tvoj element `p` mora imati otvorilsku oznaku. Otvorilske oznake imaju ovu sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector("p"));
```

Vaš element `p` bi trebalo da ima zatvarajući tag. Zatvarni tagovi imaju ovu sintaksu: `</elementName>`.

```js
assert.match(code, /<\/p\s*\>/);
```

Sadržaj vašeg elementa za `p` treba da bude `Learn the skills to become a full-stack developer`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer\s*\<\/p\s*\>/i);
```

Vaš element za `p` mora biti ispod vašeg elementa za `h2`.

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full-stack developer
--fcc-editable-region--
```
