---
id: 6823e637c1c0ed56f781b4fc
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Konačno, JavaScript čini vašu web stranicu interaktivnom — omogućava vam da kažete stranici šta da radi kada neko klikne na dugme, pošalje formular ili mnogo drugih stvari.

Za poslednji korak radionice, dodaj još jedan element za `h3` i `p` na stranicu koja opisuje JavaScript.

Prvo, dodaj element `h3` sa tekstom:

```md
Introduction to JavaScript
```

Zatim, ispod elementa sa `h3`, dodajte element sa `p` koji sadrži sledeći tekst:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

Treba da dodate treći element od `h3` na stranicu.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

Tekst vašeg elementa za `h3` treba da bude `Introduction to JavaScript`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

Moraš dodati četvrti element za `p` na stranicu.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Vaše polje za `p` trebalo bi imati tekst od `JavaScript adds interactivity to a webpage`

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
