---
id: 6823e036ea4b71553558c01b
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Napomena da svako zaglavlje sekcije izgleda malo drugačije u sažetku. Isto važi i za pasuse.

Sledeće, dodajte još jednu komponentu `p` koja prikazuje sledeći tekst:

```md
HTML represents the content and structure of a webpage
```

# --hints--

Trebalo bi da imaš drugi atribut od `p`.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Vaš element ``p`` mora imati tekst od ``HTML represents the content and structure of a webpage``. Možete kopirati tekst iz uputstava.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
