---
id: 682ba2318000b62f179bdf04
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML sastoji se od elemenata. Prvi ćete koristiti element `h1`:

```html
<h1>Welcome to freeCodeCamp</h1>
```

Počinje sa otvarajućom oznakom (`<h1>`), završava se sa zatvarajućom oznakom (`</h1>`), i sadrži tekst koji će biti prikazan između ovih oznaka.

Promenite vaš tekst od `Welcome to freeCodeCamp` u element za `h1` dodavanjem otvarajućeg taga ispred njega i zatvarajućeg taga iza njega.

# --hints--

Vaš element `h1` bi trebalo da ima oznaku za otvaranje `<h1>`.

```js
assert.exists(document.querySelector("h1"));
```

Tvoja komponenta `h1` mora imati zatvarajući tag `</h1>`.

```js
assert.match(code, /<\/h1\s*\>/);
```

Vaš element od `h1` mora da izgleda ovako: `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
