---
id: 682cd206883fc7b25eb539c5
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Imajte na umu da se HTML koji unesete u editor prikazuje u prozoru za pregled. U ovoj radionici ćete deo po deo praviti HTML za stranicu nastavnog plana.

Ispod vašeg elementa `h1`, napišite sledeće na prazan red:

```md
Full-Stack Curriculum
```

# --hints--

Trebalo bi ti da imaš tekst `Full-Stack Curriculum` u svom editoru. Pažljivo proveri pravopis.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
