---
id: 683921c4769fd23dbadec2fe
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Kada HTML definiše strukturu i sadržaj web stranice, CSS se koristi za dodavanje stila — stvari kao što su boje, fontovi, razmak i raspored.

Ispod `h3` koji si dodao tamo, dodaj još jedan element paragrafa sa tekstom:

```md
CSS is used to style a webpage
```

# --hints--

Trebalo bi da dodate drugi element `h3` na stranicu.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

Tekst vašeg elementa za `h3` bi trebalo da bude `Introduction to CSS`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

Treba da dodaš treći element od `p` na stranici.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Tvoje polje za `p` bi trebalo da sadrži tekst `CSS is used to style a webpage`. Možete kopirati tekst iz uputstava.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--
<h3>Introduction to CSS</h3>

--fcc-editable-region--
```
