---
id: 6564642ce403fe29cc23fee7
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Modul Python je fajl sa kodom kreiran za izvršavanje specifične funkcije. Standardna biblioteka Python ima mnogo modula koje možete import i koristiti u vašem kodu. Ovo možete postići korišćenjem izraza `import` praćenog imenom modula.

Pokrenite ovaj projekat korišćenjem-import modula za `string`.

# --hints--

Unapaswa ku-import moduli ya `string`.

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
