---
id: 6564683821b2ee3174e7250e
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Možete dobiti oglašene usluge unutar ugrađenog modula putem notacije tačke. Notacija tačke funkcioniše dodavanjem tačke praćene imenom usluge na ime modula. Na primer, ovako dobijate konstantu za `ascii_lowercase`:

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

U ovom projektu, koristićeš drugačiju konstantu iz modula `string`.

Objavi novi parametar nazvan `letters` i postavi vrednost `string.ascii_letters` za ovaj parametar.

# --hints--

Trebalo bi da objavite parametar koji se zove `letters`.

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

Trebalo bi da postavite vrednost `string.ascii_letters` za vaš parametar `letters`.

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
