---
id: 6564dd65c3c2fa873a83d213
title: Korak 47
challengeType: 20
dashedName: step-47
---

# --description--

U klasi objekata, možete kombinovati različite standarde tako što ćete napisati jedan standard iza drugog unutar uglastih zagrada (bez bilo kakvih dodatnih slova). Na primer: `[a-d3-6]` je kombinacija od `[a-d]` i `[3-6]`.

Sada, promenite poslednji regularni izraz kako bi odgovarao bilo kom znaku koji nije slovo ni broj. Spojite nizove `a-z`, `A-Z`, i `0-9` u jednu klasu objekata i dodajte `^` kao prvi znak da ograniči taj izraz.

# --hints--

Morate dovršiti svoj uobičajeni četvorostavni unos kako bi odgovarao bilo kom nealfanumeričkom karakteru.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*r("|')\\d\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*r("|')\[a-z\]\2\s*\)\s*,\s*\(\s*uppercase\s*,\s*r("|')\[A-Z\]\3\s*\)\s*,\s*\(\s*special_chars\s*,\s*r("|')\[\^a-zA-Z0-9\]\4\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = r'\.'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
