---
id: 6564883669b5af4b69f794cc
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Vaš kriterijum za `all_characters` je sekvenca karaktera sastavljena od svih malih i velikih slova, svih deset cifara i nekoliko specijalnih znakova.

Ispred toga, dodajte komentar koji kaže `Combine all characters`.

# --hints--

Trebalo bi da dodate taj komentar malo iznad vašeg parametra `all_characters`.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
