---
id: 65648f4b2281ba50051ae39c
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Sledeće, napiši petlju za `for` koristeći `i` kao parametar petlje. Koristi funkciju `range()` da ponovi do vrednosti `length`.

Unutar ciklusa, koristi operator sabiranja da doda slučajni karakter od `all_characters` na trenutnu vrednost `password`. Koristi funkciju `choice()` iz modula `secrets` za to.

# --hints--

Trebalo bi da napišete krug koji ponavlja nad `range(length)`, na mestu `for`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;    
    assert(function_body.match(/for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

Trebalo bi da koristiš operaciju `+=` za dodavanje nasumičnog slova iz `all_characters` trenutnoj vrednosti `password`.

```js
({ test: () =>
  {
    const generate_pwd = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_pwd;    
    assert(function_body.match(/^(\s*)for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:\s*^\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--    
    password = ''
    # Generate password
    
--fcc-editable-region--
```
