---
id: 6564c67db26c417561ab510d
title: Korak 31
challengeType: 20
dashedName: step-31
---

# --description--

Kao što vidite iz rezultata, vaš običan tekst sada odgovara prvom unutar sekvence karaktera `l`.

U vašem šablonu, možete dodati modifikator nakon slova kako biste specificirali koliko puta se to slovo treba ponoviti. Na primer, modifikator `+` znači da se treba ponoviti jednom ili više puta.

Dodajte widget za `+` na svoj dizajn.

# --hints--

Treba da promeniš svoju varijablu `pattern` u `re.compile('l+')`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
