---
id: 65649a631440b757f9f7b428
title: Korak 28
challengeType: 20
dashedName: step-28
---

# --description--

Regularni izraz, ili regex, je obrazac koji se koristi za poređenje specifične kombinacije znakova unutar niza karaktera. To je validna alternativa za uslove `if`/`else` kada vam je potrebno da poređujete ili tražite obrasce unutar niza karaktera radi validacije, promene karaktera i sl.

Funkcija ``compile()`` iz modula ``re`` prikuplja niz znakova koji je prosleđen kao argument, što predstavlja običnu literalnu vrednost koja se može koristiti i drugim metodama ``re``.

Deklarišite novi parametar za `pattern` i postavite vrednost `re.compile('i')` za ovaj parametar.

# --hints--

Morate deklarisati parametar `pattern`.

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

Trebalo bi da staviš `re.compile('i')` u svoj parametar `pattern`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
