---
id: 65649122c7f77f519aaf0975
title: Korak 20
challengeType: 20
dashedName: step-20
---

# --description--

Konačno, deklarišite parametar ``new_password`` i postavite mu vrednost na rezultat poziva ``generate_password``. Proširite ``8`` kao argument u vaš poziv ``generate_password``.

# --hints--

Treba ti da pozoveš `generate_password` kada prolaziš kroz `8` kao argument.

```js
({ test: () => assert.match(code, /generate_password\s*\(\s*8\s*\)/) })
```

Trebalo bi da postaviš vrednost `generate_password(8)` za parametar `new_password`.

```js
({
  test: () => {
    assert(__userGlobals.has("new_password"));
    assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*\)/m);
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--

--fcc-editable-region--
```
