---
id: 6564d75a923d21815caaa445
title: Korak 43
challengeType: 20
dashedName: step-43
---

# --description--

Ako je potrebno da se uporedi karakter koji ima posebno značenje u obrascu, kao što su `.` ili `+`, možete ga "escape" postavljanjem povratnika (backslash-a), `\`. Na primer, ovo dovodi do upoređivanja doslovnog karaktera za dodavanje: `\+`.

Promeni `pattern` ili odgovara na jednu stvarnu tačku.

# --hints--

Trebalo bi da promeniš svoju razliku od `pattern` u `\.`.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '.+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
