---
id: 6564cbb6311a40783b3f5de6
title: Korak 36
challengeType: 20
dashedName: step-36
---

# --description--

Klasa znakova vam takođe omogućava prikaz raspona odgovarajućih znakova. Morate navesti početni znak i krajnji znak koji su označeni znakiem križa, `-`. Znakovi prate redosled Unicode-a.

Promenite svoju varijaciju od `pattern` da odgovara bilo kom slovu `t` koje je prethodevano malim slovom u varijaciji `quote`. Koristite raspon slova od `a` do `z` za to.

# --hints--

Trebalo bi da promeniš svoju varijaciju od `pattern` kako bi odgovarala bilo kom slovu `t` u `quote` koje je prethodno malim slovom, koristeći klasu objekata `[a-z]`.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
