---
id: 656469e62377ab34a5bcb8d1
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Objavi dva nova kriterijuma koja se zovu `digits` i `symbols` na vrednost `string.digits` na `string.punctuation` za njih, u tom redosledu.

# --hints--

Trebalo bi da objavite parametar nazvan `digits`.

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

Trebalo bi da postaviš vrednost `string.digits` za svoj parametar `digits`.

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

Trebalo bi da objaviš parametar nazvan `symbols`.

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

Trebalo bi postaviti vrednost `string.punctuation` za vaš parametar `symbols`.

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
