---
id: 6564b8c9349bd76dc037967b
title: Korak 29
challengeType: 20
dashedName: step-29
---

# --description--

Funkcija `search()` iz modula `re` analizira niz karaktera poslat kao argument kako bi pronašla prvo mesto gde se obrazac redovne izraznosti poklapa sa tim nizom.

Deklariši parametar nazvan `quote` i postavi vrednost na niz znakova `'Not all those who wander are lost.'` u ovaj parametar. Zatim, ispiši rezultat od `pattern.search(quote)`.

# --hints--

Trebalo bi da imate kriterijum `quote`.

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

Treba da postaviš vrednost sekvence slova dobijene u tvoj novi parametar `quote`.

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

Trebalo bi da odštampate `pattern.search(quote)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
