---
id: 6564d97a5ef99783216a6229
title: Korak 44
challengeType: 20
dashedName: step-44
---

# --description--

Python pruža poseban tip sekvence karaktera nazvan sirova žica (*raw string*). Sirova sekvenca karaktera počinje sa prefiksom `r`. Glavna razlika u odnosu na običnu sekvencu karaktera leži u načinu na koji se obrađuje znak kosog ukočenja (backslash): u sirovoj sekvenci, kosog ukočenja se prepoznaje kao literalan karakter umesto da bude znak za bježanje. Kada pišete običan izraz/kod, korišćenje sirove sekvence karaktera je dobra praksa, jer često koristi mnogo karaktera iz `\`.

Promeni svoju sekvencu slova od `pattern` u sirovu sekvencu slova počevši sa prefiksom `r`.

# --hints--

Trebalo bi da promeniš svoj parametar `pattern` u `r'\.'`.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
