---
id: 65646ffeaed2d238c562a014
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Uobičajena praksa je stavljanje deklaracija `import` iznad vašeg koda. Štaviše, ako postoji više deklaracija `import`, organizujte ih po abecednom redosledu radi poboljšanja čitljivosti.

Iznad vašeg koda, import modul za `random`.

# --hints--

Trebalo bi da uvezete modul `random`.

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
