---
id: 656475bbf1c2573de1d2c69c
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Svaki put kada se kod izvršava, trebalo bi da vidite neočekivano slovo iz sekvence karaktera od `all_characters`. Ovo je upravo ono što želite postići kreiranjem nepredvidljive lozinke.

Međutim, algoritam koji zavisi od ``random`` čini nasumične brojeve nepredvidivim. Stoga, iako je modul ``random`` pogodan za mnoge obične aplikacije, ne može se koristiti za svrhe šifrovanja, zbog njegove predvidive prirode.

Umesto uvoza `random`, import modula za `secrets`. Zatim promenite poziv `print()` ili koristite `secrets.choice(all_characters)`.

# --hints--

Trebalo bi da uvezete modul `secrets` umesto modula `random`.

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

Trebalo bi da promeniš svoj poziv `print()` ili odštampaš `secrets.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
