---
id: 656470d517833a39bb8b5608
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Modul `random` ima generator pseudo-nasumičnih brojeva. Većina njegovih funkcija zavisi od funkcije `random()`, koja vraća decimalni broj u opsegu između `0.0` (uključujući) i `1.0` (isključujući).

Izvrši poziv funkcije `random()` iz modula `random` i ispiši rezultat.

# --hints--

Trebalo bi da ispišete rezultate poziva `random.random()`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
