---
id: 65648e4c5b316c4ec5e4fddc
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Na samom dnu vaše funkcije, deklarujte izmene za `password` i postavite vrednost praznog niza karaktera za ove izmene.

# --hints--

Trebalo bi da imaš promene `password` unutar svoje funkcije sa vrednošću praznog niza karaktera. Obrati pažnju na usklađenost linija.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/password\s*=\s*("|')\1/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--

--fcc-editable-region--
```
