---
id: 656485a8a3496d4a36b1496a
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Ovo su tri različitosti koje čine slova koja mogu biti izabrana prilikom kreiranja lozinke.

Blizu njih, dodaj komentar koji kaže `Define the possible characters for the password`.

# --hints--

Treba da dodate te komentare o maloj razlici od `letters`.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
