---
id: 656477845006313fbfea0ad1
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Iako je uticaj može izgledati slično kao `random.choice()`, `secrets`, ono vam garantuje najsigurniji izvor slučajnosti koji vaš operativni sistem može da pruži.

Sada, obriši svoje dva poziva za `print()`.

# --hints--

Treba da obrišeš svoj poziv za `print()` dva puta.

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
