---
id: 656474314bf37d3c83a85143
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Funkcija `choice()` iz modula `random` koja uzima sekvencu i vraća bilo koji element te sekvence nasumično.

Promeni svoj poziv `print()` koristeći funkciju `choice()` preko `all_characters` kao argument.

# --hints--

Treba da promeniš svoj trenutni poziv `print()` ili da odštampaš `random.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
