---
id: 6565012a2564509d40a90048
title: Korak 56
challengeType: 20
dashedName: step-56
---

# --description--

Uvukušen sa brojem elemenata u listi koji ga vraća funkcija `findall()`.

Proverite vaš postojeći poziv za `findall()` u odnosu na funkciju `len()`.

# --hints--

Trebalo bi da pozoveš funkciju `len()` sa `re.findall(pattern, password)` kao argumentom.

```js
({ test: () => assert.match(code, /^(\s*)for.+:\s*^\1\s{4}len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
--fcc-editable-region--         
        # Check constraints
        for constraint, pattern in constraints:
            re.findall(pattern, password)
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
