---
id: 6565c13fdb798865c161d8f8
title: Korak 65
challengeType: 20
dashedName: step-65
---

# --description--

Sada je vreme da testiraš svoju funkciju. Otvori detalje poslednje dve linije u svom kodu i promeni poziv funkcije koji prosleđuje argument 5. Koristi `8` za dužinu i `1` za ostale četiri praga.

# --hints--

Trebalo bi da pozoveš `generate_password` sa navedenim argumentima.

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*,\s*1\s*,\s*1\s*,\s*1\s*,\s*1\s*\)/m) })
```

Trebalo bi da objaviš svoju razliku od `new_password`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*new_password\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
        
        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break

    return password
    
--fcc-editable-region--
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
```
