---
id: 656ef54070c72161e6feb90f
title: Korak 22
challengeType: 20
dashedName: step-22
---

# --description--

Izgleda da je sve u redu, ali bi bilo dobro imati način za proveru da li generisana lozinka prati određene kriterijume. Na primer, minimalni zahtev za specijalne znakove, brojeve ili velika/mala slova. To će biti osigurano uskoro.

Trenutno, ostavi komentare na poslednje dve linije tvog koda.

# --hints--

Trebalo bi da promeniš poslednje dve linije svog koda u komentare.

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
