---
id: 656ef89b4b486b7a16077864
title: Korak 26
challengeType: 20
dashedName: step-26
---

# --description--

Tuple je druga struktura podataka ugrađena u Python. Tuple je veoma sličan listi, ali se kreira korišćenjem običnih zagrada `()`, umesto uglastih zagrada. Takođe, tuple je nepromenljiv (immutable), za razliku od liste.

```py
my_tuple = ('larch', 1, True)
```

Vaš spisak `constraints` će skladištiti tupel. Prvi element svakog tupa biće parametar ograničenja.

Promeni listu raspoloživosti `constraints` dodajući tuple u vaš list. Koristi `nums` kao prvi element i niz praznih karaktera kao drugi element.

# --hints--

Trebalo bi da dodaš tupel u listu `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

Trebalo bi da dodaš `(nums, '')` u listu `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
