---
id: 656f0d6bc60b58b9777a36c9
title: Korak 34
challengeType: 20
dashedName: step-34
---

# --description--

Klasa stavki je prikazana u kvadratnim zagradama i odgovara jednoj karakteristici među onima pomenutim unutar tih zagrada. Na primer, `ma[dnt]` odgovara `mad`, `man`, ili `mat`.

Promeni svoj format ili prati `w` praćen `h` ili `a`.

# --hints--

Vaša struktura bi trebalo da koristi klasu objekata koja odgovara `w`, praćenu sa `h` ili `a`.

```js
({ test: () => assert(runPython(`
    pattern == "w[ha]" or pattern == "w[ah]"    
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
