---
id: 656f04b20397d9a574cc7eb2
title: Korak 32
challengeType: 20
dashedName: step-32
---

# --description--

Možete dobiti iste rezultate bez korišćenja funkcije `compile()`. Promenite vaše izmene od `pattern` u niz stvarnih slova `'l+'`. Zatim, promenite poziv `print()` kako biste ispisali `re.search(pattern, quote)`.

# --hints--

Treba da promeniš svoje izmene od `pattern` u niz stvarnih slova `'l+'`.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "l+") })
```

Trebalo bi da odštampate `re.search(pattern, quote)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*re\.search\s*\(\s*pattern\s*,\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l+')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
