---
id: 6577319039f4f7de9251b822
title: Korak 61
challengeType: 20
dashedName: step-61
---

# --description--

Umjesto korišćenja stružnog kruga i parametra za izračunavanje, možeš dobiti isti rezultat na različite načine, što ćeš implementirati u sljedećim koracima.

`all()` je funkcija napravljena unutar Python koja vraća `True` ako su svi elementi unutar nekog iterabilnog objekta potvrđeni da su `True`. Inače, vraća `False`.

Izmeni svoj postojeći tok `for` sa dve informacije za `if` na jednu informaciju od `if`. Pod uslovom `if`, koristi poziv funkcije `all()` i prosledi prazan niz kao argument funkciji.

# --hints--

Morate da promenite svoj postojeći ciklus od `for` sa dve informacije za `if` u jednu informaciju za `if`.

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

Novi uslov za `if` mora biti `all([])`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

Trebalo bi da imaš `break` unutar tela `if` novog.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
            
        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
