---
id: 6577333feab1e8e927014f03
title: Korak 63
challengeType: 20
dashedName: step-63
---

# --description--

Imati `all([expression for i in iterable])` znači da se kreira nova lista evaluiranjem `expression` za svaki `i` u `iterable`. Nakon što funkcija `all()` izvrši kreiranu listu, ta lista se automatski briše jer više nije potrebna.

Memorija se može uštedeti korišćenjem generatora izraza. Generator izraz prati sintaksu liste komprehensije, ali koristi obične zagrade umesto uglastih zagrada.

Pretvori svoju *list comprehension*-u u literalni izraz uklanjanjem uglastih zagrada.

# --hints--

Treba da promeniš svoju *list comprehension* u string izraz uklanjanjem kvadratnih zagrada.

```js
({ test: () => assert.match(code, /^(\s*)if\s+all\s*\(\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        if all(
            [
                constraint <= len(re.findall(pattern, password))
                for constraint, pattern in constraints
            ]
        ):
--fcc-editable-region--
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
