---
id: 657db4cb77190e33a20e852a
title: Korak 69
challengeType: 20
dashedName: step-69
---

# --description--

Kada kombinujete argumente sa podrazumevanim vrednostima i argumente po ključu, možete proslediti manje argumenata od onih koje funkcija zahteva. Argumenti koji nisu eksplicitno prosleđeni u pozivu funkcije uzimaće svoje zadane vrednosti.

Promeni svoj poziv `generate_password()` ili samo primi `length=8`.

# --hints--

Treba da proslediš samo `length=8` u svoj poziv `generate_password()`.

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*length\s*=\s*8\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length=16, nums=1, special_chars=1, uppercase=1, lowercase=1):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break
    
    return password
    
--fcc-editable-region--
new_password = generate_password(nums=1, length=8, special_chars=1, uppercase=1, lowercase=1)
print(new_password)
--fcc-editable-region--
```
