---
id: 69b279bae0f20eda63163cfe
title: Korak 20
challengeType: 1
dashedName: step-20
---

# --description--

Pre nego što testirate funkciju `mergeSort()`, morate kreirati bazni slučaj koji zaustavlja izvršavanje funkcije kada je dužina `array` manja ili jednaka `1`.

Ovaj bazni slučaj će zaustaviti rekurzivni poziv. Bez njega, operacija spajanja (merge sort) bi nastavila da radi čak i kada je lista sortirana ili nema elemenata u njoj.

Neposredno nakon deklaracije funkcije, kreirajte izjavu `if` koja proverava da li je `array.length` manja ili jednaka 1. Unutar tela izjave `if`, dodajte izjavu `return` za zaustavljanje izvršavanja funkcije.

# --hints--

You should create an `if` statement with the condition `array.length <= 1`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

You should have a `return` statement in the body of the `if` statement.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
