---
id: 69b279bae0f20eda63163cf2
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Zatim, morate kreirati petlju koja se izvršava sve dok postoje elementi u oba `leftPart` i `rightPart`.

Za to, kreirajte petlju `while` sa dve uslova: jedan koji proverava da li je `leftArrayIndex` manji od dužine `leftPart`, i drugi uslov koji proverava da li je `rightArrayIndex` manji od dužine `rightPart`.

# --hints--

You should create a `while` loop that checks if `leftArrayIndex` is less than the length of `leftPart` and if `rightArrayIndex` is less than the length of `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
