---
id: 69b2713ff6c21fed2d3ad88b
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

Merge sort je sortirajući algoritam koji koristi princip razdvajanja i osvajanja za sortiranje kolekcija podataka. To jest, on "razdvaja" kolekciju na manje poddelove, a "osvaja" poddelove tako što ih nezavisno sortira, a zatim spaja sortirane poddelove.

Počnite kreiranjem funkcije nazvane `mergeSort` koja prima parametar `array`.

# --hints--

You should declare a function named `mergeSort` with a single `array` parameter.

```js
assert.isFunction(mergeSort);
const params = __helpers.getFunctionParams(mergeSort.toString());
assert.lengthOf(params, 1);
assert.deepEqual(params[0].name, "array");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
