---
id: 69b279bae0f20eda63163cef
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Sada kada ste podelili listu `array` na dve odvojene liste, nastavićete da delite svaku listu sve dok svaki element ne bude sam u svojoj listi. Lista sa jednim brojem je uvek sortirana.

Da biste to uradili, rekurzivno pozovite `mergeSort` unutar vaše funkcije i prosledite `leftPart` kao argument poziva.

# --hints--

You should call `mergeSort(leftPart)` inside the same `mergeSort` function.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
