---
id: 69b279bae0f20eda63163cf4
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Kada uslov `if` evaluira na `true`, to znači da je element u listi `leftPart` manji od elementa sa kojim se poredi u listi `rightPart`.

U tom slučaju, možete staviti element pronađen na mestu `leftArrayIndex` u `leftPart` unutar sortirane niz.

Unutar bloka `if`, dodelite `leftPart[leftArrayIndex]` na `array[sortedIndex]`.

# --hints--

You should assign `leftPart[leftArrayIndex]` to `array[sortedIndex]` inside the `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
