---
id: 69b278519ef46b2fd8bf1546
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Već imate vrednost `middlePoint`. Koristite je da razdelite niz na dve polovine i sačuvate levi deo spreman za sortiranje.

Koristite sintaksu 'slice' da izvučete levu polovinu od `array` i dodelite je promenljivoj nazvanoj `leftPart`.

# --hints--

You should declare a `leftPart` variable inside the `mergeSort` function.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

You should assign `array.slice(0, middlePoint)` to `leftPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
