---
id: 69b279bae0f20eda63163cf3
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Unutar petlje, počnite da poredite trenutne elemente iz svake polovine kako biste prvo spojili manju vrednost.

Dodajte blok `if` koji proverava da li je trenutni element u `leftPart` manji od trenutnog elementa u `rightPart`.

# --hints--

The `if` condition should check if `leftPart[leftArrayIndex]` is less than `rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
