---
id: 69b279bae0f20eda63163cf9
title: Korak 15
challengeType: 1
dashedName: step-15
---

# --description--

Petlja `while` koju ste kreirali poredi jedan element iz `leftPart` sa drugim u `rightPart`, a zatim dodaje manji element glavnoj `array` listi.

Ona će nastaviti ovu operaciju dok ne ostane elemenata za poređenje. Ali `leftPart` može i dalje imati elemente dok `rightPart` nema, i obrnuto.

Kreirajte još jednu petlju `while` koja se izvršava kada je `leftArrayIndex` manji od dužine `leftPart`. U sledećim koracima, koristićete je za kopiranje preostalih elemenata iz `leftPart` u listu `array`.

# --hints--

You should create a `while` loop with the condition `leftArrayIndex < leftPart.length`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
