---
id: 69b278d82f0e514f32acb724
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Nakon izvlačenja levog dela, i dalje vam je potreban desni deo da biste dovršili korak spajanja.

Koristite sintaksu za sečenje/izdvajanje da izvučete desni deo od `array` i dodelite ga promenljivoj nazvanoj `rightPart`.

# --hints--

You should have a `rightPart` variable inside the `mergeSort` function.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

You should use the slice syntax to extract the right half of the array and assign it to the `rightPart` variable.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
