---
id: 69b2762b17d23fbee40a6c00
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Algoritam merž sort izvodi tri akcije:

- Podeli neuređenu sekvencu stavki na poddelove
- Sortiraj stavke u poddelovima
- Spoji sortirane poddelove

Ovo se dešava rekursivno sve dok se poddelovi ne spoje u kompletnu sortiranu sekvencu. Počnite sa podele sekvence.

Prvo, unutar funkcije `mergeSort`, deklariš varijablu `middlePoint` koristeći `const` i dodeli joj `Math.floor(array.length / 2)` kako bi mogao da podeliš sekvencu na pola pre rekursije.

# --hints--

You should declare a `middlePoint` variable and assign `Math.floor(array.length / 2)` to it.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
