---
id: 69b279bae0f20eda63163cf6
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

U prethodnom koraku, dodelili ste element iz `leftPart` spisku `array` jer je bio manji. Ali to neće uvek biti slučaj. U nekim poređenjima, element sa desne strane može biti manji.

Kreirajte klauzulu `else` koja se izvršava kada element u `leftPart` nije manji od elementa u `rightPart`.

Unutar bloka `else`, dodelite `rightPart[rightArrayIndex]` u `array[sortedIndex]`.

# --hints--

You should create an `else` block that contains `array[sortedIndex] = rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
