---
id: 69b279bae0f20eda63163cf1
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Sada je vreme da sortirate i spojite liste (`leftPart` i `rightPart`) u originalnu `array`.

To možete uraditi poređenjem elemenata na obe liste, a spajanjem manjeg elementa u glavnu listu. Ova poređenja ćete izvršiti za sve indekse u `leftPart` i `rightPart`.

Kreirajte tri varijable: `leftArrayIndex`, `rightArrayIndex` i `sortedIndex`, i postavite njihove vrednosti na `0`. Ove varijable će vam pomoći da pratite svaki indeks tokom procesa sortiranja.

# --hints--

You should declare a `leftArrayIndex` variable initialized to `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

You should declare a `rightArrayIndex` variable initialized to `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

You should declare a `sortedIndex` variable initialized to `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);

--fcc-editable-region--
  
--fcc-editable-region--
}
```
