---
id: 686ccc2c8b967e17ab18d593
title: Implementacija metode biseksije
challengeType: 27
dashedName: implement-the-bisection-method
---

# --description--

Bisekciona metoda, poznata takođe kao binarna pretraga, koristi binarnu pretragu za pronalaženje korena realno vrednostne funkcije. Ona radi smanjivanjem intervala unutar kojeg se nalazi koren, sve dok ne dođe do vrednosti unutar postavljene tolerancije.

Na primer, ako je tolerancija `0.01`, metoda biseksije će nastaviti da deli taj opseg na pola dok razlika između gornje i donje granice nije mala ili je jednaka `0.01`.

U ovoj laboratoriji, implementiraćete funkciju koja koristi bisekcionu metodu za pronalaženje kvadratnog korena broja.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratorijum.

Priče o korisniku:

1. Treba definisati funkciju pod nazivom `square_root_bisection` sa tri parametra:
    - Broj za koji želite da pronađete kvadratni koren.
    - Tolerancija, koja je prihvatljiva granica greške za rezultat. Treba postaviti vrednost tolerancije kao podrazumevanu.
    - Maksimalan broj iteracija za izvršavanje. Treba postaviti broj iteracija kao podrazumevano.

1. Funkcija `square_root_bisection` treba da:
    - Uzrokuje `ValueError` sa porukom `Square root of negative number is not defined in real numbers` ako je prosleđeni broj negativan.
    - Za brojeve `0` i `1`, ispišite poruku: `The square root of [number] is [number]` i vratite sam broj kao kvadratni koren.
    - Za bilo koji drugi pozitivan broj, ispišite približan kvadratni koren sa porukom: `The square root of [square_target] is approximately [root]` i vratite izračunatu vrednost korena.
    - Ako nijedna vrednost ne zadovoljava uslov tolerancije, ispišite poruku o neuspehu: `Failed to converge within [maximum] iterations` i vratite `None`.

**Napomena**: Ne možete import bilo koji modul za ovu laboratoriju.

# --hints--

Ne možete bilo koji modul import.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_imports()) == 0`)) })
```

Trebalo bi da imaš funkciju pod nazivom `square_root_bisection`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("square_root_bisection")`)) })
```

Tvoja funkcija `square_root_bisection` mora imati tri parametra.

```js
({ test: () => runPython(`
    import inspect 
    sig = inspect.signature(square_root_bisection)
    assert len(sig.parameters) == 3
`) })
```

Trebalo bi da postavite podrazumevanu vrednost za toleranciju i maksimalan broj ciklusa.

```js
({ test: () => runPython(`
try:
  import inspect 
  sig = inspect.signature(square_root_bisection)
  assert len(sig.parameters) == 3
  square_root_bisection(4)
except TypeError:
  assert False
`) })
```

Vaša funkcija za `square_root_bisection` treba da izazove `ValueError` sa porukom `Square root of negative number is not defined in real numbers` kada je prosleđeni broj negativan.

```js
({ test: () => runPython(`
try:
  square_root_bisection(-6)
except ValueError as e:
  assert str(e) == "Square root of negative number is not defined in real numbers"
else:
  assert False
`) })
```

`square_root_bisection(0)` treba da vrati `0`.

```js
({ test: () => runPython(`assert square_root_bisection(0) == 0`) })
```

`square_root_bisection(0)` treba da ispiše `The square root of 0 is 0`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(0)
assert "The square root of 0 is 0" in _out
`) })
```

`square_root_bisection(0.001, 1e-7, 50)` trebalo bi da vrati broj između `0.03162267660168379` i `0.031622876601683794`.

```js
({ test: () => runPython(`assert 0.03162267660168379 <= square_root_bisection(0.001, 1e-7, 50) <= 0.031622876601683794`) })
```

`square_root_bisection(0.001, 1e-7, 50)` treba da odštampa `The square root of 0.001 is approximately X`, gde je `X` broj između `0.03162267660168379` i `0.031622876601683794`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.001, 1e-7, 50)

assert 0.03162267660168379 <= _root <= 0.031622876601683794
assert f"The square root of 0.001 is approximately {_root}" in _out
`) })
```

`square_root_bisection(0.25, 1e-7, 50)` treba da vrati broj između `0.4999999` i `0.5000001`.

```js
({ test: () => runPython(`assert 0.4999999 <= square_root_bisection(0.25, 1e-7, 50) <= 0.5000001`) })
```

`square_root_bisection(0.25, 1e-7, 50)` mora da prikaže `The square root of 0.25 is approximately X`, gde je `X` broj između `0.4999999` i `0.5000001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.25, 1e-7, 50)

assert 0.4999999 <= _root <= 0.5000001
assert f"The square root of 0.25 is approximately {_root}" in _out
`) })
```

`square_root_bisection(1)` trebalo bi da vrati `1`.

```js
({ test: () => runPython(`assert square_root_bisection(1) == 1`) })
```

`square_root_bisection(1)` trebalo bi da odštampa `The square root of 1 is 1`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(1)
assert "The square root of 1 is 1" in _out
`) })
```

`square_root_bisection(81, 1e-3, 50)` treba da vrati broj između `8.999` i `9.001`.

```js
({ test: () => runPython(`assert 8.999 <= square_root_bisection(81, 1e-3, 50) <= 9.001`) })
```

`square_root_bisection(81, 1e-3, 50)` mora da isprikaže `The square root of 81 is approximately X`, gde `X` je broj između `8.999` i `9.001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(81, 1e-3, 50)

assert 8.999 <= _root <= 9.001
assert f"The square root of 81 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-3, 100)` mora da vrati broj između `14.999` i `15.001`.

```js
({ test: () => runPython(`assert 14.999 <= square_root_bisection(225, 1e-3, 100) <= 15.001`) })
```

`square_root_bisection(225, 1e-3, 100)` treba da ispiše `The square root of 225 is approximately X`, gde `X` je broj između `14.999` i `15.001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-3, 100)

assert 14.999 <= _root <= 15.001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-5, 100)` treba da vrati broj između `14.99999` i `15.00001`.

```js
({ test: () => runPython(`assert 14.99999 <= square_root_bisection(225, 1e-5, 100) <= 15.00001`) })
```

`square_root_bisection(225, 1e-5, 100)` treba da ispiše `The square root of 225 is approximately X`, gde `X` je broj između `14.99999` i `15.00001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-5, 100)

assert 14.99999 <= _root <= 15.00001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 100)` treba da vrati broj između `14.9999999` i `15.0000001`.

```js
({ test: () => runPython(`assert 14.9999999 <= square_root_bisection(225, 1e-7, 100) <= 15.0000001`) })
```

`square_root_bisection(225, 1e-7, 100)` treba da ispiše `The square root of 225 is approximately X`, gde `X` je broj između `14.9999999` i `15.0000001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-7, 100)

assert 14.9999999 <= _root <= 15.0000001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 10)` trebalo bi da vrati `None`.

```js
({ test: () => runPython(`assert square_root_bisection(225, 1e-7, 10) is None`) })
```

`square_root_bisection(225, 1e-7, 10)` treba da odštampa `Failed to converge within 10 iterations`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(225, 1e-7, 10)
assert "Failed to converge within 10 iterations" in _out
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = square_target if square_target < 1 else 1
        high = 1 if square_target < 1 else square_target
        root = None

        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if high - low <= tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

        if root is None:
            print(f"Failed to converge within {max_iterations} iterations")

        else:   
            print(f'The square root of {square_target} is approximately {root}')

    return root
```
