---
id: 680b3ef395479b0e449ecb6e
title: Implementacija algoritma Selection Sort
challengeType: 27
dashedName: implement-selection-sort-algorithm
---

# --description--

Selection Sort je poznat algoritam za sortiranje koji se uči na mnogim kursevima iz nauke o računaru.

Ovaj algoritam radi tako što stalno traži najmanji element iz neuređenog dela liste i zamjenjuje ga sa prvim neuređenim elementom. Počinje izborom minimalne vrednosti u cijeloj listi i zamjenjivanjem tog elementa sa prvim elementom. Zatim pređe na drugu poziciju, tražeći najmanju vrednost među preostalim neuređenim elementima i zamjenjujući je sa drugim elementom. Ovaj proces se nastavlja, prolazeći kroz listu po jedan element, sve dok cijela lista nije sortirana.

Selection sort ima kvadratnu vremensku složenost (quadratic) u najboljem, prosečnom i najgorem slučaju. Prostorna složenost će biti konstantna `O(1)` jer se sortiranje vrši na mestu i količina memorije koja se koristi je konstantna bez obzira na veličinu liste.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba da definišete funkciju pod nazivom `selection_sort`.
1. Vaša funkcija `selection_sort` treba da ima jedan parametar koji predstavlja listu stavki.
1. Vaša funkcija `selection_sort` treba da uzme listu i sortira stavke na mestu od najmanjeg do najvećeg.
1. Vaša funkcija `selection_sort` treba da modifikuje ulaznu listu na mestu i vrati je nakon sortiranja.
1. Vaša funkcija `selection_sort` treba da prati algoritam selection sort, menjajući najmanji element iz nesortiranog dela liste sa prvim nesortiranim elementom.
1. Vaša funkcija `selection_sort` ne sme da vrši nepotrebne zamene kada je najmanji element već na ispravnoj poziciji.
1. Vaša funkcija `selection_sort` ne sme da koristi ugrađenu metodu iz `sort()` niti funkciju `sorted()`.

# --hints--

Trebalo bi da imaš funkciju nazvanu `selection_sort`.

```js
({test: () => {
    runPython(`
assert _Node(_code).has_function('selection_sort')
    `)
}})
```

Vaša funkcija za `selection_sort` treba da ima jedan parametar.

```js
(
    {test: () => {
        runPython(`
            import inspect
            sig = inspect.signature(selection_sort)
            assert len(sig.parameters) == 1
        `)
    }}
)
```

Ne smete da uvozite bilo koje module ili koristite ugrađene metode sortiranja u svoj kod.

```js
(
    {
        test: () => runPython(`
            assert len(_Node(_code).find_imports()) == 0
            assert not _Node(_code).block_has_call("sort")
            assert not _Node(_code).block_has_call("sorted")
        `)
    }
)
```

`selection_sort` Trebalo bi da vrati isti spisak kao i ulazni spisak.

```js
(
    {
        test: () => runPython(`
            input_list = [32, 0, 88, 1, 203]
            assert selection_sort(input_list) is input_list
        `)
    }
)
```

`selection_sort` Treba da promenite listu unosa na tom mestu. Ne smete da koristite bilo koji način za dodavanje ili uklanjanje stavki iz liste.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)

                def append(self, *args, **kwargs):
                    self._record.append(f'Appending to list: {args}, {kwargs}')
                    super().append(*args, **kwargs)

                def insert(self, *args, **kwargs):
                    self._record.append(f'Inserting to list: {args}, {kwargs}')
                    super().insert(*args, **kwargs)

                def pop(self, *args, **kwargs):
                    self._record.append(f'Popping from list: {args}, {kwargs}')
                    super().pop(*args, **kwargs)

                def remove(self, *args, **kwargs):
                    self._record.append(f'Removing from list: {args}, {kwargs}')
                    super().remove(*args, **kwargs)


            list_to_sort = CustomList([32, 0, 88, 1, 203])

            selection_sort(list_to_sort)
            assert all(
                'Setting item' in record
                for record in list_to_sort._record
            )
        `)
    }
)
```

Tvoja funkcija ``selection_sort`` bi trebalo da prati algoritam selection sort, menjajući minimalnu vrednost u neuređenom delu liste sa prvim neuređenim elementom. Izbegavaj nepotrebne zamene kada je minimalna vrednost već na odgovarajućoj poziciji.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)


            list_to_sort = CustomList([33, 1, 89, 2, 67, 245])

            swap_pairs = [
                ((0, 1, 33), (1, 33, 1)),
                ((1, 2, 33), (3, 33, 2)),
                ((2, 33, 89), (3, 89, 33)),
                ((3, 67, 89), (4, 89, 67)),
            ]

            selection_sort(list_to_sort)
            record = list_to_sort._record

            actual_pairs = [
                (record[offset * 2], record[offset * 2 + 1])
                for offset, _ in enumerate(record[::2])
            ]
            assert len(swap_pairs) == len(actual_pairs)

            for expected_pair, actual_pair in zip(swap_pairs, actual_pairs):
                assert (
                    set(f'Setting item {key} to {value}. Previous value: {old}' for key, value, old in expected_pair)
                    == set(actual_pair)
                )
        `)
    }
)
```

`selection_sort([33, 1, 89, 2, 67, 245])` trebalo bi da vrati `[1, 2, 33, 67, 89, 245]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([33, 1, 89, 2, 67, 245]) == [1, 2, 33, 67, 89, 245]
            
        `)
    }
)
```

`selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3])` treba da vrati `[3, 5, 12, 15, 16, 23, 72, 99, 567]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3]) == [3, 5, 12, 15, 16, 23, 72, 99, 567]
            
        `)
    }
)
```

`selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` treba da vrati `[1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]) == [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
            
        `)
    }
)
```

Vaša funkcija za `selection_sort` treba da sortira tačno bilo koju listu brojeva.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([42, 17, 93, 8, 61, 29]) == [8, 17, 29, 42, 61, 93]
            assert selection_sort([11, 4, 78, 23, 55, 198, 65, 90, 2]) == [2, 4, 11, 23, 55, 65, 78, 90, 198]
            assert selection_sort([9, 27, 3, 7, 101, 66, 34, 52, 87, 42, 12, 29]) == [3, 7, 9, 12, 27, 29, 34, 42, 52, 66, 87, 101]
            assert selection_sort([5, 14, 33, 77, 2, 18, 92, 1, 100, 45, 73, 64, 28, 56]) == [1, 2, 5, 14, 18, 28, 33, 45, 56, 64, 73, 77, 92, 100]
            
        `)
    }
)     
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def selection_sort(nums):
    for i, _ in enumerate(nums):
        min_index = i

        for j in range(i + 1, len(nums)):
            if nums[j] < nums[min_index]:
                min_index = j

        if min_index != i:
            nums[i], nums[min_index] = nums[min_index], nums[i]

    return nums

```
