---
id: 5f332a88dc25a0fd25c7687a
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Ime restorana je `CAMPER CAFE`. Stoga, dodajte element `h1` unutar vašeg elementa `main`. Dajte ime restorana velikim slovima kako bi bilo dobro vidljivo.

# --hints--

Trebalo bi da imaš etiketu za otvaranje `<h1>`.

```js
assert.match(code, /<h1>/i);
```

Trebalo bi da imaš oznaku za zatvaranje `</h1>`.

```js
assert.match(code, /<\/h1>/i);
```

Ne smete da menjate vaš postojeći element `main`.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Vaš element `h1` treba biti postavljen unutar vašeg elementa `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Vaše polje za `h1` treba da sadrži tekst od `CAMPER CAFE`.

```js
assert.equal(document.querySelector("h1")?.innerText.trim(), "CAMPER CAFE");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
