---
id: 5f332b23c2045fb843337579
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Da biste obavestili goste da je restoran osnovan godine `2020`, dodajte polje `p` ispod polja `h1` sa tekstom `Est. 2020`.

# --hints--

Trebalo bi da imaš etiketu za otvaranje `<p>`.

```js
assert.match(code, /<p>/i);
```

Treba vam zatvarački tag od `</p>`.

```js
assert.match(code, /<\/p>/i);
```

Ne smete da menjate vaš deo `h1` koji se nalazi. Uverite se da niste izbrisali oznaku za zatvaranje.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

Vaša komponenta `p` treba da bude ispod vaše komponente `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

Vaše polje `p` treba da sadrži tekst `Est. 2020`.

```js
assert.equal(document.querySelector("p")?.innerText.trim(), "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
