---
id: 5f33294a6af5e9188dbdb8f3
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Na ovom radincu, ćete naučiti osnove CSS (Cascading Style Sheets) za izradu menija kafe.

Počnimo dodavanjem nekog sadržaja menija. Dodajte element `main` unutar postojećeg elementa `body`. Na kraju, biće informacije o cenama za kafu i užine koje nudi kafe.

# --hints--

Vaš kod treba da ima otvarajući tag od `<main>`.

```js
assert.match(code, /<main>/i);
```

Vaš kod bi trebalo da ima zatvorni tag/oznaku od `</main>`.

```js
assert.match(code, /<\/main>/i);
```

Ne smete da menjate svoj element za `body`. Proverite da slučajno ne brišete svoj zatvarajući tag.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Etiketa za `main` treba da bude unutar etikete za `body`.

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
