---
id: 5f344fbc22624a2976425065
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Kreiraj element sa `h2` unutar elementa `section` i sa tekstom `Coffee`.

# --hints--

Treba ti da imaš oznaku za otvaranje `<h2>`.

```js
assert.match(code, /<h2\s*>/i);
```

Trebalo bi da ima oznaku za zatvaranje `</h2>`.

```js
assert.match(code, /<\/h2\s*>/i);
```

Ne smete da menjate vašu postojeću komponentu `section`. Uverite se da niste izbrisali oznaku za zatvaranje.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

Vaša komponenta za `h2` mora da bude unutar vaše komponente `section`.

```js
const h2 = document.querySelector('h2');
assert.equal(h2.parentElement.tagName, 'SECTION');
```

Vaše polje za `h2` treba da ima tekst od `Coffee`.

```js
const h2 = document.querySelector('h2');
assert.equal(h2?.innerText.trim(), 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
