---
id: 5f3477aefa51bfc29327200b
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Dodali ste tri elementa pisanjem CSS unutar oznaka `style`. Ovo radi, ali pošto će biti mnogo više stilova, bolje je sve stilove staviti u zaseban fajl i povezati ga.

Različita datoteka za `styles.css` je kreirana za vas. Možete prebacivati između datoteka koristeći kartice na vrhu uređivača.

Počnite ponovo da pišete stilove koje ste kreirali u fajlu `styles.css`. Uverite se da ne uključujete otvarajuće i zatvarajuće tagove za `style`.

# --hints--

Tvoja datoteka `styles.css` bi trebalo da ima biranje tipa `h1, h2, p`.

```js
assert.match(code.slice(code.indexOf('</html>')).replace(/[\s\n]*/g, ""), __helpers.concatRegex(__helpers.permutateRegex(['h1', 'h2', 'p'], { elementsSeparator: ',' }), /{/));
```Vaš selektor treba da postavi atribut za `text-align` na `center`.

```js
assert.match(code.slice(code.indexOf('</html>')), /text-align:\s*center;?/);
```

Твој код не сме бити са ликовима `<style>`.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
