---
id: 5f3477cbcb6ba47918c1da92
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Da bi se dizajn stranice prikazao isto na mobilnim telefonima kao što je na računaru ili laptopu, morate dodati element sa posebnom karakteristikom `meta` koji ima posebnu karakteristiku `content`.

Naučite se o viewport komponenti `meta` iz prethodnih lekcija.

Evo primera:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

Vaš kod treba da ima dva elementa od `meta`.

```js
assert.lengthOf(code.match(/<meta.*?\/?>/g),2);
```

Vaš element za `meta` mora imati svojstvo `name` sa vrednošću `viewport`.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/name=('|")viewport\1/) || meta[1]?.outerHTML?.match(/name=('|")viewport\1/));
```

Vaš element za `meta` mora imati svojstvo `content` sa vrednošću `width=device-width, initial-scale=1.0`.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
--fcc-editable-region--
    
--fcc-editable-region--
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
