---
id: 5f3477ae9675db8bb7655b30
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

U prethodnom koraku, koristili ste <dfn>selektor tipa</dfn> da biste kreirali stil za element `h1`. Postavite sadržaj elemenata `h2` i `p` u sredinu dodajući novi selektor tipa za svaki od njih na postojeći element `style`.

# --hints--

Ne smijete mijenjati postojeći selektor za `h1`.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasH1);
```

Ne smete dodati novu oznaku za `style`. Dodajte nove pravila za CSS na postojeću oznaku `style`.

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

Trebalo bi da dodate novi selektor za `h2`.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert.exists(hasH2);
```

Trebalo bi da dodate novi selektor za `p`.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert.exists(hasP);
```

Vaš element `h1` mora imati `text-align` iz `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

Vaša komponenta za `h2` mora da ima `text-align` od `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert.equal(h2TextAlign, 'center');
```

Tvoj element za `p` mora imati `text-align` za `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert.equal(pTextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1 {
        text-align: center;
      }
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

