---
id: 5f3477ae34c1239cafe128be
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Sada imate tri tipa selektora sa istim stilom. Možete dodati istu grupu stilova na više elemenata kreiranjem liste selektora. Svaki selektor je razdvojen zarezom kao što je ovo:

```css
selector1, selector2 {
  property: value;
}
```

Uklonite tri postojeći tipa prekidača i zamenite ih sa jednim spiskom prekidača koji centriraju tekst za elemente `h1`, `h2`, i `p`.

# --hints--

Trebalo bi da koristiš isti selektor za sva tri elementa, `h1, h2, p`. Proveri da pratiš taj redosled.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert.exists(hasSelector);
```

Trebalo bi da imaš samo jedan izbor u svom polju `style`.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert.equal(selectors.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

