---
id: 5f344fad8bf01691e71a30eb
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Sada imate malu kontrolu nad prikazom i izgledom vašeg sadržaja. Da biste to promenili, dodajte element `style` unutar elementa `head`.

# --hints--

Vaš kod treba da ima oznaku za otključavanje od `<style>`.

```js
assert.match(code, /<style\s*>/i);
```

Tvoj kod mora imati zatvorni tag `</style>`.

```js
assert.match(code, /<\/style\s*>/);
```

Element vašeg `style` mora biti postavljen unutar elementa vašeg `head`.

```js
assert.match(code, /<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
