---
id: 5f3477cb2e27333b1ab2b955
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Sada je potrebno uključiti fajl `styles.css`, ili ponovo koristiti stilove. Unutar elementa `head`, dodajte element `link`. Postavite atribut elementu `rel` sa vrednošću `"stylesheet"` i atribut elementu `href` sa vrednošću `"styles.css"`.

**Napomena:** Element sa `link` je element bez zatvarajućeg taga, što znači da nema zatvarački tag. Elementi bez zatvarajućeg taga treba da se pišu kao `<link>` umesto `<link></link>`.

# --hints--

Tvoj kod mora imati komponentu od `link`.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

Element `link` je element koji nema zatvorni tag, ne bi trebalo da ima zatvorni tag `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Ne smete da menjate vaš deo sa sadržajem `head` koji se nalazi. Proverite da niste izbrisali oznaku za zatvaranje.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

Trebalo bi da ima jedan element `link`.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

Tvoj element `link` treba da bude unutar tvog elementa `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Vaš element od `link` mora imati svojstvo elementa `rel` sa vrednošću `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Element vašeg `link` treba da ima svojstvo polja `href` sa vrednošću `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
