---
id: 5f344fc1520b6719f2e35605
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Biće dve sekcije na meniju, jedna za kafu i jedna za užine. Dodajte element `section` unutar elementa `main` kako biste imali mesto za sve dostupne kafe.

# --hints--

Trebalo bi da imaš oznaku za otvaranje `<section>`.

```js
assert.match(code, /<section\s*>/i);
```

Trebalo bi da imaš oznaku za zatvaranje `</section>`.

```js
assert.match(code, /<\/section\s*>/i);
```

Ne smete da menjate vašu sekciju `main` koja se nalazi. Uverite se da niste izbrisali oznaku za zatvaranje.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

Vaš element `section` treba da bude unutar vašeg elementa `main`.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert.strictEqual(section.parentElement, main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
