---
id: 5f344f9c805cd193c33d829c
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

U prethodnim lekcijama, naučili ste kako da dodate atribute i vrednosti za `CSS` kao što su:

```css
element {
 property: value;
}
```

Umetni sadržaj elementa `h1` u sredinu postavljanjem njegove atribute/svojstva `text-align` na vrednost `center`.

# --hints--

Trebalo bi da imaš birač za `h1` u svom polju `style`.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasSelector);
```

Tvoja vrednost za `text-align` treba da bude `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert.isTrue(hasTextAlign);
```

Tvoj `h1` treba da postavi svojstvo `text-align` na `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
