---
id: 5f356ed60a5decd94ab66986
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Komentari u CSS izgledaju ovako:

```css
/* comment here */
```

U vašoj stilskoj listi, ukomentariš red koji pokazuje svojstvo i vrednost `background-color`, da bi video efekat stilizovanja samo elementa `#menu`. Ovo će ponovo učiniti pozadinu belom.

# --hints--

Trebalo bi da ostaviš komentar na liniji `background-color: burlywood;` u tvom CSS.

```js
assert.match(code, /\/\*\s*background-color:\s*burlywood;?\s*\*\//i);
```

`body` Tvoj treba da ima belu pozadinu.

```js
const body = document.querySelector("body");
const compStyles = window.getComputedStyle(body);
const backgroundColor = compStyles?.getPropertyValue('background-color');
assert.equal(backgroundColor, "rgba(0, 0, 0, 0)")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
--fcc-editable-region--
  background-color: burlywood;
--fcc-editable-region--
}

h1, h2, p {
  text-align: center;
}

#menu {
  width: 300px;
}
```

