---
id: 5f356ed6199b0cdef1d2be8f
title: Korak 22
challengeType: 0
dashedName: step-22
---

# --description--

Do sada koristiš tipovne selektore i ID (id) za stilizovanje elemenata. Međutim, uobičajenije je koristiti različite selektore za stilizovanje vaših elemenata.

Naučite kako da radite sa biračkim listama `class` u prednaprednim kursevima kao što je ovo:

```css
.class-name {
  styles
}
```

Promijeni postojeći konstruktor `#menu` u konstruktor klase objekata, zamjenom `#menu` sa klasom objekata nazvanom `.menu`.

# --hints--

Treba da imaš selektor klase `.menu`.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert.exists(hasMenu);
```

Ne bi trebalo da imaš selektor za `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert.notExists(hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
--fcc-editable-region--
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
