---
id: 5f35e5c44359872a137bd98f
title: Korak 24
challengeType: 0
dashedName: step-24
---

# --description--

Pošto je glavni proizvod kave koji se prodaje kafa, možete koristiti sliku zrna kafe kao pozadinu stranice.

Uklonite komentare unutar izbornika tipa `body`. Nakon toga, dodajte atribut `background-image` i postavite mu vrednost na `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

# --hints--

Treba da uklonite funkciju `background-color` koja je objavljena kao komentar.

```js
assert.notMatch(code, /\/\*\s*background-color:\s*burlywood;?\s*\*\//i);
```

Vaš birački list `body` ne bi trebalo da sadrži nikakve komentare.

```js
assert.notMatch(code, /body\s*{\s*\/\*/i);
```

Treba da postaviš atribut `background-image` na `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert.isTrue(hasBackground)
```

Element vašeg `body` treba da ima sliku pejzaža zrna kafe.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');

assert.equal(bodyBackground, `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
--fcc-editable-region--
  /*
  background-color: burlywood;
  */
--fcc-editable-region--
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

