---
id: 5f3c866dd0d0275f01d4d847
title: Korak 35
challengeType: 0
dashedName: step-35
---

# --description--

U redu, to nije radilo. Formatiranje elemenata `p` kao što su `inline-block` i stavljanje ih na različite linije kreira dodatni razmak sa desne strane prvog elementa `p`, uzrokujući da drugi element sklizne na sledeću liniju.

Jedan način da to popravimo je smanjiti širinu svakog elementa od `p` malo ispod `50%`. Stoga, promenite vrednost `width` u `49%` za svaku klasu objekata da biste videli šta se dešava.

# --hints--

Treba da postaviš atribut `width` na `49%` u svom biraču `.flavor`.

```js
const flavorWidth = new __helpers.CSSHelp(document).getStyle('.flavor')?.getPropertyValue('width');
assert.equal(flavorWidth, '49%');
```

Treba da postaviš atribut `width` kao `49%` u svoj selektor `.price`.

```js
const priceWidth = new __helpers.CSSHelp(document).getStyle('.price')?.getPropertyValue('width');
assert.equal(priceWidth, '49%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

--fcc-editable-region--
.flavor {
  text-align: left;
  width: 50%;
}

.price {
  text-align: right;
  width: 50%;
}
--fcc-editable-region--
```

