---
id: 5f3c866d5414453fc2d7b480
title: Korak 27
challengeType: 0
dashedName: step-27
---

# --description--

Počnite ispod postojećeg para za kafu/cienu, dodajte sledeću kafu i cenu koristeći elemente `article` koji imaju dva interne komponente od `p` unutar svakog.

```md
Caramel Macchiato 3.75
Pumpkin Spice 3.50
Hazelnut 4.00
Mocha 4.50
```

Kao što je bilo ranije, tekst prvog polja `p` treba da bude aroma kafe i tekst drugog polja `p` treba da bude cena.

# --hints--

Trebalo bi da ima pet komponenti `article`.

```js
assert.lengthOf(document.querySelectorAll('article'),5);
```

Svaki element `article` mora imati dva elementa `p`.

```js
const articles = document.querySelectorAll('article');
assert.lengthOf(articles?.[0]?.children, 2);
assert.lengthOf(articles?.[1]?.children, 2);
assert.lengthOf(articles?.[2]?.children, 2);
assert.lengthOf(articles?.[3]?.children, 2);
assert.lengthOf(articles?.[4]?.children, 2);
```

Tvoj prvi element od `article` bi trebalo da ima komponente `p` koje sadrže tekst `French Vanilla` i `3.00`.

```js
const children = document.querySelector('article')?.children;
assert.match(children?.[0]?.innerText.trim(), /French Vanilla/i);
assert.match(children?.[1]?.innerText.trim(), /3\.00/i);
```

Tvoj drugi element od `article` treba da ima elemente `p` koji sadrže tekst `Caramel Macchiato` i `3.75`.

```js
const children = document.querySelectorAll('article')?.[1]?.children;
assert.match(children?.[0]?.innerText.trim(), /Caramel Macchiato/i);
assert.match(children?.[1]?.innerText.trim(), /3\.75/i);
```

Treći element vašeg `article` mora imati elemente `p` sa tekstom `Pumpkin Spice` i `3.50`.

```js
const children = document.querySelectorAll('article')?.[2]?.children;
assert.match(children?.[0]?.innerText.trim(), /Pumpkin Spice/i);
assert.match(children?.[1]?.innerText.trim(), /3\.50/i);
```

Tvoj četvrti element za `article` treba da ima elemente sa `p` koji sadrže tekst od `Hazelnut` i `4.00`.

```js
const children = document.querySelectorAll('article')?.[3]?.children;
assert.match(children?.[0]?.innerText.trim(), /Hazelnut/i);
assert.match(children?.[1]?.innerText.trim(), /4\.00/i);
```

Tvoj peti element od `article` treba da ima komponente od `p` sa tekstom `Mocha` i `4.50`.

```js
const children = document.querySelectorAll('article')?.[4]?.children;
assert.match(children?.[0]?.innerText.trim(), /Mocha/i);
assert.match(children?.[1]?.innerText.trim(), /4\.50/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
            <p>French Vanilla</p>
            <p>3.00</p>
          </article>
--fcc-editable-region--
          
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
