---
id: 6143869bb45bd85e3b1926aa
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Unutar vašeg `.heading` elementa, kreirajte `header` element sa `class` podešenim na `hero`.

U tom elementu, kreirajte `img` element sa `src` podešenim na `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`, `alt` podešenim na `freecodecamp logo`, i `class` podešenim na `hero-img`.

Atribut `loading` na `img` elementu može se podesiti na `lazy` kako bi se pretraživaču reklo da ne dohvati resurs slike dok mu nije potreban (kao što je slučaj kada korisnik skroluje sliku u vidokrug). Kao dodatna korist, elementi sa opuštenim učitavanjem neće biti učitani sve dok nisu učitani elementi bez opuštenog učitavanja – to znači da korisnici sa sporom internet vezom mogu pregledati sadržaj vaše stranice bez čekanja na učitavanje slika.

Dajte vašem novom `img` elementu atribut `loading` podešen na `lazy`.

# --hints--

You should create a `header` element.

```js
assert.exists(document.querySelector('header'));
```

Your `header` element should be within your `.heading` element.

```js
assert.equal(document.querySelector('header')?.parentElement?.className, 'heading');
```

Your `header` element should have the `class` set to `hero`.

```js
assert.equal(document.querySelector('header')?.className, 'hero');
```

You should create an `img` element.

```js
assert.exists(document.querySelector('img'));
```

Your `img` element should be within your `header` element.

```js
assert.equal(document.querySelector('img')?.parentElement?.localName, 'header');
```

Your `img` element should have the `src` set to `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`.

```js
assert.equal(document.querySelector('img')?.getAttribute('src'), 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

Your `img` element should have the `alt` set to `freecodecamp logo`.

```js
assert.equal(document.querySelector('img')?.getAttribute('alt'), 'freecodecamp logo');
```

Your `img` element should have the `loading` attribute set to `lazy`.

```js
assert.equal(document.querySelector('img')?.getAttribute('loading'), 'lazy');
```

Your `img` element should have the `class` set to `hero-img`.

```js
assert.equal(document.querySelector('img')?.className, 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
