---
id: 614387cbefeeba5f3654a291
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Nakon vašeg `header` elementa, kreirajte `div` sa `class` podešenim na `author`.

Unutar tog `div`, kreirajte `p` element sa `class` podešenim na `author-name` i date mu tekst `By freeCodeCamp`. Obuhvatite deo `freeCodeCamp` u `a` element sa `href` podešenim na `https://freecodecamp.org`, a `target` podešenim na `_blank`.

Ispod toga, dodajte drugi `p` element sa klasom `publish-date` i tekstom `March 7, 2019`.

# --hints--

You should create a new `div` element.

```js
assert.exists(document.querySelector('div'));
```

Your `div` element should come after your `header` element.

```js
assert.equal(document.querySelector('div')?.previousElementSibling?.localName, 'header');
```

Your `div` element should have the `class` set to `author`.

```js
assert.equal(document.querySelector('div')?.className, 'author');
```

You should create two new `p` elements.

```js
assert.lengthOf(document.querySelectorAll('p'), 3);
```

Your two new `p` elements should be within your `div` element.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

Your first new `p` element should have a `class` set to `author-name`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.className, 'author-name');
```

Your first new `p` element should have the text `By freeCodeCamp`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.innerText, 'By freeCodeCamp');
```

Your second new `p` element should have a `class` set to `publish-date`.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className, 'publish-date');
```

Your second new `p` element should have the text `March 7, 2019`.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText, 'March 7, 2019');
```

You should create a new `a` element.

```js
assert.exists(document.querySelector('a'));
```

Your `a` element should be within your first new `p` element.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName, 'a');
```

Your `a` element should have the `href` set to `https://freecodecamp.org`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href'), 'https://freecodecamp.org');
```

Your `a` element should have the `target` set to `_blank`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target'), '_blank');
```

Your `a` element should surround the text `freeCodeCamp`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText, 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
