---
id: 68e97fe79367ad7b5dd6c9cd
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Na ovom koraku dodajte element za `h1` sa tekstom od `XYZ Bookstore`.

# --hints--

Trebalo bi da imaš element `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Sadržaj vašeg elementa `h1` bi trebalo da bude `XYZ Bookstore`. Proverite da li ima grešaka u pravopisu ili kapitalizacije.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

Tvoj element `h1` treba da bude unutar elementa `body`.

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
