---
id: 68ebdbacdd3fa474132cc975
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Ispod polja za `h1`, dodaj polje `p` koje sadrži ovaj tekst: `Browse our collection of amazing books!`.

# --hints--

Trebalo bi da imaš element `p`.

```js
assert.exists(document.querySelector('p'));
```

Tekst vašeg polja za `p` trebalo bi da bude: `Browse our collection of amazing books!`.

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

Vaša komponenta `p` treba da bude ispod vaše komponente `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
