---
id: 68eca5c0065c82d256d29ca3
title: Korak 17
challengeType: 0
dashedName: step-17
---

# --description--

Na drugoj kartici, dodaj element `button` sa svojstvom `class` postavljeno na `btn` i tekst od `Buy Now`.

Oba elementa od `button` sada dele `class`, što znači da mogu biti stilizovani kompatibilnim stilom zajedno.

# --hints--

Trebalo bi da imaš element `button` postavljen unutar drugog elementa sa klasom `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

Funkcija za vašu drugu karticu, `button`, treba da ima atribut `class`.

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

Atribut `button` vaše druge kartice treba da ima `class` sa vrednošću `btn`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

Tekst polja `button` vaše druge kartice treba da bude `Buy Now`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
