---
id: 68eca11e99e3c5c894ca9d69
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Element za `button` se koristi za kreiranje klikabilnih dugmadi na veb stranici. Dugmad su interaktivni elementi koje korisnik može kliknuti kako bi izvršio akcije.

Dodaj element za `button` unutar elementa sa `class` od `card`, dodaj dugme atribut/svojstvo od `class` postavljeno na `btn`, i tekst sa `Buy Now`.

# --hints--

Trebalo bi da imaš element sa `button` postavljenim unutar elementa sa klasom `card`.

```js
assert.exists(document.querySelector('.card button'));
```

Vaš atribut za `button` treba da ima svojstvo od `class`.

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

Vaše polje za `button` mora imati `class` sa vrednošću `btn`.

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

Tekst vašeg polja za `button` bi trebalo da bude `Buy Now`.

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
