---
id: 68eca2a795b333ca5fee30a8
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Sada kreiraj drugu knjigu karticu. Dodaj još jedan element za `div` sa svojstvom `class` ako ima vrednost `card`. Pogledaj kako možeš ponovo da koristiš ime klase objekata za više elemenata kako bi koristio dosledan stil.

# --hints--

Trebalo bi da ima dva elementa `div` sa karakteristikom `class` sa vrednošću `card`.

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

Drugi element sa klasom `card` mora biti unutar elementa sa klasom `card-container`.

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

Dva elementa sa klasom `card` treba da budu braća i sestre iste klase.

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```
