---
id: 68ec6cd5b7e8f5a8f7319e32
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Element ``div`` služi kao kontejner za povezivanje drugih elemenata od `HTML`. Bolje je koristiti element ``div`` kada želite povezati elemente od `HTML` koji će učestvovati u skupu stilova/tema `CSS`.

Ispod elementa `p`, dodaj element `div`. `div` Ovo će biti kontejner za vaše kartice knjiga.

**Napomena**: Ovaj radionica ne izvršava CSS. Strukturirane klase i komponente su važne za stilove CSS, ali u ovoj radionici se koriste samo za struktuiranje i povezivanje sadržaja. Naučićete kako stilovi funkcionišu u sledećem modulu.

# --hints--

Trebalo bi da imaš element `div`.

```js
assert.exists(document.querySelector('div'));
```

Vaš element `div` treba da bude ispod vašeg elementa `p`.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
