---
id: 68ec6d9a315221aa31e54816
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Svojstvo elementa `class` koristi se za identifikovanje jednog ili više elemenata u svrhu stilizovanja. Za razliku od svojstva elementa `id`, imena klasa elemenata ne moraju biti jedinstvena: više elemenata može da deli jednu klasu.

Evo primera:

```html
<p class="example">example paragraph</p>
```

Dodaj svojstvo elemenata `class` na tvoj element `div` i postavi mu vrednost na `card-container`.

# --hints--

Element [`div`] treba da ima svojstvo elementa [`class`] .

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

Tvoj element za `div` mora imati svojstvo elementa `class` sa vrednošću `card-container`.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  <div>
--fcc-editable-region--
    
  </div>
</body>

</html>
```
