---
id: 68ec9332a9b5b2b32487bd00
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Atribut elementa `id` dodaje jedinstveni identifikator za element HTML. Svaki `id` mora biti jedinstven na stranici i treba ga koristiti samo jednom.

Vrednosti za `id` ne smeju sadržati razmake i moraju da sadrže samo slova, brojeve, donje podkreškinge i crtice.

Evo primera:

```html
<p id="para">example paragraph</p>
```

Dodaj atribut elementu `id` za tvoj element sa klasom `card` i postavi mu vrednost na `sally-adventure-book`.

# --hints--

Element vaš sa klasom `card` mora imati svojstvo elementa `id`.

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

Element vašeg klase `card` mora imati `id` sa vrednošću `sally-adventure-book`.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">
  --fcc-editable-region--
      
    </div>
  </div>
</body>

</html>
```
