---
id: 68ec99e478211dc578699944
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Ispod elementa ``h2`` u prvom elementu sa klasom ``card``, dodajte element ``p`` sa sledećim tekstom:

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

Trebalo bi da imaš element `p` postavljen unutar elementa sa klasom `card`.

```js
assert.exists(document.querySelector('.card p'));
```

Tekst vašeg elementa `p` treba da bude `This is an epic story of Sally and her dog Rex as they navigate through other worlds.`

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

Vaš element za `p` mora biti ispod vašeg elementa za `h2`.

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
