---
id: 68eca6364c5616d393389a30
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Napomena, element HTML izgleda ovako:

```html
<element attribute="value">
    inner text
</element>
```

Ispod elementa sa klasom `card-container`, dodaj novi element za `p` sa ovim tekstom:

```md
Review your selections and continue to checkout.
```

Ispod elementa `p`, kreiraj element `div` sa svojstvom `class` postavljeno na `btn-container`. Ovaj kontejner će povezati vaše elemente navigacionog dugmeta.

# --hints--

Novi element za `p` treba da bude ispod elementa sa klasom `card-container`.

```js
assert.exists(document.querySelector('.card-container + p'));
```

Tekst vaše nove funkcije za `p` bi trebalo da bude `Review your selections and continue to checkout.`

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

Trebalo bi da imaš element sa atributom `div` sa svojstvom `class` koje je `btn-container`.

```js
assert.exists(document.querySelector('div.btn-container'));
```

Vaš element `div` sa svojstvom `class` koji treba da bude `btn-container` mora da sledi nakon novog elementa `p`.

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
