---
id: 68eca6e5a6759ed4ea0034dc
title: Korak 19
challengeType: 0
dashedName: step-19
---

# --description--

Unutar elementa sa klasom `btn-container`, dodajte dva elementa za `button`:

Prvi dugme:

- Id: `view-cart-btn`
- Klasa predmeta: `btn`
- Tekstovi: `View Cart`

Taster za drugu:

- ID: `checkout-btn`
- Klasa stavke: `btn`
- Tekst: `Checkout`

Čestitam! Uspešno ste napravili strukturu stranice knjižare koristeći div-ove, klase elemenata i ID-jeve za organizovanje vašeg sadržaja.

# --hints--

Trebalo bi da imaš dva elementa od `button` postavljena unutar elementa `.btn-container`.

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

Tvoj prvi dugme unutar elementa sa klasom `btn-container` treba da ima atribut elementa `id` sa vrednošću `view-cart-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

Prvi taster unutar elementa sa klasom `btn-container` mora imati atribut elementa `class` sa vrednošću `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

Tekst vašeg prvog dugmeta treba da bude `View Cart`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

Tvoje drugo dugme unutar elementa sa klasom `btn-container` mora imati atribut `id` sa vrednošću `checkout-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

Tvoje drugo dugme unutar elementa sa `class="btn-container"` mora imati atribut elementa od `class` sa vrednošću `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

Tekst tvog drugog dugmeta treba da bude `Checkout`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
  <p>Review your selections and continue to checkout.</p>
  
  <div class="btn-container">
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>

  <div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
  </div>

  <p>Review your selections and continue to checkout.</p>

  <div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
  </div>
</body>

</html>
```
