---
id: 69134d14150d51f8066a3246
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Trenutno, poboljša strukturu vašeg dokumenta za HTML kako bi osigurao da je vaša stranica pravilno napisana.

Unutar elementa `head`, dodaj element `<meta charset="UTF-8">`.

Konačno, dodajte element `body` ispod sekcije `head`. Ovde će se postaviti sav vidljivi sadržaj stranice.

# --hints--

Vaš kod treba da sadrži element `<meta charset="UTF-8">`.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

Vaš kod mora da uključuje otvarajuće i zatvarajuće oznake za `<body>`.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

Vaš element za `body` treba da bude ispod elementa za `head`.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
  <title>XYZ Bookstore Page</title>
</head>

--fcc-editable-region--
</html>
```
