---
id: 69134b6cb43aa5f254950a10
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, kreiraćete stranicu knjižare tako što ćete napraviti kartice knjiga koje prikazuju informacije o različitim knjigama. Naučićete kako da strukturirate sadržaj koristeći elemente ``div``, klase, i ID-ove.

Počnite sa vašom stranicom knjižare kreiranjem osnovnog dizajna HTML.

Dodaj izjavu za `<!DOCTYPE html>` i komponente od `html` i `head`.

Dodajte atribut `lang` elementu `html` i postavite ga na `"en"`.

# --hints--

Vaš kod treba da sadrži deklaraciju za `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Vaš kod bi trebalo da sadrži otvarajuće i zatvarajuće oznake za `<html>`.

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

Vaše svojstvo za `html` trebalo bi imati svojstvo elementa `lang`, postavljeno na `"en"`.

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

Vaš kod bi trebalo da sadrži otvarajuće i zatvarajuće tagove za `<head>`.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

Element vašeg `head` treba da bude smešten unutar elementa `html`.

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
