---
id: 691f85085787a8c4ca589fbb
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Dodaj element `title` unutar elementa `head`.

Postavi naslov stranice na `XYZ Bookstore Page`.

# --hints--

Vaš kod mora da uključuje otvarajuće i zatvarajuće oznake za `<title>`.

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

Tvoja glava mora biti `XYZ Bookstore Page`.

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

Vaš element za `title` treba da bude postavljen unutar elementa `head`.

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
--fcc-editable-region-- 
</head>

</html>
```
