---
id: 67fccc2463253b213a000142
title: Izgradi medijsku ličnost
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

U prethodnim časovima, je uvedeno sa funkcijama i komponentama `audio` i `video`. U ovom laboratorijumu, ćete izgraditi medijski plejer koji će prikazivati pesmu `audio` i `video` zajedno sa transkriptom teksta.

Za element ``audio``, trebaće vam da uključite element ``source`` koji se koristi za određivanje korišćenog medija.

Evo primera:

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

Funkcija `source` se takođe može koristiti u funkciji `video`, kao što je sledeće:

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da ima element `h1` za glavni naslov stranice.
2. Trebalo bi da ima tri elementa `section`.
3. Unutar elementa `section`, trebalo bi da ima element `h2` za naslov reprodukovane pesme.
4. Ispod elementa `h2`, trebalo bi da ima element `audio` sa svojstvima `controls` i `aria-label`.
5. Unutar `audio` elementa trebalo bi da postoji `source` element sa svojstvom `src` koje predstavlja audio fajl i svojstvom `type`. Slobodno možete koristiti ovaj audio URL ako želite: `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. Unutar drugog elementa `section`, trebalo bi da ima element `h2` za naslov reprodukovanog videa.
7. Ispod elementa `h2`, trebalo bi da ima element `video` sa svojstvima `controls`, `width` i `aria-label`.
8. Unutar elementa `video`, trebalo bi da ima element `source` sa svojstvom `src` koje predstavlja video fajl i svojstvom `type`. Slobodno možete koristiti ovaj video URL ako želite: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`
9. Ispod elementa `source`, trebalo bi da ima element `track` sa svojstvom `src` koje predstavlja titlove (podnaslov), i svojstvima `kind`, `srclang` i `label`. Slobodno možete koristiti ovaj transkript URL ako želite: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.vtt`
10. Unutar trećeg elementa `section`, trebalo bi da ima element `h2` za naslov sekcije poput "Transkript".
11. Ispod elementa `h2`, trebalo bi da ima element `p` sa transkriptom videa.

# --hints--

Trebalo bi da imaš element `h1`.

```js
assert.isNotNull(document.querySelector("h1"));
```

`h1` Vaše polje mora da sadrži tekst koji predstavlja glavni naslov stranice.

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```

Trebalo bi da imaš tri elementa za `section`.

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

Prva `section` treba da ima `h2` element.

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

`h2` trebalo bi da ima tekst koji predstavlja naslov pesme koja se reprodukuje.

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

Prva `section` treba da ima `audio` element.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

Tvoj element za `audio` treba da ima svojstvo od `controls`.

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

Vaš element `audio` treba da ima svojstvo `aria-label` sa tekstom koji opisuje zvuk.

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

Vaša komponenta `audio` treba da ima komponentu `source`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

Vaš element za `source` treba da ima svojstvo `src` koje predstavlja audio fajl.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

Vaše polje `source` mora imati svojstvo `type` sa vrednošću koja određuje tip zvuka.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

`section` treba da ima karakteristiku od `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

`h2` Trebalo bi da ima tekst koji predstavlja naslov videa koji se reprodukuje.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

`section` druga mora imati karakteristiku od `video`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

Vaš element za `video` trebalo bi da ima svojstvo od `controls`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

Vaš element za `video` mora imati svojstvo `width` sa vrednošću koja određuje širinu videa.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

Element za `video` treba da ima element za `source`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

Kipengele cha `source` treba da ima svojstvo `src` koje predstavlja video fajl.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

Element `source` treba da ima atribut `type` sa vrednošću koja određuje tip videa.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

Tvoja karakteristika od `video` treba da ima karakteristiku od `track`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

Tvoj element `track` treba da ima svojstvo `src` koje predstavlja fajl sa tekstom.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

Vaš element za `track` mora imati atribut `kind` sa tekstom koji opisuje tip fajla.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

Tvoj deo `track` treba da ima svojstvo `srclang` sa tekstom koji opisuje jezik citata.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

Vaše polje za `track` trebalo bi da ima svojstvo `label` sa tekstom koji opisuje jezik citata.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

`section` treći treba da ima funkcionalnost `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

`h2` Trebalo bi da sadrži tekst koji predstavlja zaglavlje sekcije.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

`section` treći bi trebalo da ima element od `p` za tekstualnu kopiju.

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4">
      
      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
