---
id: 67ed03ac474c48692f41749e
title: Izgraditi hash tabelu
challengeType: 27
dashedName: build-a-hash-table
saveSubmissionToDB: true
---

# --description--

U ovom okruženju, kreiraćete hash tabelu od nule. Hash tabela je struktura podataka koja skladišti parove ključ-vrednost. Ona radi uzimanjem ključa kao ulaza i zatim izvršavanjem hashovanja po specifičnoj funkciji za hashovanje.

Za potrebe ovog laboratorijuma, hash funkcija će biti jednostavna: sabraće Unicode vrednosti svakog karaktera u ključu. Haš vrednost će se koristiti kao stvarni ključ za skladištenje povezane vrednosti. Ista haš vrednost će takođe se koristiti za dohvaćanje i brisanje vrednosti povezane sa ključem.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi definisati klasu objekata nazvanu `HashTable`, koja ima svojstvo atributa `collection` postavljeno na prazan rečnik kada se kreira nova instanca od `HashTable`. Rečnik za `collection` treba da skladišti parove ključ-vrednost na osnovu haš vrednosti ključa.

2. Klasa `HashTable` treba da ima četiri primjera metode: `hash`, `add`, `remove`, na `lookup`.

3. Put `hash` treba da bude:

- Uzimanje sekvence znakova kao parametar.
- Vraćanje haš vrednosti izračunate kao zbir Unicode (ASCII) vrednosti svakog znaka u sekvenci znakova. Možete koristiti funkciju `ord` za ovaj proračun.

4. Metoda za `add` treba da:
    - Prihvati dva argumenta koji predstavljaju par ključ-vrednost, i izračuna heš ključa.
    - Koristi izračunatu vrednost heša kao ključ za skladištenje rečnika koji sadrži par ključ-vrednost unutar rečnika `collection`.
    - Ako više ključeva proizvede istu vrednost heša, njihovi parovi ključ-vrednost treba da se sačuvaju u ugnježdene rečnike koji su ispod te vrednosti heša.

5. Putanja za `remove` bi trebalo da bude:

- Uzmi ključ kao argument i izračunaj njegov heš.
- Proveri da li ključ postoji u kolekciji.
- Obriši odgovarajući par ključ-vrednost iz hash tabele.
- Ako ključ ne postoji u kolekciji, ne bi trebalo ni da vrati grešku niti da obriše bilo šta.

6. Putanja za `lookup` treba da bude:

- Uzeti ključ kao argument.
- Izračunati heš ključa i vratiti vrednost sačuvanu u tabeli hešova.
- Ako ključ nije u kolekciji, treba da vrati `None`.

# --hints--

Trebalo bi da definišete klasu `HashTable`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_class("HashTable")
    `)) 
})
```

Kada se kreira nova instanca klase `HashTable`, njeno svojstvo elemenata `collection` treba da se postavi na prazan rečnik.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.collection == {}
    `) 
})
```

Darasa za `HashTable` trebalo bi da ima metodu `hash`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('hash')`))}})
```

Putanja `hash` treba da primi sekvencu karaktera kao parametar.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.hash)
    len(sig.parameters) == 1
  `))
})
```

Funkcija `hash` treba da primi sekvencu karaktera kao argument i vrati zbir Unicode vrednosti svakog karaktera u toj sekvenci.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("fcc")
        assert hash_result == 300

        # prevent hardcoding
        assert ht.hash("golf") == 424
        assert ht.hash("read") == 412
`) })
```

Klasa `HashTable` trebalo bi da ima metodu za `add`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('add')`))}})
```

Funkcija `add` treba da primi ključ i vrednost kao parametre.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.add)
    len(sig.parameters) == 2
  `))
})
```

Klasa za `HashTable` bi trebalo da ima metodu za `remove`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('remove')`))}})
```

Putanja za `remove` mora da uzme ključ kao parametar.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.remove)
    len(sig.parameters) == 1
  `))
})
```

Kada pokušate da izbrišete nepostojeći ključ u kolekciji, ne bi trebalo da generiše grešku niti da obriše bilo šta.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("rose", "flower")
        index = ht.hash("rose")
        original = ht.collection.copy()

        ht.remove("tulip")
        ht.remove("sore")

        assert ht.collection == original
        assert "rose" in ht.collection[index]
        assert "tulip" not in ht.collection.get(index, {})
    `) 
})
```

Ako više ključa ima istu vrednost hasha, metoda `remove` bi trebalo da obriše samo odgovarajući par ključ-vrednost, a ne ceo rečnik u tom objektu.

```js
({
  test: () => runPython(`
    ht = HashTable()
    ht.add("rose", "flower")
    ht.add("sore", "pain")  # "rose" and "sore" both hash to the same index

    index = ht.hash("rose")
    ht.remove("rose")

    assert index in ht.collection
    assert "rose" not in ht.collection[index]
    assert "sore" in ht.collection[index]
    assert ht.collection[index]["sore"] == "pain"
  `)
})
```

Klasa za `HashTable` treba da ima metodu `lookup`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('lookup')`))}})
```

Putanja za `lookup` treba da uzme ključ kao parametar.

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.lookup)
    len(sig.parameters) == 1
  `))
})
```

`HashTable().hash('golf')` treba da vrati `424`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("golf")
        assert hash_result == 424

        # prevent hardcoding
        assert ht.hash("dear") == 412
        assert ht.hash("cat") == 312
    `) 
})
```

`HashTable().add('golf', 'sport')` treba dodati par ključ-vrednost u kolekciju za ključ `424`.

```js

({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")
    expected_value = {424: {'golf': 'sport'}}

    assert ht.collection == expected_value
    `) 
})
```

`HashTable().add('dear', 'friend')` na `HashTable().add('read', 'book')` treba da dodaju sve parove ključ-vrednost u kolekciju, kao što su, na primer, interni strukturirani rečnici `412`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("dear", "friend")
    ht.add("read", "book")

    expected_value = {
        "dear": "friend",
        "read": "book"
    }

    assert ht.collection.get(412) == expected_value
    `) 
})
```

Kada je ključ u hash tabeli, metoda `remove()` bi trebalo da obriše dati ključ i njegovu odgovarajuću vrednost iz kolekcije.

```js
({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")

    expected_value_before_removal = {
        "golf": "sport"
    }

    index = ht.hash("golf")

    assert ht.collection.get(index) == expected_value_before_removal

    ht.remove("golf")

    assert "golf" not in ht.collection.get(index, {})
    `) 
})
```

Kada par ključ-vrednost za `'golf', 'sport'` postoji u hash tabeli, `HashTable().lookup('golf')` bi trebalo da vrati `sport`.

```js
({ 
    test: () => runPython(`
    ht = HashTable()

    ht.add("golf", "sport")

    expected_value = "sport"
    assert ht.lookup("golf") == expected_value
    `) 
})

```

Kada par ključ-vrednost za `'golf', 'sport'` nije u kolekciji, `HashTable().lookup('golf')` bi trebalo da vrati `None`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.lookup("golf") is None
`) })

```

Kada je ključ `'fcc'` u kolekciji, `HashTable().lookup('cfc')` bi trebalo da vrati `None`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("fcc", "coding")
        assert ht.lookup("cfc") is None
`) })
```

Kada dodate `('rose', 'flower')` u hash tabelu, njegovo svojstvo elementa `collection` bi trebalo da izgleda ovako: `{ 441: { 'rose': 'flower' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("rose", "flower")

    expected_value = {
        "rose": "flower"
    }

    assert ht.collection.get(441) == expected_value

    # prevent hardcoding

    ht2 = HashTable()

    ht2.add("kebab", "food")
    expected_value = {
        "kebab": "food"
    }

    assert ht2.collection.get(501) == expected_value
    `)
})
```

Kada dodate ključ koji generiše istu vrednost hasha kao postojeći ključ, kao `fcc` i `cfc`, `collection`, trebalo bi da izgleda ovako: `{ 300: { 'fcc': 'coding', 'cfc':  'chemical' }}`.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("fcc", "coding")
    ht.add("cfc", "chemical")

    expected_value = {
        "fcc": "coding",
        "cfc": "chemical"
    }

    assert ht.collection.get(300) == expected_value

    # prevent hardcoding

    ht2 = HashTable()
    ht2.add("cat", "animal")
    ht2.add("act", "verb")

    expected_value = {
        "cat": "animal",
        "act": "verb"
    }

    assert ht2.collection.get(312) == expected_value
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class HashTable:
    def __init__(self):
        self.collection = {}

    def hash(self, string):
        hashed = 0
        for char in string:
            hashed += ord(char)  
        return hashed

    def add(self, key, val):
        the_hash = self.hash(key)
        if the_hash not in self.collection:
            self.collection[the_hash] = {}
        self.collection[the_hash][key] = val

    def remove(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            del self.collection[the_hash][key]
            if not self.collection[the_hash]: 
                del self.collection[the_hash]

    def lookup(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            return self.collection[the_hash][key]
        return None
```
