---
id: 66d75dd0aa65a71600dc669b
title: Napravi program za upravljanje inventarom/skupljinom
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

U ovoj laboratoriji, kreiraćeš program za upravljanje zalihama koji će ti omogućiti da dodaješ, izmeniš, pretražuješ i uklanjaš artikle iz zaliha. Koristićeš niz objekata da predstavljaš svoj inventar, gde će svaki objekat imati ime i količinu kao ključ.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba deklarisati prazan niz pod nazivom `inventory` koji će čuvati proizvode sa ključem `name` i vrednošću u obliku malih slova, i ključ `quantity` i vrednošću celog broja.
1. Treba kreirati funkciju pod nazivom `findProductIndex` koja prima naziv proizvoda kao argument i vraća referencu na odgovarajući proizvod unutar niza `inventory`. Funkcija mora uvek da koristi mala slova naziva proizvoda za izvršavanje pretrage. Ako proizvod nije pronađen, funkcija treba da vrati `-1`.
1. Treba kreirati funkciju pod nazivom `addProduct` koja prima proizvod kao argument.
1. Ako je proizvod već u inventaru, funkcija `addProduct` treba da ažurira vrednost `quantity` dodavanjem količine prosleđene funkciji na trenutni `quantity` i da u konzolu napiše naziv proizvoda praćen razmakom i `quantity updated`.
1. Ako proizvod nije u inventaru, funkcija `addProduct` treba da doda proizvod u niz `inventory` i da u konzolu napiše naziv proizvoda praćen razmakom i `added to inventory`.
1. Treba kreirati funkciju pod nazivom `removeProduct` koja prima naziv proizvoda i količinu kao argumente.
1. Funkcija `removeProduct` treba da oduzme prosleđenu količinu od odgovarajućeg proizvoda u inventaru i da u konzolu napiše niz znakova `Remaining <product-name> pieces: <product-quantity>`, gde se `<product-name>` mora zameniti nazivom proizvoda, a `<product-quantity>` mora zameniti količinom proizvoda.
1. Ako je preostala količina nula, `removeProduct` treba da ukloni proizvod iz inventara. Ako je količina u inventaru manja od potrebne za prodaju, funkcija `removeProduct` treba da napiše niz znakova `Not enough <product-name> available, remaining pieces: <product-quantity>` u konzolu.
1. Ako proizvod koji se uklanja nije u inventaru, funkcija `removeProduct` treba da napiše `<product-name> not found` u konzolu.

**Napomena:** Da biste spriječili konflikte, sačuvajte samo sadržaj naveden u pričama korisnika prilikom pokretanja testova.

# --hints--

Trebalo bi da objavite prazan red nazvan `inventory`.

```js
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

Trebalo bi da imaš funkciju nazvanu `findProductIndex`.

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")` mora da vrati reprezentaciju objekta koji ima `name` jednak `"flour"` unutar redova `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")` treba da vrati svojstvo objekta koje je `name` jednako `"flour"` unutar reda `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

`findProductIndex("Flour")` treba da vrati `-1` kada nema ništa sa `name` jednako `"flour"` pronađeno u redu od `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

Trebalo bi da imaš funkciju nazvanu `addProduct`.

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})` mora da doda `5` na vrednost od `quantity` objekta koji ima `name` jednak `"flour"` u redu `inventory`.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`addProduct({name: "FLOUR", quantity: 5})` bi trebalo da napiše `flour quantity updated` kada nešto sa `name` jednako `"flour"` se pronađe u redu od `inventory`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

`addProduct({name: "FLOUR", quantity: 5})` treba da gura `{name: "flour", quantity: 5}` na red od `inventory` kada nema ništa sa `name` jednako `"flour"` pronađeno u računu.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

`addProduct({name: "FLOUR", quantity: 5})` treba da napiše `flour added to inventory` kada nema ničega sa `name` jednako `"flour"` dostupno u računu.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

Treba da kreiraš funkciju pod nazivom `removeProduct`.

```js
assert.isFunction(removeProduct);
```

`removeProduct("FLOUR", 5)` treba da napiše `flour not found` kada nema ništa sa `name` jednako kao `"flour"` dostupno u redu `inventory`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)` treba da izvuče `5` iz vrednosti `quantity` objekta koji ima `name` jednak `"flour"` unutar niza `inventory`.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)` treba da se napiše `Remaining flour pieces: 15` u konzoli, kada je `inventory` jednako `[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

Ako je `quantity` nakon oduzimanja nula, `removeProduct` treba da ukloni stavku proizvoda iz računa.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)` treba da napiše `Not enough flour available, remaining pieces: 5` kada `inventory` je jednako `[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// Example usage

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```
