---
id: 689f7996426c1534fa5bea66
title: Napravi prevodilac za golf bodove
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

U igri golfa, svaki ručak ima `par`, što znači prosečan broj `strokes` koji je golfaš očekivan da postigne kako bi spustio loptu u ručak i završio igru. U zavisnosti od udaljenosti visoko ili nisko od `par`, tvoji rezultati za `strokes` imaju drugačije nadimakće.

U ovoj laboratoriji, ćeš napisati funkciju koja menja `par` i `strokes` u njihovo nadimakće.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratorijum.

Priče o korisniku:

1. Treba da kreirate funkciju pod nazivom `golfScore`.
1. `golfScore` treba da prihvati dva numerička argumenta: par terena i broj udaraca.
1. `golfScore` treba da vrati string.
1. `golfScore` treba da vrati `"Hole-in-one!"` ako je `strokes` jednako `1`.
1. `golfScore` treba da vrati `"Eagle"` ako je `strokes` manje ili jednako `par - 2`.
1. `golfScore` treba da vrati `"Birdie"` ako je `strokes` jednako `par - 1`.
1. `golfScore` treba da vrati `"Par"` ako je `strokes` jednako `par`.
1. `golfScore` treba da vrati `"Bogey"` ako je `strokes` jednako `par + 1`.
1. `golfScore` treba da vrati `"Double Bogey"` ako je `strokes` jednako `par + 2`.
1. `golfScore` treba da vrati `"Go Home!"` ako je `strokes` veće ili jednako `par + 3`.

# --hints--

Trebalo bi da kreirate funkciju koja se zove `golfScore`.

```js
assert.isFunction(golfScore);
```

`golfScore` treba da uzme dva parametra.

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` treba da vrati sekvencu slova

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)` treba da vrati sekvencu slova `Hole-in-one!`

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)` treba da vrati sekvencu slova `Hole-in-one!`

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)` trebalo bi da vrati sekvencu slova `Hole-in-one!`

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)` bi trebalo da vrati niz slova `Hole-in-one!`

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)` treba da vrati niz znakova `Eagle`

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` treba da vrati niz znakova `Eagle`

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)` treba da vrati sekvencu karaktera `Birdie`

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)` treba da vrati sekvencu slova `Birdie`

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)` treba da vrati niz znakova `Birdie`

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)` treba da vrati niz karaktera `Par`

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)` treba da vrati sekvencu karaktera `Par`

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)` treba da vrati niz znakova `Par`

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)` treba vratiti niz znakova `Bogey`

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)` trebalo bi vratiti niz znakova `Bogey`

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)` trebalo bi da vrati niz znakova `Bogey`

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)` trebalo bi da vrati sekvencu karaktera `Double Bogey`

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)` trebalo bi da vrati niz znakova `Double Bogey`

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)` trebalo bi da vrati sekvencu karaktera `Double Bogey`

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)` treba da vrati niz karaktera `Go Home!`

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)` treba da vrati niz karaktera `Go Home!`

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)` treba da vrati niz znakova `Go Home!`

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
