---
id: 68e6698b29e562d21b1d260c
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, kreći ćete spisak glavnih veb pregledača. Pružen vam je HTML šablon.

Počnite dodavanjem naslova na vašu stranicu koji piše `List of Major Web Browsers` koristeći element `h1` unutar elementa `body`.

# --hints--

You should have an `h1` element inside the `body` element.

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

Your `h1` element should contain the text `List of Major Web Browsers`.

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
        
    --fcc-editable-region--
    </body>
</html>
```
