---
id: 68ecfa1003d88990bb96a33d
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Liste opisa sastavljene su od liste pojmova i detalja. Element `dt` predstavlja opisni pojam, dok element `dd` predstavlja detalj opisa.

Prvi pretraživač koji ćete dodati u svoju listu opisa biće za Google Chrome.

Unutar vašeg elementa `dl`, dodajte element `dt` koji sadrži tekst `Google Chrome`.

Ispod vašeg elementa `dt`, dodajte element `dd` sa sledećim tekstom:

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

You should have a `dt` element inside of your `dl` element.

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

Your `dt` element should have the text `Google Chrome`.

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

You should have a `dd` element below your `dt` element.

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

Your `dd` element should contain the text `This is a free web browser developed by Google and first released in 2008.`

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```

