---
id: 68da6138d70960c197ec5eec
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Pre nego što izvršite glavni algoritam, morate validirati ulaz. Parametar `pairs` mora biti ceo broj, jer predstavlja broj standardnih parova zagrada za generisanje.

Dodajte deklaraciju `if` na početak vaše funkcije kako biste proverili da li je `pairs` pun broj. Koristite funkciju `isinstance()` za to.

Ako je situacija istinita, ponovi sekvencu slova `The number of pairs should be an integer`.

# --hints--

Trebalo bi da imaš deklaraciju `if` u funkciji `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

Vaša deklaracija za `if` treba da proveri da li `pairs` nije kompletan broj koristeći `not isinstance(pairs, int)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

Trebalo bi da vratite niz znakova `The number of pairs should be an integer` kada unos nije kompletan broj.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
--fcc-editable-region--
    
--fcc-editable-region--
    
    return []
```
