---
id: 68da6138d70960c197ec5ef1
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Sada moraš utvrditi koliko je dugo sastavio kompletan set standardnih nosača. Kompletan set ima dužinu jednaku dvostrukom broju parova (jer svaki par daje otvarajući i zatvarajući nosač).

Unutar vašeg ciklusa od `while`, dodajte deklaraciju `if` ili proverite da li `len(current) == 2 * pairs`. Kada je ovaj uslov tačan, dodajte `current` u listu `result`.

# --hints--

Trebalo bi da imaš deklaraciju `if` unutar svog ciklusa `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

Vaša izjava za `if` bi trebalo da proveri kao `len(current) == 2 * pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

Trebalo bi da dodaš `current` u listu `result` unutar tvog novog izjava `if`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
--fcc-editable-region--
        
--fcc-editable-region--

    return result

print(gen_parentheses(1))
```
