---
id: 68da6138d70960c197ec5eed
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Stoga, morate potvrditi da je broj parova najmanje jedan, jer ne možete kreirati kombinaciju standardnih zagrada sa nula ili negativan broj parova.

Dodaj još jednu deklaraciju za `if` da bi proverio da li je `pairs` manje od `1`. Ako je ovaj uslov istinit, vrati niz znakova `The number of pairs should be at least 1`.

# --hints--

Trebalo bi da imaš drugi deklarativni element za `if` u svojoj funkciji `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

Vaša druga deklaracija od `if` mora da proveri da li `pairs` je manje od `1`.

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

Treba da vratiš niz karaktera `The number of pairs should be at least 1` kada je `pairs` manji od `1`.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
--fcc-editable-region--
    
--fcc-editable-region--
    
    return []
```
