---
id: 68da6138d70960c197ec5ef2
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Ako trenutna sekvenca karaktera još nije kompletirana, morate istražiti moguće sledeće stanja. Dodajte `else` u svoj izraz `if`.

Unutar bloka `else`, obradićete dodavanje početnih vitičastih zagrada. Dodajte izjavu `if` da proverite da li `opens_used < pairs`. Ovo osigurava da ne koristite više početnih vitičastih zagrada nego što je dozvoljeno.

Ako je ovaj uslov tačan, dodaj novi tuple na `queue`: `(current + '(', opens_used + 1, closes_used)`. Ovo predstavlja stanje nakon dodavanja standardnih otvarajućih zagrada.

# --hints--

Trebalo bi da imaš deo `else` za izraz `if` unutar konteksta `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse
`) })
```

Trebalo bi da imaš izraz `if` unutar bloka `else` kada proveravaš da li je `opens_used < pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("opens_used < pairs")
`) })
```

Morate dodati `(current + '(', opens_used + 1, closes_used)` u folder kada je uslov tačan.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("queue.append((current + '(', opens_used + 1, closes_used))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
--fcc-editable-region--
        
--fcc-editable-region--
    
    return result

print(gen_parentheses(1))
```
