---
id: 68da6138d70960c197ec5ef3
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Sada morate obraditi dodavanje zatvarajućih zagrada. Glavno pravilo je da možete dodati zatvarajuću zagradu samo ako osigurava balans, to jest, mora postojati više otvarajućih nego zatvarajućih zagrada.

Još unutar bloka `else`, dodaj drugi stavak od `if` ili proveriti da li `closes_used < opens_used`.

Ako je ovaj uslov istinit, dodaj još jedan novi tuple u `queue`: `(current + ')', opens_used, closes_used + 1)`. Ovo predstavlja stanje nakon dodavanja standardnih zatvarajućih zagrada.

# --hints--

Moraš imati drugu deklaraciju od `if` unutar tvog bloka `else`.

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

Vaša druga izjava konfigurisana unutar `if` mora proveriti da li je kao `closes_used < opens_used`.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

Treba da dodaš `(current + ')', opens_used, closes_used + 1)` u fasciklu kada je uslov tačan.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
--fcc-editable-region--
            
--fcc-editable-region--

    return result

print(gen_parentheses(1))
```
