---
id: 68da604b8b6012a33e2344fb
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovoj radionici, ćete implementirati funkciju koja generiše sve validne kombinacije binarnih zagrada koristeći metod širine prvenog pretraživanja (BFS). Na primer, validna kombinacija binarnih zagrada za par dve je `(())` i `()()`.

Počnite kreiranjem funkcije nazvane `gen_parentheses` sa jednim parametrom `pairs`. Za sada, vratite prazan spisak iz te funkcije.

# --hints--

Trebalo bi da definišete funkciju nazvanu `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

Tvoja funkcija `gen_parentheses` bi trebalo da ima jedan parametar nazvan `pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

Tvoja funkcija treba da vrati prazan niz/listu.

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
