---
id: 68da6138d70960c197ec5eee
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Sada ćete inicijalizovati strukturu podataka za skladištenje vaših rezultata. Kreirajte parametar nazvan `result` i počnite sa praznom listom. Ova lista će čuvati sve važeće kombinacije redovnih zagrada koje generišete.

Ispravite svoje podatke za return ili vratite `result` umesto prazne liste.

# --hints--

Trebalo bi da deklariš parametar nazvan `result` u svojoj funkciji `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

Trebalo bi da inicijalizujete `result` kao prazan spisak.

```js
({ test: () => runPython(`
_result = _Node(_code).find_function("gen_parentheses").find_variable("result")
assert _result.is_equivalent("result = []") or _result.is_equivalent("result = list()")
`) })
```

Trebalo bi da vratiš `result` na kraju svoje funkcije.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
--fcc-editable-region--
    
    return []
--fcc-editable-region--
```
