---
id: 68da6138d70960c197ec5ef0
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Sada ćete implementirati glavni BFS ciklus. Kreirajte ciklus za `while` koji nastavlja projekat `queue` haijajaja (kada se `queue` procenjuje da je `True`).

Unutar te petlje, ispisi `queue`.

# --hints--

Trebalo bi da imaš ciklus `while` u svojoj funkciji `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

Tvoj ciklus od `while` mora imati uslov `queue`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

Trebalo bi da ispisaš `queue` unutar svoje petlje od `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    
--fcc-editable-region--

    return result
```
