---
id: 68da6138d70960c197ec5eef
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Korišćenjem metode pretraživanja po širini (BFS), koristićeš red za praćenje različitih stanja kada kreiraš kombinacije standardnih struktura. Svako stanje će biti predstavljeno kao tupel sa tri elementa:

- Trenutna sekvenca karaktera
- Broj korišćenih otvarajućih uglastih zagrada do sada
- Broj korišćenih zatvarajućih uglastih zagrada do sada

Definiši parametar nazvan `queue` i počni sa listom koja sadrži jedan tupel: `('', 0, 0)`. Ovo predstavlja početno stanje sa praznim nizom znakova i korišćenim nultim zagradama.

# --hints--

Treba da deklariš parametar nazvan `queue` u tvojoj funkciji `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

Trebalo bi da inicijalizujete `queue` sa listom koja ima tupel `('', 0, 0)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
--fcc-editable-region--
    
--fcc-editable-region--
    result = []
    return result
```
