---
id: 68e00c5349cbddfef6dbc053
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Unutar vaše petlje ``while``, koristite ``queue.pop(0)`` da uklonite i dobijete prvi element iz redova. Ovo implementira ponašanje "prvo uđe, prvo izađe" (FIFO), poznato u BFS.

Otvori ovaj tuple sa tri parametra: `current`, `opens_used`, i `closes_used`.

# --hints--

Trebalo bi da koristiš `queue.pop(0)` da dobiješ prvi element iz liste.

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

Trebalo bi da otvorite te rezultate u tri parametra: `current`, `opens_used`, i `closes_used`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
--fcc-editable-region--
        
--fcc-editable-region--

    return result
```
