---
id: 587d774c367417b2b2512a9d
title: Znajte kada treba ostaviti alt tekst praznim
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9P4t2'
forumTopicId: 301019
dashedName: know-when-alt-text-should-be-left-blank
---

# --description--

U prethodnom zadatku, naučili ste da je uključivanje `alt` atributa na `img` tagovima obavezno. Međutim, ponekad su slike već opisane kroz okolni tekstualni sadržaj, ili se koriste samo za dekoracione svrhe. U ovim slučajevima, alt tekst može izgledati suvišno ili nepotreban.

Kada je slika već opisana tekstualnim sadržajem ili ne dodaje značenje na stranici, `img` i dalje treba `alt` atribut, ali može biti postavljen kao prazan string. Evo primera:

```html
<img src="visualDecoration.jpeg" alt="">
```

Dekorativne slike obično padaju ispod 'dekoracije' kategorije. Međutim, često se koriste za CSS pravila, i zato nisu deo markup oznaka koje čitač ekrana procesira.

**Napomena:** Za slike sa opisima, možda i dalje želite da uključite alt tekst jer pomaže pretraživačima da indeksiraju sadržaj slike.

# --instructions--

Camper Cat je napisao kod za stranicu dizajna za blog deo svog sajta. Planira da doda vizuelne pauze između svoje dva članka sa dekorativnom slikom samurajskih mačeva. Dodajte `alt` atribut na `img` tag i postavite ga kao prazan string. (Primetite da slika u `src` nema stvarnu putanju do fajla - ne brinite što nema mačeva koji se prikazuju.)

# --hints--

Vaš `img` tag treba da ima `alt` atribut.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

`alt` atribut treba da bude postavljen kao prazan string.

```js
assert.isEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg" alt="">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>