---
id: 587d774c367417b2b2512a9c
title: Dodajte tekstualnu alternativu slikama za pristupačnost osobama sa oštećenjem vida
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7VfD'
forumTopicId: 16628
dashedName: add-a-text-alternative-to-images-for-visually-impaired-accessibility
---

# --description--

Možda ste već videli `alt` atribut na `img` tagu u prethodnim zadacima. `alt` tekst opisuje sadržaj slike i pruža tekstualnu alternativu za tu sliku. `alt` atribut pomaže u slučajevima kada slika nije učitana ili korisnik ne može da je vidi. Pretraživači takođe koriste ovaj atribut da bi razumeli šta slika sadrži kako bi je mogli da uključe u rezultate pretrage. Evo primera:

```html
<img src="importantLogo.jpeg" alt="Company logo">
```

Osobe sa oštećenjem vida zavise od čitača ekrana da bi pretvorile sadržaj na webu u audio interfejs. Oni neće dobiti informacije ako su prikazane samo vizuelno. Za slike, čitači ekrana mogu da pristupe `alt` atributu i pročitaju sadržaj da bi pružili važne informacije.

Dobar `alt` tekst pruža kratak opis slike. Trebalo bi uvek da uključite `alt` atribut na vaše slike. Prema HTML5 specifikaciji, ovo se sada smatra obaveznim.

# --instructions--

Camper Cat je programer i pravi ninja, koji gradi sajt za deljenje svog znanja. Slika profila koju želi da koristi pokazuje njegove veštine i treba da bude cenjena od strane svih posetilaca sajta. Dodajte `alt` atribut na `img` tag koji opisuje da Camper Cat radi karate. (Slika u `src` nema stvarnu putanju do fajla, tako da trebalo da vidite kako se `alt` tekst prikazuje.)

# --hints--

Vaš `img` tag treba da ima `alt` atribut koji nije prazan.

```js
assert.isNotEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<img src="doingKarateWow.jpeg">
```

# --solutions--

```html
<img src="doingKarateWow.jpeg" alt="Someone doing karate">