---
id: 587d7787367417b2b2512aa1
title: Poboljšajte navigaciju čitača ekrana sa header landmark elementom
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB76vtv'
forumTopicId: 301023
dashedName: make-screen-reader-navigation-easier-with-the-header-landmark
---

# --description--

Još jedan HTML5 element koji dodaje semantički značaj i poboljšava pristupačnost je `header` tag. Koristi se za obuhvat informacija za uvod ili navigacione linkove za svoje roditeljski element, i dobro radi za okruženje sadržaja koji se ponavljaju na vrhu na više stranica.

`header` deli landmark element koji ste videli sa `main`, omogućavajući tehnologijama za pomoć da brzo navigiraju kroz taj sadržaj.

**Napomena:** `header` je namenjen za korišćenje unutar `body` tela vašeg HTML dokumenta. Razlikuje se od `head` elementa, koji sadrži naslov stranice, meta informacije, itd.

# --instructions--

Camper Cat piše odlične članke o obuci za ninje, i želi da doda stranicu za njih na svoj sajt. Promenite `div` tag na vrhu koji trenutno sadrži `h1` u `header` tag umesto njega.

# --hints--

Vaš kod treba da ima jedan `header` element.

```js
assert.lengthOf(document.querySelectorAll('header'),1);
```

Vaši `header` elementi treba da omotaju `h1`.

```js
const header = document.querySelector('header');
const children = header?.querySelectorAll(`:scope ${'h1'}`);
assert.lengthOf(children , 1);
```

Vaš kod ne treba da ima nijedan `div` element.

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

Vaš `header` element treba da ima closing tag.

```js
assert.isTrue(code.match(/<\/header>/g)?.length === code.match(/<header>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>

  <div>
    <h1>Training with Camper Cat</h1>
  </div>


  <main>
    <section id="stealth">
      <h2>Stealth & Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>

  <header>
    <h1>Training with Camper Cat</h1>
  </header>


  <main>
    <section id="stealth">
      <h2>Stealth & Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>