---
id: 587d778f367417b2b2512aae
title: Dajte značenje linkovima korišćenjem opisnog teksta linka
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437DcV'
forumTopicId: 301013
dashedName: give-links-meaning-by-using-descriptive-link-text
---

# --description--

Korisnici čitača ekrana imaju različite opcije za tipove sadržaja koje njihov uređaj čita. Ove opcije uključuju preskakanje do (ili preko) elemenata navigacije, preskakanje do glavnog sadržaja, ili dobijanje pregleda stranice iz naslova sekcija. Jedna opcija je da čuju samo linkove koji su na stranici.

Čitači ekrana to rade tako što čitaju tekst linka, ili ono što je između `a` tagova. Imati listu linkova "kliknite ovde" ili "pročitajte više" ne pomaže. Umesto toga, koristite kratak, ali opisni tekst unutar `a` tagova da biste pružili više značenja za ove korisnike.

# --instructions--

Tekst linka koji Camper Cat koristi nema dovoljno opisa bez konteksta. Pomerite `a` tagove da omotaju tekst `information about batteries` umesto `Click here`.

# --hints--

Vaš kod treba da pomeri `a` tagove da omotaju tekst `information about batteries` umesto `Click here`.

```js
assert.match(document.querySelector('a')?.textContent, /^(information about batteries)$/g);
```

`a` element treba da ima `href` atribut sa praznim string vrednošću `""`.

```js
assert.isEmpty(document.querySelector('a')?.getAttribute('href'));
```

`a` element treba da ima closing tag.

```js
assert.isTrue(code.match(/<\/a>/g)?.length === code.match(/<a href=(''|"")>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. <a href="">Click here</a> for information about batteries</p>
  </article>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. Click here for <a href="">information about batteries</a></p>
  </article>
</body>