---
id: 587d778a367417b2b2512aa5
title: Poboljšajte pristupačnost grafikona pomoću figure elementa
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMqtE'
forumTopicId: 301015
dashedName: improve-chart-accessibility-with-the-figure-element
---

# --description--

HTML5 je predstavio `figure` i `figcaption` elemente koji su povezani. Kada se koriste zajedno, ovi elementi omotavaju predstavljanje slike (kao što je slika, dijagram, ili grafikon) zajedno sa njegovim opisom slike. Omotavanje ovih elemenata zajedno pruža dvostruki pristup pristupačnosti povezujući po značenju povezani sadržaj i pružajući tekstualnu alternativu koja opisuje `figure`.

Za predstavljanje podataka kao što su grafikoni, opis slike može biti korišćen za prikazivanje sažetog trenda ili zaključka za korisnike sa oštećenjem vida. Naredni izazovi se tiču kako da izbacite verziju tabele podataka grafikona sa ekrana (koristeći CSS) za korisnike čitača ekrana.

Evo primera - zapamtite da `figcaption` ide unutar `figure` tagova i može biti povezan sa drugim elementima:

```html
<figure>
  <img src="roundhouseDestruction.jpeg" alt="Photo of Camper Cat executing a roundhouse kick">
  <br>
  <figcaption>
    Master Camper Cat demonstrates proper form of a roundhouse kick.
  </figcaption>
</figure>
```

# --instructions--

Camper Cat naporno radi na izradi grafikonu sa grupisanim stubcima koji prikazuje vreme po nedelji za obuku u potaji, borbu i oružje. Pomozite mu da poboljša strukturu svoje stranice promenom `div` taga koji je koristio u `figure` tag, i `p` taga koji omotava opis slike u `figcaption` tag.

# --hints--

Vaš kod treba da ima jedan `figure` element.

```js
assert.lengthOf(document.querySelectorAll('figure') , 1);
```

Vaš kod treba da ima jedan `figcaption` element.

```js
assert.lengthOf(document.querySelectorAll('figcaption') , 1);
```

Vaš kod ne treba da ima `div` elemente.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

Vaš kod ne treba da ima `p` elemente.

```js
assert.lengthOf(document.querySelectorAll('p') , 0);
```

`figcaption` treba da bude dete `figure` elementa.

```js
const figure = document.querySelector('figure');
const children = figure?.querySelectorAll(`:scope ${'figcaption'}`);
assert.lengthOf(children, 1);
```

Vaš `figure` element treba da ima closing tag.

```js
assert.isTrue(
    code.match(/<\/figure>/g)?.length === code.match(/<figure>/g)?.length
);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth & Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>

      <!-- Only change code below this line -->
      <div>
        <!-- Stacked bar chart will go here -->
        <br>
        <p>Breakdown per week of time to spend training in stealth, combat, and weapons.</p>
      </div>
      <!-- Only change code above this line -->

    </section>
    <section id="stealth">
      <h2>Stealth & Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth & Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>
      <figure>
        <!-- Stacked bar chart will go here -->
        <br>
        <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
      </figure>
    </section>
    <section id="stealth">
      <h2>Stealth & Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>