---
id: 587d778f367417b2b2512aad
title: >-
  Izbegnite probleme sa slepilom za boje pažljivim izborom boja koje prenose informacije
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437as3'
forumTopicId: 301011
dashedName: >-
  avoid-colorblindness-issues-by-carefully-choosing-colors-that-convey-information
---

# --description--

Postoji više tipova slepila za boje. Oni mogu da variraju od smanjene osetljivosti na određene talasne dužine svetlosti do potpune nemogućnosti razlikovanja boja. Najčešći oblik je smanjena osetljivost na zelene nijanse.

Na primer, ako su dve slične nijanse zelene boje i boja teksta i boja pozadine vašeg sadržaja, korisnik sa slepilom za boje možda neće moći da ih razlikuje. Slične boje možete posmatrati kao susedne na krugu boja, a takve kombinacije treba izbegavati kada boja nosi važnu informaciju.

**Napomena:** Neki online alati za izbor boja uključuju vizuelne simulacije kako boje izgledaju osobama sa različitim tipovima slepila za boje. Ovo su korisni resursi zajedno sa online kalkulatorima za proveru odnosa kontrasta boja.

# --instructions--

Camper Cat isprobava različite stilove za važno dugme, ali žuta (`#FFFF33`) `background-color` i zelena (`#33FF33`) `color` boja teksta su susedne boje na krugu boja i gotovo se ne razlikuju za neke korisnike sa slepilom za boje. (Njihova slična osvetljenost takođe ne prolazi proveru odnosa kontrasta boja.) Promenite `color` boju teksta u tamnoplavu (`#003366`) da biste rešili oba problema.

# --hints--

Vaš kod treba da promeni `color` teksta za `button` u tamnoplavu boju.

```js
const button = document.querySelector('button');
const buttonColor = window.getComputedStyle(button).color; 
assert.equal(buttonColor, 'rgb(0, 51, 102)');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  button {
    color: #33FF33;
    background-color: #FFFF33;
    font-size: 14px;
    padding: 10px;
  }
  </style>
</head>
<body>
  <header>
    <h1>Danger!</h1>
  </header>
  <button>Delete Internet</button>
</body>
```

# --solutions--

```html
<head>
  <style>
    button {
      color: #003366;
      background-color: #FFFF33;
      font-size: 14px;
      padding: 10px;
    }
  </style>
</head>
<body>
  <header>
    <h1>Danger!</h1>
  </header>
  <button>Delete Internet</button>
</body>
```
