---
id: 587d778b367417b2b2512aa8
title: Dodajte pristupačni birač datuma
challengeType: 0
videoUrl: 'https://scrimba.com/c/cR3bRbCV'
forumTopicId: 301008
dashedName: add-an-accessible-date-picker
---

# --description--

Forme često uključuju `input` element, koji se može koristiti za kreiranje različitih kontrola forme. `type` atribut na ovom elementu pokazuje koji tip input elementa će biti kreiran.

Možda ste već videli `text` i `submit` tipove inputa u prethodnim zadacima, i HTML5 je predstavio `date` opciju za specificiranje input polja za datum. Prema podršci pretraživača, date picker se pojavljuje na input polju kada je u fokusu, što čini popunjavanje forme lakšim za sve korisnike.

Za starije pretraživače, tip će biti `text` po default-u, tako da je korisno prikazati korisnicima očekivani format datuma u `label` ili `placeholder` tekstu za upozorenje.

Evo primera:

```html
<label for="input1">Enter a date:</label>
<input type="date" id="input1" name="input1">
```

# --instructions--

Camper Cat organizuje Mortal Kombat turnir i želi da pita svog takmičara koji je najbolji datum za održavanje. Dodajte `input` element sa `type` atributom sa vrednošću `date`, `id` atributom sa vrednošću `pickdate`, i `name` atributom sa vrednošću `date`.

# --hints--

Vaš kod treba da doda jedan `input` element za polje za izbor datuma.

```js
assert.lengthOf(document.querySelectorAll('input'), 2);
```

Vaš `input` element treba da ima `type` atribut sa vrednošću `date`.

```js
assert.equal(document.querySelector('input')?.getAttribute('type'), 'date');
```

Vaš `input` element treba da ima `id` atribut sa vrednošću `pickdate`.

```js
assert.equal(document.querySelector('input')?.getAttribute('id'),'pickdate');
```

Vaš `input` element treba da ima `name` atribut sa vrednošću `date`.

```js
assert.equal(document.querySelector('input')?.getAttribute('name'), 'date');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>

        <!-- Only change code below this line -->



        <!-- Only change code above this line -->

        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>
        <input type="date" id="pickdate" name="date">
        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>