---
id: 587d778a367417b2b2512aa6
title: Poboljšajte pristupačnost polja forme pomoću label elementa
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMMAN'
forumTopicId: 301016
dashedName: improve-form-field-accessibility-with-the-label-element
---

# --description--

Poboljšanje pristupačnosti korišćenjem semantičkih HTML oznaka uključuje korišćenje imena labela i odgovarajućih atributa. Naredni izazovi će se baviti nekim važnim slučajevima korišćenja atributa u formama.

`label` element omotava tekst specifičnog kontrolnog elementa forme, obično ime ili oznaku opcije. Ovo povezuje značenje sa elementom i čini formu lakšom za čitanje. Atribut `for` na `label` elementu eksplicitno povezuje `label` sa kontrolom forme i koristi ga čitač ekrana.

Naučili ste o radio dugmadima i njihovim labelama u lekciji osnovni HTML. U toj lekciji, omotali smo element inputa radio dugmeta unutar `label` elementa zajedno sa tekstom labele kako bi tekst mogao da se klikne. Drugi način da postignete ovo je korišćenjem `for` atributa, kako je objašnjeno u ovoj lekciji.

Vrednost `for` atributa mora biti ista kao vrednost `id` atributa kontrolnog elementa forme. Evo primera:

```html
<form>
  <label for="name">Name:</label>
  <input type="text" id="name" name="name">
</form>
```

# --instructions--

Camper Cat očekuje veliko interesovanje za svoje blog postove sa zanimljivim mislima i želi da doda formu za prijavu na email. Dodajte `for` atribut na `label` za email koji odgovara `id` na njegovom `input` elementu.

# --hints--

Vaš kod treba da ima `for` atribut na label elementu koji nije prazan.

```js
assert.isNotEmpty(document.querySelector('label')?.getAttribute('for'));
```

Vrednost vašeg `for` atributa treba da odgovara vrednosti `id` na email input elementu.

```js
assert.equal(document.querySelector('label')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label>Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>