---
id: 587d7788367417b2b2512aa2
title: Olakšajte navigaciju čitača ekrana sa nav landmark elementom
challengeType: 0
videoUrl: 'https://scrimba.com/c/czVwWSv'
forumTopicId: 301024
dashedName: make-screen-reader-navigation-easier-with-the-nav-landmark
---

# --description--

`nav` element je još jedan HTML5 element sa semantičkim markup oznakama dizajniranim za lakšu navigaciju čitača ekrana. Ovaj tag je namenjen za omatanje glavnih navigacionih linkova na vašoj stranici.

Ako ima ponavljajućih linkova sajta na dnu stranice, nije ih nužno omotati `nav` tagom. Korišćenje `footer` elementa (koji će biti obrađen u sledećem zadatku) je dovoljno.

# --instructions--

Camper Cat je ubacio navigacione linkove na vrhu svoje stranice za učenje, ali ih je omotao `div` elementom. Promenite `div` u `nav` element da poboljšate pristupačnost na vašoj stranici.

# --hints--

Vaš kod treba da ima jedan `nav` element.

```js
assert.lengthOf(document.querySelectorAll('nav') , 1);
```

Vaši `nav` elementi treba da omotaju `ul` i njegove elemente liste.

```js
const nav = document.querySelector('nav');
const children = nav?.querySelectorAll(`:scope ${'ul'}`);
assert.lengthOf(children,1);
```

Vaš kod ne treba da ima nijedan `div` element.

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

Vaš `nav` element treba da ima closing tag.

```js
assert.isTrue(code.match(/<\/nav>/g)?.length === code.match(/<nav>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training with Camper Cat</h1>

    <div>
      <ul>
        <li><a href="#stealth">Stealth & Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </div>

  </header>
  <main>
    <section id="stealth">
      <h2>Stealth & Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training with Camper Cat</h1>

    <nav>
      <ul>
        <li><a href="#stealth">Stealth & Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>

  </header>
  <main>
    <section id="stealth">
      <h2>Stealth & Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>