---
id: 587d7790367417b2b2512ab0
title: Koristite tabindex da dodate fokus tastature elementu
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMDHW'
forumTopicId: 301027
dashedName: use-tabindex-to-add-keyboard-focus-to-an-element
---

# --description--

HTML atribut `tabindex` ima tri različite funkcije koje se odnose na fokus tastature elementu. Kada je na tagu, on označava da element može primiti fokus. Njegova vrednost (pozitivan ceo broj, negativan ili nula) određuje njegovo ponašanje.

Neki elementi, kao što su linkovi i kontrole formi, automatski primaju fokus tastature kada korisnik klikne tab kroz stranicu. To je u istom redosledu kao što se elementi pojavljuju u HTML kodu izvora. Isto ponašanje može biti dodeljeno i drugim elementima, kao što su `div`, `span`, i `p`, stavljanjem `tabindex="0"` atributa na njih. Evo primera:

```html
<div tabindex="0">I need keyboard focus!</div>
```

**Napomena:** Negativna vrednost `tabindex` (obično -1) označava da element može primiti fokus, ali joj nije dostupan preko tastature. Ovaj pristup se obično koristi za donošenje fokusa sadržaju programski (kao kada se `div` koji se koristi za pop-up prozor aktivira), a nalazi se van opsega ovih izazova.

# --instructions--

Camper Cat je kreirale novi anketu da prikupi informacije od svojih korisnika. On zna da input polja automatski primaju fokus tastature, ali želi da osigura da njegovi korisnici sa tastaturom ostanu fokusirani na upustva kada tab-uju kroz stvari. Dodajte `tabindex` atribut na `p` tag koji drži upustva formi i postavite njegovu vrednost na `0`. Bonus - koriščenje `tabindex` omogućava `li` pseudo-klasu CSS `:focus` da radi sa `p` tagom.

# --hints--

Vaš kod trebao bi da doda `tabindex` atribut na `p` tag koji sadrži upustva za formu.

```js
assert.isNotNull(document.querySelector('p')?.getAttribute('tabindex'));
```

Vaš kod trebao bi da postavi `tabindex` atribut na `p` tagu na vrednost 0.

```js
assert.equal(document.querySelector('p')?.getAttribute('tabindex'), '0');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p>Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p tabindex="0">Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
