---
id: 587d7789367417b2b2512aa4
title: Poboljšajte dostupnost audio sadržaja sa audio elementom
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJVkcZ'
forumTopicId: 301014
dashedName: improve-accessibility-of-audio-content-with-the-audio-element
---

# --description--

HTML5 `audio` element ima semantičko značenje kada omotava audio ili sadržaj audio strima u vašim oznakama. Audio sadržaj takođe treba tekstualni ekvivalent kako bi bio dostupan osobama koje nisu gluvo ili imaju teškoće sa sluhom. To se može uraditi sa tekstom blizu na stranici ili linkove ka transkripciji.

`audio` tag podržava `controls` atribut. Ovo prikazuje osnovne kontrole preglednika za reprodukciju, pauzu itd., i podržava funkcionalnost tastature. To je boolean atribut, što znači da mu nije potrebna vrednost - samo njegove prisustvo na tagu aktivira tu postavku.

Evde je primer:

```html
<audio id="meowClip" controls>
  <source src="audio/meow.mp3" type="audio/mpeg">
  <source src="audio/meow.ogg" type="audio/ogg">
</audio>
```

**Napomena:** Medijski sadržaj obično ima i vizuelne i slušne komponente. Trebaju im uparene slike i tekstualne verzije kako bi se osobe sa vizuelnim i/ili slušnim invaliditetom mogle uključiti. Obično nije zadatak web programera da kreira opis slike ili tekstualnu verziju, ali mora biti svestan kako ih uključiti.

# --instructions--

Vreme je za predah od Camper Cat-a i upoznavanje njegovog saputnika Zersiax-a (@zersiax), zagovornika pristupačnosti i korisnika čitača ekrana. Da bi čuo kako njegov čitač ekrana radi, dodaj element `audio` nakon elementa `p`. Uključi atribut `controls`. Zatim unutar elementa `audio` dodaj element `source` sa atributom `src` postavljenim na `https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3` i atributom `type` postavljenim na `"audio/mpeg"`.

**Napomena:** Audio klip može da se čini brz i teško razumljiv, ali to je uobičajena brzina za korisnike čitača ekrana.

# --hints--

Vaš kod treba da ima jedan `audio` element.

```js
assert.lengthOf(document.querySelectorAll('audio'),1);
```

Vaš `audio` element treba da ima zatvarajući tag.

```js
assert.match(code,/<audio.*>[\s\S]*<\/audio>/g);
assert.lengthOf(code.match(/<\/audio>/g),1);
```

`audio` tag trebao bi da ima `controls` atribut.

```js
assert.exists(document.querySelector('audio')?.getAttribute('controls'));
```

Vaš kod trebao bi da ima jedan `source` element.

```js
assert.lengthOf(document.querySelectorAll('source'), 1);
```

Vaš `source` element trebao bi biti unutar `audio` tagova.

```js
const audio = document.querySelector('audio');
const children = audio.querySelectorAll(`:scope ${'source'}`);
assert.lengthOf(children,1);
```

Vrednost `src` atributa na `source` tagu trebala bi da se tačno poklapa sa linkom navedenim u uputstvima.

```js
assert.equal(
  document.querySelector('source')?.getAttribute('src'),
  'https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3'
);
```

Vaš kod trebao bi da uključi `type` atribut na `source` tagu sa vrednošću audio/mpeg.

```js
assert.equal(document.querySelector('source')?.getAttribute('type'), 'audio/mpeg');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>



  </main>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>
    <audio controls>
      <source src="https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3" type="audio/mpeg"/>
    </audio>
  </main>
</body>
```
