---
id: 587d7790367417b2b2512aaf
title: Poboljšajte navigaciju linkova pomoću HTML access keys
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQvmaTp'
forumTopicId: 301021
dashedName: make-links-navigable-with-html-access-keys
---

# --description--

HTML pruža `accesskey` atribut za postavljanje prečica na tastaturi za aktiviranje ili fokusiranje određenog elementa. Dodavanje `accesskey` atributa može olakšati navigaciju za korisnike koji koriste samo tastaturu.

HTML5 dozvoljava ovaj atribut da se koristi na bilo kom elementu, ali je posebno važan kada se koristi na interaktivnim elementima. To uključuje linkove, dugmad, i kontrole forme.

Evo primera:

```html
<button accesskey="b">Important Button</button>
```

# --instructions--

Camper Cat želi da linkovi koji okružuju dva naslova blog postova imaju prečice na tastaturi kako bi korisnici njegovog sajta mogli brzo da naviguju do punog teksta. Dodajte `accesskey` atribut na oba linka i postavite prvi na `g` (za Garfield) i drugi na `c` (za Chuck Norris).

# --hints--

Vaš kod treba da doda `accesskey` atribut na `a` element sa `id` vrednošću `first`.

```js
assert.isTrue(document.querySelector('#first')?.hasAttribute('accesskey'));
```

Vaš kod treba da doda `accesskey` atribut na `a` element sa `id` vrednošću `second`.

```js
assert.isTrue(document.querySelector('#second')?.hasAttribute('accesskey'));
```

Vaš kod treba da postavi `accesskey` atribut na `a` element sa `id` vrednošću `first` na `g`. Zapamtite da velika i mala slova su važna.

```js
assert.equal(document.querySelector('#first')?.getAttribute('accesskey'), 'g');
```

Vaš kod treba da postavi `accesskey` atribut na `a` element sa `id` vrednošću `second` na `c`. Zapamtite da velika i mala slova su važna.

```js
assert.equal(document.querySelector('#second')?.getAttribute('accesskey'), 'c');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" accesskey="g" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" accesskey="c" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>