---
id: 587d774e367417b2b2512aa0
title: Omotajte sadržaj u article element
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp79S3'
forumTopicId: 301029
dashedName: wrap-content-in-the-article-element
---

# --description--

`article` je jedan od novih HTML5 elemenata koji dodaju semantički značaj vašim markup oznakama. `article` je element za deljenje sekcija i koristi se za nezavisne, samostalne sadržaje. Ovaj tag je posebno pogodan za blog unose, postove na forumima, ili vesti članke.

Da biste proverili da li je sadržaj dovoljno nezavisan, obično je korisno postaviti dva pitanja. Pitajte se da li, ako uklonite sav kontekst koji ga okružuje, da li bi sadržaj i dalje imao smisao? Takođe, za tekstualne sadržaje, da li bi sadržaj ostao smislen ako bi bio u RSS toku?

Imajte u vidu da ljudi koji koriste asistivne tehnologije zavise od dobro strukturisanih i semantički značajnih markup oznaka za bolje razumevanje vašeg rada.

**Napomena:** `section` je takođe novi element u HTML5, i ima blago drugačiji semantički značaj od `article`. `article` je za nezavisne sadržaje, dok je `section` za grupisanje povezanih sadržaja u kontekstu. Oni se mogu koristiti jedan unutar drugog, u zavisnosti od potreba. Na primer, ako je knjiga `article`, onda je svaki poglavlje `section`. Kada nema veze među grupama sadržaja, koristite `div`.

`<div>` - grupše sadržaj
`<section>` - grupše povezane sadržaje
`<article>` - grupše nezavisne, samostalne sadržaje

# --instructions--

Camper Cat je koristio `article` oznake za obavješćivanje postova na svojoj blog stranici, ali je zaboravio da ih koristi na najgornjem postu. Promenite `div` oznaku da koristi `article` oznaku umesto nje.

# --hints--

Vaš kod treba da ima tri `article` oznake.

```js
assert.lengthOf(document.querySelectorAll('article') ,3);
```

Vaš kod ne treba da ima nijednu `div` oznaku.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <div>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </div>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>