---
id: 587d7790367417b2b2512ab1
title: Koristi tabindex da odrediš redosled fokusa tastature za više elemenata
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzRRcb'
forumTopicId: 301028
dashedName: use-tabindex-to-specify-the-order-of-keyboard-focus-for-several-elements
---

# --description--

Svojstvo `tabindex` takođe definiše eksplicitni redosled tabovanja za elemente. Ovo se dešava kada je vrednost svojstva postavljena na pozitivan broj (1 ili veće).

Postavljanje `tabindex="1"` dovodi fokus tastature na taj element prvi. Zatim se prolazi kroz definisane vrednosti `tabindex` (2, 3, itd.) pre nego što se pređe na podrazumevane elemente i one sa `tabindex="0"`.

Važno je napomenuti da kada se redosled taba postavi ovako, on nadjačava podrazumevani redosled (koji proizlazi iz HTML strukture). To može zbuniti korisnike koji očekuju da navigacija počne od vrha stranice. Ova tehnika može biti korisna u određenim situacijama, ali zbog pristupačnosti je koristite oprezno.

Evo primera:

```html
<div tabindex="1">I get keyboard focus, and I get it first!</div>
```

```html
<div tabindex="2">I get keyboard focus, and I get it second!</div>
```

# --instructions--

Camper Cat ima polje za pretragu na svojoj stranici sa Inspirativnim citatima koje želi da postavi u gornji desni ugao koristeći CSS. Želi da kontrole obrasca — `input` za pretragu i `input` za slanje — budu prva dva elementa u redosledu taba. Dodaj atribut `tabindex` sa vrednošću `1` na `input` sa id-jem `search`, i atribut `tabindex` sa vrednošću `2` na `input` sa id-jem `submit`.

Još jedna stvar za razmotriti je da neki pregledači mogu umetnuti fokus unutar vašeg redosleda taba kada je element dodat klikom. Element dodat na stranicu osigurava da ćeš uvek početi na početku svog redosleda taba.

# --hints--

Tvoj kod treba da doda atribut `tabindex` na `input` element sa id-jem `search`.

```js
assert.isNotNull(document.querySelector('#search')?.getAttribute('tabindex'));
```

Tvoj kod treba da doda atribut `tabindex` na `input` element sa id-jem `submit`.

```js
assert.isNotNull(document.querySelector('#submit')?.getAttribute('tabindex'));
```

Tvoj kod treba da postavi vrednost `tabindex` na `input` elementu sa id-jem `search` na `1`.

```js
assert.equal(document.querySelector('#search')?.getAttribute('tabindex'), '1');
```

Tvoj kod treba da postavi vrednost `tabindex` na `input` elementu sa id-jem `submit` na `2`.

```js
assert.equal(document.querySelector('#submit')?.getAttribute('tabindex'), '2');
```

# --seed--

## --seed-contents--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input type="search" name="search" id="search">
    <input type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```

# --solutions--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input tabindex="1" type="search" name="search" id="search">
    <input tabindex="2" type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```
