---
id: 68814221fc4752e17b9696bd
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Kada korisnik klikne na belešku, izvrši izmene, a zatim klikne van beleške, biće vam potreban način da dobijete tu tek prepravljenu verziju. Ovo možete postići korišćenjem `blur` događaja. Ovaj događaj se aktivira kada element izgubi fokus. Evo primera:

```js
element.addEventListener("blur", () => {
  // do something here
});
```

Priložite metod `addEventListener` na varijablu `noteEl`. Prvi argument treba da bude događaj `"blur"`, a drugi argument treba da bude strelica funkcija (arrow function). Unutar tela te strelice funkcije, koristite `const` da kreirate novu varijablu nazvanu `newContent` i dodelite joj vrednost `noteEl.innerHTML`.

`innerHTML` se ovde koristi kako bi pomogao u očuvanju razmaka i formatiranja za belešku.

# --hints--

You should attach an `addEventListener` method to the `noteEl` variable.

```js
assert.match(code, /noteEl\.addEventListener/);
```

Your `addEventListener` method should have a first argument of the `"blur"` event.

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],/);
```

Inside of the body of the arrow function, you should use `const` to create a new variable called `newContent` and assign it the value of `noteEl.innerHTML`.

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],\s*\(\)\s*=>\s*{[^}]*const\s+newContent\s*=\s*noteEl\.innerHTML[^}]*}\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>
    
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const noteEl = document.getElementById("note");
const statusEl = document.getElementById("status");

let currentContent = "";

--fcc-editable-region--
  
--fcc-editable-region--

window.addEventListener("DOMContentLoaded", () => {
  currentContent = noteEl.textContent;
});
```
