---
id: 669e81368e52b3a5c35a2dc5
title: Kreiraj stranicu za kolekciju videa
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba imati element `main` kao jedini potomak elementa `body`.
1. Treba imati element `h1` sa temom vaše stranice.
1. Treba imati uvodni paragraf koji opisuje temu vaše stranice ispod vašeg elementa `h1`.
1. Treba imati tri elementa `section` ispod vašeg prvog paragrafa.
1. Svaki deo treba da ima element `h2`, paragraf i element `iframe`, u ovom redosledu.
1. Tri elementa `iframe` treba da imaju atribut `src` postavljen za validno video snimanje.
1. Svaki element `iframe` takođe treba da ima atribut `title` koji opisuje postavljen sadržaj, i atribute `height` i `width` za postavljanje elementa na odgovarajuću veličinu.

# --hints--

Trebalo bi da imaš element `main` unutar svog elementa `body`.

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

Tvoj element `main` mora biti jedino dete od elementa `body`.

```js
assert.equal(document.querySelector('body').children.length, 1);
```

Unutar elementa `main` treba da imaš element `h1` koji sadrži naslov tvoje stranice.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Trebalo bi da imate uvodni paragraf za temu vaše stranice ispod vašeg elementa `h1`.

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

Trebalo bi da imaš tri elementa `section` ispod tvog prvog elementa `p`.

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

Svaki element `section` mora početi sa elementom `h2` koji služi kao naslov tog dela.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

Svaki element sa `section` mora imati element `p` da identifikuje video sadržaj kao svoje drugo dete.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

Svaka instanca `section` mora imati instancu `iframe` kao svog trećeg potomka.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

Vaš prvi `iframe` treba da ima svojstvo `src` postavljeno za validni video.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

Tvoj prvi element od `iframe` treba da ima svojstvo `title` za opisivanje postavljenog sadržaja.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

Vaš element `iframe`` prvi treba da ima svojstvo od ``height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

Vaš prvi element `iframe` mora imati atribut `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

Vaš drugi element `iframe` mora imati atribut `src` podešen za validan video.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

Drugi vaš element `iframe` treba da ima atribut od `title` za opisivanje postavljenog sadržaja.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

Vaš element za `iframe`, drugi, treba da ima svojstvo od `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

Vaša druga komponenta od `iframe` mora imati svojstvo `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

Vaš treći element `iframe` mora imati svojstvo `src` postavljeno za validan video.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

Vaš treći element `iframe` treba da ima atribut `title` za opisivanje postavljenog sadržaja.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

Vaš treći element za `iframe` treba da ima svojstvo od `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

Vaš treći element `iframe` bi trebalo da ima svojstvo `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Front-End Web Development</h1>
        <p>Front-End Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Front-end
            developers ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
