---
id: a6b0bb188d873cb2c8729495
title: Konverzija specijalnih znakova u HTML entitete
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

Ovaj radionik se bavi konvertovanjem specijalnih znakova u nizu karaktera u njihove odgovarajuće entitete HTML.

**Cilj**: Popunite priče korisnika ispod i prođite sve testove kako biste završili ovaj rad.

**Priče korisnika:**

1. Trebalo bi da ima funkciju `convertHTML` koja prima niz karaktera kao argument.
2. Funkcija `convertHTML` treba da vrati novi niz karaktera tako što će zamijeniti specifične karaktere u argumentnom nizu sa njihovim odgovarajućim entitetima HTML.

- `&` mora biti zamenjeno sa `&amp;`.
   - `<` mora biti zamenjeno sa `&lt;`.
   - `>` mora biti zamenjeno sa `&gt;`.
   - `"` mora biti zamenjeno sa `&quot;`.
   - `'` mora biti zamenjeno sa `&apos;`.

# --hints--

Trebalo bi da imaš funkciju `convertHTML`.

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` trebalo bi vratiti sekvencu karaktera `Dolce &amp; Gabbana`.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` treba da vrati sekvencu karaktera `Hamburgers &lt; Pizza &lt; Tacos`.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` treba da vrati sekvencu karaktera `Sixty &gt; twelve`.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` treba da vrati sekvencu karaktera `Stuff in &quot;quotation marks&quot;`.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` bi trebalo da vrati sekvencu slova `Schindler&apos;s List`.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` trebalo bi da vrati sekvencu karaktera `&lt;&gt;`.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` trebalo bi vratiti niz znakova `abc`.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
