---
id: 69babdf3f896711c11516373
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Tokom svakog simulacionog ticka:

- Određeni broj posetilaca stiže na kapiju (iz njegovog `queue`).
- Kapija obrađuje posetioca na osnovu svog `capacity`.
- Ako stigne više posetilaca nego što kapija može da obradi, neki će ostati (preliv).

Sada ćete kreirati funkciju koja obrađuje ovu logiku za jednu kapiju u jednom tikotu.

Kreirajte praznu funkciju pod nazivom `processGateFlow` koja prima dva parametra:

- `gate`: Objekt kapije koji se obrađuje.
- `tickIndex`: Trenutna pozicija simulačkog tikota u nizu `queue`.

# --hints--

Trebalo bi da kreirate funkciju pod nazivom `processGateFlow`.

```js
assert.isFunction(processGateFlow);
```

Vaša funkcija `processGateFlow` trebalo bi da ima parametre `gate` i `tickIndex`.

```js
const regex = __helpers.functionRegex('processGateFlow', ['gate', 'tickIndex']);
assert.match(__helpers.removeJSComments(processGateFlow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
