---
id: 69bac067cd15c4685d07906d
title: Korak 15
challengeType: 1
dashedName: step-15
---

# --description--

Sada kada imate funkcije za obradu pojedinačnih tikova i rukovanje prelivom, vreme je da kreirate kontrolersku funkciju za jednu kapiju. Ova funkcija će obrađivati sve akcije za jednu kapiju tokom jednog tikova simulacije.

Kreirajte praznu funkciju pod nazivom `handleGateAtTick` sa parametrima:

- `gates`: Puni niz objekata kapija.
- `gate`: Trenutna obrađivana kapija.
- `tickIndex`: Trenutni indeks tikova za simulaciju.
- `throughputSummary`: Objekt koji prati ukupan broj obrađenih posetioca po kapiji.

# --hints--

Trebalo bi da kreirate funkciju pod nazivom `handleGateAtTick`.

```js
assert.isFunction(handleGateAtTick);
```

Funkcija vašeg `handleGateAtTick` treba da ima parametre `gates`, `gate`, `tickIndex` i `throughputSummary`.

```js
const regex = __helpers.functionRegex('handleGateAtTick', ['gates', 'gate', 'tickIndex', 'throughputSummary']);
assert.match(__helpers.removeJSComments(handleGateAtTick.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
