---
id: 69baafd6d683dd9921b75db3
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovom radionici, koristićete JavaScript da simulirate protok posetilaca na muzičkom festivalu.

Radit ćete sa dva skupa podataka koji opisuju ulazne kapije festivala tokom jutarnjih i noćnih smena:

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

Svaki objekat kapije sadrži sledeće svojstva:

- `id`: String koji identifikuje kapiju.
- `capacity`: Broj posetilaca koje kapija može obraditi po tiku.
- `queue`: Niz brojeva koji predstavlja koliko posetilaca stiže na kapiju tokom određenog tik-a.

Da biste počeli, kreirajte praznu funkciju pod nazivom `initializeThroughput` koja prima parametar pod nazivom `gates`.

# --hints--

Trebalo bi da kreirate funkciju pod imenom `initializeThroughput`.

```js
assert.isFunction(initializeThroughput);
```

Funkcija `initializeThroughput` bi trebalo da ima parametar `gates`.

```js
const regex = __helpers.functionRegex('initializeThroughput', ['gates']);
assert.match(__helpers.removeJSComments(initializeThroughput.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
