---
id: 69babfd65c4c4632a0c76f63
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Kada kapija ne može da obradi sve prisutne osobe tokom tika, neke prisutne osobe ostaju u redu. Da bi se rešio ovaj preliv, kreiraćete funkciju koja preusmerava prisutne osobe koje su u prelivu na drugu kapiju.

Kreirajte praznu funkciju pod nazivom `rerouteOverflow` sa sledećim parametrima:

- `gates`: Puni niz objekata kapija.
- `currentGate`: Kapija gde je došlo do prelivanja.
- `tickIndex`: Trenutna pozicija tik-a u redu.
- `overflowAmount`: Broj prisutnih osoba koje nije moguće obraditi.

# --hints--

Trebalo bi da kreirate funkciju pod imenom `rerouteOverflow`.

```js
assert.isFunction(rerouteOverflow);
```

Funkcija vašeg `rerouteOverflow` treba da ima parametre `gates`, `currentGate`, `tickIndex` i `overflowAmount`.

```js
const regex = __helpers.functionRegex('rerouteOverflow', ['gates', 'currentGate', 'tickIndex', 'overflowAmount']);
assert.match(__helpers.removeJSComments(rerouteOverflow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

--fcc-editable-region--

--fcc-editable-region--
```
