---
id: 66a83e5e491625454b6f62c3
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Formularima se sastoji od `inputs` gde korisnici mogu uneti svoje podatke. Možete grupisati povezane ulaze zajedno koristeći element `fieldset`.

Evo primera korišćenja elementa `fieldset`:

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

Unutar vašeg elementa `form`, dodajte element `fieldset`.

# --hints--

You should have an opening `fieldset` tag.

```js
assert.match(code, /<fieldset>/i);
```

You should have a closing `fieldset` tag.

```js
assert.match(code, /<\/fieldset>/i);
```

Your `fieldset` element should be inside the `form` element.

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">        
--fcc-editable-region--
        
--fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
