---
id: 66a84111965a0c46df6bbd0a
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

U prethodnim lekcijama, naučili ste kako da povežete element `label` sa elementom `input` na ovaj način:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

Atribut `for` na elementu `label` treba da odgovara atributu `id` na elementu `input`. Ovo se naziva eksplicitna asocijacija.

Ispod vašeg elementa `legend`, dodajte element `label` sa tekstom `Name (required):`. Postavite njegov atribut `for` na vrednost `"full-name"`.

Zatim ispod vašeg elementa `label`, dodajte element `input` bez atributa. U sledećim koracima, dodaćete potrebne atribute.

# --hints--

You should have an opening `label` tag.

```js
assert.match(code, /<label.*>/);
```

You should have a closing `label` tag.

```js
assert.match(code, /<\/label>/);
```

Your `label` element should be placed below your `legend` inside the `fieldset` element.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

Your `label` element should have a `for` attribute with the value of `"full-name"`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

Your `label` element should have the text of `Name (required):`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

You should have an `input` element below your `label` element.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
