---
id: 66a83601cd819e37f0dccd14
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, veštaćete se sa radom sa HTML formularima kreiranjem obrasca za povratne informacije o hotelu.

Za uvodni tekst, želite da prikažete glavni naslov praćen kratkom napomenom o ostavljanju povratnih informacija. 

Unutar vašeg `body` elementa, dodajte `header` element.

Unutar `header` elementa, dodajte `h1` element. Tekst `h1` elementa treba da bude `Hotel Feedback Form`.

Ispod vašeg `h1` elementa, dodajte `p` element. Tekst `p` elementa treba da bude `Thank you for staying with us. Please provide feedback on your recent stay.`

# --hints--

You should have an opening `header` tag.

```js
assert.match(code, /<header>/i);
```

You should have a closing `header` tag.

```js
assert.match(code, /<\/header>/i);
```

You should have an opening `h1` tag.

```js
assert.match(code, /<h1>/i);
```

You should have a closing `h1` tag.

```js
assert.match(code, /<\/h1>/i);
```

Your `h1` element should have the text `Hotel Feedback Form`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

Your `h1` element should be inside the `header` element.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

You should have an opening `p` tag.

```js
assert.match(code, /<p>/i);
```

You should have a closing `p` tag.

```js
assert.match(code, /<\/p>/i);
```

Your paragraph element should have the text `Thank you for staying with us. Please provide feedback on your recent stay.`

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

Your paragraph element should be inside your `header`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

Your paragraph element should come after your `h1`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
