---
id: 66a83bdcf425e7446900b7c4
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

U prethodnim lekcijama, naučili ste kako da radite sa elementom `form` na ovaj način:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

Atribut `action` koristi se za određivanje gde bi podaci forme trebalo da budu poslati kada se forma pošalje.

Atribut `method` koristi se za određivanje HTTP metode koja treba da se koristi prilikom slanja podataka forme. Najčešće metode su `GET` i `POST`. 

**NAPOMENA**: Naučićete o tome kako rade HTTP metode u kasnijim modulima.

Unutar vašeg `main` elementa, dodajte `form` element sa atributom `action` podešenim na `"https://hotel-feedback.freecodecamp.org"` i atributom `method` podešenim na `"POST"`.

# --hints--

You should have an opening `form` tag.

```js
assert.match(code, /<form.*>/i);
```

You should have a closing `form` tag.

```js
assert.match(code, /<\/form>/i);
```

Your `form` element should have an `action` attribute set to `"https://hotel-feedback.freecodecamp.org"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

Your `form` element should have a `method` attribute set to `"POST"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
