---
id: 66a937e74920ba68ebe5e86d
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

U prethodnim lekcijama, naučili ste kako da radite sa `placeholder` i `required` atributima na ovaj način:

```html
<input type="text" placeholder="e.g., John Doe" required>
```

Za vaš postojeći element `input`, dodajte atribut `placeholder` sa vrednošću `"e.g., John Doe"`.

Takođe, dodajte atribut `required` elementu `input`.

# --hints--

You should not remove the existing `input` element.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Your `input` should have a `placeholder` attribute with the value of `"e.g., John Doe"`.

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

Your `input` should have a `required` attribute.

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
