---
id: 66a9419e2d18476c645ce693
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Broj `input` se koristi za kreiranje numeričkog polja za unos.

Evo primera numeričkog polja za unos:

```html
<input type="number" id="age" name="age" min="18" max="100">
```

Atributi `min` i `max` koriste se za postavljanje minimalnih i maksimalnih vrednosti koje mogu biti unete u polje za unos.

Ispod vašeg elementa `label`, dodajte unos sa atributom `type` podešenim na `"number"` i `id` od `"age"`.

Atribut `name` treba da bude podešen na `"age"`, atribut `min` treba da bude podešen na `"3"`, a atribut `max` treba da bude podešen na `"100"`.

# --hints--

You should have a number `input` below your `label` element.

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

Your number `input` should have an `id` of `age`.

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

Your number `input` should have a `name` attribute set to `age`.

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

Your number `input` should have a `min` attribute set to `"3"`.

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

Your number `input` should have a `max` attribute set to `"100"`.

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
      --fcc-editable-region--
          
      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
