---
id: 66a93730719e1f68410cce54
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Atribut `name` se koristi za identifikaciju podataka forme nakon što su poslati na server.

Evo primera kako koristiti atribut `name`:

```html
<input type="email" name="email">
```

Dodajte atribut `name` na element `input` sa vrednošću `"name"`.

# --hints--

You should not remove the existing `input` element.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

You should not create a new `input` element.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Tvoj element `input` treba da ima atribut `name` sa vrednošću `"name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
