---
id: 66a93c95bc58e26a8fe95818
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Zatim, dodajte `input` sa tipom `"email"` ispod vaše e-pošte `label`. Ovaj `input` treba da ima atribut `id`, postavljen na vrednost `"email"`. Takođe, postavite atribut `name` na vrednost `"email"`.

Ovaj `input` je takođe neophodan, pa se pobrinite da dodate atribut `required`.

Na kraju, dodajte atribut `placeholder`, postavljen na `"example@email.com"`.

# --hints--

You should have an `input` element below your email `label`.

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

Your `input` element should have a `type` of `"email"`.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

Your `input` element should have an `id` attribute set to `"email"`.

```js
assert.isNotNull(document.querySelector('input#email'));
```

Your `input` element should have a `name` attribute set to `"email"`.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

Your `input` element should have a `required` attribute.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

Your `input` element should have a `placeholder` attribute set to `"example@email.com"`.
  
```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
