---
id: 66a9521bc70162712caf118d
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Ako želite da korisnici izaberu jednu opciju iz liste opcija, možete koristiti skup radijumskih dugmadi (radio buttona).

Evo primera sa dva radijumska dugmeta:

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

U ovom primeru, radijumska dugmad su grupisana korišćenjem iste vrednosti atributa ``name``. To znači da se u bilo kom trenutku može izabrati samo jedno radijumsko dugme.

Ispod vašeg elementa ``legend``, dodajte dugme ``radio`` sa ``id`` podešenim na ``"yes-option"``, atributom ``name`` podešenim na ``"hotel-stay"``, i atributom ``value`` podešenim na ``"yes"``.

Ispod vašeg dugmeta ``radio``, dodajte element ``label`` sa tekstom ``Yes`` i atributom ``for`` podešenim na ``"yes-option"``.

# --hints--

You should have a `radio` button below your `legend` element.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

Your `radio` button should have an `id` attribute set to `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

Your `radio` button should have a `name` attribute set to `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

Your radio button should have a `value` attribute set to `"yes"`.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.value,
  'yes'
);
```

You should have a `label` element below your `radio` button.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

Your `label` element should have a `for` attribute set to `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

Your `label` element should have the text `Yes`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
