---
id: 67a51d3a8a6fe123b77b0c6e
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Elementi `input` mogu imati atribut `size`. Ovaj atribut definiše broj znakova koji bi trebalo da bude vidljiv dok korisnik unosi podatke u polje za unos. Vrednost `size` treba da bude nenegativan ceo broj veći od nule. Ako `size` nije navedeno, ili je navedeno sa nevažećom vrednošću, polje za unos će imati podrazumevanu širinu postavljenu od strane pregledača.

```html
<label for="lastName">Last Name:</label>
<input id="lastName" name="lastName" type="text" size="10" />
```

Dajte poljima za ime i e-mail atribut `size` sa vrednošću `"20"`.

# --hints--

You should give the name `input` a `size` attribute with a value of `"20"`.

```js
const label1 = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(label1?.getAttribute('size'), '20');
```

You should give the email `input` a `size` attribute with a value of `"20"`.

```js
const label2 = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(label2?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
