---
id: 699f1d72e906e4616155290d
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, ćete nastaviti da učite o specifikacijama tipova i kako funkcionišu sa funkcijama za izgradnju sigurnog kalkulatora za tipove.

Počnite kreiranjem funkcije nazvana `square` koja prima parametar i vraća kvadrat broja.

# --hints--

Trebalo bi da imaš funkciju koja se zove `square`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

Tvoja funkcija `square` treba da ima jedan parametar.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

Tvoja funkcija `square` treba da vrati kvadrat broja.

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
