---
id: 699f52220de88c8091032076
title: Korak 18
challengeType: 1
dashedName: step-18
---

# --description--

Kada otvorite konzolu, rezultat je `NaN` jer drugi argument za `exponent` nije bio dostavljen. Potrebno je nekoliko ažuriranja za funkciju `raiseTo` koje će ukloniti ovaj problem.

Prva promena je vraćanje `base ** 2`.

U sledećem koraku, videćete kako da rukujete opcionim parametrom za `exponent`.

# --hints--

Tvoja funkcija za `raiseTo` treba da vrati broj.

```js
assert.isNumber(raiseTo(4));
```

Tvoja funkcija za `raiseTo` treba da vrati vrednost `base` koja je dohvaćena pomoću snage/moći `2`.

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
