---
id: 699f3f636aa840b7201aa37c
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Kada otvorite konzolu, trebalo bi da vidite poruku o grešci za TypeScript sledeću:

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

Pošto je parametar jasno definisan kao broj, TypeScript ima mogućnost detekcije greške prilikom pokušaja prosleđivanja argumenta tipa niz znakova funkciji za `square`.

Promeni argument koji se prosleđuje funkciji `square` u broj `5` kako bi se ispravila greška i videli tačni rezultati na konzoli.

# --hints--

Poziv funkcije `square` treba da ima argument `5` umesto `"something"`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```
