---
id: 699f52fac083214b6bf51de6
title: Korak 19
challengeType: 1
dashedName: step-19
---

# --description--

Na ovom zadnjem koraku, moraćete da dodate proveru uslova za parametar `exponent` u funkciju `raiseTo`.

Ako se argument `exponent` ukloni, funkcija mora da vrati `base` koji je podignut sa snagom `exponent`. Ako se argument `exponent` ne ukloni, mora da vrati `base ** 2`.

Sa ovim poslednjim korakom, vaš uređaj za izračunavanje sa tipom sigurnosti je završen!

# --hints--

Vaša funkcija za `raiseTo` treba da vrati `base ** exponent` kada se argument/parametar `exponent` dostavi.

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

Tvoja funkcija za `raiseTo` mora vratiti `base ** 2` kada argument/parametar `exponent` nije dostupan.

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {

  return base ** 2;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }

  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
