---
id: 699f3738921fb79a651ded56
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Sada je vreme da pozoveš svoju funkciju.

Kreiraj parametar nazvan `result` i postavi mu vrednost na poziv funkcije `square` sa argumentom `5`.

Nakon ovoga, unesite vrednost `result` u konzolu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da imaš parametar nazvan `result`.

```js
assert.isDefined(result);
```

Vaše polje za `result` mora biti broj.

```js
assert.isNumber(result);
```

Tvoj parametar za `result` treba da bude kvadrat od `5`.

```js
assert.strictEqual(result, 25);
```

Trebalo bi da napišete parametar `result` u konzoli.

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
