---
id: 699f41b8da20f816b5c7dece
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Ako otvorite konzolu, bi trebalo da vidite sledeću poruku o grešci:

```md
Expected 1 arguments, but got 2.
```

Da biste rešili tu grešku, ažurirajte svoj poziv funkcije `square` ili uzmite samo jedan argument od `5`.

# --hints--

Vaš poziv funkcije `square` treba da ima samo jedan argument od `5`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```
