---
id: 699f4639ae0016caaf81b7c6
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Sada je vreme da testiraš svoju funkcionalnost.

Počnite kreiranjem parametra nazvanog `avgResult` i postavite mu vrednost za poziv funkcije `getAverage`. Argument tog poziva funkcije treba da bude `[2, 14, 26, 8]`.

Zatim unesite vrednost `avgResult` u konzolu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da imaš parametar nazvan `avgResult`.

```js
assert.isDefined(avgResult);
```

Treba da staviš `getAverage([2, 14, 26, 8])` u svoj parametar `avgResult`.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

Vaš parametar za `avgResult` mora biti broj.

```js
assert.isNumber(avgResult);
```

Treba da napišete `avgResult` u konzoli.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
