---
id: 699f412114c87221f16a99e3
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Još jedna prednost korišćenja TypeScript je što detektuje greške ako prosledite previše ili premalo argumenata za poziv funkcije.

Da biste ovo prikazali, ažurirajte vaš poziv funkcije za `square` ili uzmite dva argumenta za `5` i `10`.

# --hints--

Vaš poziv funkcije za `square` mora imati dva argumenta od `5` i `10`.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
