---
id: 699f513e01eb1768cc5df726
title: Korak 17
challengeType: 1
dashedName: step-17
---

# --description--

Sada je parametar za `exponent` postavljen kao opciono, možete pozvati funkciju samo sa jednim argumentom.

Ispravite vaš poziv funkcije za `raiseTo` ili unesite broj `3` samo.

# --hints--

Vaša funkcija za `raiseTo` treba da bude pozvana sa samo jednim argumentom od `3`.

```js
assert.match(
  code,
  /raiseTo\s*\(\s*3\s*\)/
);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}

--fcc-editable-region--
const raisedResult = raiseTo(2, 3);
--fcc-editable-region--
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
