---
id: 699f38937b053e4679878ff1
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Kada radiš sa JavaScript normalno, ništa ti ne sprečava da pozoveš funkciju `square` koristeći sekvencu karaktera ili niz podataka, što će rezultirati nepredvidivim rezultatima.

Da biste ovo prikazali, promenite poziv vaše funkcije koristeći niz znakova `"something"` umesto broja `5`.

U ovom kursu ćete naučiti kako da koristite TypeScript da povećate tipsku sigurnost u vašem kodu i sprečite greške poput ovih.

# --hints--

Vaš poziv funkcije za `square` mora imati argument sa `"something"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
