---
id: 699f434e294a9e11cf4f65f0
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Sada je vreme da dodate drugog kalkulator asistenta u vaš alatni projekat.

Počnite kreiranjem funkcije koja se zove `getAverage` i koja prima jedan parametar. Vaša funkcija treba da vrati prosek niza brojeva.

U narednim koracima, dodaćete informacije o tipu ovoj funkciji.

# --hints--

Trebalo bi da imaš funkciju pod nazivom `getAverage`.

```js
assert.isFunction(getAverage);
```

Vaša funkcija treba da primi jedan parametar.

```js
assert.strictEqual(getAverage.length, 1);
```

Funkcija mora da vrati broj.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Vaša funkcija mora da vrati tačne rezultate za bilo koji niz brojeva.

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
