---
id: 699f49669ea51a31d2ba89d8
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

Kao što ste zapamtili iz prethodnih lekcija, možete eksplicitno definisati tip povratne vrednosti za funkciju kao ovo:

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

U ovom primeru, tip povratne vrednosti funkcije `circleArea` je `number`. Ako pokušate da vratite različiti tip, TypeScript će prikazati grešku.

Ispravite svoju funkciju od `getAverage` ili eksplicitno postavite tip povratne vrednosti na `number`.

# --hints--

Vaša funkcija za `getAverage` treba da ima tip povratne vrednosti od `number`.

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
