---
id: 699f3cb8d057f29954205d7d
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Kao što se sećate iz prethodnih lekcija, možete dodati tipne opise za parametre funkcije na ovaj način:

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

U ovom primeru, TypeScript će prikazati grešku ako pokušate da pozovete funkciju `add` sa argumentima koji nisu brojevi.

Sahihisha svoj operater `square` ili uključite opis tipa za parametar `num`. Tip treba da bude `number`.

# --hints--

Trebalo bi da dodate tip opis za parametar `num` u funkciju `square`. Tip treba da bude `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
