---
id: 699f4ec939f10da62fb56c3e
title: Korak 15
challengeType: 1
dashedName: step-15
---

# --description--

Sada je vreme za testiranje vaše funkcionalnosti.

Definiši parametar nazvan `raisedResult` i postavi mu vrednost kao rezultat pozivanja `raiseTo` sa argumentima `2` i `3`.

Zatim, upišite `raisedResult` u konzolu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da imate parametar koji se zove `raisedResult`.

```js
assert.isDefined(raisedResult);
```

Vaše polje za `raisedResult` mora biti broj.

```js
assert.isNumber(raisedResult);
```

Treba da napišete `raisedResult` u konzolu.

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
