---
id: 699f4b903bf5427a599f791d
title: Korak 14
challengeType: 1
dashedName: step-14
---

# --description--

U nekoliko prethodnih koraka, kreiraćete još jednu funkciju kako biste nastavili da vežbate opisivanje tipa za funkcije.

Počnite kreiranjem funkcije nazvane `raiseTo` koja prima dva parametra nazvana `base` i `exponent`. Oba parametra moraju biti tipa `number`. Funkcija treba da vrati rezultat podizanja `base` sa snagom `exponent`. Tip povratne vrednosti funkcije takođe treba da bude prikazan kao `number`.

# --hints--

Trebalo bi da imaš funkciju `raiseTo`.

```js
assert.isFunction(raiseTo);
```

Funkcija vašeg `raiseTo` treba imati dva parametra, `base` i `exponent`.

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*(?::\s*\w+\s*)?\s*=>/
);
```

Vaši parametri za `base` i `exponent` moraju sve biti tipa `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

Tvoja funkcija `raiseTo` treba da vrati broj.

```js
assert.isNumber(raiseTo(2, 3));
```

Tvoja funkcija za `raiseTo` treba da vrati rezultat podizanja `base` sa snagom `exponent`.

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

Vaša funkcija za `raiseTo` treba da ima opis tipa povratne vrednosti od `number`.

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
