---
id: 699f4861dd5b5f164e1bda4a
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Još jedna prednost korišćenja TypeScript je osiguravanje da su implementirani i da su to oni tipovi povratnih vrednosti koje očekujete. Ponekad TypeScript može odrediti tip povratne vrednosti funkcije, ali ga možete eksplicitno definisati.

Da biste ovo prikazali, promenite tip povratne vrednosti funkcije `getAverage` u niz karaktera `"something"`.

# --hints--

Tvoja funkcija `getAverage` bi trebalo da vrati niz znakova.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

Vaša funkcija za `getAverage` treba da vrati niz znakova `"something"`.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
