---
id: 699f4a3e7b4c6bfa538cc4f7
title: Korak 13
challengeType: 1
dashedName: step-13
---

# --description--

Kada otvorite konzolu, videćete sledeću poruku o grešci:

```md
Type 'string' is not assignable to type 'number'.
```

Da biste rešili ovu kosu, morate zameniti return sa `sum / numbers.length`.

# --hints--

Tvoja funkcija za `getAverage` treba da vrati broj.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Funkcija `getAverage` treba da vrati prosek brojeva prisutnih u redu podataka.

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```
