---
id: 699f458e47fb88c756fdba81
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Trenutno, atribut `numbers` ima nejasnu vrstu `any`. Da biste ovo ispravili, možete postaviti jasnu vrstu tog atributa kao niz podataka brojeva.

Zapamtite iz prethodnih lekcija, da možete definisati tip numeričkog niza podataka kao što sledi:

```ts
const numbers: number[] = [1, 2, 3];
```

Ispravite svoju funkciju vašeg `getAverage` postavljanjem eksplicitnog tipa parametra `numbers` kao numeričkog niza podataka.

# --hints--

Vaša funkcija ``getAverage`` treba da ima parametar nazvan ``numbers``, koji je eksplicitno definisan kao numerički tip podataka.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
