---
id: 69a5f35669099ed52f8563b1
title: Napravi program za pametno dopunjavanje spajalice
challengeType: 26
dashedName: lab-smart-pantry-restocker
---

# --description--

U ovoj laboratoriji, ćeš izgraditi mali program za upravljanje pantrijom koristeći osnovne koncepte od JavaScript kao nizove podataka, objekte, cikluse i uslove.

Osigurati prijem tereta sa namirnicama iz ormarića, odlučujući šta uraditi sa svakim predmetom i organizujući stvari za skladištenje.

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

Red `rawData` ima sekvencu karaktera sa komponentama povezanim znakiem cevi (pipe) `|` za format `sku|name|qty|expires|zone`, gde je `zone` opcionalno.

**Priče korisnika:**

1. Moraš implementirati funkciju ``parseShipment(rawData)`` koja prima niz karakter sekvenci i vraća niz objekata sa atributima/svojstvima ``{ sku, name, qty, expires, zone }``.

- Vrednosti `sku` koje se ponavljaju u opterećenju treba ignorisati.
- Kada deo `zone` nije dostavljen, trebalo bi ga postaviti na `"general"`.
- Vrednost `qty` treba da se konvertuje u broj.

2. Treba izvršiti funkciju `planRestock(pantry, shipment)` koja poredi trenutni inventar sa dolaznim teretom i vraća listu akcija za tip `{ type, item }`, gde je `type` jedan od `"restock"`, `"discard"`, ili `"donate"`, a `item` je analizirani artikal tereta.

Parametar `pantry` je niz elemenata sa istom strukturom kao analizirani element opterećenja (`{ sku, name, qty, expires, zone }`).

- Ako artikal tereta ima `qty` od `0` ili manje, tip akcije `type` bi trebalo da bude `"discard"`, bez obzira da li je artikal u skladištu.
   - Inače, ako `sku` artikala tereta već postoji u skladištu, tip akcije `type` bi trebalo da bude `"restock"`.
   - Inače (ako `sku` artikala tereta nije u skladištu), tip akcije `type` bi trebalo da bude `"donate"`.

3. Treba da izvršite funkciju `groupByZone(actions)` koja prikuplja akcije u lokacije skladištenja na osnovu svojstva `zone` svakog elementa. Funkcija bi trebalo da vrati objekat tipa objekat, gde je svaki ključ ime lokacije, a vrednost niz akcija povezanih sa tom lokacijom. Na primer, ako akcije imaju stavke sa lokacijama `"fridge"` i `"pantry"`, rezultat bi trebalo da bude `{ fridge: [...], pantry: [...] }`.

4. Trebalo bi da izvršite funkciju `clonePantry(pantry)` koja vraća duboku kopiju panteona/skladišta kako bi promene sortiranja ne promenile originalnu listu. Duboka kopija znači kreiranje novog niza sa novim stavkama, tako da menjanje kopije ne menja originalno skladište.

5. Treba koristiti sve funkcije zajedno za obradu opterećenja i zapisivanje konačnog prikupljenog rezultata na konzolu.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da definišete funkciju koja se zove `parseShipment` i koja prima jedan red niz znakova pod nazivom `rawData` kao argument.

``` js
assert.isFunction(parseShipment);
assert.lengthOf(parseShipment, 1);
```

Vaša funkcija za `parseShipment` treba da promeni sekvencu karaktera opterećenja u redu u stavke sa atributima: `sku`, `name`, `qty`, `expires`, i `zone`.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);

assert.isArray(result);
assert.isObject(result[0]);
assert.deepEqual(result[0], {
  sku: "A10",
  name: "Tomatoes",
  qty: 5,
  expires: "2027-01-01",
  zone: "fridge"
});
```

SKU koji se ponavljaju u teretu treba ih zanemariti.

``` js
const rawData = [
  "C32|Eggs|3|2027-01-01",
  "C32|Eggs|3|2027-01-01"
];

const result = parseShipment(rawData);

assert.lengthOf(result, 1);
assert.strictEqual(result[0].sku, "C32");
```

Kada `zone` nije u sekvenci opterećenja, treba ih postaviti na `general`.

``` js
const rawData = ["B21|Bananas|10|2027-01-01"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].zone, "general");
```

Vrednost `qty` mora biti promenjena u broj.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].qty, 5);
assert.isNumber(result[0].qty);
```

Treba da definiš funkciju nazvanu `planRestock` koja prima dva parametra: `pantry` i `shipment`.

``` js
assert.isFunction(planRestock);
assert.lengthOf(planRestock, 2);
```

Funkcija za `planRestock` treba da vrati niz akcija sa svojstvima od `type` i `item`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 6, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.isArray(result);
assert.isObject(result[0]);
assert.property(result[0], "type");
assert.property(result[0], "item");
```

Predmeti sa količinom `0` ili manje u teretu bi trebalo da kreiraju akciju tipa `discard`, bez obzira na to da li je predmet u ormanu.

``` js
const pantry = [{ sku: "D43", name: "Pineapples", qty: 2, expires: "2027-01-01", zone: "general" }];
const shipment = [{ sku: "D43", name: "Pineapples", qty: 0, expires: "2027-01-01", zone: "general" },
                  { sku: "E54", name: "Peppers", qty: -1, expires: "2027-01-01", zone: "fridge" }
                ];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "discard");
assert.strictEqual(result[1].type, "discard");
```

Kada je artikl tereta već u skladištu (SKU jednak `sku`), akcija bi trebalo da bude `restock`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "restock");
```

Ako artikal tereta nije u ormanu, akcija bi trebalo da bude `"donate"`.

``` js
const pantry = [];
const shipment = [{ sku: "C32", name: "Eggs", qty: 3, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "donate");
```

Trebalo bi da definišete funkciju nazvanu `groupByZone` koja prima jedan parametar nazvan `actions`.

``` js
assert.isFunction(groupByZone);
assert.lengthOf(groupByZone, 1);
```

Tvoja funkcija za `groupByZone` treba da vrati prikupljene akcije u skladu sa svojstvom `zone` svakog objekta.

``` js

const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "restock", item: { sku: "B1", zone: "pantry" } }, 
];

const result = groupByZone(actions);

assert.isObject(result);
assert.isArray(result.fridge);
assert.isArray(result.pantry);
```

Vaša funkcija za `groupByZone` mora da prikuplja akcije sa ispravnim sadržajem i brojem.

``` js
const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "donate",  item: { sku: "B1", zone: "fridge" } },
  { type: "discard", item: { sku: "C1", zone: "pantry" } },
];

const result = groupByZone(actions);

assert.lengthOf(result.fridge, 2);
assert.lengthOf(result.pantry, 1);
assert.strictEqual(result.fridge[0].item.sku, "A1");
assert.strictEqual(result.pantry[0].type, "discard");
```

Trebalo bi da definišete funkciju nazvanu `clonePantry` koja prima jedan parametar nazvan `pantry`.

``` js
assert.isFunction(clonePantry);
assert.lengthOf(clonePantry, 1);
```

`clonePantry` trebalo bi da vrati novi red umesto prvobitnog reda police.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.isArray(pantryClone);
assert.notStrictEqual(pantryClone, pantry);
```

Predmeti u panteografskom prostoru, koji su kopirani, takođe bi trebali biti novi predmeti (duboka kopija).

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.notStrictEqual(pantryClone[0], pantry[0]);
assert.deepEqual(pantryClone[0], pantry[0]);

pantryClone[0].qty = 100;
assert.strictEqual(pantry[0].qty, 4);
```

Radnici moraju da rade zajedno kako bi obradili opterećenje i prikupili sve dešavajuće događaje.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];

const shipment = parseShipment(rawData);
const pantryClone = clonePantry(pantry);
const actions = planRestock(pantryClone, shipment);
const grouped = groupByZone(actions);

assert.isObject(grouped);
assert.strictEqual(grouped.fridge[0].item.qty, 5);
```

Trebalo bi da unesete prikupljene aktivnosti prema lokacijama za skladištenje.

``` js
const groupedLog = getLogs().find(log => {
  if (!log || typeof log !== "object" || Array.isArray(log)) return false;

  return Object.values(log).some(zoneActions => {
    if (!Array.isArray(zoneActions)) return false;

    return zoneActions.some(action =>
      action &&
      ["restock", "discard", "donate"].includes(action.type) &&
      action.item &&
      typeof action.item.sku === "string"
    );
  });
});

assert.isOk(groupedLog, "Log the grouped actions object to the console.");
```

# --seed--

## --seed-contents--

```js
const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01",
  "B21|Bananas|10|2027-01-01",
  "C32|Eggs|3|2027-01-01|fridge",
  "C32|Eggs|3|2027-01-01",
  "D43|Pineapples|0|2027-01-01",
  "E54|Peppers|-1|2027-01-01|fridge"
];
```

# --solutions--

```js

function parseShipment(rawData) {

  // Early return if rawData is not an Array
  if (!Array.isArray(rawData)) return [];

  const result = [];
  const seenSKUs = [];

  for (let line of rawData) {

    const parts = line.split("|");

    const sku = parts[0].trim();

    if (seenSKUs.includes(sku)) continue; // Avoid duplicates

    const name = parts[1].trim();
    const qty = Number(parts[2].trim()); // Parse qty into a number
    const expires = parts[3].trim();
    const zone = ( parts[4] ) ? parts[4].trim() : "general";

    result.push({
      sku: sku,
      name: name,
      qty: qty,
      expires: expires,
      zone: zone
    });

    seenSKUs.push(sku);

  }

  return result;
}

function planRestock(pantry, shipment) {
  const actions = [];

  for (let item of shipment) {

    var exists = false;

    for (let stock of pantry)
    {
      if (stock.sku === item.sku)
      {
        exists = true;
      }
    }

    if (item.qty > 0) {
      if (exists) {
        actions.push({ type: "restock", item: item });
      } else{
        actions.push({ type: "donate", item: item });
      }
    } else {
      actions.push({ type: "discard", item: item });
    }
  }

  return actions;
}

function groupByZone(actions) {

  const byZone = {};

  for (const action of actions) {

    const zone = action.item.zone;

    if (!byZone[zone]) {
      byZone[zone] = [];
    }

    byZone[zone].push(action);
  }

  return byZone;

}

function clonePantry(pantry) {
  const pantryClone = [];

  for (let item of pantry) {
    pantryClone.push({
      sku: item.sku,
      name: item.name,
      qty: item.qty,
      expires: item.expires,
      zone: item.zone
    });
  }

  return pantryClone;
}

const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01", // Restock existing item
  "B21|Bananas|10|2027-01-01", // Donate new item without zone
  "C32|Eggs|3|2027-01-01|fridge", // Donate to a defined zone
  "C32|Eggs|3|2027-01-01", // Duplicated SKU in shipment
  "D43|Pineapples|0|2027-01-01", // Discard with quantity of 0
  "E54|Peppers|-1|2027-01-01|fridge" // Discard even if it's not in pantry
];

const shipment = parseShipment(rawData);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

// Show items expanded in zones
console.log(grouped);

```
