---
id: 690ddb5e6ee94eb73ed172b0
title: Popravite formular za doprinose
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

Lokalna organizacija za pomoć napravila je sajt sa formularom za donacije, ali postoje nekoliko problema koje je potrebno ispraviti. Formular nije lako dostupan i ima sintaksičku grešku od HTML.

Tvoj zadatak je da ispraviš sve greške kako bi forma ispravno radila i bila pristupačna svim korisnicima. Ispuni zahteve iz korisničkih priča ispod i klikni na „Proveri svoj kod“ da vidiš da li si uklonio sve greške.

**Priče korisnika:**

1. Elementi `input` su elementi koji nemaju zatvorenu oznaku/opis. Uklonite sve zatvarajuće oznake za `</input>` iz formulara.
2. Dodajte elemente `label` u svako polje unosa formulara kako bi korisnici znali za šta je svako polje. Tekst etikete treba da odgovara onome što se nalazi pored svakog polja unosa.
3. Tip unosa `Email Address:` treba da bude `email` umesto `text`.
4. Morate povezati svaki element `label` sa odgovarajućim elementom `input` koristeći atribut `for` na `label` i odgovarajući atribut `id` na `input`.
5. Dodajte atribut `required` za tekstualna, e-mail i numerička polja unosa (ali ne za checkbox ili dugme za slanje) kako biste osigurali da korisnici unose potrebne podatke.

# --hints--

Potrebno je da imate zatvarajući tagove za `</input>` u vašem kodu.

```js
assert.notMatch(code, /<\/input>/);
```

Trebalo bi da imaš četiri elementa `input` u svom formularu.

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

Trebalo bi da imaš četiri elementa `label` u tvom formularu.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

`label` Vaš prvi treba da ima tekst od `Full Name:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

Tvoj prvi element `input` treba da ima atribut `required`.

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

Tvoj drugi element `label` treba da sadrži tekst `Email Address:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

Tvoj drugi element `input` treba da ima atribut `required`.

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

`label` Treći treba da ima tekst sa `Donation Amount ($):`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

Tvoj treći element `input` treba da ima atribut `required`.

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

`label` Tvoj četvrti deo treba da ima tekst od `Subscribe`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

Trebalo bi da povežete svaki element `input` sa elementom `label` koristeći atribute `for` i `id`.

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

`input` vaša e-mail adresa mora imati tip `email`.

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

Vaše polje za izbor/selektovanje ``input`` ne bi trebalo da ima svojstvo ``required``.

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

Tvoj dugme za slanje sa `button` ne bi trebalo da ima atribut `required`.

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
