---
id: 66f17db06803d11a1bd19a20
title: Igra jenga kratkog testa
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da kreirate niz podataka pod nazivom `questions`.
1. Niz podataka `questions` treba da sadrži najmanje pet elemenata, od kojih svaki mora imati ključeve `category`, `question`, `choices` i `answer`.
1. Ključ `category` mora imati string vrednost koja predstavlja grupu pitanja.
1. Ključ `question` mora imati string vrednost koja predstavlja pitanje.
1. Ključ `choices` mora imati vrednost niza podataka sa tri stringa, koji predstavljaju alternativne odgovore na pitanje.
1. Ključ `answer` mora imati string vrednost koja predstavlja tačan odgovor na pitanje. Takođe, vrednost `answer` treba da bude uključena u niz podataka `choices`.
1. Trebalo bi da postoji funkcija pod nazivom `getRandomQuestion` koja prima niz podataka pitanja kao argument i vraća nasumični objekat pitanja iz tog niza.
1. Trebalo bi da postoji funkcija pod nazivom `getRandomComputerChoice` koja prima niz dostupnih opcija kao argument i vraća nasumičan odgovor za izabrano pitanje.
1. Trebalo bi da postoji funkcija pod nazivom `getResults` koja prima objekat pitanja kao prvi argument i opciju računara kao drugi argument. Funkcija treba da vrati `The computer's choice is correct!` ako je odgovor tačan. U suprotnom, vraća `The computer's choice is wrong. The correct answer is: <correct-answer>`, gde je `<correct-answer>` vrednost tačnog odgovora za izabrano pitanje.

# --hints--

Treba da kreirate red podataka nazvan `questions`.

```js
assert.isArray(questions);
```

Podatni red za `questions` mora imati najmanje pet stavki, svaki sa ključevima `category`, `question`, `choices`, i `answer`.

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

Ključ za `category` treba da ima vrednost alfanumeričke sekvence koja predstavlja grupu pitanja.

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => {
    assert.isString(category);
    assert.isNotEmpty(category);
});
```

Ključ `question` mora imati vrednost niza slova koji predstavlja pitanje.

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
    assert.isNotEmpty(question.replaceAll('?', ''));
})
```

Ključ za `choices` treba da ima vrednost niza podataka sa sekvencom od tri različita slova i svako.

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    assert.lengthOf(new Set(choices), 3);
    choices.forEach(choice => {
        assert.isString(choice);
        assert.isNotEmpty(choice);
    });
})
```

Funguo ya `answer` mora imati vrednost sekvence slova.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => {
    assert.isString(answer);
    assert.isNotEmpty(answer);
});
```

Vrednost `answer` treba da bude uključena u red podataka `choices`.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

Trebalo bi da imaš funkciju nazvanu `getRandomQuestion` koja prima niz podataka pitanja kao argument i vraća nasumično pitanje iz tog niza.

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

Trebalo bi da imaš funkciju nazvanu `getRandomComputerChoice` koja prima niz dostupnih opcija kao argument i vraća slučajni odgovor na odabrano pitanje.

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

Trebalo bi da imaš funkciju nazvanu `getResults`.

```js
assert.isFunction(getResults);
```

Tvoja funkcija ``getResults`` treba da primi objekat pitanja kao prvi parametar i opciju kompjuter/računara kao drugi parametar.

```js
const testQuestion = questions[0];
const wrongChoice = questions[0].choices.find(choice => choice !== questions[0].answer);
assert.equal(getResults(testQuestion, testQuestion.answer), `The computer's choice is correct!`);
assert.equal(getResults(testQuestion, wrongChoice), `The computer's choice is wrong. The correct answer is: ${testQuestion.answer}`);
```

Ako opcija računara odgovara odgovoru, `getResults` treba da vrati `The computer's choice is correct!`

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

Ako opcija računara ne odgovara odgovoru, `getResults` mora vratiti `The computer's choice is wrong. The correct answer is: <correct-answer>`, gde je `<correct-answer>` vrednost tačnog odgovora za izabrano pitanje.

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

Tvoja funkcija `getResults` treba da koristi potpuno poređenje jednakosti, a ne porediti deo niza znakova.

```js
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Hamburger"}, "Ham"), `The computer's choice is wrong. The correct answer is: Hamburger`);
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Ham"}, "Hamburger"), `The computer's choice is wrong. The correct answer is: Ham`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
