---
id: 587d8259367417b2b2512c85
title: Implementacija algoritma Selection Sort
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

U ovom laboratorijumu ćete implementirati algoritma Selection Sort. Selection Sort radi tako što bira minimalnu vrednost u nizu i razmenjuje njenu poziciju sa prvim elementom u nizu. Zatim počinje od druge pozicije, birajući najmanju vrednost u preostalom nizu i razmenjujući njenu poziciju sa drugim elementom. Nastavlja da iterira kroz niz i razmenjuje elemente sve dok ne dođe do kraja niza. Sada je niz sortiran. Selection Sort ima vremensku složenost O(n^2) u svim slučajevima.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da ima funkciju za `selectionSort`.
2. Funkcija `selectionSort` treba da uzme niz podataka i vrati niz podataka sa istim elementima, ali sortirane od najmanjeg do najvećeg.

# --hints--

`selectionSort` Trebalo bi da bude funkcija.

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` treba da vrati nepromenjen red podataka, osim po redosledu.

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` treba da vrati sortirovan red podataka (od manjeg do većeg).

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort` ne sme da koristi metod ugrađen u `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
