---
id: 68cabc534d2e741f33a3fe1d
title: Građivač pratioca statistike lika u igri
challengeType: 27
dashedName: lab-game-character-stats
---

# --description--

U ovoj laboratoriji, napravićete alat za praćenje statistike lika. Program će vam omogućiti da kreirate lika sa specifičnim atributima, modifikujete te atribute i dobijete trenutne statistike lika.

**Cilj:** Popunite priče korisnika ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Unda darasa linaloitwa `GameCharacter` linalowakilisha mhusika wa mchezo na kusimamia takwimu za mhusika.
2. Wakati darasa hilo linapotengenezwa, kitu kipya cha `GameCharacter` kinapaswa kuwa na sifa zifuatazo:
- `_name` iwe na mfuatano wa herufi uliotolewa wakati wa kuanzisha.
    - `_health` iwe na `100`.
    - `_mana` iwe na `50`.
    - `_level` iwe na `1`.
3. Kreiraj svojstvo `name` koje omogućava samo čitanje imena lika.
4. Kwa sifa ya `health`:
     - Tafsiri getter itoe afya ya sasa.
     - Tafsiri setter isiruhusu afya kuwekwa chini ya `0`, na iweze kufikia kiwango cha juu cha `100`.
5. Kwa sifa ya `mana`:
     - Tafsiri getter itoe mana ya sasa.
     - Tafsiri setter isiruhusu mana kuwekwa chini ya `0`, na iweze kufikia kiwango cha juu cha `50`.
6. Unda getter kwa `level` kurudisha kiwango cha sasa cha mhusika.
7. Tafsiri njia (method) liitwalo `level_up` ambayo:
    - Inaongeza kiwango cha mhusika kwa 1.
Vraća zdravlje na `100` i manu na `50` koristeći njihove settere.
Štampa poruku u formatu `<name> leveled up to <level>!` (gde `<name>` i `<level>` treba zameniti imenom lika i novim nivoom, redom).
8. Prevedi metodu `__str__` koja vraća formatirani string koji uključuje:
    - Jina la mhusika.
    - Kiwango cha mhusika.
    - Afya ya sasa ya mhusika.
    - Mana ya sasa ya mhusika.
Na primer, lik po imenu `Kratos`, odmah nakon kreiranja, treba da bude predstavljen na sledeći način:

   ```md
   Name: Kratos
   Level: 1
   Health: 100
   Mana: 50
   ```

## Primer upotrebe

```python
hero = GameCharacter('Kratos') # Creates a new character named Kratos
print(hero)  # Displays the character's stats

hero.health -= 30  # Decreases health by 30
hero.mana -= 10    # Decreases mana by 10
print(hero)  # Displays the updated stats

hero.level_up()  # Levels up the character
print(hero)  # Displays the stats after leveling up
```

# --hints--

Trebalo bi da imaš klasu objekata `GameCharacter`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("GameCharacter")`)) })
```

Trebalo bi da imaš putanju `__init__` u klasi `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__init__")`) 
})
```

Putanja za `__init__` treba da ima dva parametra, gde je prvi `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.__init__).parameters)
    assert params[0] == "self"
    assert len(params) == 2
    `) 
})
```

Treba da postavite `self._name` kao niz znakova koji je generisan tokom inicijalizacije unutar putanje `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._name == "Kratos"
    `) 
})
```

Trebalo bi da staviš `100` u `self._health` unutar puta `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._health == 100
    `) 
})
```

Trebalo bi da staviš `50` u `self._mana` unutar puta `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._mana == 50
    `) 
})
```

Trebalo bi da staviš `1` u `self._level` unutar puta `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._level == 1
    `) 
})
```

Trebalo bi da imaš putanju `name` u klasi `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("name")`) 
})
```

Putanja za `name` treba da ima samo jedan parametar, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_args("self")`) 
})
```

Moraš vratiti `self._name` u putanji od `name`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_return("self._name")`)
})
```

Putanja za `name` bi trebalo da ima dekorator za `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_decorators("property")`) 
})
```

Trebalo bi da imaš putanju `health` u klasi `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("health")`) 
})
```

Putanja za `health` treba da ima samo jedan parametar, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_args("self")`) 
})
```

Treba ti vratiti `self._health` na putu `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_return("self._health")`)
})
```

Putanj za `health` treba da ima dekoraciju od `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_decorators("property")`) 
})
```

Treba da napraviš setter metod za `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1]`) 
})
```

Setter putanja za `health` mora imati dekorator od `@health.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_decorators("health.setter")`) 
})
```

Setter metoda za `health` treba da ima dva parametra, a prvi je `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.health.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Trebalo bi postaviti `self._health` na `0` ako je vrednost koja se pruža za seter od `health` manja od `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = -1
    assert char.health == 0
    char.health = -10
    assert char.health == 0
`)})
```

Trebalo bi postaviti `self._health` na `100` ako je vrijednost postavljena za setter `health` veća od `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 101
    assert char.health == 100
    char.health = 200
    assert char.health == 100
`)})
```

Trebalo bi da postavite `self._health` kao vrednost koja se pruža za setera `health` kada je vrednost između `0` i `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 95
    assert char.health == 95
    char.health = 1
    assert char.health == 1
    char.health = 0
    assert char.health == 0
    char.health = 100
    assert char.health == 100
`)})
```

Trebalo bi da imaš putanju `mana` u klasi `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("mana")`) 
})
```

Putanja za `mana` mora imati samo jedan parametar, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_args("self")`) 
})
```

Treba ti vratiti `self._mana` na putu `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_return("self._mana")`)
})
```

Putanja za `mana` trebalo bi da ima dekoraciju od `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_decorators("property")`) 
})
```

Treba da kreiraš setter metod za `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1]`) 
})
```

Setter metoda za `mana` treba da ima dekorator od `@mana.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1].has_decorators("mana.setter")`) 
})
```

Setter metoda za `mana` treba da ima dva parametra, od kojih je prvi `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.mana.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Trebalo bi postaviti `self._mana` na `0` ako je vrijednost dodeljena setteru za `mana` manja od `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = -1
    assert char.mana == 0
    char.mana = -10
    assert char.mana == 0
`)})
```

Morate postaviti `self._mana` na `50` ako je vrednost prosleđena seteru za `mana` veća od `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 51
    assert char.mana == 50
    char.mana = 100
    assert char.mana == 50
`)})
```

Trebalo bi da postaviš `self._mana` na vrednost koja je data za setter od `mana` ako je vrednost između `0` i `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 25
    assert char.mana == 25
    char.mana = 0
    assert char.mana == 0
    char.mana = 50
    assert char.mana == 50
`)})
```

Trebalo bi da imaš putanju `level` u klasi `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level")`) 
})
```

Putanja za `level` treba da ima samo jedan parametar, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_args("self")`) 
})
```

Treba ti vratiti `self._level` putem `level`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_return("self._level")`)
})
```

Putanja za `level` treba da ima dekoraciju od `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_decorators("property")`) 
})
```

Trebalo bi da imaš putanju `level_up` u klasi `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level_up")`) 
})
```

Putanja za `level_up` treba da ima samo jedan parametar, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_args("self")`) 
})
```

Treba da dodaš `self._level` na `1` u putanji od `level_up`.

```js
({ 
  test: () => runPython(`
    node = _Node(_code).find_class("GameCharacter").find_function("level_up")
    assert (
        node.has_stmt("self._level += 1") or
        node.has_stmt("self._level = self._level + 1")
    )
`)
})
```

Trebalo bi da staviš `self.health` u `100` na mestu `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.health = 100")`) 
})
```

Treba ti postaviti `self.mana` u `50` putanju `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.mana = 50")`) 
})
```

Metoda za `level_up` treba da ispiše `<name> leveled up to <level>!` (gde `<name>` i `<level>` moraju biti zamenjeni imenom karaktera i novim nivoom, po odvojeno).

```js
({
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.level_up()
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
            
    assert "test leveled up to 2!" in output
`)
})
```

Trebalo bi da imaš putanju `__str__` u klasi `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__str__")`) 
})
```

Putanja za `__str__` treba da ima samo jedan parametar, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("__str__").has_args("self")`) 
})
```

Tvoja putanja ``__str__`` bi trebalo da vrati niz slova sa statističkim podacima o karakteru koristeći navedeni format.

```js
({ 
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.mana = 15
    gc.health = 77
    print(gc)
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
    expected = """Name: test
    Level: 1
    Health: 77
    Mana: 15"""
    assert expected in output
`)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class GameCharacter:
    def __init__(self, name):
        self._name = name
        self._health = 100
        self._mana = 50
        self._level = 1

    @property
    def name(self):
        return self._name

    @property
    def health(self):
        return self._health

    @health.setter
    def health(self, value):
        if value < 0:
            self._health = 0
        elif value > 100:
            self._health = 100
        else:
            self._health = value

    @property
    def mana(self):
        return self._mana

    @mana.setter
    def mana(self, value):
        if value < 0:
            self._mana = 0
        elif value > 50:
            self._mana = 50
        else:
            self._mana = value

    @property
    def level(self):
        return self._level

    def level_up(self):
        self._level += 1
        self.health = 100
        self.mana = 50
        print(f"{self._name} leveled up to {self._level}!")

    def __str__(self):
        return (
            f"Name: {self._name}\n"
            f"Level: {self._level}\n"
            f"Health: {self._health}\n"
            f"Mana: {self._mana}"
        )
```
