---
id: 5dc1798ff86c76b9248c6eb3
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Elementi naslova od `h1` do `h6` koriste se za prikaz važnosti sadržaja ispod njih. Manji broj ukazuje na veću važnost, pa elementi od `h2` imaju manju važnost od elemenata od `h1`.

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

Koristi samo jedan element `h1` po stranici i stavi manje važne podnaslove ispod naslova sekcija sa većom važnosti.

Ispod elementa `h1`, dodajte element `h2` koji sadrži ovaj tekst:

`Cat Photos`

# --hints--

Tvoj element `h1` treba da ima početnu oznaku. Početne oznake koriste sintaksu `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Vaš element od `h1` treba da ima zatvarajući tag. Zatvorni tagovi su `/` odmah nakon karaktera `<`.

```js
assert(code.match(/<\/h1\>/));
```

Treba ti samo jedan element od `h1`. Ukloni višak.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

Tekst tvog polja `h1` bi trebalo da bude `CatPhotoApp`. Ostavio si tekst ili postoje pravopisne greške.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Oznake `h2` treba da budu napisane malim slovima. Po konvenciji, sve HTML oznake pišu se malim slovima.

```js
assert.notMatch(code, /<\/?H2>/);
```

Vaš element za `h2` mora imati otvarajući tag. Otvarajući tagovi imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelector('h2'));
```

Vaš element za `h2` mora imati zatvarajući tag. Zatvarajući tagovi imaju `/` odmah nakon karaktera/slovo `<`.

```js
assert(code.match(/<\/h2\>/));
```

Sadržaj vašeg elementa za `h2` treba da bude `Cat Photos`. Postavite tekst `Cat Photos` samo između otvarajućih i zatvarajućih oznaka za `h2`.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

Element vašeg `h2` treba da bude ispod elementa `h1`. Element `h1` ima veću važnost i treba da bude iznad elementa `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    
--fcc-editable-region--
  </body>
</html>
```
