---
id: 5dc174fcf86c76b9248c6eb2
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Elementi za HTML imaju otvarajuće oznake kao `<h1>` i zatvarajuće oznake kao `</h1>`. Sadržaj koji element prikazuje postavlja se između njegovih otvarajućih i zatvarajućih oznaka.

Promeni tekst elementa `h1` ispod sa `Hello World` na `CatPhotoApp` i pregledaj promene u pretraživačkom prikazu.

Kada završiš, pritisni dugme „Check Your Code“ da vidiš/proveriš da li je ispravno.

# --hints--

Tekst `CatPhotoApp` treba da bude u kodu. Možda želiš da proveriš svoju pravopisnost.

```js
assert(code.match(/catphotoapp/i));
```

Vaš element za `h1` mora imati otvorenu oznaku. Otvorene oznake imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Etikete za `h1` bi trebalo da budu u malim slovima. Prema ugovoru, sve etikete za HTML su napisane malim slovima.

```js
assert.notMatch(code, /<\/?H1>/);
```

Vaš element za `h1` mora imati zatvarajuću oznaku. Zatvarajuće oznake su `/` odmah nakon slova `<`.

```js
assert(code.match(/<\/h1\>/));
```

Poseduje više od jednog elementa za `h1`. Uklonite dodatni element od `h1`.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Sadržaj vašeg elementa `h1` treba da bude `CatPhotoApp`. Možda ste ostavili tekst, postoje greške u pravopisu ili nije između otvarajućih i zatvarajućih oznaka elementa `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```
