---
id: 5dc17dc8f86c76b9248c6eb5
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Ostavljanje komentara omogućava vam da ostavite poruku bez uticaja na način na koji pregledač prikazuje stranicu. Takođe vam omogućava da učinite kod neaktivnim. Komentari u HTML počinju sa `<!--`, imaju bilo koje linije teksta i završavaju se sa `-->`.

Ovo je primer komentara sa `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Ostavite komentare o komponenti `p` za ovaj tekst:

`TODO: Add link to cat photos`

# --hints--

Vaši komentari moraju početi sa `<!--`. Nedostaje vam jedna ili više slova koja označavaju početak komentara.

```js
assert(code.match(/<!--/));
```

Vaši komentari bi trebalo da se završe sa `-->`. Nedostaje vam jedno ili više slova koja označavaju kraj komentara.

```js
assert(code.match(/-->/));
```

Kod ne bi trebalo da ima dodatne karaktere za otvaranje/zatvaranje komentara. Imate višak `<!--` ili `-->` koji je vidljiv u pretraživaču.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

Vaši komentari bi trebalo da budu iznad elementa `p`. Postavili ste ga u pogrešnom redosledu.

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

Tvoj komentar treba da sadrži tekst `TODO: Add link to cat photos`.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
