---
id: 5dc17d3bf86c76b9248c6eb4
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Element `p` se koristi za kreiranje tekstualnog pasusa na veb stranici. Kreirajte element `p` ispod vašeg elementa `h2` koristeći sledeći tekst:

`Everyone loves cute cats online!`

# --hints--

Tvoj element `p` bi trebalo da ima oznaku otvaranja. Oznake otvaranja imaju sledeću sintaksu: `<elementName>`.

```js
assert(document.querySelector('p'));
```

Tvoji etikete za `p` moraju biti u malim slovima. Po ugovoru, sve etikete za HTML se pišu malim slovima.

```js
assert.notMatch(code, /<\/?P>/);
```

Vaša komponenta za `p` bi trebalo da ima zatvarajući tag. Zatvarajući tagovi su `/` odmah nakon slova/karaktera `<`.

```js
assert(code.match(/<\/p\>/));
```

Tekst vašeg polja za `p` treba da bude `Everyone loves cute cats online!`. Ostavili ste tekst ili postoje greške u kucanju.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Vaš element za `p` treba da bude ispod elementa za `h2`. Postavili ste ga u pogrešan redosled.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>
    
--fcc-editable-region--
  </body>
</html>
```
