---
id: 5dc24614f86c76b9248c6ebd
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Možete se povezati sa drugom stranicom koristeći anker polje (`a`).

Ovo je primer povezivanja sa `https://www.freecodecamp.org`:

```html
<a href="https://www.freecodecamp.org"></a>
```

Dodajte anker element nakon paragrafa koji se povezuje sa `https://freecatphotoapp.com`. Trenutno, link neće biti vidljiv u početnom prikazu.

# --hints--

Vaš anker element (`a`) mora imati oznaku za otvaranje. Oznake za otvaranje imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Vaš anker element (`a`) treba da ima zatvarajući tag. Zatvarajući tagovi imaju `/` odmah nakon karaktera od `<`.

```js
assert(code.match(/<\/a\>/));
```

Tvoj element sidra (`a`) treba da bude ispod elementa `p`. Postavio si ih pogrešnim redosledom.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

Vaš jangkorni element (`a`) nema atribut `href`. Uverite se da postoji razmak nakon imena otvarajuće oznake i/ili da postoji razmak prije svih imena atributa.

```js
assert(document.querySelector('a').hasAttribute('href'));
```

Vaš element sidra (`a`) treba da vodi na `https://freecatphotoapp.com`. Ili nedostaje URL ili postoji greška u njegovom pisanju.

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

Iako si postavio atribut `href` elementa sidra (`a`) na ispravnu vezu, preporučuje se da vrednosti atributa uvek budu pod navodnicima.

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
