---
id: 5dc23f9bf86c76b9248c6eba
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Možete dodati slike na vaš sajt koristeći element `img`. Elementi `img` imaju otvarajući tag bez zatvarajućeg taga. Element bez zatvarajućeg taga poznat je kao <dfn>prazan element</dfn>.

Dodajte element za `img` ispod elementa `p`. Do sada, slika neće biti prikazana u pretraživaču.

# --hints--

Vaša komponenta za `img` mora imati oznaku za otvaranje. Oznake za otvaranje imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelector('img'));
```

Tvoj element za `img` ne bi trebalo da ima zatvarajući atribut/tag `</img>`.

```js
assert(!code.match(/<\/img\>/));
```

Trebalo bi da ima samo jedan element od `img`. Uklonite sve viškove.

```js
assert(document.querySelectorAll('img').length === 1);
```

Tvoj element za `img` treba da bude ispod elementa za `p`. Postavio si ga u pogrešnom redosledu.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
