---
id: 5dc2385ff86c76b9248c6eb7
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 ima elemente koji identifikuju različite sekcije sadržaja. Ovi elementi čine HTML lakšim za čitanje i pomažu u Optimizaciji pretraživača (SEO) i Pristupačnosti.

Element `main` se koristi za predstavljanje glavnog sadržaja tela dokumenta HTML. Sadržaj unutar elementa `main` mora biti jedinstven za taj dokument i ne sme se ponavljati u drugim delovima tog dokumenta.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Pronađite glavni deo ove stranice dodavanjem otvarajućeg oznaka `<main>` pre elemenata `h1`, i zatvarajućeg oznaka `</main>` nakon elemenata `p`.

# --hints--

Vaš element ``main`` treba da ima otvarajuću oznaku. Otvarajući tagovi imaju ovu sintaksu: ``<elementName>``.

```js
assert(document.querySelector('main'));
```

Vaši oznake za `main` bi trebalo da budu u malim slovima. Prema ugovoru, sve oznake HTML su napisane malim slovima.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0];
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

Tvoj element `main` treba da ima završnu oznaku. Završne oznake imaju znak `/` odmah posle znaka `<`.

```js
assert(code.match(/<\/main\>/));
```

Početna oznaka elementa `main` treba da bude ispod početne oznake elementa `body`. Postavio si ih pogrešnim redosledom.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Otvarajući tag tvog elementa ``main`` mora biti iznad elementa ``h1``. Postavili ste ih u pogrešnom redosledu.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H1'));
```

Završna oznaka elementa `main` treba da bude ispod elementa `p`. Postavio si ih pogrešnim redosledom.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode.contains(pNode));
assert.match(pNode.textContent.toLowerCase(), /everyone loves cute cats online!/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
  
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
