---
id: 5dc24165f86c76b9248c6ebc
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Svi elementi za `img` moraju imati atribut elementa `alt`. Tekst atributa elementa `alt` koristi se od strane čitača ekrana za poboljšanje pristupačnosti i prikazuje se ako slika ne učita.

Ovo je primer elementa `img` sa svojstvom elementa `alt`:

```html
<img src="cat.jpg" alt="A cat">
```

Unutar elementa `img`, dodaj atribut elementu `alt` sa ovim tekstom:

`A cute orange cat lying on its back`

# --hints--

Tvoj kod bi trebalo da ima komponentu `img`. Ukloni komponentu `img` iz prethodnog koraka.

```js
assert(document.querySelector('img'));
```

Vaš element `img` nema svojstvo elemenata `alt`. Uverite se da postoji razmak nakon imena otvarajuće oznake i/ili da postoji razmak pre svih imena svojstava elementa.

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

Vrednost atributa `alt` elementa `img` razlikuje se od `A cute orange cat lying on its back`. Uverite se da je vrednost atributa `alt` okružena navodnicima.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
