---
id: 5dc24073f86c76b9248c6ebb
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Svojstva elemenata za HTML su specijalne reči koje se koriste unutar elementnog atributa za kontrolisanje ponašanja tog elementa. Svojstvo elementa `src` u elementu `img` opisuje URL slike (mesto gde se slika nalazi).

Ovo je primer elementa `img` sa atributom elementa `src` koji se odnosi na logo freeCodeCamp-a:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

Unutar postojećeg elementa `img`, dodajte svojstvo elementu `src` sa ovim URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Tvoj kod bi trebalo da ima element `img`. Možda ste uklonili element `img` ili niste obuhvatili vrednost atributa elementa `src` sa navodnici.

```js
assert(document.querySelector('img'));
```

Vaš element ``img`` bi trebalo da ima atribut elementa ``src``. Možda ste izostavili taj atribut ili postoji greška u kucanju. Proverite da postoji razmak između imena elementa i imena atributa elementa.

```js
assert(document.querySelector('img').src);
```

Atribut `src` elementa `img` treba da bude postavljen na `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`. Možda si izostavio URL ili napravio grešku u pisanju. Velika i mala slova u URL-u su važna.

```js
assert.strictEqual(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Iako ste postavili atribut elementa `img` za element `src` za URL tačno, savetuje se da vrednost atributa elementa uvek obuhvatite navodnicima.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

Element vašeg `img` treba da se zatvori pomoću `>` ili `/>`.

```js
assert.match(code, /<img\s+src\s*=\s*("|')\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\s*\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```
