---
id: 5ddbd81294d8ddc1510a8e56
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Inline sadržaj mora biti između otvarajućeg i zatvarajućeg taga jangkošnog elementa (`a`).

Ovo je primer poveznog elementa sa tekstom `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Dodajte tekst za sidro `cat photos` u polje sidra. Ovo će biti tekst veze/link teksta.

# --hints--

Vaš anker element (`a`) treba da ima otvarajući tag. Otvoreći tagovi imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Vaš ankorni element (`a`) treba da ima zatvarajući tag. Zatvarajući tagovi se nalaze `/` odmah nakon slova/karaktera `<`.

```js
assert(code.match(/<\/a\>/));
```

Tekst vašeg ankora elementa (`a`) treba da bude `cat photos`. Osigurajte da umetnete tekst veze između otvarajućeg i zatvarajućeg taga ankora elementa (`a`).

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
