---
id: 5dfa2407b521be39a3de7be1
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Da biste otvorili linkove u novoj kartici, možete koristiti atribut elementa `target` na elementu sidroznaka taga (`a`).

Svojstvo elementa `target` opisuje gde će se povezani dokument otvoriti. `target="_blank"` otvara povezani dokument u novoj karti ili prozoru.

Ovo je osnovna sintaksa za element `a` koji ima svojstvo elementa `target`:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Dodajte atribut elementu `target` sa vrednošću `_blank` na atribut otvaranja elementa sidra (anchor tag) (`a`), kako bi se veza otvorila u novoj karti.

# --hints--

Vaš element `p` bi trebalo da ima komponentu ankora oznake (`a`) koja sadrži tekst `cat photos`. Možda ste ga izbrisali ili postoje pravopisne greške.

```js
const anchor = document.querySelectorAll('p > a');
assert.isNotEmpty(anchor);
assert.strictEqual(anchor[1]?.innerText?.toLowerCase().replace(/\s+/g, ' '), 'cat photos')
```

Vaš anker tag element (`a`) nema atribut elemenata `target`. Proverite da postoji razmak nakon imena otvarajuće oznake i/ili da postoji razmak pre svih elemenata atributa.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

Vrednost atributa elementa `target` treba da bude `_blank`. Možda ste zaboravili vrednost ili postoji greška u kucanju. Zapamtite da atributi elementa moraju biti obavijeni navodnicima.

```js
assert.strictEqual(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
