---
id: 5dfa3589eacea3f48c6300ae
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

U okviru drugog elementa od `section`, dodaj novi element sa `h2` koji sadrži tekst `Cat Lists`.

# --hints--

Vaš element za `section` treba da ima oznaku otvaranja. Oznake otvaranja imaju ovu sintaksu: `<elementName>`.

```js
assert(
  document.querySelectorAll('section').length === 2 &&
    code.match(/<\/section>/g).length === 2
);
```

Vaš element `h2` treba da ima otvorajući tag. Otvarajući tagovi imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelectorAll('h2').length === 2);
```

Vaš element za `h2` mora imati zatvarajući tag. Zatvarajući tagovi imaju `/` odmah nakon karaktera/slovo `<`.

```js
assert(code.match(/<\/h2\>/g).length === 2);
```

Tvoj drugi element `h2` mora biti direktno iznad zatvarajuće oznake drugog elementa `section`. Nije na pravilnoj poziciji.

```js
const secondSection = document.querySelectorAll('section')[1];
assert(secondSection.lastElementChild.nodeName === 'H2');
```

Drugi element `h2` treba da sadrži tekst `Cat Lists`. Nedostaje tekst ili postoje greške u kucanju.

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase() === 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
