---
id: 5dfa37b9eacea3f48c6300b0
title: Korak 22
challengeType: 0
dashedName: step-22
---

# --description--

Za kreiranje neuređene liste stavki, možete koristiti funkciju `ul`.

Nakon elementa sa `h3` koji sadrži tekst `Things cats love:`, dodajte element neuređenog spiska (`ul`). Zapamtite da trenutno ništa neće biti prikazano.

# --hints--

Vaš element za `ul` treba da ima aktivacionu etiketu. Aktivacione etikete imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelector('ul'));
```

Vaš element `ul` mora imati zatvorni tag. Zatvorni tagovi imaju `/` odmah nakon karaktera `<`.

```js
assert(code.match(/<\/ul>/));
```

Element ``ul`` mora biti iznad zatvarajućeg taga drugog elementa ``section``.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')?.[1]?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName , 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
