---
id: 5dfb6a35eacea3f48c6300b4
title: Korak 26
challengeType: 0
dashedName: step-26
---

# --description--

Element za opis slike (`figcaption`) se koristi za dodavanje opisa slika koji opisuju sliku sadržanu u elementu `figure`.

Ovde je primer elementa ``figcaption`` sa opisima za ``A cute cat``:

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

Nakon slike postavljene unutar elementa `figure`, dodaj element `figcaption` sa tekstom postavljenim na:

`Cats love lasagna.`

# --hints--

Tvoj element `figcaption` treba da ima početnu oznaku. Početne oznake imaju sledeću sintaksu: `<elementName>`.

```js
assert(document.querySelector('figcaption'));
```

Vaš element za `figcaption` mora imati zatvarajući tag. Zatvarajući tagovi su `/` odmah nakon karaktera `<`.

```js
assert(code.match(/<\/figcaption\>/));
```

Element za `figcaption` treba biti postavljen unutar elementa `figure`.

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

Komponenta lasagne za `img` treba da bude postavljena unutar komponente za `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Element `figcaption` koji je postavljen unutar elementa `figure` mora biti ispod elementa `img`. Postavili ste ga u pogrešnom redosledu.

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

Sadržaj vašeg elementa `figcaption` bi trebalo da bude `Cats love lasagna.`. Ostavili ste tekst ili postoje greške u pravopisu.

```js
assert(
  document.querySelector('figcaption').innerText.trim().match(/^Cats love lasagna\.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          
        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
