---
id: 5dfb655eeacea3f48c6300b3
title: Korak 25
challengeType: 0
dashedName: step-25
---

# --description--

Funkcija `figure` prikazuje samostalni sadržaj i omogućit će vam da uključite slike i opise slika.

Postavite sliku koju ste nedavno dodali unutar elementa `figure`.

# --hints--

Tvoj element `figure` mora imati labelu za otvaranje. Labeli za otvaranje imaju sledeću sintaksu: `<elementName>`.

```js
assert(document.querySelector('figure'));
```

Tvoj element `figure` treba da ima završnu oznaku. Završne oznake imaju znak `/` odmah posle znaka `<`.

```js
assert(code.match(/<\/figure\>/));
```

Element za `figure` mora biti neposredno iznad zatvarajuće oznake drugog elementa `section`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.lastElementChild?.nodeName, 'FIGURE');
```

Element `img` za lazanju treba da bude smešten unutar elementa `figure`.

```js
assert.strictEqual(document.querySelector('figure > img')?.src.toLowerCase(), 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
      
        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
      
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
