---
id: 5ef9b03c81a63668521804ea
title: Korak 67
challengeType: 0
dashedName: step-67
---

# --description--

Napomena da se sve što ste do sada postavili na stranicu nalazi unutar elementa `body`. Svi elementi sadržaja stranice koji treba da se prikažu na stranici nalaze unutar elementa `body`. Međutim, druge važne informacije nalaze se unutar elementa `head`.

Element `head` se koristi za čuvanje meta podataka o dokumentu, kao što su njegov naslov, veze stilskih listova i skripte. Meta podaci su informacije o stranici koje nisu direktno prikazane na stranici.

Dodajte funkcionalnost za `head` iznad funkcionalnosti za `body`.

# --hints--

Izbrisali ste element za `body`, ili mu nedostaje otvarajući ili zatvarajući tag.

```js
assert(document.querySelector('body') && code.match(/<\/body>/));
```

Element vašeg `head` mora imati oznaku za otvaranje. Oznake za otvaranje imaju sledeću sintaksu: `<elementName>`.

```js
assert(code.match(/\<head\>/));
```

Tvoj element za `head` bi trebalo da ima zatvarajući tag. Zatvarajući tagovi su `/` odmah nakon karaktera `<`.

```js
assert(code.match(/\<\/head\>/));
```

Kipengele chako cha `head` treba da bude iznad oznake otvaranja komponente za `body`. Postavio/la si ga negde drugde.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/\<\/head\>\<body\>/));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>

  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving" checked> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic">Energetic</label>
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
