---
id: 5ef9b03c81a63668521804df
title: Korak 48
challengeType: 0
dashedName: step-48
---

# --description--

Atribut elementa `id` koristi se za identifikovanje posebnih elemenata od HTML. Svaka vrednost atributa elementa `id` mora biti jedinstvena u poređenju sa svim drugim vrednostima od `id` za celu stranicu.

Evo primera elementa `input` sa atributom `id`:

```html
<input id="email">
```

Dodaj atribut `id` sa vrednošću `indoor` radio dugmetu. Kada element ima više atributa, njihov redosled nije važan.

# --hints--

Tvoje radio dugme i dalje treba da bude između početne i završne oznake elementa `label`.

```js
const labelChildNodes = [...document.querySelectorAll('form > label')?.[0]?.childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

Vaše radio dugme treba da ima atribut `id`. Uverite se da postoji razmak nakon naziva otvarajuće oznake i/ili razmak ispred svih naziva atributa.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('id'));
```

Vaše radio dugme treba da ima atribut `type`. Uverite se da postoji razmak nakon naziva otvarajuće oznake i/ili razmak ispred svih naziva atributa.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Tvoj radio dugme mora imati atribut svojstvo `id` sa vrednošću `indoor`. Izostavili ste vrednost ili napravili pravopisnu grešku. Zapamtite da atributi svojstava moraju biti okruženi navodnicima.

```js
assert.match(document.querySelector('input')?.id, /^indoor$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input type="radio"> Indoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
