---
id: 5ef9b03c81a63668521804d6
title: Korak 37
challengeType: 0
dashedName: step-37
---

# --description--

Sada ćete dodati online formular za prikupljanje informacija od korisnika.

Kipengele cha `form` se koristi za dobijanje informacija od korisnika kao što su njihovo ime, e-mail adresa i drugi detalji.

Nakon zaglavlja sekcije za `Cat Form`, dodaj element za `form`.

# --hints--

Vaš element `form` mora imati otvarajući i zatvarajući tag. Možda vam nedostaje jedan ili oba zahtevana taga, ili ih postavljate u pogrešan redosled.

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

Etikete za vaš element sa `form` nisu u pravilnom redosledu.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

Element ``form`` koji je postavljen unutar elementa ``section`` treba da bude ispod elementa ``h2``. Imate element ``h2`` i element ``form`` u pogređem redosledu.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

Element `form` ne sme imati sadržaj. Uklonite bilo koje elemente od HTML ili tekst između oznaka (tagova) elementa `form`.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
