---
id: 5ef9b03c81a63668521804db
title: Korak 43
challengeType: 0
dashedName: step-43
---

# --description--

Da biste sprečili korisnika da pošalje vašu formu kada nedostaju važne informacije, morate dodati svojstvo elementa `required` na element `input`.

Ovo je primer polja za unos sa atributom elementa `required`:

```html
<input type="text" name="firstName" required>
```

Nije potrebno dodeljivati vrednost atributu elementa `required`. Umesto toga, samo dodajte reč `required` u element `input`, osiguravši da postoji razmak između njega i drugih atributa.

# --hints--

Izbrisali ste svoj element `input` ili ima neispravnu sintaksu. Sve vrednosti atributa polja moraju biti okružene navodnicima.

```js
assert(document.querySelectorAll('input').length);
```

`form` trebao bi imati samo element `input`. Izbrišite druge elemente od HTML ili tekst unutar elementa `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Tvoj element `input` treba da ima atribut `required`. Potrebno je samo da dodaš reč `required` unutar početne oznake elementa `input`.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('required'));
```

Vrednost ne bi trebalo da se daje za atribut elementa `required`.

```js
assert.strictEqual(document.querySelector('input')?.getAttribute('required'), '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
