---
id: 5ef9b03c81a63668521804d1
title: Korak 28
challengeType: 0
dashedName: step-28
---

# --description--

Nakon elementa sa `figure`, dodajte drugi element sa `h3` sa tekstom:

`Top 3 things cats hate:`

# --hints--

Element `h3` treba da se nalazi neposredno iznad zatvarajuće oznake drugog elementa `section`. Uverite se da ima i otvarajuću i zatvarajuću oznaku.

```js
assert(
  document.querySelectorAll('main > section')[1].lastElementChild.nodeName ===
    'H3' && code.match(/<\/h3\>/g).length === 2
);
```

Novi element `h3` treba da sadrži tekst `Top 3 things cats hate:`. Obavezno uključite dvotačku na kraju teksta.

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase()
    .replace(/\s+/g, ' ') === 'top 3 things cats hate:'
);
```

Element `figure` treba da bude iznad novog elementa `h3`. Možda si slučajno obrisao element `figure`.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')[1]
  .lastElementChild;
assert(
  secondSectionLastElemNode.nodeName === 'H3' &&
    secondSectionLastElemNode.previousElementSibling.nodeName === 'FIGURE'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
