---
id: 5ef9b03c81a63668521804e2
title: Korak 56
challengeType: 0
dashedName: step-56
---

# --description--

Forme često koriste polja za višestruki izbor za pitanja koja mogu imati više od jednog odgovora. Element `input` sa svojstvom elementa `type` postavljeno na `checkbox` stvara polje za višestruki izbor.

Ispod elementa od `legend` koji ste postavili, dodajte `input` sa atributom elementa `type` koji je postavljen na `checkbox` i čiji tekst je:

` Loving`

# --hints--

Komponenta za `input` za vaš padajući okvir treba da ima etiketu za otvaranje, ali ne i etiketu za zatvaranje.

```js
assert.isTrue(document.querySelector('fieldset > input') && !code.match(/<\/input\>/g));
```

Trebalo je da dodate samo jedno polje za unos u svoj okvir za izbor. Uklonite sve suvišnosti.

```js
assert.lengthOf(document.querySelectorAll('fieldset > input'), 1);
```

Tvoja nova sekcija `input` treba da bude ispod sekcije `legend` sa tekstom `What's your cat's personality?`. Postavila si je u pogređem redosledu.

```js
const existingLegendElem = document.querySelectorAll('fieldset > legend')?.[1];
assert.strictEqual(
  existingLegendElem?.nextElementSibling?.nodeName,
  'INPUT'
);
```

Tvoj novi element `input` nema svojstvo elementa `type`. Proveri da postoji razmak nakon imena oznake za otvaranje.

```js
assert.isTrue(document.querySelector('fieldset > input')?.hasAttribute('type'));
```

Tvoj novi element za `input` treba da ima atribut elementa `type` sa vrednošću `checkbox`. Izostavili ste vrednost ili ste napravili greške u kucanju. Zapamtite da vrednosti atributa elementa moraju biti okružene navodnicima.

```js
assert.match(document.querySelector('fieldset > input')?.getAttribute('type'),/^checkbox$/i);
```

Iako ste postavili atribut elementa `input` novog elementa `type` da bude `checkbox`, savetuje se uvek izvlačenje vrednosti atributa elementa i navodnika.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

Tekst ` Loving` treba da bude postavljen direktno sa desne strane vašeg padajućeg menija. Proverite da postoji razmak između elementa i teksta. Ostavili ste tekst ili ste napravili pravopisne greške.

```js
const checkboxInputElem = document.querySelector('input[type="checkbox"]');
assert.match(
  checkboxInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Loving/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>
            
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
