---
id: 5ef9b03c81a63668521804d9
title: Korak 42
challengeType: 0
dashedName: step-42
---

# --description--

Tekst naslovnika se koristi da pruži korisnicima smernice o kakvoj vrsti informacija koje treba uneti u polje za unos.

Ovo je primer elementa `input` sa zamjennikom postavljenim na `e.g. Jane Doe`:

```html
<input type="text" placeholder="e.g. Jane Doe">
```

Dodajte tekst zamene `cat photo URL` u vaše polje `input`.

# --hints--

Izbrisali ste svoj element `input` ili ima neispravna sintaksa. Sve vrednosti atributa polja moraju biti obavijene navodnicima.

```js
assert(document.querySelectorAll('input').length);
```

`form` Vaše bi trebalo da ima samo element od `input`. Izbrišite druge elemente od HTML ili tekst unutar elemenata `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Vaš element ``input`` nema atribut ``placeholder``. Osigurajte da postoji razmak nakon imena otvarajućeg taga i/ili da postoji razmak pre svih atributnih imena elementa.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('placeholder'));
```

Vaš element `input` treba da ima atribut `placeholder` sa vrednošću `cat photo URL`. Vrednost je izostavljena ili postoje greške u kucanju. Zapamtite da vrednosti atributa elemenata moraju biti obavijene navodnicima.

```js
assert(
  document.querySelector('input')
    ?.getAttribute('placeholder')
    .replace(/\s+/g, ' ')
    .match(/^cat photo URL$/i)
);
```

Iako ste postavili atribut ``input`` elementa ``placeholder`` kao ``cat photo URL``, savetuje se uvek da izbacujete vrednost atributa elementa i navodnike.

```js
assert(!/\<\s*input\s+placeholder\s*=\s*cat\s+photo\s+url/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
