---
id: 5ef9b03c81a63668521804da
title: Korak 44
challengeType: 0
dashedName: step-44
---

# --description--

Element `button` se koristi za kreiranje klikabilnog dugmeta.

Dodaj element/polje `button` sa tekstom `Submit` ispod elementa `input`. Standardno ponašanje kliktanja na dugme za formular bez ikakvog atributa je slanje (submitovanje) formulara na lokaciju navedenu u atributu `action` formulara.

# --hints--

Vaš element `button` mora imati otvarajući tag. Otvarajući tagovi imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelector('button'));
```

Vaš element sa `button` treba da ima zatvarajući tag. Zatvarajući tagovi su `/` odmah nakon slova `<`.

```js
assert(code.match(/<\/button\>/));
```

Tekst vašeg elementa `button` treba da bude `Submit`. Tekst je prazan ili ste napravili grešku u pravopisu.

```js
assert(document.querySelector('button').innerText.toLowerCase() === 'submit');
```

Tvoj element `button` bi trebalo da bude ispod elementa `input`. Postavili ste ga u pogrešnom redosledu.

```js
const collection = [...document.querySelectorAll('input, button')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('INPUT') < collection.indexOf('BUTTON'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
