---
id: 5ef9b03c81a63668521804d7
title: Korak 38
challengeType: 0
dashedName: step-38
---

# --description--

Atribut elementa `action` pokazuje gde podaci forme treba da budu poslati.

Ovo je primer elementa `form` sa svojstvom elementa `action`:

```html
<form action="/submit-url"></form>
```

U ovom primeru, `action="/submit-url"` kažem pretraživaču da podaci forme bi trebalo poslati na putanju `/submit-url`.

Dodaj svojstvo elementa `action` sa vrednošću `https://freecatphotoapp.com/submit-cat-photo` na polje `form`.

# --hints--

Vaš element `form` mora imati otvarajući i zatvarajući tag u ispravnom redosledu. Možda vam nedostaje jedan ili oba neophodna taga, ili ih imate u pogrešnom redosledu.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Vaš element `form`, koji je postavljen unutar elementa `section`, treba da bude ispod elementa `h2`. Imate element `h2` i element `form` u pogrešnom redosledu.

```js
assert.equal(document.querySelector('form')?.previousElementSibling?.nodeName, 'H2');
```

Element vašeg `form` ne bi trebalo da sadrži nikakav sadržaj. Uklonite sve komponente HTML ili tekst između oznaka elementa `form`.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

Vaš element `form` nedostaje atribut elementa `action`. Uverite se da postoji razmak nakon imena otvarajuće oznake i/ili da postoji razmak pre svih imena atributa elemenata.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

Vaš element od `form` mora imati svojstvo elementa `action` sa vrednošću `https://freecatphotoapp.com/submit-cat-photo`.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

Iako ste postavili atribut elementa `action` za URL tačno, savetuje se uvek da obuhvatite vrednost atributa elementa i navodnike.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
