---
id: 5ef9b03c81a63668521804d8
title: Korak 39
challengeType: 0
dashedName: step-39
---

# --description--

Ovaj element `input` omogućava nekoliko načina prikupljanja podataka sa web forme. Kao elementi `img`, elementi `input` su <dfn>element bez sadržaja</dfn> i ne zahtevaju zatvarajući label-e.

Postavi unutar elementa `input`, unutar elementa `form`.

# --hints--

Vaš element ``form`` mora imati otvarajući i zatvarajući tag u pravilnom redosledu. Možda ste propustili jedan ili oba neophodna taga, ili ste ih postavili u pogrešan redosled.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Mekanizam otvaranja vašeg elementa sa `form` mora imati samo atribut elementa `action`. Uklonite bilo šta drugo što ste napisali unutar njega.

```js
assert([...document.querySelector('form').attributes].length < 2);
```

Treba da kreiraš element za `input`. Proveri sintaksu.

```js
assert(document.querySelector('input'));
```

Tvoj element `input` treba da ima oznaku za otvaranje, ali ne i oznaku za zatvaranje.

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

Element za `input` treba da se postavi unutar elementa `form`.

```js
assert(document.querySelector('form > input'));
```

`form` treba da sadrži samo element `input` unutar sebe. Uklonite bilo koji element HTML ili tekst između oznaka polja `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
