---
id: 5ef9b03c81a63668521804dc
title: Korak 46
challengeType: 0
dashedName: step-46
---

# --description--

Možete koristiti radio dugmad za pitanja u kojima želite samo jedan odgovor od više opcija.

Evo primera radio dugmeta sa tekstom `cat`:

```html
<input type="radio"> cat
```

Zapamti da je element `input` prazan element.

Pre polja za unos teksta, dodajte radio dugme sa tekstom postavljenim kao:

`Indoor`

# --hints--

Treba da kreirate element `input` za vaše radio dugme. Proverite sintaksu.

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

Vaš element `input` treba da ima otvarajući tag, ali ne i zatvarajući tag.

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

Trebalo bi da dodaš samo jedan element `input` na tvoj radio dugme. Ukloni sve suvišnosti.

```js
assert(document.querySelectorAll('form > input'), 2);
```

Tvoj novi element `input` treba da bude iznad postojećeg elementa `input` čiji je atribut `type` postavljen na `text`. Postavio si ih pogrešnim redosledom.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Tvoja nova karakteristika `input` nema atribut `type`. Proveri da postoji razmak nakon imena oznake za otvaranje.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Vaša nova funkcija za `input` treba da ima samo jednu karakteristiku od `type`. Uklonite sve dodatke.

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Tvoj novi element `input` treba da ima atribut `type` sa vrednošću `radio`. Vrednost nedostaje ili je pogrešno napisana. Zapamti da vrednosti atributa treba da budu pod navodnicima.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

Iako ste postavili atribut `input` novog elementa `type` na `radio`, savetuje se da uvek obuhvatite vrednost atributa navodnicima.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

Tekst dugmeta `radio` za/od ` Indoor` treba da bude postavljen posle njega umesto pre njega.

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

Tekst ` Indoor` treba da bude smešten direktno desno od vašeg dugmeta `radio`. Ostavili ste tekst ili postoje greške u pravopisu.

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
 
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
