---
id: 5ef9b03c81a63668521804d2
title: Korak 29
challengeType: 0
dashedName: step-29
---

# --description--

Kod uređenog spiska (`ol`) je isti kao neuređeni spisak, ali stavke u uređenom spisku dobijaju brojeve kada se prikazuju.

Ispod elementa `h3`, dodaj numerisana listu sa ovim tri stavke liste:

`flea treatment`
`thunder`
`other cats`

# --hints--

Vaš element `ol` mora imati otvorenu oznaku. Otvarajuće oznake imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelector('ol'));
```

Vaš element za `ol` mora imati zatvarnu oznaku. Zatvarne oznake imaju `/` odmah nakon slova `<`.

```js
assert(code.match(/<\/ol>/));
```

Element `ol` bi trebalo da bude iznad zatvarajuće oznake drugog elementa `section`. Postavili ste ga u pogrešan redosled.

```js
assert.equal(document.querySelectorAll('main > section')?.[1]?.lastElementChild?.nodeName, 'OL');
```

Tri elementa za `li` treba da budu postavljeni unutar elementa `ol`.

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

Trebalo bi da imaš tri elementa sa `li` koji sadrže tekst `flea treatment`, `thunder` i `other cats` u bilo kom redosledu.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

Trebalo bi da ima samo jedan element `ol`.

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
