---
id: 5ef9b03c81a63668521804de
title: Korak 50
challengeType: 0
dashedName: step-50
---

# --description--

Molimo obratite pažnju da se oba radijumska dugmeta mogu izabrati istovremeno. Da bi se izabiranjem jednog radijumskog dugmeta ne odabrilo drugo automatski, oba dugmeta moraju imati atribut elementa `name` sa istom vrednošću.

Evo primera dva radio dugmeta sa istim atributom `name`:

```html
<input type="radio" name="meal"> Breakfast
<input type="radio" name="meal"> Lunch
```

Dodaj atribut `name` sa vrednošću `indoor-outdoor` za oba radio dugmeta.

# --hints--

Oba radio dugmeta bi trebalo još uvek da se postave između otvarajućeg i zatvarajućeg labela elementa `label`.

```js
const labelChildNodes = [...document.querySelectorAll('form > label')].map(
  (node) => node.childNodes
);
assert(
  labelChildNodes.filter((childNode) => childNode[0].nodeName === 'INPUT')
    .length === 2
);
```

Oba radijska dugmeta moraju imati atribut `name`. Proverite da postoji razmak nakon imena otvarajuće oznake i/ili da postoji razmak pre svih imena atributa.

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(radioButtons.every((btn) => btn.hasAttribute('name')));
```

Sve dve radio opcije moraju imati atribut elementa `type`. Osigurajte da postoji razmak nakon imena oznake za otvaranje i/ili da postoji razmak prije svih imena atributa elemenata.

```js
const radioButtons = [...document.querySelectorAll('input')];
assert(radioButtons.every((btn) => btn.hasAttribute('type')));
```

Oba radio dugmeta treba da imaju atribut `name` sa vrednošću `indoor-outdoor`. Možda si izostavio tu vrednost ili napravio grešku u pisanju. Zapamti da vrednosti atributa treba da budu pod navodnicima.

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(
  radioButtons.every((btn) =>
    btn.getAttribute('name').match(/^indoor-outdoor$/)
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
          <label><input id="outdoor" type="radio"> Outdoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
