---
id: 5ef9b03c81a63668521804d0
title: Korak 27
challengeType: 0
dashedName: step-27
---

# --description--

Da bi naglasio određenu reč ili izraz, možeš koristiti element `em`.

Sisistiza neno `love` u delu `figcaption` pokrivanjem sa elementom naglaska `em`.

# --hints--

Vaš element naglašavanja za `em` treba da ima otvorajući tag. Otvarajući tagovi imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelector('em'));
```

Tvoj element `em` treba da ima završnu oznaku. Završne oznake imaju znak `/` odmah posle znaka `<`.

```js
assert(code.match(/<\/em\>/));
```

Umefuta kipengele cha `figcaption` au kinakosa lebo ya kufungua au ya kufunga.

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

Tvoj element `em` treba da obuhvata tekst `love`. Možda nedostaje taj tekst ili postoji greška u pisanju.

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

Tekst elementa `figcaption` treba da bude `Cats love lasagna`. Proveri pravopis i uveri se da postoje odgovarajući razmaci oko početne i završne oznake elementa `em`.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
