---
id: 5ef9b03c81a63668521804d4
title: Korak 34
challengeType: 0
dashedName: step-34
---

# --description--

Funkcija `strong` se koristi za pokazivanje da je određeni tekst od velike važnosti ili hitan.

U dijelu koji ste dodali, `figcaption`, prikažite da je `hate` od velike važnosti za okruživanje ga u elementu `strong`.

# --hints--

Tvoj element za `strong` treba da ima oznaku za proširivanje. Oznake za proširivanje imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelector('strong'));
```

Vaš element za `strong` treba da ima zatvarajući tag. Zatvarajući tagovi su `/` odmah nakon karaktera `<`.

```js
assert(code.match(/<\/strong\>/));
```

Vaš element za `strong` treba da obuhvati reč `hate` u tekstu `Cats hate other cats.`. Ostavili ste tekst ili postoje greške u pravopisu.

```js
assert(
  document
    .querySelectorAll('figcaption')[1]
    .querySelector('strong')
    .innerText.trim().toLowerCase() === 'hate'
);
```

Tekst za `figcaption` treba da bude `Cats hate other cats.` Proverite pravopisne greške i osigurajte da su potrebni razmaci blizu otvarajućih i zatvarajućih tagova elementa za `strong`.

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert(
  secondFigCaption &&
    secondFigCaption.innerText
      .replace(/\s+/gi, ' ')
      .trim()
      .match(/cats hate other cats\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
--fcc-editable-region--
        
          <figcaption>Cats hate other cats.</figcaption>  
        
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
