---
id: 5efb23e70dc218d6c85f89b1
title: Korak 40
challengeType: 0
dashedName: step-40
---

# --description--

Postoji mnogo vrsta unosa koje možete kreirati koristeći atribut elementa `type`. Možete lako kreirati polje za lozinku, dugme za resetovanje, ili kontrolu koja omogućava korisniku da izabere fajl sa svog računara.

Kreiraj tekstualni deo da dobije unos teksta od korisnika dodavanjem atributa elementu `type` sa vrednošću `text` na elementu `input`.

# --hints--

Element vašeg koda `input` je uklonjen, ili ima neispravnu sintaksu. Ako ste dodali atribute elementima, uverite se da su njihove vrednosti obuhvaćene navodnicima.

```js
assert(document.querySelectorAll('input').length);
```

Tvoj element `form` treba da ima samo element `input`. Obriši dodatne elemente HTML ili tekst između oznaka (labela) elementa `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Element ``input`` nema svojstvo elementa ``type`` sa vrednošću ``text``. Osigurajte da postoji razmak nakon imena otvarajuće oznake i/ili da postoji razmak pre svih imena svojstava elemenata.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Element vašeg `input` treba da ima svojstvo elementa `type` sa vrednošću `text`. Izostavili ste vrednost ili postoje greške u kucanju. Zapamtite da vrednosti svojstava elemenata moraju biti okružene navodnicima.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type').replace(/\s+/g, ' '),
  /^text$/i
);
```

Iako ste postavili atribut elementa `input` od elementa `type` na `text`, savetuje se da uvek koristite vrednost atributa elementa i navodnike.

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
