---
id: 5efb2c990dc218d6c85f89b2
title: Korak 45
challengeType: 0
dashedName: step-45
---

# --description--

Iako ste postavili svoje dugme ispod polja za unos teksta, ona se pojavljuje pored svake stavke na stranici. Ovo je zato što `input` i `button` su <dfn>elementi sa jedne linije</dfn>, koji se ne prikazuju na novim redovima.

Dugme koje ste postavili za slanje forme podrazumevano. Međutim, oslanjanje na podrazumevano ponašanje može izazvati zbunjenost. Dodajte atribut elementa `type` sa vrednošću `submit` na `button` kako biste ga definisali kao dugme za slanje.

# --hints--

Vaš element `button` mora imati otvarajući tag. Otvarajući tagovi imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelector('button'));
```

Vaš element sa `button` treba da ima zatvarajući tag. Zatvarajući tagovi su `/` odmah nakon slova `<`.

```js
assert(code.match(/<\/button\>/));
```

Element vašeg `button` nema svojstvo elementa `type`. Proverite da postoji razmak nakon imena otvarajuće oznake.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('type'));
```

Tvoj element `button` treba da ima svojstvo elementa `type` sa vrednošću `submit`. Možda ste ostavili vrednost ili napravili pravopisnu grešku. Zapamtite da vrednosti svojstava elemenata moraju biti okružene navodnicima.

```js
assert.match(
  document.querySelector('button')?.getAttribute('type'),
  /^submit$/i
);
```

Iako ste postavili atribut elementa ``button`` elementa ``type`` na ``submit``, savetuje se uvek da vrednost atributa elementa obuhvatite navodnicima.

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
