---
id: 5efc4f528d6a74d05e68af74
title: Korak 58
challengeType: 0
dashedName: step-58
---

# --description--

Postoji drugi način za povezivanje teksta elementa `input` sa samim tim elementom. Možete umetnuti taj tekst unutar elementa `label` i dodati atribut elementa `for` sa vrednošću jednaku atributu elementa `input` iz elementa `id`.

Ako postoji funkcija/karakteristika `input`, kao što je prikazano ispod:

```html
<input id="breakfast" type="radio" name="meal" value="breakfast">
```

Primer elementa za `label` povezan sa ovim elementom `input` je:

```html
<label for="breakfast">Breakfast</label>
```

Povežite tekst od `Loving` i padajući meni za sortiranje unutar teksta `Loving` isključivo u elementu `label`, i dodelite mu odgovarajuće svojstvo elementa `for`.

# --hints--

Morate osigurati da je element za padajući meni i dalje prisutan i da niste napravili nikakve izmene unutar njega. Mora ostati onako kako je bio u originalnom kodu.

```js
const checkboxElementRegex = /\s*<input\s+id\s*=\s*"loving"\s*type\s*=\s*"checkbox"\s*>/;
assert.match(code, checkboxElementRegex);
```

Vaš izborni okvir bi trebalo da ima svojstvo atributa `id` sa vrednošću `loving`. Očekivalo se `--fcc-expected--`, ali je pronađeno `--fcc-actual--`.

```js
const checkboxElementId = document.querySelector('input[type="checkbox"]')?.getAttribute('id');
assert.strictEqual(checkboxElementId, 'loving');
```

Treba da dodate novi element `label` nakon izbornog okvira. Uverite se da ima oznaku za otvaranje i oznaku za zatvaranje.

```js
assert.lengthOf(document.querySelectorAll('label'), 3);
assert.lengthOf(code.match(/<\/label\>/g), 3);
```

Vaš element `label` treba da bude nakon vaše selektne kutije.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElement = checkboxElement?.nextElementSibling?.nodeName;
assert.strictEqual(checkboxSiblingElement, 'LABEL');
```

Vaš element za `label` treba da ima tekst od `Loving`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementText = checkboxElement?.nextElementSibling?.innerText;

assert.strictEqual(checkboxSiblingElementText, 'Loving');
```

Tekst etikete `Loving` treba da bude samo unutar tvog novog elementa `label`. Osigurajte se da ne ponavljate taj tekst izvan elementa `label`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingNode = checkboxElement?.nextSibling?.nodeValue ?? '';
const labelElement = checkboxElement?.nextElementSibling;
const labelSiblingNode = labelElement?.nextSibling?.nodeValue ?? '';

assert.notInclude(checkboxSiblingNode, 'Loving');
assert.notInclude(labelSiblingNode, 'Loving');
```

Vašem elementu `label` nedostaje atribut `for`. Uverite se da postoji razmak između naziva otvarajuće oznake `label` i atributa `for`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.isNotNull(checkboxSiblingElementAttr);
```

Vaše svojstvo od `label` trebalo bi imati atribut svojstva `for` sa vrednošću `loving`. Očekivalo se `--fcc-expected--`, ali je pronađeno `--fcc-actual--`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.strictEqual(checkboxSiblingElementAttr, 'loving');
```

Treba da postoji razmak između vaše opcije za odabir i vašeg novog elementa/komponente `label`.

```js
assert.match(code, />\s+<label\s+for\s*=\s*('|")loving/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
