---
id: 5f07c98cdb9413cbd4b16750
title: Korak 19
challengeType: 0
dashedName: step-19
---

# --description--

Vreme je da se doda novi deo. Dodaj element `section`, drugi, ispod postojećeg elementa `section`.

# --hints--

Vaš element za `section` treba da ima oznaku otvaranja. Oznake otvaranja imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 2);
```

Treba ti dodati samo jednu oznaku za otvaranje `section`. Molimo te ukloni sve višak sadržaje.

```js
assert(document.querySelectorAll('section').length === 2);
```

Vaš element za `section` treba da ima zatvarajući tag. Zatvarajući tagovi su `/` odmah nakon slova `<`.

```js
assert(code.match(/<\/section>/g).length >= 2);
```

Treba da dodate samo jednu zatvarajuću oznaku `section`. Molimo vas da uklonite sve dodatke/višak elemente.

```js
assert(code.match(/<\/section>/g).length === 2);
```

Drugi element `section` ne bi trebalo da se postavi unutar prvog elementa `section`.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

Oba elementa od `section` bi trebalo da budu između otvarajućeg i zatvarajućeg tag-a elemenata `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
