---
id: 5f07be6ef7412fbad0c5626b
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Kada dodajete bilo kakav novi sadržaj, trebalo bi da koristite element `section` za razdvajanje sadržaja slika mačaka od budućeg sadržaja.

Element `section` se koristi za definisanje sekcije u dokumentu, kao što su poglavlja, naslovi, podnaslovi ili bilo koji drugi deo dokumenta. To je semantički element koji pomaže SEO-u i pristupačnosti.

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

Uklonite vaše komponente od `h2`, `p`, i anker tag (`a`) na rasporedu unutar komponente `section`.

# --hints--

Vaš element `section` treba da ima oznaku za otvaranje. Oznake za otvaranje imaju sledeću sintaksu: `<elementName>`.

```js
assert(document.querySelector('section'));
```

Vaša komponenta za `section` treba da ima zatvarajući tag. Zatvarajući tagovi se pojavljuju `/` odmah nakon karaktera `<`.

```js
assert(code.match(/<\/section\s*>/i));
```

Onemogućen element `section` treba da bude između otvarajućeg i zatvarajućeg oznaka elementa `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

Postojeći elementi od `h2`, dva elementa `p`, i sidreni tag (`a`) bi trebalo da budu između otvarajućeg i zatvarajućeg oznake za element `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4);
```

Element od `h1` ne sme biti postavljen unutar elementa `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```
