---
id: 5f07fb1579dc934717801375
title: Korak 35
challengeType: 0
dashedName: step-35
---

# --description--

Vreme je da dodate novi deo. Dodajte treći element za `section` ispod drugog elementa za `section`.

# --hints--

Vaš element za `section` treba da ima oznaku otvaranja. Oznake otvaranja imaju ovu sintaksu: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 3);
```

Trebalo bi da dodaš samo jednu oznaku za otvaranje `section`. Molimo te, ukloni sve suvišnosti.

```js
assert(document.querySelectorAll('section').length === 3);
```

Element vašeg `section` treba da ima zatvarajući tag. Zatvarajući tagovi su `/` odmah nakon slova `<`.

```js
assert(code.match(/<\/section>/g).length >= 3);
```

Treba da dodate samo jednu zatvarajuću oznaku `section`. Molimo vas da uklonite sve dodatke.

```js
assert(code.match(/<\/section>/g).length === 3);
```

Svi elementi iz `section` moraju biti između otvarajućeg i zatvarajućeg oznake elementa `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const sectionElemsFound = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(sectionElemsFound.length === 3);
```

Poslednji element `section` ne bi trebalo da ima sadržaj. Uklonite bilo koji element HTML ili tekst unutar elementa `section`.

```js
assert.lengthOf(document.querySelectorAll('main > section')?.[2]?.children, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
