---
id: 5f1a89f1190aff21ae42105a
title: Korak 62
challengeType: 0
dashedName: step-62
---

# --description--

Kao radio dugmići, podaci forme za izabrane čekotaksuše su parovi atributa elemenata za `name` / `value`. Iako je atribut `value` opcionalan, najbolja praksa je da se uključi zajedno sa čekotaksušima ili radio dugmićima na stranici.

Dodajte atribut `value` svakoj čekirtaži. Za jednostavnost, postavite atribut `value` svake čekirtaže da ima istu vrednost kao njen atribut `id`.

# --hints--

Sve tri odabrane kutije moraju imati atribut `value`. Osigurajte da postoji razmak nakon imena otvarajuće oznake i/ili da postoji razmak prije svih atributa elemenata.

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')];
assert(checkboxes.every((checkbox) => checkbox.hasAttribute('value')));
```

Svojstvo `value` za kontrolu izbora `Loving` treba da se postavi na `loving`. Izostavili ste vrednost ili postoje greške u pravopisu. Zapamtite da vrednosti svojstava elemenata treba da budu okružene navodnicima.

```js
const lovingCheckbox = document.querySelector('#loving');
assert(lovingCheckbox.getAttribute('value').match(/^loving$/));
```

Svojstvo `value` za čekiricu `Lazy` treba biti postavljeno na `lazy`. Ostavili ste vrednost ili postoje greške u kucanju. Zapamtite da vrednosti svojstava elemenata treba biti okružene navodnicima.

```js
const lazyCheckbox = document.querySelector('#lazy');
assert(lazyCheckbox.getAttribute('value').match(/^lazy$/));
```

Svojstvo `value` za čekboks (checkbox) `Energetic` treba da se postavi na `energetic`. Nedostaje vrednost ili postoje greške u kucanju. Zapamtite da vrednosti svojstava elementa moraju biti obavijene navodnicima.

```js
const energeticCheckbox = document.querySelector('#energetic');
assert(energeticCheckbox.getAttribute('value').match(/^energetic$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality"> <label for="energetic"> Energetic</label>
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
