---
id: 671141d8e32fe934c26fa1be
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Dodajte tekst `See more ` pre jangkora elementa i ` in our gallery` posle jangkora elementa.

# --hints--

Trebalo bi da koristite postojeću sidro funkciju, ne kreirajte novu funkciju.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Tekst unutar anchor elementa ne sme da se menja, mora biti `cat photos`.

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

Trebalo bi da imaš `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` u svom kodu.

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
