---
id: 7cf9b03d81a65668421804c3
title: Korak 41
challengeType: 0
dashedName: step-41
---

# --description--

Ako su podaci forme prisutni na mestu definisanom u atributa elementa `action`, morate postaviti atribut elementa `name` u tekstualno polje i uneti vrednost koja predstavlja podnesene podatke.

Ovo je primer elementa za `input` sa svojstvom elementa `name`:

```html
<input type="text" name="name">
```

Dodajte svojstvo elementa `name` sa vrednošću `catphotourl` u vaše tekstualno polje.

# --hints--

Izbrisali ste svoj element `input` ili ima neispravnu sintaksu. Sve vrednosti atributa elemenata moraju biti obavijene navodnicima.

```js
assert(document.querySelectorAll('input').length);
```

`form` treba da ima samo element `input`. Obrišite druge komponente od HTML ili tekstove sadržane unutar elemenata `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
  document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Vaš element `input` nema atribut elementa `name`. Osigurajte da postoji razmak nakon imena otvarajućeg taga i/ili da postoji razmak pre svih naziva atributa elemenata.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('name'));
```

Element `input` mora imati atribut elementa `name` sa vrednošću `catphotourl`. Ostavili ste vrednost ili postoje greške u kucanju. Zapamtite da vrednosti atributa elemenata moraju biti obuhvaćene navodnicima.

```js
assert.match(
  document.querySelector('input')?.getAttribute('name'),
  /^catphotourl$/i
);
```

Iako ste postavili atribut elementa ``input`` elementa ``name`` na ``catphotourl``, savetuje se da uvek obuhvatite vrednost atributa elementa i navodnike.

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
