---
id: 662a6bc12cde72c32fb526f0
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

<dfn>Interfejs</dfn> je kao šema klase objekata. Interfejs sadrži set metoda i atributa koje klasa mora da implementira.

Počnite sa ovim projektom deklarisanjem prazne klase pod nazivom `Equation`. Koristićete ovu klasu za definisanje interfejsa, opšte mape za izračunavanje.

# --hints--

Trebalo bi da definišete novu klasu objekata nazvanu `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
