---
id: 662bde88dc84f1e249801b1a
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` stoji za *Osnovnu apstraktnu klasu*. Klasa `ABC` vam omogućava da pretvorite regularnu klasu u apstraktnu klasu, koja je klasa korišćena kao šablon (blueprint) za stvarnu klasu.

Promenite svoju deklaraciju za `import` da biste uvezli samo klasu `ABC` iz modula `abc`. Možete uvesti određeni element `x` iz modula `y` prateći strukturu import `from y import x`.

Zatim, promeni svoju klasu `Equation` u apstraktnu klasu tako što ćeš je učiniti da nasleđuje od `ABC`.

# --hints--

Trebalo bi da uvezete `ABC` iz modula `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

Darasa lako la `Equation` linapaswa kurithi kutoka `ABC`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass
    
eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
