---
id: 662bd8260da84bdd5feae419
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Klasa ``LinearEquation`` treba da implementira, umesto samo da nasledi sve metode deklarisane unutar klase ``Equation``, koja bi trebalo da se koristi kao interfejs.

Trenutno, klasa `Equation` je samo roditeljska klasa za `LinearEquation`. U narednim koracima, naučićete kako ga promeniti u službeni interfejs.

Trenutno, kreiraj instancu od `Equation` i postavi njenu vrednost u parametar `eq`, a instancu od `LinearEquation` i postavi njenu vrednost u parametar `lin_eq`.

# --hints--

Treba da deklariš parametar `eq` i postaviš mu vrednost kao `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

Trebalo bi da deklarišete parametar `lin_eq` i postavite mu vrednost kao što je `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
