---
id: 662bdd364bf2cde1487922a9
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Za razliku od drugih programskih jezika, Python ne implementira interfejse u svom osnovnom jeziku, ali standardna biblioteka za Python vam omogućava da definišete interfejse na jednostavan način.

Za ovaj projekat, koristićeš alate iz modula `abc`. Stoga, import ovaj modul u svom kodu.

# --hints--

Trebalo bi da koristiš modul import, `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
    

eq = Equation()
lin_eq = LinearEquation()
```
