---
id: 662bd552e1c1d2db1b88ba47
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Sada, definiši još jednu klasu koja se naziva `LinearEquation` i neka nasleđuje od `Equation`. Koristićeš ovu klasu za predstavljanje jednačine linije.

# --hints--

Trebalo bi da definišete klasu objekata koja se zove `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

Tvoja klasa za `LinearEquation` mora da nasleduje od klase za `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
