---
id: 662bd456896f16d9bd03f1a6
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Unutar klase `Equation`, definiš dve nove metode primera sa imenima `solve` i `analyze`.

# --hints--

Treba da definišete metod sa imenom `solve` unutar klase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

Tvoja putanja ``solve`` bi trebalo da uzme jedan parametar, ``self``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

Trebalo bi da definišete metod pod nazivom `analyze` unutar klase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

Tvoj put ``analyze`` treba da uzme jedan parametar, ``self``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
