---
id: 662fbcef5f05e1b84f541a0c
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Svaki element jednačine će biti inicijalizovan prosleđivanjem više argumenata, kao što su parametri jednačine, počevši od opsega n za \\( x \\) do opsega nule, uključujući mogući opseg sa vrednošću `0`.

Na primer, `LinearEquation(4, 5)` će predstavljati jednačinu \\( 4x + 5 = 0 \\), gde `4` je nivo prvog stepena (veće ovde) i `5` je nivo nule.

Morate proveriti da je prosleđen tačan broj argumenata za inicijalizovanje objekta jednačine.

Unutar metode `__init__`, kreiraj deklaraciju `if` da proveri da li dužina `args` je različita od broja standarda koje jednačina treba da ima (`degree + 1`). Ako jeste, ukloni `TypeError` i koristi sledeći niz znakova za izdavanje specifične poruke: `f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"`.

Zatim, ispravite grešku prosljeđivanjem `2` i `3` inicijalizovanjem `lin_eq`.

# --hints--

Treba da napraviš deklaraciju `if` koja proverava da li je broj korišćenih kriterijuma za uspostavljanje jednačine različit od `degree + 1`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

Unutar nove `if` naredbe treba da podigneš `TypeError` i koristiš dati tekst kao prilagođenu poruku o grešci.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

Trebalo bi da prođete kroz `2` i `3` da biste aktivirali `lin_eq`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        
    
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation()
--fcc-editable-region--
```
