---
id: 662fa2e2cf27c09f21f4f5d0
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

U Python, tipovi podataka se prepoznaju tokom izvršavanja koda (runtime). Stoga, nema potrebe da navodiš tip podatka parametra kada ga deklariraš. Međutim, možeš definisati tip podatka parametra kako bi se specificiralo da će sačuvati određeni tip podataka korišćenjem `variable: <data type> = value` ili samo `variable: <data type>`. Zapamti da interpretator Python ne proverava tipove korišćene za definisanje parametara, i obično ti je potreban eksterni alat za to.

Unutar klase `Equation`, deklariraj svojstvo elementa klase `degree`. Ne dodeljuj mu vrednost. Umesto toga koristi annotaciju tipa `int` da pokažeš da će sačuvati ceo broj unutar stvarnih klasa.

Kasnije ćete koristiti ovaj atribut klase kao deo procesa validacije argumenata koji se šalju za inicijalizaciju objekata jednačina.

# --hints--

Deklariši atribut klase pod nazivom `degree` i postavi njegovu oznaku tipa na `int` unutar klase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
