---
id: 662f96576ef178927de87975
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Ako se prepozna kao apstraktan način, način mora biti ukrašen dekorativnim elementom `@abstractmethod`.

<dfn>Dekoratori</dfn> koriste se u Python da bi promenili ponašanje funkcije. Evo primera kako koristiti dekorator nazvan `spam`:

```py
@spam
def foo():
    pass
```

Promenite svoju deklaraciju od import ili unesete dekorator za `abstractmethod` i na oba puta za `solve` i `analyze` za klasu `Equation`. Ovo će uzrokovati dve greške.

Kada klasa nasleđuje od `ABC` i ima apstraktnu metodu, ta klasa ne može biti instancirana. Zbog toga, obrišite primer `Equation` da biste uklonili tu grešku.

Druga greška se javlja zato što klasa `LinearEquation` mora da implementira sve apstraktne metode definisane u interfejsu. Uverite se da ih implementirate unutar klase `LinearEquation`. Nije vam dozvoljeno da koristite dekorator `abstractmethod` u stvarnoj klasi.

# --hints--

Treba ti da dovedeš `abstractmethod` iz modula `abc`.

```js
({ test: () => assert(runPython(`
_Node(_code).has_import("from abc import ABC, abstractmethod") or \\
_Node(_code).has_import("from abc import abstractmethod, ABC") or \\
(_Node(_code).has_import("from abc import abstractmethod") and _Node(_code).has_import("from abc import ABC"))
`)) })
```

Trebalo bi da se ukrasi na `@abstractmethod` putem `solve` unutar klase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_decorators("abstractmethod")`)) })
```

Trebalo bi da se ukrasi na `@abstractmethod` put putem `analyze` unutar klase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_decorators("abstractmethod")`)) })
```

Treba da napraviš metod koji se zove `solve` unutar klase `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("solve")`)) })
```

Tvoja putanja ``solve`` treba da uzme jedan parametar, ``self``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_args("self")`)) })
```

Kreiraj metod pod nazivom `analyze` unutar klase `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("analyze")`)) })
```

Tvoja putanja ``solve`` mora uzeti jedan parametar, ``self``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC


class Equation(ABC):
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
