---
id: 662fc3eba556a6bf800d48c1
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Funkcija izgrađena unutar `isinstance()` koja prima dva argumenta i vraća booleanski vrednost, pokazujući da li je nešto što je prosleđeno kao prvi argument instanca klase objekata prosleđenih kao drugi argument.

```py
isinstance(7, int) # True
```

Još jedna stvar koju želite da proverite je da je svaki argument broj. Nakon vašeg `if`, kreirajte petlju `for` koja iterira preko `args` i proverite da li argument u trenutnoj petlji nije instanca od `int` ili `float`. Koristite funkciju `isinstance()` i prosledite tuple koji sadrži `int` i `float` kao drugi argument.

Ako je argument broj, prikaži `TypeError` ako kaže `"Coefficients must be of type 'int' or 'float'"`.

# --hints--

Trebalo bi ti da napraviš krug `for` koji se ponavlja na `args` nakon tvog izjave `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_iter().is_equivalent("args")`)) })
```

Trebalo bi da napravite deklaraciju za `if` koja proverava da li trenutni koeficijent nije primer od `int` ili `float` unutar opsega `for`.

```js
({ test: () => assert(runPython(`
var = str(_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_vars())
cond1 = f'not isinstance({var}, (int, float))'
cond2 = f'not isinstance({var}, (float, int))'
if_stmt = _Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
if_stmt.is_equivalent(cond1) or if_stmt.is_equivalent(cond2)
`)) })
```

Trebalo bi da koristite niz znakova navedenog za prikaz `TypeError` unutar deklaracije `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
