---
id: 662f6d7c92381a3049e4c987
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Interfejs nije samo potreban za definisanje apstraktnih metoda, već može da implementira metode koje će nasleđivati konkretne klase.

Pre izvršenja `solve` na `analyze` unutar klase `Equation`, definišite metodu za `__init__`. Ne koristite nikakav dekorator na njemu.

# --hints--

Treba da definišete metodu `__init__` u vašoj klasi `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

Tvoja putanja ``__init__`` mora uzeti jedan parametar, ``self``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
