---
id: 6639fdcc701833a54c364211
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Nakon validacije vrednosti, morate ih sačuvati u atributu primernog elementa. Koristite dictionary comprehension da kreirate rječnik čiji je ključ klasa vrednosti, a odgovarajuća vrijednost je ta vrednost, zatim postavite tu vrijednost u atribut nazvan `coefficients`.

Na primer, predmet od `LinearEquation` objavljen od strane `2` i `4` mora imati sledeće svojstvo `coefficients`: `{1: 2, 0: 4}`, jer se `2` odnosi na prvi stepen/nivo `x` i `4` se odnosi na nulti stepen/nivo `x`.

Kreirajte parove ključ-vrednost u vašem novom rječniku, prateći isti format kao što je u `args`.

# --hints--

Trebalo bi da objaviš funkciju pod nazivom `coefficients` unutar svog puta `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_variable("self.coefficients")`)) })
```

Trebalo bi da koristiš generisanje rečnika kako bi sačuvao svoje standarde.

```js
({ test: () => runPython(`
import ast
node = _Node(_code).find_class("Equation").find_function("__init__").find_variable("self.coefficients")
assert isinstance(node.tree.value, ast.DictComp)
`) })
```

Vaše svojstvo `coefficients` treba da bude rečnik koji sadrži parove ključ-vrednost u strukturi na nivou klase. Imajte na umu da pratite isti redosled gde su sačuvani parametri unutar `args`.

```js
({ test: () => runPython(`
actual1 = list(LinearEquation(1, 6).coefficients.items())
expected1 = list({1: 1, 0: 6}.items())
actual2 = list(LinearEquation(-3.5, 0).coefficients.items())
expected2 = list({1: -3.5, 0: 0}.items())
assert actual1 == expected1
assert actual2 == expected2
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
--fcc-editable-region--
        
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass

    
lin_eq = LinearEquation(2, 3)
```
