---
id: 6639f947d3a1818c9322c64a
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Poslednji korak validacije nivoa je proveriti da li apsolutni maksimum nije jednak nuli. Zapamtite da apsolutni maksimum treba da se prosledi kao prvi argument prilikom inicijalizacije nečega.

Dodajte izjavu `if`, a uklonite `ValueError` koristeći sledeću sekvencu slova da biste dobili posebnu poruku: `'Highest degree coefficient must be different from zero'`.

# --hints--

Treba da kreirate deklaraciju `if` koja proverava da li je prvi uslov koji je prošao za uspostavljanje jednačine jednak nuli.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_conditions()[0]
cond.is_equivalent("args[0] == 0") or cond.is_equivalent("0 == args[0]") or cond.is_equivalent("not args[0]")
`)) })
```

Treba da pružite `ValueError` unutar nove deklaracije od `if` i koristite priloženi niz karaktera da vratite specifičnu poruku o grešci.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_bodies()[0].has_stmt("raise ValueError('Highest degree coefficient must be different from zero')")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
