---
id: 663a22ba7420c4d2f7fd2aec
title: Korak 25
challengeType: 20
dashedName: step-25
---

# --description--

Vreme je da se primeni metoda `solve`. Ako postoji linearna jednačina u formi \\( ax + b = 0 \\), rešenje je \\(x = -\frac{b}{a}\\).

Izvuci vrednosti koeficijenata sačuvanih u svojstvu elementa `coefficients` i sačuvajte ih u parametre `a` i `b`. Zapamtite da ćete morati da koristite metodu `.values()`.

Zatim, deklariši parametar `x`, postavi mu vrednost na rešenje jednačine i vrati iz metode `solve`.

# --hints--

Treba da pristupi sačuvane vrednosti unutar atributa elementa `coefficients` i sačuvaš ih u parametre `a` i `b`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("a, b = self.coefficients.values()")`)) })
```

Treba da deklariš parametar nazvan `x` i dodeli mu vrednost rešenja linearnog jednačine.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("x = -b/a")`)) })
```

Trebalo bi da vratiš `x` putem svog `solve`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_return("x")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
--fcc-editable-region--
    def solve(self):
        pass
--fcc-editable-region--
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
