---
id: 663b83a28943e6aa6275a514
title: Korak 19
challengeType: 20
dashedName: step-19
---

# --description--

Dok si unutar klase `Equation`, deklariraj metodu `__str__` da pruži tačnu reprezentaciju sekvence karaktera za objekte jednačina koje ćeš kreirati.

Trenutno, unutar puta `__str__`, deklariš parametar `terms` i postavi mu vrednost prazne liste. Koristićeš ovaj parametar za skladištenje svakog člana (koeficijent moći puta) tvog jednačine.

Zatim, deklariraj parametar `equation_string`, postavi mu vrednost na rezultat spajanja elemenata iz liste `terms` koristeći razmak. Na kraju, vrati `equation_string`.

# --hints--

Trebalo bi da objaviš putanju `__str__` unutar klase `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

Tvoj put ``__str__`` treba da uzme jedan parametar, ``self``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

Morate deklarisati parametar `terms` i postaviti vrednost prazne liste unutar putanje `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

Trebalo bi da objaviš parametar `equation_string` i postaviš mu vrednost rezultata spajanja elemenata koji se nalaze unutar `terms` za poziciju unutar puta `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

Trebalo bi da vratiš `equation_string` putem svog `__str__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
