---
id: 663c9f31306353460da54542
title: Korak 23
challengeType: 20
dashedName: step-23
---

# --description--

Kao što možete da vidite, oznaka `+` nije prikazana u rezultatima. Oznaka broja se podrazumevano prikazuje samo ako je negativna. Da biste promenili ovo ponašanje, možete napisati dve tačke nakon izrasčunske ekspresije koja treba da se proceni unutar zagrada vaše f-string komponente, i navesti opciju `+`. Ovo će vam omogućiti da prikažete oznaku za pozitivne i negativne brojeve.

Promeni sekvencu karaktera u vašim dva uslovna bloka dodavanjem `:+` unutar zagrade nakon `coefficient`.

# --hints--

Trebalo bi da promenite niz slova kako biste dodali na listu `terms` unutar vašeg izjava/deklaracije `if` u `f'{coefficient:+}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_bodies()[0].is_equivalent("terms.append(f'{coefficient:+}')")`)) })
```

Trebalo bi da promeniš niz karaktera kako bi uvrstio `terms` u svoj deo `elif`, tako da postane `f'{coefficient:+}x'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[1].find_bodies()[1].is_equivalent("terms.append(f'{coefficient:+}x')")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient}')
            elif n == 1:
                terms.append(f'{coefficient}x')                
        equation_string = ' '.join(terms)        
        return equation_string
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
