---
id: 6675aaf418b41157f6ccd692
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Vreme je za povratak na putanju `__init__`. U zavisnosti od tipa jednačine, moraćete da prosledite različit broj argumenata prilikom inicijalizacije.

Dodaj drugi parametar `args` u tu funkciju i koristi operater `*` da ga učini da prima različit broj argumenata.

# --hints--

Tvoj put ``__init__`` treba da uzme dva parametra: ``self``, i ``*args``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
