---
id: 6675a38a8b535e4ff3274520
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Funkcija `hasattr` koja je ugrađena uzima objekat kao svoj prvi argument i niz znakova koji predstavlja ime svojstva elementa kao svoj drugi argument. Vraća vrednost tačno ili netačno da pokaže da li objekat poseduje pomenuto svojstvo elementa.

Sada ćete koristiti metodu `__init_subclass__` da proverite da li podklasa ima svojstvo elementa `degree` tokom njegove inicijalizacije.

Kreiraj deklaraciju za `if` da proveri da li `cls` nema atribut od `degree`. Ako je tako, prikaži `AttributeError` i koristi niz karaktera `f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"` da generiše specifičnu poruku.

Nakon toga, ispravi grešku koja je nastala u terminalu deklarisanjem atributa člana klase ``degree`` unutar klase ``LinearEquation``. Ovaj atribut treba da predstavlja stepen jednačine, koji je moć najvišeg eksponenta za `\\( x \\)`. Stoga, postavi vrednost celog broja ``1`` za atribut člana ``degree``.

# --hints--

Trebalo bi kreirati deklaraciju za `if` koja proverava da li `cls` nema svojstvo elementa `degree` unutar puta `__init_subclass__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_conditions()[0].is_equivalent("not hasattr(cls, 'degree')")`)) })
```

Trebalo bi da prikažete `AttributeError` koristeći niz slova dat unutar vašeg izraza `if`.

```js
({ test: () => runPython(`
raise_stmt = 'raise AttributeError(f"Cannot create \\'{cls.__name__}\\' class: missing required attribute \\'degree\\'")'
node = _Node(_code).find_class("Equation").find_function("__init_subclass__").find_ifs()[0].find_bodies()[0]
assert node.has_stmt(raise_stmt) 
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        pass

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    
--fcc-editable-region--
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
