---
id: 66759e32b88fb5459b1e0234
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Metoda `__init_subclass__` se poziva svaki put kada se kreira klasa objekata i omogućava personalizaciju podklasa. Ova metoda prima argument nazvan po kontraktu `cls` (koji predstavlja "klasu objekata"), koji prosleđuje novu podklasu.

Napravi put za `__init_subclass__` u svoj razred `Equation` i sa dodaj parametar `cls`.

# --hints--

Trebalo bi da napraviš putanju za `__init_subclass__` sa parametrom `cls` u svojoj klasi `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
