---
id: 66c06d618d075c7f7f1b890a
title: Izgradi proročnika sudbine
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

U ovoj laboratoriji, definisaćete pet sudbina i nasumično će izabrati jednu kako bi prikazao sudbinu korisniku.

**Cilj:** Popunite priče korisnika ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Morate definisati pet parametara `fortune1`, `fortune2`, `fortune3`, `fortune4`, i na `fortune5` sa vrednošću niza karaktera koji izaberete. Možete koristiti opcije ispod ako želite:

- `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. Treba da izaberete nasumičan sretan broj između 1 i 5, uključujući ih, i unesete tu vrednost u parametar `randomNumber`.

3. Trebalo bi da kreirate parametar `selectedFortune` i postavite odgovarajuće stanje u skladu sa ovim pravilima:

- Ako je `randomNumber` jednako 1, postavi vrednost `fortune1` u `selectedFortune`.
- Ako je `randomNumber` jednako 2, postavi vrednost `fortune2` u `selectedFortune`.
- Ako je `randomNumber` jednako 3, postavi vrednost `fortune3` u `selectedFortune`.
- Ako je `randomNumber` jednako 4, postavi vrednost `fortune4` u `selectedFortune`.
- Ako je `randomNumber` jednako 5, postavi vrednost `fortune5` u `selectedFortune`.

4. Treba da ukucate `selectedFortune` u konzolu.

# --hints--

Trebalo bi da inicijalizujete `fortune1` sa vrednošću sekvence slova.

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

Trebalo bi da inicijalizujete `fortune2` sa string vrednošću.

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

Trebalo bi da inicijalizujete `fortune3` sa vrednošću sekvence slova.

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

Trebalo bi da inicijalizujete `fortune4` sa vrednošću sekvence slova.

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

Treba da inicijalizujete `fortune5` sa vrednošću niza karaktera.

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

Trebalo bi da koristiš metodu `Math.random()` za kreiranje slučajnog broja sreće.

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

Treba da generiš slučajni broj između 1 i 5, uključujući, i da mu dodeli vrednost parametru `randomNumber`.

```js
assert.isNotNull(randomNumber);
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length]
  const userCode__test_code__ = new Function(`${code};return randomNumber`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__()
    assert.strictEqual(result, i + 1);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Trebalo bi da imaš kriterijum `selectedFortune` čija je vrednost definisana na osnovu vrednosti `randomNumber`.

```js
assert.isNotNull(selectedFortune);
```

`randomNumber` mora da odgovara njegovoj vrednosti. Na primer, ako je `randomNumber` 1, `selectedFortune` mora biti jednak `fortune1` i tako dalje.

```js
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length];
  const fortunes__test_code__ = [fortune1, fortune2, fortune3, fortune4, fortune5,];
  const userCode__test_code__ = new Function(`${code};return selectedFortune`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__();
    assert.strictEqual(result, fortunes__test_code__[i]);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Treba da unesete `selectedFortune` u konzolu.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
