---
id: 66da326c02141df538f29ba5
title: Kreiraj stranicu za plaćanje
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Mora da postoji komponenta `h1` sa tekstom `Checkout`.
1. Treba da postoji dva elementa `section` odmah nakon komponente `h1`.
1. Mora da postoji komponenta `h2` sa tekstom `Your Cart` u prvom delu.
1. U prvom delu mora postojati slika proizvoda sa odgovarajućim alternativnim tekstom. Možete koristiti ovu sliku ako želite: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. Mora da postoji komponenta `h2` sa tekstom `Payment Information` u drugom delu.
1. Mora da postoji komponenta `form` u drugom delu.
1. U vašem formularu mora postojati polje sa `id` i `name` od `card-name`, i `type` od `text`, zajedno sa povezanim `label`.
1. U vašem formularu mora postojati polje sa `id` i `name` od `card-number`, i `type` od `text`, zajedno sa povezanim `label`.
1. Mora da postoji najmanje dva elementa `input` sa karakteristikom `required`.
1. Treba uključiti komponentu `span` sa tekstom `*` i `aria-hidden` postavljenu kao `true` unutar komponente `label` za svako potrebno polje, kako bi bile jasno vidljive potrebne sekcije.
1. Mora da postoji komponenta `p` sa tekstom pomoći koji opisuje zahtevani format broja kartice, postavljen odmah nakon ulaznog polja za broj kartice. `p` treba da ima `id` od `card-number-help` i da se referiše na polje broja kartice koristeći `aria-describedby`.

# --hints--

Trebalo bi da imaš element sa tekstom `h1` koji sadrži tekst `Checkout`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

Trebalo bi da imate samo jedan element `h1` na svojoj stranici.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Trebalo bi da ima najmanje dva elementa `section` nakon elementa `h1`.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

Treba da ima element sa `h2` koji sadrži tekst `Your Cart` unutar prve sekcije.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

Trebalo bi ti da imaš barem jednu sliku sa alternativnim tekstom u tvom prvom delu.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

Treba da imate element `h2` sa tekstom `Payment Information` unutar drugog odeljka.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

Trebalo bi da imaš element `form` unutar drugog elementa, `section`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

Trebalo bi da imaš polje `input` sa `id` i `name` za `card-name` unutar tvog formulara.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

Vaš unos imena kartice mora imati `type` od `text`.

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

Trebalo bi da imaš polje sa `input` koje sadrži `id` i `name` za `card-number` unutar tvog formulara.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

Unos vašeg broja kartice mora imati `type` od `text`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

Svi elementi od `input` koji nisu `type` iz `submit` moraju imati element `label` koji sa njima odgovara.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
});
```

Trebalo bi da imaš najmanje dva elementa sa karakteristikom `required` iz `input`.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

Treba da uključite element `span` koji sadrži tekst `*` unutar polja `label` za svaki neophodan unos.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

Vaše komponente za `span` moraju imati `aria-hidden` postavljeno na `true`.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

Trebalo bi da imate polje `p` sa `id` za `card-number-help` odmah nakon unosa broja kartice.

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
const nextElement = cardNumberInput.parentElement?.tagName === 'LABEL'
  ? cardNumberInput.parentElement.nextElementSibling
  : cardNumberInput.nextElementSibling;
assert.equal(nextElement, cardNumberHelp);
```

Vaši podaci za pomoć sa brojem kartice ne smeju biti prazni.

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

Unos vašeg broja kartice mora imati `aria-describedby` podešen na `card-number-help`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
