---
id: a3bfc1673c0526e06d3ac698
title: Konstruktor kalkulatora za sabiranje osnovnih brojeva
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

U ovoj laboratoriji, ćeš napraviti kalkulator koji prima brojeve i vraća sumu svih premijnih brojeva koji su manji ili jednaki tom broju.

**Cilj**: Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da ima funkciju `sumPrimes` koja prima brojeve kao argumente.
1. Funkcija `sumPrimes` treba da vrati zbir svih osnovnih brojeva manjih ili jednakih datoj/unesenom broju.
1. Ako je ulazni broj manji od `2`, funkcija treba da vrati `0`.

# --hints--

Trebalo bi da imaš funkciju `sumPrimes`.

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` trebalo bi da vrati `17`.

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` trebalo bi da vrati `10`.

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` treba da vrati `2`.

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` trebalo bi vratiti `0`.

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` treba da vrati `73156`.

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
