---
id: a3566b1109230028080c9345
title: Napravi algoritam za sabiranje svih brojeva
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

U ovoj laboratoriji ćeš morati da napraviš algoritam za sabiranje svih brojeva.

**Cilj:** Popunite korisničke priče ispod i prođite sve testove da biste završili ovu laboratoriju.

**Priče korisnika:**

1. Trebalo bi da postoji funkcija pod nazivom `sumAll` koja prima niz od dve brojeve.
1. `sumAll([n, m])` treba da vrati zbir `n` i `m` zajedno sa zbirom svih brojeva između njih. Manji broj neće uvek biti prvi. Na primer, `sumAll([4,1])` treba da vrati `10` jer je zbir svih brojeva između `1` i `4` (uključujući sve) `10`.

# --hints--

Trebalo bi da imaš funkciju koja se zove `sumAll`.

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])` Trebalo bi da vrati broj.

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])` treba da vrati `10`.

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` trebalo bi da vrati `10`.

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` treba da vrati `45`.

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` treba da vrati `45`.

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;
  
  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
