---
id: 68d275dd800f404d22a07564
title: Implementiraj algoritam dubinskog pretraživanja prvo
challengeType: 27
dashedName: lab-depth-first-search
---

# --description--

U ovoj laboratoriji, implementiraćeš algoritma za pretraživanje grafa koji se zove <dfn>dubinska pretraga</dfn>.

Wakati <dfn>široko-prva pretraga</dfn> traži korak po dužini ivica od izvornog čvora, <dfn>duboko-prva pretraga</dfn> počinje spuštanjem po putanji ivica sve dok ne stigne do mesta gde može da stigne.

Kada stigne do kraja puta, pretraga se vraća na poslednji čvor koji ima neposetašenu ivicu i nastavlja sa pretragom.

Za razliku od pretrage po širini, svaki put kada se čvor posjeti, on ne posjećuje sve svoje susjedstvo. Umjesto toga, počinje posjećivanjem jednog susjeda i nastavlja nizak tom putanjom sve dok ne dođe do tačke gdje nema drugih čvorova za posjetu na toj putanji.

Da biste implementirali ovaj algoritam, moraćete koristiti stack (struktura podataka gde je poslednji dodati element prvi uklonjen, prateći princip <dfn>Last-In-First-Out</dfn>). Stack je koristan u algoritmima za pretragu dubine jer kada dodajete susjednike na stack, želite da posetite najnovije dodate susjednike prvo i uklonite ih sa steka.

Jednostavan rezultat ovog algoritma je lista čvorova koji su dostignivi iz određenog čvora. Stoga, takođe ćete morati da pratite čvorove koje posjećujete.

**Cilj**: Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da postoji funkcija nazvana `dfs`.
1. Funkcija `dfs` treba da primi dva argumenta:
   - Asilinearna matrica susedstva.
   - Oznaka čvora (node label), koja je vrednost broja čvora između `0` i `n - 1`, gde `n` predstavlja ukupan broj čvorova u grafu.
1. Funkcija `dfs` treba da implementira algoritam dubinskog pretraživanja kako bi prvo vratila listu svih dostupnih čvorova od čvora koji joj je prosleđen.

# --hints--

Trebalo bi da imaš funkciju nazvanu `dfs` koja prima dva argumenta.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs)
  sig = inspect.signature(dfs)
  assert len(sig.parameters) == 2
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` treba da vrati listu koja sadrži `1`, `2`, `3`, na `0`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {1, 2, 3, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` trebalo bi da vrati listu koja sadrži `1`, `2`, `3`, i na `0`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {3, 2, 1, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` treba da vrati `[3]`.

```js
({ test: () => runPython(`
  assert dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3) == [3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` trebalo bi da vrati listu koja sadrži `3` i `2`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)
  assert result == [3, 2] or result == [2, 3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` trebalo bi da vrati listu koja sadrži `0` i `1`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)
  assert result == [0, 1] or result == [1, 0]
`) })
```

Funkcija `dfs` treba da vrati ispravne rezultate.

```js
({ test: () => runPython(`
  import random

  def solve(g, s):
      v, q = set(), [s]
      while q:
          u = q.pop()
          if u not in v:
              v.add(u)
              q.extend([i for i, x in enumerate(g[u]) if x == 1 and i not in v])
      return v

  random.seed(0)

  for _ in range(10):
      n = 6
      mat = [[0]*n for _ in range(n)]
      for i in range(n):
          for j in range(i+1, n):
              if random.random() > 0.5:
                  mat[i][j] = mat[j][i] = 1

      start = random.randint(0, n-1)
      assert set(dfs(mat, start)) == solve(mat, start), "Random graph check failed"
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs(graph, root):
    stack = []
    temp_v = None
    visited = []
    temp_v_neighbors = []
    stack.append(root)
    while stack:
        temp_v = stack.pop()
        if temp_v not in visited:
            visited.append(temp_v)
            temp_v_neighbors = graph[temp_v]
            for n, is_neighbor in enumerate(temp_v_neighbors):
                if is_neighbor == 1:
                    stack.append(n)
    return visited
```
