---
id: 0095d0567b6817114989a1dd
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, ćete implementirati algoritam najkraćeg puta koristeći JavaScript. Napisaćete funkciju koja računa najkraći put između čvorova u težiteljenom grafu i vraća rastojanje zajedno sa pređene staze.

Na primer, uzimajući u obzir graf gde su gradovi povezani putevima različitih udaljenosti, algoritam će pronaći najkraći put od jednog grada do drugog. Ako želite putovati od Grada A do Grada D, algoritam može otkriti da je putovanje A ⇨ B ⇨ C ⇨ D (ukupno: km 15) kraći nego direktan put A ⇨ D (km 20).

Da biste počeli, definisati parametar nazvan `INF` i postaviti mu vrednost `Infinity`. Kasnije ćete ga koristiti da pokažete da ne postoji direktna veza između dve čvora (node).

# --hints--

Morate definisati parametar koji se zove `INF`.

```js
assert.isDefined(INF);
```

Vaš parametar za `INF` treba da ima vrednost `Infinity`.

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
