---
id: 69c918841987bf16c1d1bac3
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

Unutarici ciklus od `for`, potrebno je proveriti da li trenutni čvor nije posetao i da li je bliži od najboljeg koji ste do sada pronašli.

Izjava za `if` koja proverava da li `nodeNo` nije posetano od strane `distances[nodeNo]` i da je manje od `minDistance`.

Ostavite telo prazno za sada.

# --hints--

Mora imati deklaraciju `if` sa uslovom od `!visited[nodeNo] && distances[nodeNo] < minDistance` unutar unutrašnjeg ciklusa `for`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*!visited\s*\[\s*nodeNo\s*\]\s*&&\s*distances\s*\[\s*nodeNo\s*\]\s*<\s*minDistance\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
