---
id: 69c918841987bf16c1d1bac7
title: Korak 16
challengeType: 1
dashedName: step-16
---

# --description--

Sada nakon što je trenutni čvor označen kao posetjen, moraš proveriti sve njegove susede da vidiš da li ih možeš dostići efikasnije preko tog čvora.

Nakon `visited[current] = true`, dodaj susedski krug `for` koji prolazi kroz svaki `nodeNo` od `0` do `n`. Napomena, ovo je drugi interni krug na istom nivou kao prvi.

U ovom ciklusu, deklariši parametar nazvan `distance` sa vrednošću `matrix[current][nodeNo]`. Ovo ti daje težinu ivice između trenutnog čvora i njegovog suseda.

# --hints--

Treba da dodate susjedni krug `for` nakon objavljivanja trenutnog čvora kao posjećenog.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

Ciklus susjeda `for` bi trebalo da se ponovi od `0` do `n`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

Trebalo bi da imaš kriterijum `distance` sa vrednošću `matrix[current][nodeNo]` unutar ciklusa suseda.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
