---
id: 69c918841987bf16c1d1bac5
title: Korak 14
challengeType: 1
dashedName: step-14
---

# --description--

Nakon analize svih čvorova radi pronalaženja najbližeg neposetjenog, morate obraditi situaciju u kojoj nije dostupan nijedan validan čvor. Ovo se dešava kada sve preostale čvorove nema.

Nakon unutrašnje petlje `for`, ali i dalje unutar spoljašnje petlje, dodajte deklaraciju `if` koja proverava da li je `current` potpuno jednako `-1`. Ako je to tačno, koristite `break` unutar bloka `if` za ranije prekidanje spoljašnje petlje.

# --hints--

Trebalo bi dodati izjavu za `if` koja proverava kao `current === -1`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)/
);
```

Trebalo bi da koristiš `break` unutar bloka `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)[\s\S]*?break/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
