---
id: 69c918841987bf16c1d1babd
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Potrebno je da se prati najkraća udaljenost poznata od početnog čvora do svakog drugog čvora. Početno, pretpostavlja se da je svaki čvor na beskonačnoj udaljenosti.

U JavaScript, možete kreirati kolonu podataka popunjenu vrednostima koristeći `new Array(n).fill(value)`. Primer: `new Array(3).fill(0)` stvara `[0, 0, 0]`.

Unutar funkcije `shortestPath`, kreiraj parametar koji se zove `distances` i postavi vrednost na `new Array(n).fill(INF)`.

# --hints--

Treba da kreiraš parametar koji se zove `distances` unutar `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

Vaš parametar za `distances` mora biti postavljen na `new Array(n).fill(INF)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
