---
id: 69c918841987bf16c1d1bac4
title: Korak 13
challengeType: 1
dashedName: step-13
---

# --description--

Ako je uslov tačan, trenutni čvor je najbolji kandidat koji do sada nije posetjen. Morate ispraviti vaše parametre za praćenje kako biste to prikazali.

U izjavi `if`, ažuriraj `minDistance` u `distances[nodeNo]` i `current` u `nodeNo`.

# --hints--

Treba da staviš `minDistance` u `distances[nodeNo]` unutar dela `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /minDistance\s*=\s*distances\s*\[\s*nodeNo\s*\]/
);
```

Trebalo bi da staviš `current` u `nodeNo` unutar sekcije `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /current\s*=\s*nodeNo/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
