---
id: 69c918841987bf16c1d1bac1
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Sledeće ćete dodati glavnu petlju koja pokreće algoritam. Pokreće se za svaki čvor u grafu, izabirući najbliži neposetačeni čvor svaki put.

Unutar funkcije `shortestPath`, dodaj petlju `for` koja radi `n` puta, koristeći `i` kao varijablu petlje.

Na početku tela kruga, koristeći `let`, deklarisati dva parametra `minDistance` i `current` koja su postavljena za `INF` i `-1` respektivno. Ovo će pratiti najbliži neposetačeni čvor dostupan tokom svakog ciklusa.

# --hints--

Treba da ima krug od `for` unutar `shortestPath` koji se izvršava jednom `n`. Ne zaboravi da staviš `i` i `let` unutra.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

Trebalo bi da imaš kriterijum od `minDistance` postavljen za `INF` unutar ciklusa.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

Trebalo bi da imate parametar `current` postavljen za `-1` unutar kruga.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
