---
id: 69c918841987bf16c1d1babf
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Takođe treba pratiti stvarnu putanju do svakog čvora. Sačuvati ćete listu putanja gde je svaki element niz podataka indikatora čvora koji predstavlja pređenu putanju.

Početno, putanja svakog čvora će sadržati samo taj čvor. Možete kreirati ovu strukturu koristeći `Array.from()`:

```js
Array.from({ length: n }, (_, i) => [i]);
```

Ovo stvara red elemenata `n`, gde je element koji se nalazi u primeru `i` to `[i]`.

Još unutar funkcije ``shortestPath``, definiši parametar nazvan ``paths`` i postavi vrednost ``Array.from({ length: n }, (_, i) => [i])``.

# --hints--

Trebalo bi da kreirate parametar `paths` unutar funkcije `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

Vaš parametar `paths` mora biti inicijalizovan korišćenjem `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

Tvoj parametar `paths` mora koristiti `{ length: n }` kao prvi argument za `Array.from`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

Drugi argument za `Array.from` treba da bude poziv za vraćanje `(_, i) => [i]`, gde je `_` placeholder za neupotrebljen element i `i` je primer. Trebalo bi da vrati `[i]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
