---
id: 69c918841987bf16c1d1babb
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Sada ćete kreirati glavnu funkciju koja obavlja sve zadatke.

Funkcija nazvana `shortestPath` koja prima tri parametra: `matrix`, `startNode`, i `targetNode`.

U narednim koracima, kreiraćeš logiku ove funkcije.

# --hints--

Trebalo bi da kreirate funkciju nazvanu `shortestPath`.

```js
assert.isFunction(shortestPath);
```

Tvoja funkcija za `shortestPath` mora da ima parametre `matrix`, `startNode`, na `targetNode`.

```js
const params = __helpers.getFunctionParams(shortestPath.toString());
const names = params.map(p => p.name);
assert.include(names, 'matrix');
assert.include(names, 'startNode');
assert.include(names, 'targetNode');
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--

--fcc-editable-region--
```
