---
id: 69c918841987bf16c1d1bacb
title: Korak 20
challengeType: 1
dashedName: step-20
---

# --description--

Kada se glavni ciklus završi, mora da utvrdi koji čvor ili čvorovi treba da prikažu njegove rezultate.

Ako je dostupan specifičan `targetNode`, prikažite rezultate samo za taj čvor. Inače, prikažite rezultate za sve čvorove.

U JavaScript, možete pronaći sve indikatore podatomog reda kao ponovljiv koristeći `[...Array(n).keys()]`.

Nakon spoljne petlje od `for`, kreiraj varijablu nazvanu `targets`. Koristeći trostruki uslovni izraz, ako `targetNode` nije potpuno jednak `null`, postavi vrednost `[targetNode]`, inače postavi `[...Array(n).keys()]`.

# --hints--

Trebalo bi da napraviš parametar `targets` nakon glavnog kruga.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+targets\s*=/
);
```

Tvoj kriterijum za `targets` treba da koristi ternarni motor uslova koji proverava `targetNode !== null`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /targets\s*=\s*targetNode\s*!==\s*null\s*\?/
);
```

Kada je polje dato, `targets` trebalo bi da bude `[targetNode]`, inače trebalo bi da bude `[...Array(n).keys()]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /\?\s*\[\s*targetNode\s*\]\s*:\s*\[\s*\.\.\.\s*Array\s*\(\s*n\s*\)\s*\.keys\s*\(\s*\)\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
