---
id: 69c937c73f23b1140bcd297a
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Parametar za `targetNode` je opcion. Ako nije dostavljen, funkcija će izračunati najkraći put od `startNode` do svih ostalih čvorova u grafu.

Navedite podrazumevanu vrednost za `targetNode` za parametar `null`.

# --hints--

Parametar `targetNode` njegove funkcije `shortestPath` treba da ima podrazumevanu vrednost `null`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /targetNode\s*=\s*null/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode) {

}
--fcc-editable-region--
```
