---
id: 69c918841987bf16c1d1bac8
title: Korak 17
challengeType: 1
dashedName: step-17
---

# --description--

Prije korekcije udaljenosti, morate potvrditi da je susjed pogodan za razmatranje. Mora imati stvarni rub (`distance !== INF`) i još nije posjećen.

Unutar susjednog ciklusa `for`, dodaj izjavu za `if` koja provjerava da `distance` nije sasvim jednako sa `INF` **i** da `nodeNo` nije posjećeno.

Unutar bloka ``if``, deklarisan je parametar ``newDistance`` sa vrednošću ``distances[current] + distance``. Ovo je ukupna cena dostizanja suseda preko trenutnog čvora.

# --hints--

Treba da dodaš deklaraciju `if` sa uslovom `distance !== INF && !visited[nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*distance\s*!==\s*INF\s*&&\s*!\s*visited\s*\[\s*\w+\s*\]\s*\)/
);
```

Moraš da objaviš `newDistance` sa vrednošću `distances[current] + distance` unutar bloka `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+newDistance\s*=\s*distances\s*\[\s*current\s*\]\s*\+\s*distance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
