---
id: 5d5a813321b9e3db6c106a46
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript je moćan jezik koji vam omogućava da gradite sajtove za saradnju.

*Napomena*: Za sve preostale projekte u ovom kurikulumu, biće vam potreban osnovni nivo znanja o HTML i CSS. Ako ste novi u HTML i CSS, molimo vas da pregledate [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

Da biste počeli, kreirajte svoju osnovnu strukturu od HTML koja uključuje `DOCTYPE`, `html`, `head`, i `body`, zatim dodajte oznaku `meta` za `charset`. Dodajte element/funkcionalnost za `title` i koristite tekst od `RPG - Dragon Repeller` za njega. Uključite oznaku `link` za vašu stilsku listu (stylesheet) kako biste povezali fajl od `styles.css`.

Konačno, kreiraj funkciju za `div` koja ima `id` postavljeno na `game` unutar vaše `body`.

# --hints--

Vaš kod treba da ima referencu na `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Treba da uključite razmak nakon referencije `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Trebalo bi da definišete tip dokumenta kao `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Morate zatvoriti deklaraciju za `DOCTYPE` za `>` nakon tipa.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Element vašeg `html` mora imati oznaku za otvaranje. Ne zaboravite atribut `lang`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Element vašeg `html` mora imati zatvarajući tag.

```js
assert(code.match(/<\/html\s*>/));
```

Tvoje saopštenje za `DOCTYPE` trebalo bi biti na početku tvog HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Trebalo bi da imate etiketu za otvaranje `head`.

```js
assert(code.match(/<head\s*>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Trebalo bi da imaš oznaku za otvaranje `body`.

```js
assert(code.match(/<body\s*>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Elementi za `head` i `body` bi trebalo da budu sestrinski elementi.

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

Element `head` treba da bude unutar elementa `html`.

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

Element za `body` treba da bude unutar elementa `html`.

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

Твој код мора имати елемент `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Vaše polje `meta` mora imati atribut `charset` sa vrednošću `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Vaš kod mora imati karakteristiku `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Trebalo bi da imaš oznaku za zatvaranje `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Vaš kod mora imati element `link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Trebalo bi da imaš element `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Tvoj element `div` mora imati svojstvo `id` sa vrednošću `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Tvoj element `div` mora biti unutar elementa `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
