---
id: 62a115879a6d51422652cbfc
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

U vašoj interaktivnoj igri, korisnici će moći da prate svoje statistike, kupuju oružje i bore se protiv monstrima. Pre nego što nastavite sa interaktivnim delom igre JavaScript, prvo morate kreirati komponente HTML koje će prikazivati informacije o igri.

Kreiraj četiri elementa za `div` unutar tvog elementa za `#game`. Postavi sledeće vrednosti za `id` u redosledu: `stats`, `controls`, `monsterStats`, i `text`.

# --hints--

Trebalo bi da kreiraš četiri nova elementa za `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

Trebalo bi mu dati jedan od novih elemenata iz `div` `id` na `stats`.

```js
assert.exists(document.querySelector('div#stats'));
```

Trebalo bi da mu date jedan od novih elemenata iz `div` `id` na `controls`.

```js
assert.exists(document.querySelector('div#controls'));
```

Trebalo bi da mu/izmi jedan od novih atributa `div` `id` od `monsterStats`.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

Trebalo bi mu dati jedan od novih elemenata za `div` `id` od `text`.

```js
assert.exists(document.querySelector('div#text'));
```

Morate postaviti nove komponente `div` u ispravnom redosledu.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

Trebalo bi da staviš nove elemente `div` unutar elementa `#game`.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
