---
id: 62a1166ed9a56d439c0770e7
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Unutar elementa `#stats` kreiraj tri elementa `span`. Svakom dodeli klasu `stat`. Zatim prvom elementu dodeli tekst `XP: 0`, drugom `Health: 100`, a trećem `Gold: 50`.

# --hints--

Trebalo bi da imaš tri elementa `span` unutar elementa `#stats`.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

Trebalo bi da dodate tri nova svojstva klasi objekta `stat` na `span`.

```js
assert.lengthOf(document.querySelectorAll('#stats > .stat'), 3);
```

Vaš prvi element od `.stat` mora da sadrži tekst koji je dostavljen putem `XP: 0`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

Tvoj drugi element `.stat` treba da sadrži tekst koji dolazi iz `Health: 100`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

Vaš treći deo `.stat` bi trebalo da sadrži tekst koji je dostavljen iz `Gold: 50`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
