---
id: 62a23cb9bc467a147516b500
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Za tvoj komponent sa `#controls`, kreiraj tri komponente za `button`. Prvi treba da ima `id` podešen na `button1`, i tekst od `Go to store`. Drugi treba da ima `id` podešen na `button2`, i tekst od `Go to cave`. Treći treba da ima `id` podešen na `button3`, i tekst od `Fight dragon`.

# --hints--

Trebalo bi da dodate tri elementa od `button` u vaš element `#controls`.

```js
const buttons = document.querySelectorAll('#controls > button');
assert.exists(buttons);
assert.equal(buttons.length, 3);
```

Tvoja prva dugme treba da ima `id` postavljeno na `button1`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.id, 'button1');
```

Tvoja prva dugme treba da ima tekst `Go to store`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.innerText, 'Go to store');
```

Твоа друга дугме треба да има `id` постављено на `button2`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.id, 'button2');
```

Tvoja druga dugme treba da ima tekst `Go to cave`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.innerText, 'Go to cave');
```

Treće dugme bi trebalo da ima `id` postavljeno na `button3`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.id, 'button3');
```

Tvoje treće dugme treba da ima tekst `Fight dragon`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.innerText, 'Fight dragon');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
--fcc-editable-region--
      <div id="controls">

      </div>
--fcc-editable-region--
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
