---
id: 62a2509ba163e020bb9d84ea
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Sada možeš početi da pišeš JavaScript. Počni kreiranjem elementa `script`. Ovaj element se koristi za učitavanje JavaScript-a u HTML datoteku.

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

Trebalo bi da imaš element `script`.

```js
assert.isAtLeast(document.querySelectorAll('script').length, 1);
```

Vaš element `script` treba da ima oznaku za otvaranje.

```js
assert.match(code, /<script\s*>/i);
```

Vaša komponenta za `script` mora imati zatvarajući label/tag.

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
