---
id: 62a23d1c5f1c93161f3582ae
title: Korak 19
challengeType: 0
dashedName: step-19
---

# --description--

Slično elementu `#stats`, element `#monsterStats` treba da sadrži dva elementa `span`. Dodaj im klasu `stat`. Prvom dodeli tekst `Monster Name: `, a drugom `Health: `. Nakon teksta u svakom elementu dodaj element `strong`, koji sadrži prazan element `span`. Prvom unutrašnjem elementu `span` dodeli `id="monsterName"`, a drugom `id="monsterHealth"`.

# --hints--

Vaša funkcija od `#monsterStats` bi trebalo da ima dva sastava od `span`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

Svi vaši novi atributi `span` moraju imati vrednost `class` od `stat`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

Prvi deo vašeg `span` treba da ima deo `strong` sa unutrašnjim, praznim delom `span`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Drugi vaš element za `span` mora imati polje `strong` koje sadrži ugniježđeno polje unutar `span` koje nema sadržaja.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Tvoj prvi element od `span` treba da ima tekst `Monster Name: `. Proveri da je pozicija tačna.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+/);
```

Vaš drugi element od `span` treba da sadrži tekst `Health: `. Proverite da je razmak tačan.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[1].innerHTML, /Health:\s+/);
```

Vaše komponente za `strong` i `span` treba da dolaze nakon teksta.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+<strong>/);
assert.match(spans[1].innerHTML, /Health:\s+<strong>/);
```

Tvoj prvi interni element za `span` mora imati `id` od `monsterName`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[0]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterName');
```

Tvoj drugi interni element `span` treba da ima `id` od `monsterHealth`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[1]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterHealth');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
--fcc-editable-region--
      <div id="monsterStats">

      </div>
--fcc-editable-region--
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
