---
id: 62a255dae245b52317da824a
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Jedan od najmoćnijih alata je vaš konzola za programere. U zavisnosti od vašeg pretraživača, ovo se može otvoriti klikom na `F12` ili `Ctrl+Shift+I`. Na Mac-u, možete kliknuti na `Option + ⌘ + C` i izabrati "Console". Takođe možete kliknuti na dugme "Console" koje se nalazi na vrhu prozora za prikaz kako biste videli naš ugrađeni konzolu.

Developer konzola će prikazivati greške iz tvog koda, ali je takođe možeš koristiti da vidiš vrednosti varijabli u svom kodu, što je korisno za otklanjanje grešaka.

Dodajte liniju `console.log("Hello World");` između vaših tagova za `script`. Zatim kliknite na dugme "Console" da biste otvorili konzolu. Trebalo bi da vidite tekst `"Hello World"`.


# --hints--

Moraš da dodaš liniju `console.log("Hello World");` u svoj kod.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

Vaš red za `console.log("Hello World");` treba da bude između vaših oznaka za `script`.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
