---
id: 62a257659d0d1e2456f24ba2
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Pre nego što počnete da pišete kod za svoj projekat, trebalo bi ga prebaciti u sopstvenu datoteku kako biste sve organizovali.

Uklonite svoj marker od `console.log("Hello World");`. Zatim mu dodajte vaše polje `script` koje je sada prazno, atribut polja `src` postavljeno na `./script.js`.

# --hints--

U tvom kodu ne sme biti linija `console.log("Hello World");`.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Tvoj element za `script` mora da ima svojstvo elementa `src` postavljeno na `./script.js`.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
