---
id: 62a3b41c9494f937560640ab
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Navedi još dva parametra nazvana `monsterHealth` i `inventory`.

Za vaš parametar sa `inventory`, postavite mu vrednost na niz karaktera `"stick"`.

Zapamti da si radio sa podacima (data) u prethodnom projektu:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

Trebalo bi da koristiš `let` za deklarisanje svog parametra `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Trebalo bi da koristiš `camelCase` za ime svog parametra ``monsterHealth``.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` Ne bi trebalo da ima vrednost.

```js
assert.isUndefined(monsterHealth);
```

Ne smete da postavite vrednost za vaš parametar `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

Trebalo bi da koristiš `let` za deklarisanje svog parametra `inventory`.

```js
assert.match(code, /let\s+inventory/i);
```

Vaš kriterijum za `inventory` treba i dalje biti red podataka.

```js
assert.isArray(inventory);
```

Vaš parametar za `inventory` mora imati samo jednu vrednost.

```js
assert.lengthOf(inventory, 1);
```

Vaše polje za `inventory` mora da sadrži niz slova `"stick"`.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
