---
id: 62a39f5a5790eb27c1e5d4bf
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Vaš prikaz je prebačen na vaš novi fajl sa `script.js`. Zapamtite da možete koristiti tabove na vrhu za prebacivanje između fajlova.

Dodajte svoju liniju `console.log("Hello World");` u ovaj fajl, i videće se da je prikazano na vašoj konzoli.

# --hints--

Trebalo bi da imaš liniju `console.log("Hello World");` u svom kodu.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
