---
id: 62a3a75d8466a12e009eff76
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Inicijalizuj drugi parametar nazvan `health` sa vrednošću `100`, i drugi parametar nazvan `gold` sa vrednošću `50`.

# --hints--

Trebalo bi da koristiš `let` za deklarisanje parametra nazvanog `health`.

```js
assert.match(code, /let\s+health/);
```

Trebalo bi da inicijalizujete parametar nazvan `health` sa vrednošću `100`.

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

Trebalo bi da koristite `let` za deklarisanje parametra koji se zove `gold`.

```js
assert.match(code, /let\s+gold/);
```

Trebalo bi da inicijalizujete parametar nazvan `gold` sa vrednošću `50`.

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

`health` treba da ima vrednost `100`.

```js
assert.equal(health, 100);
```

`gold` treba da ima vrednost od `50`.

```js
assert.equal(gold, 50);
```

`xp` još bi trebalo da ima vrednost od `0`.

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
