---
id: 62a3b79d520a7f3d0e25afd6
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript se integriše sa HTML-om pomoću <dfn>Document Object Model-a</dfn>, odnosno DOM-a. DOM predstavlja stablo koje opisuje HTML dokument. HTML-u možeš pristupiti preko objekta `document`, koji predstavlja ceo HTML dokument.

Jedan način za dobijanje specifičnih elemenata u vašem HTML je korišćenje metode `querySelector()`. Metoda `querySelector()` uzima selektor CSS kao argument i vraća prvi element koji odgovara tom selektoru. Na primer, da biste dobili element `<h1>` u vašem HTML, napisali ćete:

```js
let h1 = document.querySelector("h1");
```

Zapamti da je `h1` niz slova i odgovara selektoru CSS koji koristiš.

Kreiraj parametar za `button1` i koristi `querySelector()` da postavi vrednost tvog elementa koji ima `id` za `button1`. Zapamti da selektori za CSS za `id` počinju sa prefiksom `#`.

# --hints--

Treba da koristiš `let` za deklarisanje parametra `button1`.

```js
assert.match(code, /let\s+button1/);
```

Trebalo bi da koristiš `document.querySelector()`.

```js
assert.match(code, /document\.querySelector/);
```

Trebalo bi da koristiš alat `"#button1"`.

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
