---
id: 62a3a0a3c0a4b32915d26a6e
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Uklonite svoju liniju `console.log("Hello World");` kada počnete da pišete kod svog projekta.

U prethodnom projektu, naučili ste kako da radite sa parametrima i ključnom reči `let` kao što je ovo:

```js
let camperbot;
```

Takođe si naučio/la kako da definišeš parametar i vrednost kao ovaj:

```js
let age = 32;
```

Koristi glavnu reč `let` da deklariraš parametar nazvan `xp` sa vrednošću `0`, broj.

# --hints--

Ne smete imati red od `console.log("Hello World");` u vašem kodu.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Trebali biste koristiti glavnu reč `let` da deklarišete svoj parametar.

```js
assert.match(code, /let/);
```

Trebalo bi da objavite parametar nazvan `xp`.

```js
assert.match(code, /let\s+xp/);
```

`xp` mora imati vrednost od `0`.

```js
assert.equal(xp, 0);
```

Trebalo bi da inicijalizujete parametar `xp` za `0`.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
