---
id: 62a3b3eab50e193608c19fc6
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Objavi kriterijum nazvan `fighting`, ali ne slušaj ga sa bilo kojom vrednošću.

# --hints--

Trebalo bi da koristite `let` za deklarisanje parametra `fighting`.

```js
assert.match(code, /let\s+fighting/);
```

Твој параметър за `fighting` не сме да има вредност.

```js
assert.isUndefined(fighting);
```

Ne treba da postaviš vrednost za svoj parametar `fighting`.

```js
assert.match(code, /let\s+fighting\s*;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
--fcc-editable-region--
```
