---
id: 62a3a7e4f1060e2fc5ffb34b
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Definišite još jedan parametar koji se zove `currentWeaponIndex` i sa postavi njegovu vrednost na `0`.

# --hints--

Trebalo bi da koristite `let` za deklarisanje parametra pod nazivom `currentWeaponIndex`.

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

Trebalo bi da koristiš camelCase za ime svog parametra.

```js
assert.match(code, /currentWeaponIndex/);
```

Vaš parametar za `currentWeaponIndex` treba biti postavljen na `0`.

```js
assert.equal(currentWeaponIndex, 0);
```

Trebalo bi da inicijalizujete svoj parametar na `0`.

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
