---
id: 62a7bf06d2ad9d1c5024e833
title: Korak 40
challengeType: 0
dashedName: step-40
---

# --description--

Takođe ćete morati da ispravite akcije koje se izvršavaju kada se dugme ponovo klikne.

U vašoj funkciji za `goStore()`, ažurirajte svojstvo `onclick` za svako dugme kako biste pokrenuli `buyHealth`, `buyWeapon` i `goTown`, odgovarajuće.

# --hints--

Treba da koristiš tačkasto pisanje da bi dobio atribut `onclick` od `button1`.

```js
assert.match(goStore.toString(), /button1\.onclick/);
```

Ne smete koristiti `let` ili `const` da biste dobili status `onclick` od `button1`.

```js
assert.notMatch(code, /(let|const)\s+button1\.onclick/);
```

Trebalo bi da postaviš atribut `onclick` na `button1` kao `buyHealth`.

```js
assert.match(code, /button1\.onclick\s*=\s*buyHealth/);
```

Trebalo bi da postaviš atribut `onclick` na `button1` u svojoj funkciji `goStore`.

```js
assert.match(goStore.toString(), /button1\.onclick\s*=\s*buyHealth/);
```

Moraš da koristiš notaciju tačke da dobiješ svojstvo za `onclick` od `button2`.

```js
assert.match(goStore.toString(), /button2\.onclick/);
```

Ne smiješ da koristiš `let` ili `const` da bi dobio/la titulu `onclick` od `button2`.

```js
assert.notMatch(code, /(let|const)\s+button2\.onclick/);
```

Morate postaviti atribut `onclick` za `button2` na `buyWeapon`.

```js
assert.match(code, /button2\.onclick\s*=\s*buyWeapon/);
```

Trebalo bi da postaviš atribut `onclick` na `button2` u svojoj funkciji `goStore`.

```js
assert.match(goStore.toString(), /button2\.onclick\s*=\s*buyWeapon/);
```

Trebalo bi da koristite dot notaciju da biste dobili svojstvo za `onclick` od `button3`.

```js
assert.match(goStore.toString(), /button3\.onclick/);
```

Ne smeš koristiti `let` ili `const` da bi dobio/la status `onclick` od `button3`.

```js
assert.notMatch(code, /(let|const)\s+button3\.onclick/);
```

Trebalo bi postaviti svojstvo `onclick` sa `button3` na `goTown`.

```js
assert.match(code, /button3\.onclick\s*=\s*goTown/);
```

Trebalo bi da postaviš atribut `onclick` na `button3` u svojoj funkciji `goStore`.

```js
assert.match(goStore.toString(), /button3\.onclick\s*=\s*goTown/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

--fcc-editable-region--
function goStore() {
  button1.innerText = "Buy 10 health (10 gold)";
  button2.innerText = "Buy weapon (30 gold)";
  button3.innerText = "Go to town square";
}
--fcc-editable-region--

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
