---
id: 62a8d81f539f004776dd9b1e
title: Korak 110
challengeType: 0
dashedName: step-110
---

# --description--

Sada možete započeti kod za borbu protiv monstrima. Da bi vaš kod bio dobro organizovan, vaša funkcija za `fightDragon` je premještena da bude blizu drugih funkcija od `fight`.

Ispod vašeg niza podataka `weapons`, deklariš varijablu `monsters` i postavi njenu vrednost kao niz podataka. Postavi taj niz da ima tri elementa, svaki sa svojstvima `name`, `level` i `health`. Vrednosti prvog elementa treba da budu `"slime"`, `2` i `15`, u tom redosledu. Drugi element treba da ima `"fanged beast"`, `8` i `60`. Treći element treba da ima `"dragon"`, `20` i `300`.

# --hints--

Trebalo bi da koristiš `const` za deklarisanje varijable `monsters`.

```js
assert.match(code, /const\s+monsters/);
```

Varijabla `monsters` treba da bude red podataka.

```js
assert.isArray(monsters);
```

Vaša promenljiva `monsters` mora imati vrednost 3.

```js
assert.lengthOf(monsters, 3);
```

Red vaših podataka `monsters` trebalo bi da ima 3 stavke.

```js
assert(monsters.every(val => typeof val === "object"));
```

Prvi element u vašem nizu `monsters` treba da bude nešto sa karakteristikom `name` postavljeno na `"slime"`.

```js
assert.equal(monsters[0].name, "slime");
```

Prvi element u vašoj sekvenci `monsters` treba da bude nešto sa karakteristikom `level` postavljeno na `2`.

```js
assert.equal(monsters[0].level, 2);
```

Prvi element u vašem nizu `monsters` treba da bude nešto sa karakteristikom `health` postavljeno na `15`.

```js
assert.equal(monsters[0].health, 15);
```

Vrednost u drugom redu vašeg `monsters` treba da bude element sa karakteristikom `name`, postavljen na `"fanged beast"`.

```js
assert.equal(monsters[1].name, "fanged beast");
```

Vrednost polja ``monsters`` u vašem nizu treba da bude nešto sa karakteristikom ``level`` postavljeno na ``8``.

```js
assert.equal(monsters[1].level, 8);
```

Vrednost drugog elementa u vašem nizu `monsters` mora biti nešto sa karakteristikom `health` postavljeno na `60`.

```js
assert.equal(monsters[1].health, 60);
```

Treća vrednost u vašem redu od `monsters` treba da ima svojstvo `name` postavljeno na `"dragon"`.

```js
assert.equal(monsters[2].name, "dragon");
```

Treća vrednost u vašem redu od `monsters` treba da bude nešto sa karakteristikom `level` postavljeno na `20`.

```js
assert.equal(monsters[2].level, 20);
```

Treća vrednost u vašem redu od `monsters` treba da bude element sa svojstvom `health`, postavljen na `300`.

```js
assert.equal(monsters[2].health, 300);
```

Ne smete dodavati nikakve dodatne atribute svojim stavkama.

```js
assert.deepEqual(monsters, [
  {name: "slime", level: 2, health: 15},
  {name: "fanged beast", level: 8, health: 60},
  {name: "dragon", level: 20, health: 300}
])
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");
const weapons = [
  { name: 'stick', power: 5 },
  { name: 'dagger', power: 30 },
  { name: 'claw hammer', power: 50 },
  { name: 'sword', power: 100 }
];
--fcc-editable-region--

--fcc-editable-region--
const locations = [
  {
    name: "town square",
    "button text": ["Go to store", "Go to cave", "Fight dragon"],
    "button functions": [goStore, goCave, fightDragon],
    text: "You are in the town square. You see a sign that says \"Store\"."
  },
  {
    name: "store",
    "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
    "button functions": [buyHealth, buyWeapon, goTown],
    text: "You enter the store."
  },
  {
    name: "cave",
    "button text": ["Fight slime", "Fight fanged beast", "Go to town square"],
    "button functions": [fightSlime, fightBeast, goTown],
    text: "You enter the cave. You see some monsters."
  }
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {
  button1.innerText = location["button text"][0];
  button2.innerText = location["button text"][1];
  button3.innerText = location["button text"][2];
  button1.onclick = location["button functions"][0];
  button2.onclick = location["button functions"][1];
  button3.onclick = location["button functions"][2];
  text.innerText = location.text;
}

function goTown() {
  update(locations[0]);
}

function goStore() {
  update(locations[1]);
}

function goCave() {
  update(locations[2]);
}

function buyHealth() {
  if (gold >= 10) {
    gold -= 10;
    health += 10;
    goldText.innerText = gold;
    healthText.innerText = health;
  } else {
    text.innerText = "You do not have enough gold to buy health.";
  }
}

function buyWeapon() {
  if (currentWeaponIndex < weapons.length - 1) {
    if (gold >= 30) {
      gold -= 30;
      currentWeaponIndex++;
      goldText.innerText = gold;
      let newWeapon = weapons[currentWeaponIndex].name;
      text.innerText = "You now have a " + newWeapon + ".";
      inventory.push(newWeapon);
      text.innerText += " In your inventory you have: " + inventory;
    } else {
      text.innerText = "You do not have enough gold to buy a weapon.";
    }
  } else {
    text.innerText = "You already have the most powerful weapon!";
    button2.innerText = "Sell weapon for 15 gold";
    button2.onclick = sellWeapon;
  }
}

function sellWeapon() {
  if (inventory.length > 1) {
    gold += 15;
    goldText.innerText = gold;
    let currentWeapon = inventory.shift();
    text.innerText = "You sold a " + currentWeapon + ".";
    text.innerText += " In your inventory you have: " + inventory;
  } else {
    text.innerText = "Don't sell your only weapon!";
  }
}

function fightSlime() {

}

function fightBeast() {

}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
