---
id: 62b46e3a8d4be31be5af793d
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Sastavili smo se sa malim problemom. Pokušavate da izvršite pretragu elementa dugmeta na svojoj stranici, ali vaša oznaka `script` je unutar `head` vašeg HTML. To znači da se vaš kod izvršava pre nego što pregledač završi čitanje HTML, i vaša `document.querySelector()` neće pronaći to dugme – jer pregledač još uvek nije završio sa obradom.

Da biste ovo popravili, premestite vaš element `script` iznutra elementa `head`, i postavite ga na kraj vašeg elementa `body` (odmah pre zatvarajuće oznake `</body>`).

# --hints--

Твой елемент `script` не сме бити унутар твог елемента `head`.

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

Vaš element `script` mora da dođe na kraju vašeg elementa `body`.

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
