---
id: 62fc20387ef88d1d1998aac5
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Koristite `querySelector()` da dobijete druga dva elementa od `button` koristeći njihove `id`: `button2` na `button3`. Sačuvajte u parametre nazvane `button2` i `button3`. Zapamtite da koristite `const`.

# --hints--

Trebalo bi da objaviš parametar `button2` i `const`.

```js
assert.match(code, /const\s+button2/);
```

Vaš parametar za `button2` mora imati vrednost vašeg atributa za `"#button2"`.

```js
assert.match(code, /const\s+button2\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#button2\1\s*\)/);
```

Trebalo bi da objavite parametar `button3` sa `const`.

```js
assert.match(code, /const\s+button3/);
```

Parametar vašeg `button3` mora imati vrednost polja vašeg `"#button3"`.

```js
assert.match(code, /const\s+button3\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#button3\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");

--fcc-editable-region--
```
