---
id: 69aaafeaa4037197c75587cd
title: Tekeleza foleni
challengeType: 26
dashedName: implement-a-queue
---

# --description--

U ovoj laboratoriji, ćete implementirati strukturu podataka reda koristeći funkcije. Red je struktura podataka tipa Prvi-ulaz, prvi-izlaz (FIFO), gde se elementi dodaju na zadnji kraj reda i uklanjaju sa prednjeg kraja.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Mora da ima funkciju `enqueue` koja dodaje element na zadnji deo redoslika.
1. Mora da ima funkciju `dequeue` koja uklanja i vraća prednji element redoslika.
1. Mora da ima funkciju `front` koja vraća prednji element redoslika bez uklanjanja.
1. Mora da ima funkciju `size` koja vraća broj elemenata u redosliku.
1. Mora da ima funkciju `isEmpty` koja vraća `true` ako je redoslik prazan i `false` u suprotnom slučaju.

**Napomena:** Mnoge funkcije zavise od funkcionalnosti `enqueue`. Implementirajte to prvo, jer su funkcije za `dequeue`, `front`, `size` i `isEmpty` potrebne za dodavanje komponenti u red (queue).

# --hints--

Trebalo bi da imaš funkciju `enqueue`.

```js
assert.isFunction(enqueue);
```

Funkcija `enqueue` bi trebalo da doda karakteristiku iza rednog sistema/kolije.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

Trebalo bi da imate funkciju `dequeue`.

```js
assert.isFunction(dequeue);
```

Funkcija `dequeue` bi trebalo da ukloni i vrati prednji element reda.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

`dequeue` trebalo bi da vrati `undefined` kada se pozove na praznoj listi.

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue` Trebalo bi da nastavi dodavanje elemenata iza reda nakon uklanjanja elemenata.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

Različiti profili ne smeju deliti unutrašnji status.

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

Trebalo bi da imaš funkciju `front`.

```js
assert.isFunction(front);
```

Funkcija `front` treba da vrati prednji element redova bez uklanjanja.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

`front` treba vratiti `undefined` kada je pozvan na prazan spisak.

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

Trebalo bi da imaš funkciju `size`.

```js
assert.isFunction(size);
```

Funkcija `size` treba da vrati tačan broj elemenata u redu.

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

Trebalo bi da imaš funkciju `isEmpty`.

```js
assert.isFunction(isEmpty);
```

Funkcija za `isEmpty` treba da vrati `true` u praznu mapu.

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

Funkcija `isEmpty` treba da vrati `false` u neku ne-praznu mapu.

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

Funkcija za `isEmpty` treba da vrati `true` nakon što su svi elementi uklonjeni iz redova.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
