---
id: 587d78af367417b2b2512b03
title: Izradi obrazac za anketu
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301145
dashedName: build-a-survey-form
---

# --description--

**Cilj:** Napravite aplikaciju koja funkcioniše kao <a href="https://survey-form.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://survey-form.freecodecamp.rocks</a>. **Ne kopirajte ovaj demo projekat**.

**Priče korisnika:**

1. Treba da imate naslov stranice u elementu `h1` sa `id`-om `title`.
1. Treba da imate kratak opis u elementu `p` sa `id`-om `description`.
1. Treba da postoji element `form` sa `id`-om `survey-form`.
1. Unutar forme, morate uneti vaše ime u polje za unos `input` sa `id`-om `name` i `type`-om `text`.
1. Unutar forme, morate uneti vaš e-mail u polje za unos `input` sa `id`-om `email`.
1. Ako unesete nevažeći e-mail, videćete HTML5 validacioni deo.
1. Unutar forme, možete uneti broj u polje za unos `input` sa `id`-om `number`.
1. Polje za unos broja ne bi trebalo da dozvoljava unos karaktera koji nisu brojevi, bilo tako što sprečava kucanje ili prikazuje HTML5 validacioni deo (u zavisnosti od vašeg pregledača).
1. Ako unesete brojeve van opsega polja za unos broja, definisanog atributima `min` i `max`, videćete HTML5 validacioni deo.
1. Za polja za ime, e-mail i broj, možete videti odgovarajuće elemente `<label>` u formi, koji opisuju svrhu svakog polja sa sledećim ID-ovima: `id="name-label"`, `id="email-label"` i `id="number-label"`.
1. Za polja za ime, e-mail i broj, možete videti placeholder tekst koji daje opis ili uputstvo za svako polje.
1. Unutar elementa forme, treba da postoji `<select>` padajući meni sa `id`-om `dropdown` i sa najmanje dve opcije za izbor.
1. Unutar elementa forme, možete izabrati opciju iz grupe radio dugmadi koristeći atribut `name`.
1. Unutar elementa forme, možete odabrati više stavki iz niza polja za čekiranje, svako sa atributom `value`.
1. Unutar elementa forme, je prisutan polje za dodatne komentare `<textarea>`.
1. Unutar elementa forme, postoji dugme sa `id`-om `submit` za slanje svih unosa.

Popunite sve priče korisnika i prođite kroz sve testove ispod kako biste završili ovaj projekat. Dodajte svoj stil. Uživajte u pisanju koda!

**Napomena:** Obavezno dodajte `<link rel="stylesheet" href="styles.css">` u vaš HTML da povežete vaš CSS fajl.

# --hints--

Treba da imate element `h1` sa `id`-om `title`.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

Vaš `#title` ne sme biti prazan.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Treba da imate element `p` sa `id`-om `description`.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

Vaš `#description` ne sme biti prazan.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Treba da imate element `form` sa `id`-om `survey-form`.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

Treba da imate element `input` sa `id`-om `name`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Vaš `#name` treba da ima `type` postavljeno na `text`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

Vaš `#name` mora zahtevati unos.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Vaš `#name` treba da bude potomak `#survey-form`.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

Treba da imate element `input` sa `id`-om `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Vaš `#email` treba da ima `type` postavljeno na `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

Vaš `#email` mora zahtevati unos.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Vaš `#email` treba da biti potomak `#survey-form`.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

Treba da imate element `input` sa `id`-om `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Vaš `#number` treba da biti potomak `#survey-form`.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

Vaš `#number` treba da imati `type` postavljeno na `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

Vaš `#number` treba da ima atribut `min` sa brojevima.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

Vaš `#number` treba da imati atribut `max` sa brojevima.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

Treba da imate element `label` sa `id`-om `name-label`.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Treba da imate element `label` sa `id`-om `email-label`.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Treba da imate element `label` sa `id`-om `number-label`.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Vaš `#name-label` treba da sadrži tekst koji opisuje unos.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Vaš `#email-label` treba da sadrži tekst koji opisuje unos.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Vaš `#number-label` treba da sadrži tekst koji opisuje unos.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Vaš `#name-label` treba da biti potomak `#survey-form`.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

Vaš `#email-label` treba da biti potomak `#survey-form`.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

Vaš `#number-label` treba da biti potomak `#survey-form`.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

Vaš `#name` treba da imati atribut `placeholder` sa vrednošću.

```js
const el = document.getElementById