---
id: 660ef19b95d3308e7dd31bb6
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

<dfn>Konsole</dfn> vam omogućava da štampate i vidite rezultate od JavaScript. Možete poslati informacije u konzolu koristeći `console.log()`. Na primer, ovaj kod će ispisati `"Naomi"` u konzolu:

```js
let developer = "Naomi";
console.log(developer);
```

Kod iznad dobija pristup promenljivoj `developer` po njenom imenu u `console.log()`. Zapamtite da je vrednost koja se nalazi između običnih zagrada ona koja će biti štampana.

Ispišite vrednost varijable `character` u konzoli. Zatim, kliknite na dugme "Console" da biste videli konzolu JavaScript.

# --hints--

Inicijalizacija varijable `character` ne sme se menjati.

```js
assert.strictEqual(character, 'Hello');
```

Trebalo bi da dobiješ pristup `console` u svom kodu.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Trebalo bi da dobiješ kvalifikaciju `log` od `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Trebalo bi da koristite obične zagrade za pozivanje puta `.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Trebalo bi da odštampate varijablu `character` u konzolu.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
