---
id: 660ef5105b8ba095307a0e50
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Sada ponovite unos vrednosti promena `character` na konzoli. Trebalo bi da vidite sekvencu slova `"Hello"`, pa sekvencu slova `"World"`, na konzoli.

# --hints--

Trebalo bi da koristiš `console.log` drugi put.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Treba da unesete vrednost `character` u konzolu drugi put.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```

`console.log()` Nova vrednost bi trebalo da dođe nakon resetovanja.

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
