---
id: 660ef55dd468079679ee0092
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

Sada ste spremni da objavite svoj drugi parametar. Uklonite oba izjave od `console.log`, i opseg vrednosti za `character`.

Takođe uklonite vaš parametar za `secondCharacter`, ali ostavite aktivaciju za `character` nepromenjenom.

# --hints--

Nemojte imati nikakvu izjave od `console.log()` u vašem kodu.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Ne smete da postavite ponovnu vrednost za `character`.

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

Tvoj parametar za `character` mora imati vrednost `"Hello"`.

```js
assert.equal(character, "Hello");
```

Ne bi trebalo da ima parametar `secondCharacter`.

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
