---
id: 660ef6355e8f5a9e67fe5f46
title: Korak 16
challengeType: 1
dashedName: step-16
---

# --description--

Za tip podataka `number`, možete izvršiti aritmetičke operacije, kao što je sabiranje. Pokušajte da prikažete `count + 1` na konzoli.

# --hints--

Trebalo bi da dobiješ dostupnost za `console` u svom kodu.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Trebalo bi da dobiješ pristup `.log()` za `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Trebaće ti da napišeš `count + 1` u konzolu.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
