---
id: 660ee6e3a242da6bd579de69
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Jedan od najvažnijih koncepata u programiranju su parametri. <dfn>Parametar</dfn> ukazuje na posebnu memorijsku adresu koja čuva vrednost. Parametri dobijaju ime koje se može koristiti u celom vašem kodu da bi se pristupila toj vrednosti.

Deklarisanje parametra znači mu dati ime. U JavaScript, ovo se često radi korišćenjem reči `let`. Na primer, ovde je kako deklarisati parametar `hello`:

```js
let hello;
```

Ime parametra prati posebna pravila: imena mogu sadržavati slova, brojeve, znakove dolara i znakove crte ispod, ali ne smeju imati razmake i ne bi smela počinjati sa brojem.

Koristi reč `let` da deklariraš parametar koji se zove `character`.

_Napomena_: Uobičajena praksa je završiti izjavu u JavaScript sa znakom polutačke tačke. `;`

# --hints--

Trebalo bi da koristiš `let` u svom kodu.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Trebalo bi da koristiš `character` u svom kodu.

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

Trebalo bi da koristiš `let` za deklarisanje parametra `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Vaš post bi trebalo da se završi tačkom.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
