---
id: 660ef31a5be625914a0102cd
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Kada deklarišete parametar za glavnu reč `let`, možete <dfn>ponovo postaviti vrednost</dfn> (ili promeniti vrednost) tog parametra kasnije. U ovom primeru, vrednost `programmer` se menja iz `"Naomi"` u `"CamperChan"`.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Zapamtite da kada ponovo postavite vrednost parametru, ponovo ćete prikazati ključnu reč `let`.

Nakon vašeg `console.log`, postavite vrednost `"World"` za vaš parametar `character`.

# --hints--

Moraš koristiti `character` dva puta u svoj kod.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

Trebalo bi da koristiš operater dodeljivanja za ponovno postavljanje vrednosti za `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

Trebalo bi da koristiš niz znakova `"World"` u svom kodu.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

Vaš parametar za `character` mora imati vrednost `"World"`.

```js
assert.equal(character, "World");
```

Postavljanje vrednosti više ne bi trebalo da koristi `let`.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
