---
id: 660eebd83100d37862268781
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Vaš parametar `character` trenutno nema vrednost. Možete postaviti vrednost koristeći operator <dfn>vrednosti</dfn> `=`. Primer:

```js
let hello = "Hello";
```

Postavljanje vrednosti parametra tokom njegovog deklarisanja naziva se <dfn>inicijalizacija</dfn>.

Inicijalizujte svoj parametar za `character` da prosledi vrednost `"Hello"` prilikom deklarisanja.

# --hints--

Trebalo bi da koristiš operater za raspodelu vrednosti `=`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

Trebalo bi da koristiš niz znakova `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

Trebalo bi da koristiš dvostruke navodne znakove za svoju sekvencu slova `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

Trebalo bi da staviš `"Hello"` za svoj parametar `character`.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
