---
id: 660ef857f2806aa626d29d17
title: Korak 17
challengeType: 1
dashedName: step-17
---

# --description--

Možete takođe izvršiti oduzimanje (`-`), množenje (`*`) i deljenje (`/`). Molimo vas da pokušate da koristite operacije i brojeve u vašem `console.log`. Kada budete spremni da nastavite, uklonite `console.log`.

# --hints--

Ne bi trebalo da imaš `console.log()` u svoj kod.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Vaš parametar za `character` treba da ima vrednost od `"Hello"`.

```js
assert.equal(character, "Hello");
```

Vaš parametar za `count` bi trebalo da ima vrednost od `8`.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
