---
id: 660ef0f7c4b8e68ccd1f0786
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

JavaScript ima sedam primitivnih tipova podataka, od kojih je jedan niska karaktera. U JavaScript-u, <dfn>niska karaktera</dfn> predstavlja niz znakova i može biti okružena jednostrukim (`'`) ili dvostrukim (`"`) navodnicima.

Zapamtite da je niz karaktera <dfn>nepromenljiv</dfn>, što znači da, jednom kada se kreira, ne može se promeniti. Međutim, parametar i dalje može dobiti drugu vrednost.

Promeni svoju sekvencu od `"Hello"` koristeći pojedinačne navodne znakove.

# --hints--

Ne bi trebalo da imate dvostruke navodne oznake u svom kodu.

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

Trebalo bi da koristiš pojedinačne navodne znakove u svojoj sekvenci `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

Još moraš da koristiš `let` u svom kodu.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Još treba da objavi parametar `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Vaš atribut za ``character`` treba da ima sekvencu karaktera ``"Hello"`` kao svoju vrednost.

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
