---
id: 660ef5c1904955978a986a5c
title: Korak 15
challengeType: 1
dashedName: step-15
---

# --description--

Koristi `let` da objaviš varijablu `count`. Postavi vrednost <dfn>broj</dfn> `8`. Kada koristiš numeričku vrednost, ne koristi navodnike. Primer:

```js
let money = 100;
```

# --hints--

Trebalo bi da imaš ključnu reč `let` drugog u tvom kodu.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Trebalo bi da koristite `let` za deklarisanje varijable `count`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

Treba da postaviš vrednost broja `8` za svoju varijablu `count`.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
