---
id: 660f3ba3cceef11b6ba08b59
title: Korak 75
challengeType: 1
dashedName: step-75
---

# --description--

Nažalost, sada je izgubljeno ispod piramide. Ovo je zato što si ponovo napravio <dfn>jednu grešku</dfn>.

Vaš početni ciklus je bio sa vrednostima `i` od `0` do `7`, jer `count` je `8` i vaš uslov zahteva da `i` bude ispod `count`. Vaš ciklus se sada izvršava sa vrednostima `i` od `1` do `7`.

Vaš krug mora biti popravljen da bi radio kada je `i` jednako `8`, takođe. Ako pogledate svoju logiku, to znači da vaš krug treba da radi kada `i` je <dfn>manje ili jednako</dfn> `count`. Možete koristiti operator manji ili jednak od `<=` za ovo.

Ispravite uslov vašeg ciklusa kako bi radio kada je `i` manji ili jednak `count`.

# --hints--

Tvoj ciklus za `for` ne bi trebalo da proverava da li je `i` ispod `count`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Tvoj ciklus od `for` treba da proveri da li `i` je manji ili jednak `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
