---
id: 660f1cedf3676fe26122ebf6
title: Korak 38
challengeType: 1
dashedName: step-38
---

# --description--

Bi trebalo da vidite brojeve od nule do sedmice štampane na vašoj konzoli, jedan po redu. Ovo će se koristiti kao osnova za kreiranje vaše piramide.

Promeni svoju poruku (log) i poruku za pushovanje `i` u red podataka `rows`.

# --hints--

Ne bi trebalo da ima poziv ka `console.log`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Trebalo bi da pozoveš `.push()` na red podataka `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Treba da guraš `i` u red podataka `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

`.push()` bi trebalo da se pojavi u vašem ciklusu od `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
