---
id: 660f09a2694b59c3a10ee304
title: Korak 23
challengeType: 1
dashedName: step-23
---

# --description--

Trenutno, uklonite vaš prvi konzolni log i vaš prilog za `rows[rows.length - 1]`. Ostavite drugi log zapis/izjavu za `rows` za kasnije.

# --hints--

Treba da uklonite vašu izjavu za `console.log(rows[0])`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

Trebalo bi da uklonite svoj deo `rows[rows.length - 1]`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

Ne smete da uklonite jednostavnu deklaraciju `console.log(rows)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
