---
id: 660f039ff313dbb696b007ca
title: Korak 19
challengeType: 1
dashedName: step-19
---

# --description--

Kada red podataka ima vrednosti, ili <dfn>elementi</dfn>, te vrednosti su razdvojene zarezom. Evo reda podataka koji sadrži sekvencu od dva karaktera:

```js
let array = ["first", "second"];
```

Promeni svoju deklaraciju `rows` u niz podataka sa sekvencama karaktera `"Naomi"`, `"Quincy"`, na `"CamperChan"`. Redosled vrednosti u nizu podataka je važan, pa ga prate. Zapamtite da sekvenca karaktera pazi na velika i mala slova.

# --hints--

Prvi element u vašoj liniji podataka treba da bude sekvenca karaktera `"Naomi"`.

```js
assert.equal(rows[0], "Naomi");
```

Drugi element u vašem nizu podataka treba da bude sekvenca karaktera `"Quincy"`.

```js
assert.equal(rows[1], "Quincy");
```

Treće polje u vašem nizu podataka treba da bude niz karaktera `"CamperChan"`.

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
