---
id: 660f1b6e60bd9edf902c81fd
title: Korak 36
challengeType: 1
dashedName: step-36
---

# --description--

Izjava koja <dfn>iteration</dfn> će vam reći šta vaš petlja treba da uradi sa zvučnikom nakon svake petlje.

Kada postavite parametru novu vrednost, možete koristiti taj parametar da se vratite na originalnu vrednost pre resetovanja. To vam omogućava da uradite stvari kao što je dodavanje tri postojećem broju. Na primer, `bees = bees + 3` će dodati vrednost od `bees` sa tri.

Koristi tu sintaksu da promeniš niz karaktera `"iteration"` u deklaraciju za ponovno postavljanje vrednosti koja dodaje `i` sa jediničnim korakom.

# --hints--

Trebalo bi da dodate jedan na vaš atribut `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

Treba da ponovo postaviš vrednost `i + 1` za svoj parametar `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Tvoj ciklus od `for` treba da doda `i` sa `1` tokom svake iteracije.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
