---
id: 660f53ad3d39175c5d4335ac
title: Korak 117
challengeType: 1
dashedName: step-117
---

# --description--

Tvoj generator piramide je sada skoro završen, i ima više funkcionalnosti nego što si prvobitno planirao! Sledeći korak je refaktoring tvog koda.

Uklonite sve komentare, i one sa jednim redom i one sa više redova, iz vašeg koda.

# --hints--

Ne bi trebalo da imaš nikakve komentare na jednoj liniji u svom kodu.

```js
assert.notMatch(code, /\/\//);
```

Nemoj imati nikakve višelinijske komentare u svom kodu.

```js
assert.notMatch(code, /(?:\*\/|\/\*)/);
```

Nemoj imati nikakve komentare u svom kodu.

```js
assert.equal(code, __helpers.removeJSComments(code));
```

Trebali biste da uklonite kod koji je bio komentarisano korišćenjem višelinijskih komentara.

```js
assert.notMatch(code, /while/);
assert.lengthOf(code.match(/for\s*\(/g), 2)
assert.lengthOf(code.match(/rows\.push/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
