---
id: 660f2b6fd54ac1fc142804dd
title: Korak 47
challengeType: 1
dashedName: step-47
---

# --description--

Logika za kreiranje ove piramidalne strukture može biti komplikovana, što znači da je dobro vreme da se taj kod izvuče i stavi u funkciju.

<dfn>Funkcija</dfn> je deo koda koji se može ponovno koristiti u celoj vašoj aplikaciji. Funkcije se deklarišu sledećom sintaksom:

```js
function name(parameter) {

}
```

Glavna reč za `function` govori JavaScript da će parametar `name` biti funkcionalan. `parameter` je parametar koji predstavlja vrednost koja se prosleđuje unutar funkcije kada se koristi. Funkcija može imati mnogo ili malo parametara, kako god želite. Kao u opsegu/kontekstu `for`, prostor između okruglih zagrada je <dfn>telo funkcije</dfn>.

Deklariši funkciju za `padRow`. Nemoj kreirati nikakve parametre još. Telo funkcije mora biti prazno. Zapamti da moraš koristiti Camel konvenciju case za svoje imena.

# --hints--

Trebalo bi da koristiš glavnu reč `function`.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

Trebalo bi objaviti funkciju `padRow`.

```js
assert.isFunction(padRow);
```

Tvoja funkcija za `padRow()` ne bi trebalo da ima nikakvih parametara.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

Tvoja funkcija za `padRow()` mora imati prazno telo.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
