---
id: 660f4a83373de83ea101685f
title: Korak 100
challengeType: 1
dashedName: step-100
---

# --description--

Zato što ćete pokretati kolo u suprotnom smeru, vaše kolo mora da nastavi kada `i` je iznad `0`. Možete koristiti parametar <dfn>iznad </dfn> `>` za ovo.

Postavite uslov da vaš krug se izvrši kada `i` je veći od `0`.

# --hints--

Tvoj ciklus od `for` treba da se pokrene kada `i` je veći od `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; false; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
