---
id: 660f5179b3b0ca558f6b4d4f
title: Korak 113
challengeType: 1
dashedName: step-113
---

# --description--

Koristi deklaraciju `if` da proveriš da li je `inverted` tačno. Zapamti da ovde nema potrebe za korišćenjem operatera jednakosti.

# --hints--

Trebalo bi da koristiš deklaraciju `if`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(/);
```

Vaša izjava za `if` treba da proveri da li `inverted` je `true`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted/);
```

Zahtev je da `if` ne koristite nikakve operatore za poređenje.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {

  rows.unshift(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
