---
id: 660f39b444fd6f16d1e49c1f
title: Korak 73
challengeType: 1
dashedName: step-73
---

# --description--

Umesto da je potrebno prosleđivati ``i + 1`` za vaš poziv ``padRow``, možete umesto toga započeti svoj ciklus na ``1``. Ovo će vam omogućiti da kreirate ciklus koji počinje sa jednim modelom.

Ispravite svoj ciklus da počnete sa `1` umesto `0`.

# --hints--

Tvoj ciklus za `for` treba da inicijalizuje `i` za `1`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
