---
id: 660f4de78f775e480ba2e451
title: Korak 107
challengeType: 1
dashedName: step-107
---

# --description--

Metoda `.unshift()` za red podataka omogućava vam da dodate vrednost na početak reda podataka, za razliku od `.push()` koja dodaje vrednost na kraj reda podataka. `.unshift()` vraća novu dužinu korišćenog reda podataka.

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

Koristite `const` za deklarisanje varijable `unshifted`, i postavite njenu vrednost na rezultat poziva `.unshift()` u vašem redu podataka `numbers`. Prosledite `5` kao argument. Zatim odštampajte varijablu `unshifted`.

# --hints--

Treba da deklariš svoju varijablu za `unshifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

Trebalo bi da pozoveš putanju `.unshift()` u svoj red podataka `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

Trebalo bi da prosledite `5` kao argument vašem pozivu `.unshift()`.

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

Trebalo bi da staviš rezultat svog poziva na `.unshift()` u svoju varijablu `unshifted`.

```js
assert.equal(unshifted, 4);
```

Treba da ispišete vašu varijablu `unshifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
