---
id: 660f4cffb1459d45e34902d1
title: Korak 106
challengeType: 1
dashedName: step-106
---

# --description--

Možeš da izgradiš preokrenutu piramidu bez potrebe da zaokružiš "nazad" kao što si to uradio.

Da biste ovo uradili, moraćete da naučite nekoliko novih metoda za niz podataka. Počnite korišćenjem `const` da deklarišete varijablu `numbers`. Postavite joj vrednost da bude niz podataka sa elementima `1`, `2` i `3`. Zatim upišite niz podataka `numbers` u memoriju.

# --hints--

Trebalo bi da koristiš `const` za deklarisanje varijable `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

Vaša promenljiva `numbers` mora da bude red podataka.

```js
assert.isArray(numbers);
```

Redak vaših podataka za `numbers` mora imati elemente `1`, `2`, i `3` u istom redosledu.

```js
assert.deepEqual(numbers, [1,2,3]);
```

Trebaće ti da napišeš svoj red podataka `numbers` u memoriju.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
