---
id: 660f229d2dbe09ef2954a4a1
title: Korak 42
challengeType: 1
dashedName: step-42
---

# --description--

Zapamtite da ste u svom početnom ciklusu koristili operater za sabiranje `+` da biste dodali vrednost `i` za `1`.

Možete uraditi nešto slično sa vrednošću niza karaktera, spajanjem novog niza na postojeći niz. Na primer, `hello = hello + " World";` će dodati niz karaktera `" World"` na niz koji je sačuvan u varijabli `hello`. Ovo se zove <dfn>spajanje</dfn>.

U vašem ciklusu od `for...of`, koristite operator za sabiranje da spojite vrednost `row` sa vrednošću `result`.

# --hints--

Moraš koristiti operater za spajanje na tvojoj varijabli `result`.

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

Trebalo bi da povežete `row` u vašu varijablu `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

Treba da postaviš vrednost rezultata spajanja u varijablu `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
