---
id: 660f1a00ac619ddc1e259a66
title: Korak 35
challengeType: 1
dashedName: step-35
---

# --description--

<dfn>Uslov</dfn> za ciklus `for` objašnjava koliko puta bi se ciklus trebao ponoviti. Ciklus će prestati kada `condition` nije tačno.

U JavaScript, vrednost istinite ili lažne vrednosti može biti `true` ili `false`. Ovo nije niz slova - saznaćete više o ovoj razlici kasnije.

Trenutno, koristićete <dfn>operator manje od</dfn> (`<`). Ovo vam omogućava da proverite da li je vrednost na levoj strani manja od vrednosti na desnoj strani. Na primer, `count < 3` će biti izračunato kao `true` ako je `count` jednako `2`, i `false` ako je `count` jednako `4`.

Promeni svoj niz `"condition"` u uslov za proveru da li je `i` manji od `count`.

# --hints--

Trebalo bi da koristite donju operu.

```js
assert.match(__helpers.removeJSComments(code), /</);
```

Treba koristiti operater manjeg nego ako `i` je manje od `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Tvoj ciklus za `for` bi trebalo da koristi `i < count` kao uslov.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
