---
id: 660f4e74f7fd3f4a99ac2e50
title: Korak 108
challengeType: 1
dashedName: step-108
---

# --description--

Nizovi podataka takođe imaju metodu `.shift()`. Ovo će vratiti prvi element niza podataka, za razliku od `.pop()` koji vraća poslednji element. Evo primera metode `.shift()`:

```js
const numbers = [1, 2, 3];
numbers.shift();
```

Podatni red za `numbers` će biti `[2, 3]`.

Direktno ispod vašeg reda podataka `numbers`, deklariši parametar `shifted` i postavi mu vrednost na rezultat poziva `.shift()` u redu podataka `numbers`. Na sledećoj liniji, napišite parametar `shifted` u konzolu.

# --hints--

Trebalo bi da koristiš `const` za deklarisanje parametra `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

Trebalo bi da pozoveš `.shift()` u svoj red podataka `numbers`.

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

Treba da postavite vrednost rezultata vašeg poziva ``.shift()`` u svoj parametar ``shifted``.

```js
assert.equal(shifted, 1);
```

Treba da napišeš svoj parametar `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
