---
id: 660f34626216270c682e2f7b
title: Korak 66
challengeType: 1
dashedName: step-66
---

# --description--

<Poziv funkcije> Nije ti dozvoljeno da koristiš funkciju direktno. Možda to ne znaš, ali metode poput `.push()` predstavljaju pozive funkcija.

Funkcija se poziva referenciranjem imena funkcije, i dodavanjem `()`. Ovako treba da pozovete funkciju za `test`:

```js
test();
```

Zameni `character.repeat(i + 1)` u vašem pozivu `.push()` i funkcioniskim pozivom za svoju funkciju `padRow`. Ne dodajte nikakve argumente/parametre za sada.

# --hints--

Koristite `i + 1` u vašem pozivu `push`.

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

Trebalo bi da koristite `character.repeat` u vašem pozivu `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

Trebaće ti da pozoveš `padRow` u svom pozivu `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\((.+?)?\)\)/);
```

Trebalo bi da imaš argument u tvom pozivu za `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\(\s*?\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
