---
id: 660f4774e3e0df35a68bb5f2
title: Korak 90
challengeType: 1
dashedName: step-90
---

# --description--

Da biste ponovo generisali svoju piramidu, gurnite rezultat poziva `padRow` sa `done` i `count` kao argumente u vaš red podataka od `rows`, kao što ste to uradili u svom prvom ciklusu.

# --hints--

Vaš ciklus treba da pozove putanju `.push()` na vašem `rows`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(/);
```

Trebalo bi da pozoveš svoju funkciju `padRow` u svojoj putanji `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(/);
```

Treba da proslediš `done` kao prvi argument u pozivu `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done/);
```

Trebalo bi da proslediš `count` kao drugi argument u tvom pozivu od `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--

--fcc-editable-region--
  if (done === count) {
    continueLoop = false;
  } 
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
