---
id: 660f46460f9c36330ebc07d8
title: Korak 88
challengeType: 1
dashedName: step-88
---

# --description--

Opereta za jednakost može izazvati čudne ponašanja u JavaScript. Na primer, `"0" == 0` je tačno, iako je jedno niz karaktera a drugo broj.

Operater `<dfn>usawa mkali</dfn> `===`` se koristi za proveru da li su dve vrednosti jednake i da imaju isti tip. Opšćenito govoreći, ovo je operator jednakosti koji biste trebali koristiti u svakoj situaciji. Za operator stroge jednakosti, ``"0" === 0`` postaje netačno (false), jer iako imaju jednaku vrednost od nule, ne pripadaju istom tipu.

Ispravite vaš uslov za `done == count` ili koristite operator strogoj jednakosti.

# --hints--

Uslov za `if` mora koristiti strogu jednakost.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
