---
id: 660f447efc0e722f016c1be0
title: Korak 81
challengeType: 1
dashedName: step-81
---

# --description--

Tekstovi su ponovo vidljivi! Ovo je zato što ``"false"`` je niz karaktera, koji kada se evaluira kao tačno ili netačno postaje ``true``. To znači da ``"false"`` je tačna vrednost.

<dfn>Istinitost</dfn> je vrednost koja se smatra istinitom kada se proceni kao istinita ili neistinitoj. Mnoge vrednosti koje naiđete u JavaScript biće istinite vrednosti.

<dfn>Lažna vrednost</dfn> je suprotno - vrednost koja se smatra netačnom kada se procenjuje kao istinita ili neistinitost. JavaScript ima definisanu listu lažnih vrednosti. Neki od njih uključuju `false`, `0`, `""`, `null`, `undefined`, i `NaN`.

Pokušajte da promenite vaš uslov `if` u niz praznih znakova `""`, koji je lažna vrednost.

# --hints--

Tvoja rečenica sa `if` mora imati `""` kao uslov.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
