---
id: 660f4b641290da41b2cf0dd9
title: Korak 103
challengeType: 1
dashedName: step-103
---

# --description--

Kao što je slučaj sa dodavanjem, postoje različiti operateri koje možete koristiti za oduzimanje. Opereta ya <dfn>mgawo wa thamani wa kutoa</dfn> `-=` koristi vrednost dobijenu iz trenutne vrednosti parametra, a zatim vraća rezultat tom parametru.

Promeni svoje ponavljajuće izjave sa ispravnom izjavom koja koristi operater za dodelu vrednosti.

# --hints--

Tvoj ciklus od `for` ne bi trebalo da koristi `i = i - 1`.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /for\(leti=count;i>0;i=i-1\)\{rows\.push\(padRow\(i,count\)\);/);
```

Tvoj ciklus/loop `for` treba da koristi raspon vrednosti za smanjenje `i` za `1`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*-=\s*1\s*\)\s*\{\s*rows\.push\(\s*padRow\s*\(\s*i\s*,\s*count\s*\)\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i = i - 1) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
