---
id: 660f49e32001983c90b75850
title: Korak 98
challengeType: 1
dashedName: step-98
---

# --description--

Ovo je čisto kolo sa odličnim performansama. Odličan posao! Ali još ima mnogo toga za istraživanje.

Koristite višeredne komentare za dodavanje komentara u vaš ciklus `while`.

# --hints--

Vaš ciklus od `while` treba da pruži povratnu informaciju.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

Vaš krug od `while` treba da pruži povratnu informaciju.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /rows\.push/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
