---
id: 660f48e1d3682f39e81843c4
title: Korak 94
challengeType: 1
dashedName: step-94
---

# --description--

Tvoj piramidski generator još uvek funkcioniše. Međutim, postoji mogućnost ulaska u beskonačni ciklus.

Zato što samo proveravaš da li `done` nije jednako `count`, ako `done` biće **veće od** `count`, tvoj ciklus će trajati zauvek.

Proverite uslov svog ciklusa, ili proverite da li `done` je manje od ili jednako `count`.

# --hints--

Tvoj ciklus od `while` treba da proveri da li `done` je manjije ili jednako `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
