---
id: 660f34e99571070d56d2f231
title: Korak 67
challengeType: 1
dashedName: step-67
---

# --description--

Vaša funkcija `padRow` ima dva parametra koja ste opisali. Vrednost se pruža za te parametre kada se funkcija pozove.

Vrednost koju pružaš u pozivu funkcije naziva se <dfn>argument</dfn>, i prosleđuješ argumente u pozivu funkcije. Evo poziva funkcije sa `"Hello"` korišćenog kao argument:

```js
test("Hello");
```

Proveri `i + 1` i `count` kao argumente za tvoj poziv od `padRow`. Kao parametre, argumenti su razdvojeni zarezom.

# --hints--

Trebalo bi da prođeš `i + 1` za poziv `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

Trebalo bi da staviš zarez nakon tvog argumenta `i + 1`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

Trebalo bi da prosledite `count` kao vaš drugi argument.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
