---
id: 660f45ccf4ca5c31f253005a
title: Korak 86
challengeType: 1
dashedName: step-86
---

# --description--

Trenutno, ako promenite `continueLoop` u true, vaša petlja od `while` će trajati zauvek. Ovo se naziva <dfn>beskonačna petlja</dfn>, i morate biti pažljivi da izbegnete ovo. Beskonačna petlja može zamrznuti vaš sistem, zahtevajući potpuno ponovno pokretanje kako bi se izašlo.

Da biste izbegli ovo, počnite korišćenjem operatera za povećanje da biste povećali vrednost parametra `done` unutar vašeg ciklusa.

# --hints--

Vaš ciklus `while` treba da doda parametar `done`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
