---
id: 660f3dd626be3a1ffe27e5d1
title: Korak 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript takođe ima podršku za višelinijske komentare. Višelinijski komentari počinju sa `/*` i završavaju se sa `*/`.

Za razliku od komentara u jednoj liniji, komentari više linija će pokriti mnogo linija.

Koristite `/*` i `*/` da promenite vaš trenutni blok koda `for`, uključujući njegovo telo, u višelinijski komentar.

# --hints--

Trebalo bi da počnete sa višelinijskim komentarima i `/*`.

```js
assert.match(code, /\/\*/);
```

Treba da završiš više linija komentara sa `*/`.

```js
assert.match(code, /\*\//);
```

Tvoj ceo ciklus od `for` treba biti postavljen kao komentar.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
