---
id: 660f0c34aad72dc712b97624
title: Korak 27
challengeType: 1
dashedName: step-27
---

# --description--

Još jedan važan način za ovaj projekat je putanja `.pop()`. Izvlači poslednji element iz reda podataka i <dfn> vraća/ponavlja </dfn> taj element.

Kada funkcija vraća vrednost, možete to smatrati da vam vraća tu vrednost, čineći je dostupnom za upotrebu u drugim delovima vašeg koda.

Kreiraj novi parametar nazvan `popped` i postavi mu vrednost koja je rezultat `rows.pop()`. Zatim, unesi `popped` u konzolu.

# --hints--

Trebalo bi da objavite parametar nazvan `popped`.

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

Trebalo bi da koristiš `let` za deklarisanje svog parametra koji se zove `popped`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

Trebalo bi da pozoveš putanju `.pop()`.

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

Trebaće ti da pozoveš putanju `.pop()` u svoj red podataka `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

Treba da napišete svoj parametar `popped` u konzoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
