---
id: 660f3ce51f70571e1c5227c8
title: Korak 76
challengeType: 1
dashedName: step-76
---

# --description--

Komentari mogu pomoći da se objasni zašto vaš kod prolazi kroz određenu putanju, ili ostaviti napomene za kasnije samo za sebe.

U JavaScript, možete koristiti `//` da ostavite komentar u jednoj liniji vašeg koda.

Ostavi jedan red komentar na vašem kolažu `for` ili da vam se podsjeti da promijenite kod u drugi tip kolaža.

# --hints--

Trebalo bi da počnete sa komentarom na jednoj liniji kod `//`.

```js
assert.match(code, /\/\//);
```

Tvoj komentar od jedne linije mora imati najmanje pet slova.

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
