---
id: 660f415b76859a2736771607
title: Korak 78
challengeType: 1
dashedName: step-78
---

# --description--

Tvoja piramida je ponovo nestala. To je u redu – to se očekuje.

Kada kreirate svoj novi krug, morate naučiti o deklaraciji `if`. Deklaracija <dfn>`if`</dfn> vam omogućava da pokrenete deo koda samo kada se ispuni određeni uslov. Koristeći sledeću sintaksu:

```js
if (condition) {
  logic
}
```

Kreiraj deklaraciju za `if` koja ima vrednost tačno ili netačno `true` kao uslov. U telu, ispiši niz znakova `"Condition is true"`.

# --hints--

Treba da kreirate deklaraciju za `if`.

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

Vaša izjava za `if` mora imati `true` kao status.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

Telo vaše deklaracije za `if` treba da ispiše `"Condition is true"`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
