---
id: 660f1bf673487ae0bb25b900
title: Korak 37
challengeType: 1
dashedName: step-37
---

# --description--

Vaš krug sada mora biti pokren osam puta. Unutar tela kruga, ispišite vrednost kruga od `i` i vidite šta se dešava.

# --hints--

Trebalo bi da koristiš `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

Treba da unesete vrednost `i`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

Morate uneti vrednost `i` u vašu petlju `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
