---
id: 660f165270622fd4ec0da3f7
title: Korak 31
challengeType: 1
dashedName: step-31
---

# --description--

Deklarisanje parametra za ključnu reč `let` omogućava ponovno dodeljivanje vrednosti. To znači da možete kasnije promeniti `character` u potpuno drugačiju vrednost.

Za ovaj projekat, ne želite da promenite ove vrednosti parametara. Zato biste trebali umesto toga da koristite `const` da deklarišete te parametre. Parametri za `const` su specifični.

Prvo, parametar `const` ne može se ponovo postaviti kao parametar `let`. Ovaj kod će nam vratiti grešku:

```js
const firstName = "Naomi";
firstName = "Jessica";
```

Parametar `const` takođe ne može biti deklarisan bez vrednosti. Ovaj kod će vratiti grešku:

```js
const firstName;
```

Promeni svoje ključne reči `let` sa `const`.

# --hints--

Treba da koristiš `const` za deklarisanje svog parametra `character`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

Trebalo bi da koristiš `const` za deklarisanje svog parametra `count`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

Trebalo bi da koristite `const` za deklarisanje vašeg parametra `rows`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

Ne smiješ koristiti `let` u svom kodu.

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
