---
id: 660f38c34a4de6141c0c369f
title: Korak 71
challengeType: 1
dashedName: step-71
---

# --description--

Tvoj piramidalni generator sada radi kako je očekivano. Ali ovo je sjajna prilika da detaljnije pregledaš kod koji si napisao.

Inkrement nije jedini način za povećanje vrednosti u promenljivoj. Inkrement dodeljivanja vrednosti može se koristiti kao kratak način da se kaže: „uzmi početnu vrednost promenljive, dodaj ovu vrednost i ponovo stavi rezultat u promenljivu.“ Na primer, ove dve rečenice bi dale isti rezultat:

```js
test = test + 1;
test += 1;
```

Ispravite svoju ponavljajuću rečenicu u petlji `for` da koristi opseg inkrementalnih vrednosti.

# --hints--

Tvoj ciklus za `for` ne bi trebalo da koristi `i = i + 1`;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Tvoj ciklus za `for` bi trebalo da koristi dodeljiva vrednost sa `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
