---
id: 660f1e3f047bf4e403268713
title: Korak 39
challengeType: 1
dashedName: step-39
---

# --description--

Nažalost, sada ne možeš da vidiš šta tvoj ciklus radi.

Koristite `let` da deklarišete parametar `result`, i postavite mu vrednost kao prazan niz karaktera. Prazan niz karaktera se predstavlja navodnicima bez ničega između njih, kao što je `""`.

# --hints--

Trebalo bi da objaviš parametar `result`.

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

Trebalo bi da koristiš `let` za deklarisanje svog parametra `result`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

Vaš parametar za `result` treba da bude prazna sekvenca karaktera.

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
