---
id: 660f50a21fe7645252804f2b
title: Korak 112
challengeType: 1
dashedName: step-112
---

# --description--

Želite li način za prebacivanje između inverzne i standardne piramide?

Započni deklarisanjem promenljive `inverted` i postavi njenu vrednost na `true`. Ne treba da menjaš ovu promenljivu u svom kodu, ali treba da koristiš `let` kako bi naši testovi kasnije mogli da je izmene.

# --hints--

Trebalo bi da objavite parametar ``inverted`` na ``let``.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted/);
```

Trebalo bi da inicijalizujete ``inverted`` na thamani ya ``true``.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted\s*=\s*true;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
--fcc-editable-region--

--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  rows.unshift(padRow(i, count));
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
