---
id: 660f44f10ea40f300b896a5e
title: Korak 84
challengeType: 1
dashedName: step-84
---

# --description--

Sada si navikao da radiš sa deklaracijama ``if...else if...else``, možeš ih ukloniti iz svog koda.

Nakon što završiš to, koristi `let` da deklariraš parametar za `continueLoop` i zatim postavi vrednost tačno ili netačno za `false`. Zatim koristi `let` da deklariraš parametar za `done` i zatim postavi vrednost `0`.

# --hints--

Ne bi trebalo da ima dubinu od `if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

Ne bi trebalo da imaš dubinu od `else if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

Ne bi trebalo da ima dubinu od `else`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

Trebalo bi da koristiš `let` za deklarisanje parametra `continueLoop`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

Vaš parametar `continueLoop` mora imati vrednost `false`.

```js
assert.isFalse(continueLoop);
```

Trebalo bi da koristiš `let` za deklarisanje parametra `done`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

Vaš parametar za `done` mora imati vrednost `0`.

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
