---
id: 660f3915b41a441537ec9f5e
title: Korak 72
challengeType: 1
dashedName: step-72
---

# --description--

Budući da dodajete `i` na `1` samo, možete koristiti <dfn>operator za povećanje</dfn> `++`. Ovaj operator povećava vrednost parametra za 1, ažurirajući raspored vrednosti za taj parametar. Na primer, `test` će se pretvoriti u `8` ovde:

```js
let test = 7;
test++;
```

Promenite vaše podešavanje pojačanja i operaciju pojačanja za vaš ciklus osvežavanja.

# --hints--

Tvoj ciklus za `for` ne bi trebalo da koristi dodeljivanje sa `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

Tvoj ciklus za `for` bi trebalo da koristi operator sabiranja na `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
