---
id: 660f374d532dc41189cc9cc2
title: Korak 69
challengeType: 1
dashedName: step-69
---

# --description--

Sada je vreme za malo računanja. Zamislite piramidu sa tri sloja. Ako želimo da bude u sredini, biće kao što sledi:

```md
··#··
·###·
#####
```

Praznine su zamenjene medijanima ili prosečnim vrednostima radi lakšeg čitanja. Ako se ovaj obrazac nastavi, videće se da praznine u početku i na kraju niza prate određeni obrazac.

Ispravite vaš niz praznih mesta kako bi se vratilo `rowCount - rowNumber` puta.

Otvori konzolu da vidiš rezultate.

# --hints--

Trebalo bi da pozoveš `.repeat()` u svoj redosled od `" "` kako bi ponovio/ponovila `rowCount - rowNumber`.

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
