---
id: 660f0ee51d7460ce88cd248d
title: Korak 29
challengeType: 1
dashedName: step-29
---

# --description--

Da li ste očekivali da vidite `4` na konzoli? `.push()` obnavlja novu dužinu niza podataka, nakon dodavanja vrednosti koju unosite.

Važno je razumeti koju vrednost metoda vraća. Odvojite vreme da testirate `.push()` i `.pop()`. Kada budete spremni, uklonite sve pozive ka `.push()` i `.pop()`, zajedno sa vašim podacima o `console.log`.

# --hints--

Ne bi trebalo da imaš poziv za `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

Nema poziva za `.pop()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

Nema informacija o log fajlovima.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

Ne smeš da imaš parametar `popped`.

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

Ne sme da ima parametar od `pushed`.

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
