---
id: 660f2fbd45b520046cac68e8
title: Korak 65
challengeType: 1
dashedName: step-65
---

# --description--

Zapamtite da ste u prethodnom koraku naučili o vrednostima return. Funkcija može <dfn>kurudisha</dfn> vrednosti za vaš program da ih koristi na drugačiji način.

U funkciji, ključna reč `return` koristi se za određivanje vrednosti return. Na primer, ova funkcija bi vratila vrednost zadatu za prvi parametar:

```js
function name(parameter) {
  return parameter;
}
```

Koristi glavnu reč `return` da vrati vrednost varijable `character`, ponavljajući to puta `rowNumber`.

# --hints--

Trebalo bi da koristiš putanju `.repeat()`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

Trebalo bi da koristiš putanju `.repeat()` u svojoj varijabli `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

Treba da prođeš `rowNumber` za svoj poziv `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

Trebalo bi da koristiš glavnu reč `return`.

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

Trebalo bi da vratiš rezultate tvog poziva za `.repeat()`.

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
