---
id: 660f540c2176ea5dec01306d
title: Korak 118
challengeType: 1
dashedName: step-118
---

# --description--

Bravo! Isprobajte različite vrednosti za vaše parametre od `character`, `count`, i `inverted`.

Kada budete spremni da nastavite sa svojim sledećim projektom, postavite `character` na `"!"`, `count` na `10`, i `inverted` na `false` ili za nastavak.

Čestitam na završetku vašeg prvog projekta JavaScript!

# --hints--

Trebalo bi da staviš `character` umesto `"!"`.

```js
assert.equal(character, "!");
```

Trebalo bi da staviš `count` na mesto `10`.

```js
assert.equal(count, 10);
```

Trebalo bi da zamenite `inverted` sa `false`.

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
