---
id: 660f20473aef47e9b8c9afc6
title: Korak 40
challengeType: 1
dashedName: step-40
---

# --description--

Povećajte deklaraciju memorije kako biste ispisali vrednost `result`. U zavisnosti od konzole koju koristite, možda nećete videti ništa ispisano.

# --hints--

Trebalo bi da dodaš poziv `console.log`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Trebalo bi da odštampate vrednost vaše parametre `result`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
