---
id: 660f0a55847d6cc485f29ba5
title: Korak 26
challengeType: 1
dashedName: step-26
---

# --description--

Način (metoda) u JavaScript je funkcija povezana sa nekim vrednostima ili određenim objektima. Primer koji ste već naišli je metoda `.log()`, koja je deo objekta `console`.

Nizovi podataka imaju svoje metode, a prva metoda koju ćete istražiti je `.push()`. Ovo vam omogućava da "gurnete" (push-ujete) vrednost na kraj niza podataka. Evo primera dodavanja broja `12` na kraj niza podataka:

```js
array.push(12);
```

Koristite `.push()` da dodate niz karaktera `"freeCodeCamp"` na kraju vašeg reda `rows`. Dodajte ovaj kod pre `console.log` da vidite izmene koje ste napravili u svom redu podataka.

# --hints--

Moraš koristiti `.push()` u svoj kod.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

Trebalo bi da koristiš putanju `.push()` svog redova `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Treba ti da prođeš sekvencu slova `"freeCodeCamp"` putem svog `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
