---
id: 660f07d231941bc11719f664
title: Korak 21
challengeType: 1
dashedName: step-21
---

# --description--

Redovi podataka su specifični jer se smatraju da su <dfn>promenljivi</dfn>. To znači da možete promeniti vrednost u jednom polju.

Na primer, ovaj kod će postaviti broj `25` u drugo polje u redu podataka:

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

Ispravite element **tri** iz vašeg reda `rows` u broj `10`. Zatim ispišite red `rows` na vašoj konzoli.

# --hints--

Trebalo bi da koristiš podvučeni zapis na redu `rows` ponovo.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

Trebalo bi da dobijete pristup trećem elementu reda `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

Treba da koristiš operater raspodele na trećem elementu reda `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

Trebalo bi da postaviš vrednost `10` u treći element svog reda `rows`.

```js
assert.equal(rows[2], 10);
```

Treba ti da imaš drugi deo `console.log` u svom kodu.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Trebali biste napisati red `rows` u konzolu.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
