---
id: 660f359af3e32e0f1a6880b7
title: Korak 68
challengeType: 1
dashedName: step-68
---

# --description--

Sada bi trebalo da vidiš isti niz slova na svojoj konzoli. Tvoja funkcija ``padRow`` izvodi istu radnju koju si prethodno izvršio, ali je sada u samoj ponovno upotrebljivoj sekciji.

Koristite operater za dodavanje da spojite prazno mesto `" "` na početku i kraju vaše ponavljajuće sekvence slova `character`.

Zapamtite da možete koristiti operater `+` za spajanje nizova karaktera na ovaj način:

```js
" " + "string"
```

# --hints--

Trebalo bi da dodaš jedan razmak na početak vrednosti koju vraćaš.

```js
assert.match(padRow(1, 1), /^\s/);
```

Treba da dodaš jedan razmak na kraju vrednosti koju vraćaš.

```js
assert.match(padRow(1, 1), /\s$/);
```

Vaša funkcija za `padRow()` treba da vrati sekvencu `character` ponovljenu sa razmakom pre i posle te sekvence.

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
