---
id: 660f505d02b2bd513a1c3468
title: Korak 111
challengeType: 1
dashedName: step-111
---

# --description--

Vaša piramida ne prestaje biti suprotna. To je zato što dodajete novi red na kraj reda podataka.

Popravite telo vaše petlje da dodate novi red na početak podataka.

# --hints--

Trebalo bi da koristiš putanju `unshift` od `rows`.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

Trebalo bi da proslediš poziv `padRow()` kao argument na način `.unshift()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

Trebalo bi da prosledite `i` kao prvi argument svom pozivu `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

Trebalo bi da prosledite `count` kao drugi argument vašem pozivu `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
