---
id: 660f23b53db70af0f2620e78
title: Korak 43
challengeType: 1
dashedName: step-43
---

# --description--

Sada su svi tvoji brojevi vidljivi na jednoj liniji. Ovo neće raditi za kreiranje piramide.

Biće vam potrebno dodati novi red za svaki redak. Međutim, pritiskanje tastera return za umetanje razloma linije između navodnika u JavaScript će verovatno rezultovati greškom analize. Umesto toga, morate koristiti specijalnu sekvencu <dfn>escape sequences</dfn> `\n`, koja se tretira kao novi red kada se niz znakova napiše. Na primer:

```js
lineOne = lineOne + "\n" + lineTwo;
```

Koristite operator dvostrukog plusa da dodate novi red nakon vrednosti `result` i `row`.

# --hints--

Trebalo bi koristiti sekvencu za bekapanovanje `\n`. Zapamti da bi trebalo da bude niz karaktera, pa je zatvoren u navodnike.

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

Trebalo bi da povežete vaše promene iz `row` sa `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

Trebalo bi da povežete sekvencu za apstrahovanje `\n` sa vašim izmenama `row`.

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

Treba da ponovo unesete celu smesu u vaše izmene od `result`.

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
