---
id: 660f455b044d3230ed971e98
title: Korak 85
challengeType: 1
dashedName: step-85
---

# --description--

Ciklus od <dfn>`while`</dfn> će se nastaviti ponavljati sve dok `condition` više nije tačno. Ima sledeću sintaksu:

```js
while (condition) {
  logic;
}
```

Koristi tu sintaksu za deklarisanje ciklusa `while` sa `continueLoop` kao uslov. Njegovo telo treba da bude prazno.

# --hints--

Trebalo bi da koristiš kolačić `while`.

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

Vaš ciklus `while` treba da koristi `continueLoop` kao uslov.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
