---
id: 660f2eccfe3f820304af1b39
title: Korak 64
challengeType: 1
dashedName: step-64
---

# --description--

Da bi znao kako da kreiraš strukturu redova, tvoja funkcija ``padRow`` će morati da zna koji je red broj i koliko ukupno generisanih redova postoji.

Najbolji način da se ovo uradi je kreiranje funkcionalskih parametara za njih. Pružite svoj funkcionalni parametar `padRow`, `rowNumber` i `rowCount`. Više parametara se razdvajaju zarezom:

```js
function name(first, second) {

}
```

# --hints--

Vaša funkcija za `padRow` mora da ima parametar `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

Trebalo bi da dodate zarez nakon vašeg parametra `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

Tvoja funkcija za `padRow` mora imati parametar `rowCount`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
