---
id: 660f255022991ef34ed0ee88
title: Korak 44
challengeType: 1
dashedName: step-44
---

# --description--

Ispisivanje broja nije proizvelo piramidu koja izgleda dobro. Sada kada izvučete strukturirani sadržaj vašeg reda podataka za `rows`, vreme je da ispravite svoj početni ciklus.

Umesto slanja `i` u red podataka, pošalj vrednost svog parametra `character`.

# --hints--

Ne bi trebalo ponovo da guraš svoj parametar `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Trebalo bi da pritisnete svoj parametar `character`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
