---
id: 660f4a1472f8e63d76162ce5
title: Korak 99
challengeType: 1
dashedName: step-99
---

# --description--

Šta misliš da napraviš svoju piramidu okrenuvši je, ili <dfn>da ih promeniš </dfn>? Je l' vreme za pokušaj!

Počnite kreiranjem novog konteksta za `for`. Deklarišite svoj kontekst `i` i postavite mu vrednost `count`, zatim koristite boolean vrednosti (tačno ili netačno) `false` za uslove i informacije konteksta.

# --hints--

Vaš kod bi trebalo da ima ciklus od `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

Vaš ciklus od `for` treba da inicijalizuje `i` sa vrednošću `count`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count/);
```

Mzunguko wako wa `for` mora da koristi `false` kao uslov.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false/);
```

Tvoj ciklus `for` bi trebalo da koristi `false` kao ciklus.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
