---
id: 660f280dda5040f707c76b4a
title: Korak 45
challengeType: 1
dashedName: step-45
---

# --description--

Sada imate niz slova `#`, ali forma piramide još uvek nije tu. Srećom, razlika `i` predstavlja trenutni "čišćenje" broj u vašem ciklusu, što vam omogućava da ga koristite za kreiranje strukture tipa piramide.

Za postizanje ovoga, koristićete funkciju `.repeat()` dostupnu za niz znakova. Ova funkcija prima broj kao argument, specificirajući koliko puta treba ponoviti ciljani niz znakova. Na primer, korišćenjem `.repeat()` stvara se niz znakova `"Code! Code! Code!"`:

```js
const activity = "Code! ";
activity.repeat(3);
```

Koristi putanju `.repeat()` na tvojoj `character`, i sačuvaj je kao broj u `i`.

# --hints--

Trebalo bi da koristiš putanju `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

Trebalo bi da koristiš putanju `.repeat()` u tvojoj verziji `character`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

Trebalo bi da prođeš kroz svoj `i` na način `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

Trebalo bi ti da koristiš putanju `.repeat()` u putanji `.push()`

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\.repeat\(\s*i\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
