---
id: 660f18f059fe0fda192ce394
title: Korak 34
challengeType: 1
dashedName: step-34
---

# --description--

Tvoja petlja sada zahteva ispravan iterator. <dfn>Iterator</dfn> je parametar koji možete eksplicitno deklarisati u vašoj petlji od `for` da kontrolišete kako se petlja ponavlja ili prolazi kroz vašu logiku.

Je uobičajena praksa koristiti `i` kao vaš parametar prekidača u kolu. Kolo `for` omogućava vam da ovo deklarišete u standardnim zagradama `()`. Na primer, ovde je kolo `for` koje deklarše parametar `index` i postavlja njegovu vrednost na `100`.

```js
for (let index = 100; "second"; "third") {

}
```

Zameni niz znakova `"iterator"` sa izrazom `let` za parametar `i`. Postavi vrednost `0` od. Ovo će parametru `i` dati vrednost `0` za **prvi put** kada vaš ciklus prođe.

# --hints--

Trebalo bi da koristiš `let` za deklarisanje parametra `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

Trebalo bi da postaviš vrednost `0` za svoj parametar `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

Tvoj ciklus za `for` treba da inicijalizuje preključač `i` za `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
