---
id: 660f207334fabaeac3269c38
title: Korak 41
challengeType: 1
dashedName: step-41
---

# --description--

Da bi kontrolisao niz karaktera ``result``, koristiće se različita vrsta petlje. Specifično, petlja za ``for...of``, koja ponavlja svaki element u iterativnom objektu i privremeno mu dodeljuje vrednost na parametar.

Sintaksa kruga za `for...of` izgleda ovako:

```js
for (const value of iterable) {

}
```

Zapamti da možeš koristiti `const` jer taj parametar postoji samo u jednom ciklusu, a ne tokom celog ciklusa.

Napravi petlju za `for...of` koja iterira kroz niz podataka `rows`, postavljajući svaku vrednost u parametar `row`.

# --hints--

Treba da koristiš drugačiju glavnu reč od `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

Trebalo bi da objaviš parametar `row`.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

Trebalo bi da koristiš `const` za deklarisanje svog parametra `row`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

Vaš ciklus za `for...of` treba da deklarše vaš parametar `row`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

Vaš parametar za `row` mora biti preuzet iz `rows` koristeći glavnu reč `of`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

Struktura kruga vašeg `for...of` treba da bude prazna.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
