---
id: 660f0da9bf1035c9097af20a
title: Korak 28
challengeType: 1
dashedName: step-28
---

# --description--

Trebalo bi da ste videli `"freeCodeCamp"` ispisano na konzoli. Ovo je zato što `.pop()` vraća vrednost preuzetu iz niza podataka - i ranije ste dodali `"freeCodeCamp"` na kraj niza podataka.

Ali šta `.push()` vraća? Postavi vrednost `rows.push()` koju imaš u novom parametru `pushed`, i snimi ga u memoriju.

# --hints--

Trebalo bi da objaviš parametar `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

Trebalo bi da koristiš `let` za deklarisanje svog parametra `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

Treba da postaviš vrednost `rows.push("freeCodeCamp")` u svoj parametar `pushed`.

```js
assert.equal(pushed, 4);
```

Trebalo bi da pišete u memoriju vašeg parametra `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
