---
id: 660f47afe4c98536715d5fa4
title: Korak 91
challengeType: 1
dashedName: step-91
---

# --description--

Operacija <dfn>kutokufanana kwa ukali</dfn> `!==` vam omogućava da proverite da li dve vrednosti ne odgovaraju, ili nemaju isti tip. Sintaksa je ista kao za operaciju jednakosti: `value !== 4`.

Trenutno, ciklus za `while` radi samo ako `continueLoop` je tačno. Proverite status kruga za `while` da biste proverili da li `done` nije kompatibilan sa `count`.

# --hints--

Vaš ciklus za `while` treba da proveri da li `done` i `count` ne odgovaraju.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
