---
id: 660f4f79e2a82a4e92290f44
title: Korak 110
challengeType: 1
dashedName: step-110
---

# --description--

Ponekad želite da vratite kod koji je ranije bio komentarisano. To možete učiniti uklanjanjem `/*` i `*/` koji okružuju taj kod. Ovo se zove <dfn>kuondoa maoni</dfn>.

Uklonite samo komentare u vašoj prvoj petlji od `for`. Ostavite jednoredne komentare i ostale dva višelinredna komentara na mestu.

# --hints--

Trebalo bi da uklonite komentare u vašem prvom krugu od `for`.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

Ne bi trebalo da brišeš svoje komentare sa jedne linije.

```js
assert.match(code, /\/\//);
```

Trebalo bi da uklonite komentare u vašem ciklusu od `while`.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

Ne bi trebalo da uklanjaš komentare u svoj drugi ciklus od `for`.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
