---
id: 660f2a70ad6225fa503e71c3
title: Korak 46
challengeType: 1
dashedName: step-46
---

# --description--

Približava se! Trenutno, susrećate se sa onim što je poznato kao <dfn>pojedinačna greška</dfn>, problem koji se često javlja u jezicima baziranim na nuli, kao što su JavaScript.

Prvi indeks tvog niza `rows` je `0`, zbog čega svoju `for` petlju započinješ sa `i = 0`. Međutim, ponavljanje niske karaktera nula puta ne ispisuje ništa.

Da biste ovo ispravili, dodajte `1` u vrednost `i` u vašem pozivu `.repeat()`. Nemojte se vraćati na `i` kao što ste to uradili u uslovima svog ciklusa.

# --hints--

Treba da dodaš `1` u `i` na putanji `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
