---
id: 660f061d259bbebc37461080
title: Korak 20
challengeType: 1
dashedName: step-20
---

# --description--

Možete dobiti vrednost unutar niza podataka koristeći <dfn>kielezo</dfn> te vrednosti. Indeks je broj koji predstavlja poziciju vrednosti u nizu podataka, počevši od `0` za prvu vrednost.

Možete dobiti vrednost korišćenjem <dfn>korišćenje zagrada</dfn>, kao što je `array[0]`.

Koristite ``console.log`` i uglate zagrade za štampanje prve vrednosti iz vaše kolone ``rows``.

# --hints--

Trebalo bi da imaš deklaraciju `console.log()` u svom kodu.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Trebalo bi da dobiješ svoj red `rows`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

Trebalo bi da koristiš pisanje sa zagradama i svoj red `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

Trebalo bi da koristiš uglate zagrade da dođeš do prvog elementa svog redova `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

Trebalo bi da napišete prvi deo vašeg reda `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
