---
id: 660f033cf051ebb50ea3bf48
title: Korak 18
challengeType: 1
dashedName: step-18
---

# --description--

U programima, često ćete morati da radite sa velikom količinom podataka. Postoje mnoge strukture podataka koje vam mogu pomoći u organizovanju i upravljanju vašim podacima. Jedna od osnovnih struktura podataka je <dfn>red podataka</dfn>.

<dfn>Niz podataka</dfn> je neprimitivni tip podataka koji može da skladišti sekvencu vrednosti. Neprimitivni tipovi podataka se razlikuju od primitivnih tipova jer mogu da skladište složenije podatke. Primitivni tipovi, kao što su nizovi karaktera i brojeva, mogu da skladište samo jednu vrednost u jednom trenutku.

Podatni nizovi su prikazani korišćenjem uglastih zagrada (`[]`). Evo primera parametra sa vrednošću praznog niza podataka:

```js
let array = [];
```

Deklariši parametar `rows` i postavi mu vrednost na prazan niz podataka.

# --hints--

Trebalo bi da imate kriterijum za `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

Trebalo bi da koristite `let` za deklarisanje svog parametra `rows`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

Trebalo bi da postavite vrednost praznog niza podataka za vaš parametar `rows`.

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
