---
id: 660f17d4e9f227d86e834abd
title: Korak 33
challengeType: 1
dashedName: step-33
---

# --description--

Da biste kreirali piramidu, moraćete da napravite mnogo redova. Kada morate ponoviti vežbu dok ne bude ispunjen neki uslov, koristićete <dfn>mzunguko</dfn>. Postoji mnogo načina da se piše petlja.

Počeo je sa osnovnim ciklusom `for`. Ciklusi `for` koriste sledeću sintaksu:

```js
for (iterator; condition; iteration) {
  logic;
}
```

U sledećih koraka, ćete pregledati svaki deo kruga detaljno. Trenutno, kreirajte krug za `for` koji uključuje reči `"iterator"`, `"condition"` i `"iteration"` za tri dela. Postavite telo kružnika <dfn>mwili</dfn>, sekciju unutar okruglih zagrada `{}`, na prazno.


# --hints--

Trebalo bi da imaš krug `for`.

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

Prvi deo tvog kruga `for` mora biti niz znakova `"iterator"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

Drugi deo tvog kruga `for` treba da bude niz slova `"condition"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

Treći deo vašeg ciklusa za `for` mora biti niz slova `"iteration"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

Telo vašeg kruga wa `for` bi trebalo da bude prazno.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
