---
id: 660f383d4c772c12ff59904b
title: Korak 70
challengeType: 1
dashedName: step-70
---

# --description--

Možete proslediti cijeli matematički izraz kao argument. Funkcija će primiti rezultat evaluacije tog izraza. Na primjer, ova dva poziva funkcije ispod dati bi isti rezultat. (Napomena da se pravilo redosleda operacija PEMDAS—Zagrade, Eksponenti, Množenje, Deljenje, Sabiranje, Oduzimanje—primjenjuje):

```js
test(2 * 3 + 1);
test(7);
```

Ako ponovo pogledamo tu strukturu:

```md
··#··
·###·
#####
```

Ispravite vrednost `character` ili ponoviti puta `2 * rowNumber - 1`.

Ponavljanje otvaranja konzole radi pregleda ažuriranih rezultata.
*(Alternative, more direct imperative: Ponovo otvorite konzolu da biste videli ažurirane rezultate.)*

# --hints--

Tvoja funkcija mora da prosledi `2 * rowNumber - 1` za tvoj poziv `.repeat()`.

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
