---
id: 660f0f980e98e8cf77f1ce31
title: Korak 30
challengeType: 1
dashedName: step-30
---

# --description--

Promenite vaš deklarativni deo `rows` ako mu je ponovo dodeljena prazna linija podataka.

Takođe, promenite svoju sekvencu karaktera ``'Hello'`` koristeći ponovo dvostruke navodne znakove. Generalno, nije važno koji od ova dva preferirate, ali želite da imate doslednost u tom izboru celom projektu.

# --hints--

Vaš red podataka za `rows` trebao bi da bude prazan.

```js
assert.empty(rows);
```

Tvoj niz slova za `"Hello"` bi trebalo da koristi dvostruke navodne oznake.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
