---
id: 660f4934fb48f63abd5ae371
title: Korak 95
challengeType: 1
dashedName: step-95
---

# --description--

Korišćenje `done` za praćenje broja kreiranih redova je funkcionalan metod, ali možeš malo očistiti logiku.

Redovi podataka imaju posebnu osobinu od `length` koja vam omogućava da vidite koliko vrednosti, ili <dfn>elementa</dfn>, sadržanih u redu podataka. Možete doći do pristupa ovoj osobinici koristeći sintaksu kao `myArray.length`.

Imajte na umu da `rows.length` u pozivu `padRow` će vam dati jednu grešku, jer `done` se dodaje *pre* tog poziva.

Proverite svoje postavke ili proverite da li `rows.length` je manje od `count`.

# --hints--

Tvoj ciklus za `while` treba da proveri da li `rows.length` je manje od `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
