---
id: 660f17294346b7d69e79db3d
title: Korak 32
challengeType: 1
dashedName: step-32
---

# --description--

Sada si spreman da počneš sa izgradnjom svog piramidskog generatora. Tvoj parametar za `character` će biti kao temeljni kamen piramide.

`"Hello"` Možda neće dobro funkcionisati za to. Promeni vrednost `character` da bude heš vrednost `"#"`.

# --hints--

Vaš parametar za `character` mora biti znak heksa/hash.

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
