---
id: 6610c2d8d67563174fcf96dc
title: Korak 52
challengeType: 1
dashedName: step-52
---

# --description--

Ako imate vrednost koja je eksplicitno napisana u vašem kodu, kao niz znakova `"Hello!"` u svojoj funkciji, smatra se da je <dfn>hardkodirano</dfn>. Hardkodiranje vrednosti unutar funkcije može učiniti da više nije upotrebljivo kako želite.

Umesto toga, možete definisati <dfn>parametare</dfn> za funkciju. Parametri su specijalne varijable koje primaju vrednosti kada pozovete funkciju, i mogu se koristiti unutar vaše funkcije za modifikovanje rezultata funkcionalskog koda putem promena.

Da biste dodali parametar u vašu funkciju, morate dodati ime parametra unutar običnih zagrada. Na primer, ova funkcija `demo` ima parametar `name`:

```js
function demo(name) {

}
```

`name` izgleda da je važan parametar, pa nastavi i dodaj ga u svoju funkciju od `padRow`.

# --hints--

Vaša funkcija za `padRow` mora imati parametar `name`.

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
