---
id: 6610b8f6a98d25f4d485a94d
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Osnovna vrednost neinicijalizovane varijable je `undefined`. Ovo je poseban tip podataka koji predstavlja vrednost koja još uvek nema definiciju.

Još uvek možete dodeliti vrednost neinicijalizovanoj promenljivoj. Evo primera:

```js
let uninitialized;
uninitialized = "assigned";
```

Ubacite sekvencu karaktera `"Test"` za vašu izmenu `secondCharacter` ispod vaše deklaracije. Otvorite konzolu da vidite kako je vaša memorija promenjena.

# --hints--

Trebaš da inicijalizuješ `secondCharacter`. Zapamti da je inicijalizacija postavljanje vrednosti kada objaviš promenu.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

Trebalo bi da koristiš distributivnu operaciju na `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

Trebaće ti da uneseš niz slova `"Test"` za svoju transformaciju `secondCharacter`.

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
