---
id: 6610c6541c82551f95e765ab
title: Korak 58
challengeType: 1
dashedName: step-58
---

# --description--

Promene vrednosti se mogu deklarisati i unutar funkcije. Ove promene se smatraju lokalnim opsegom <dfn>lokalni opseg</dfn>, ili <dfn>blok opseg</dfn>. Varijabla deklarisana unutar funkcije može se koristiti samo unutar te funkcije. Ako pokušate da joj pristupite izvan funkcije, dobićete grešku referenciranja.

Da biste ovo videli, koristite `const` da deklarišete parametar `test` u vašoj funkciji `padRow`. Inicijalizujte sa vrednošću `"Testing"`.

Nakon toga, ispod svoje funkcije, pokušajte da napišete `test` u konzolu. Videće grešku jer nije deklarisano van opsega funkcije. Uklonite `console.log` da biste prošli test i nastavili.

# --hints--

Tvoja funkcija mora da deklariše parametar `test`.

```js
assert.match(padRow.toString(), /var\s+test/);
```

Trebalo bi da inicijalizujete `test` sa vrednošću `"Testing"`. Ne zaboravite tačku.

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

Tvoj parametar za `test` treba da dođe pre glavne reči `return`.

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
