---
id: 6610c83b52583e245a079217
title: Korak 61
challengeType: 1
dashedName: step-61
---

# --description--

Sada vrednost vaših promena `call` je `"Testing"`. Međutim, vaša funkcija više ne koristi parametar `name`.

Uklonite parametar `name` iz deklaracije vaše funkcije, a zatim uklonite niz znakova `"CamperChan"` iz poziva `padRow`.

Pia, ukloni sve `console.log` iz funkcije `padRow`.# --hints--

Tvoja funkcija `padRow` ne sme imati parametar `name`.

```js
assert.notMatch(padRow.toString(), /name/);
```

Ne smete da prosledite `"CamperChan"` u vaš poziv `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

Još moraš da pozoveš svoju funkciju `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

Ne smete imati `console.log` pre glavne reči od `return`.

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

Ne smeš imati `console.log` nakon glavne reči od `return`.

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
