---
id: 6610b9f7619764fad5fd516d
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Takođe možete postaviti vrednost jednog parametra u drugi parametar. Na primer:

```js
let first = "One";
let second = "Two";
second = first;
```

Parametar za `second` biće sa vrednošću `"One"`.

Da biste videli kako to radi, promenite vrednost `secondCharacter` od `"Test"` na vaš parametar za `character`.

Zatim otvori konzolu da vidiš šta se beleži.

# --hints--

Ne smete postaviti vrednost za `"Test"` u vaš parametar/polje `secondCharacter`.

```js
assert.notEqual(secondCharacter, "Test");
```

Trebali biste postaviti vrednost parametra `character` u vaš parametar `secondCharacter`. Ne zaboravite tačku.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

Vaš parametar za `secondCharacter` sada bi trebalo da ima vrednost od `"World"`.

```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
