---
id: 6610c424b7119919b62932f4
title: Korak 53
challengeType: 1
dashedName: step-53
---

# --description--

Funkcija ne mora da vrati vrednost koja je direktno postavljena. Može vratiti vrednost sačuvanu u parametru. Parametri su specifični parametri funkcije, pa se oni mogu vratiti i.

Promeni svoju funkciju `padRow` ili `return` parametar `name` direktno.

# --hints--

Tvoja funkcija sa `padRow` treba da vrati vrednost parametra za `name`.

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
