---
id: 6610c105bbdacc114d6cdc44
title: Korak 48
challengeType: 1
dashedName: step-48
---

# --description--

Da biste koristili funkciju, morate je pozvati. <dfn>Poziv funkcije</dfn> govori vašem programu da izvrši kod iz funkcije gde god želite da ga pozovete. Sintaksa poziva funkcije je ime funkcije praćeno standardnim zagradama. Na primer, ovaj kod definiše i poziva funkciju `test`.

```js
function test() {

}
test();
```

Pozovi svoju funkciju `padRow`.

# --hints--

Trebalo bi da pozoveš funkciju `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

Telo vašeg `padRow` treba da bude prazno.

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
