---
id: 6610c8cfe4cf4d278e35c156
title: Korak 63
challengeType: 1
dashedName: step-63
---

# --description--

Kao što je očekivano, tvoja funkcija sada ponovo vraća `undefined`. Tvoj parametar za `call` više nije potreban, stoga ukloni deklaraciju za `call` i `console.log` za parametar `call`.

# --hints--

Ne bi trebalo da imaš izjavu o `call`.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /constcall/);
```

Nemojte da pišete memoriju za vaš parametar `call`.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
