---
id: 6610b741b54b90f0c0fb3d58
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Kada su imena parametara više od jedne reči, postoje posebna pravila o tome kako treba kapitalizovati te reči. U JavaScript, korišćena konvencija je <dfn>camel case</dfn>.

Camel case znači da je prva reč u imenu napisana tamamen malim slovima, ali su sve sledeće reči napisane velikim slovom na početku svake reči. Evo primera za camel case:

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

Koristi camel case da objaviš novi atribut `secondCharacter`.

# --hints--

Trebalo bi da objavite parametar `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

Trebalo bi da koristiš `let` za deklarisanje svog parametra `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

Ne treba da postaviš vrednost za svoj parametar `secondCharacter`. Ne zaboravi tačku-zarez.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
