---
id: 6610c71600966a2191d3a64a
title: Korak 59
challengeType: 1
dashedName: step-59
---

# --description--

Vrednosti koje se vraćaju iz funkcije koriste se prilikom pozivanja te funkcije. Možete koristiti poziv funkcije direktno kao vraćenu vrednost, ili uhvatiti vraćenu vrednost u parametar. Na ovaj način, možete koristiti vrednost dodeljenu u lokalnom opsegu parametra, izvan funkcije gde je kreirana.

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

Da biste koristili vašu vrednost od `"Testing"`, vratite se iz funkcije `padRow` tako što ćete ispraviti svoju deklaraciju od `return` kako bi vratila samo parametar `test`.


# --hints--

Tvoja funkcija `padRow` treba da vrati parametar `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
