---
id: 6610c538372aa61cc0f5b122
title: Korak 57
challengeType: 1
dashedName: step-57
---

# --description--

Promene u JavaScript su dostupne u specifičnom <dfn>wigo</dfn>. Drugim rečima, mesto gde se promene najavljuju određuje gde u vašem kodu mogu biti korišćene.

Prvi opseg je globalni opseg (global scope). Promene deklarisane izvan bilo kog "bloka" kao funkcija ili `for` petlje su u <dfn> globalnom opsegu</dfn>. Vaše promene `character`, `count` i `rows` su sve u globalnom opsegu.

Kada su vaše promene u globalnom opsegu, funkcija može da ih pronađe u svojoj definiciji. Evo primera funkcije koja koristi globalne promene: `title`

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

Ovaj primer će vratiti `"Professor Naomi"`. Ažurirajte svoju funkciju `padRow` da vrati vrednost koja spaja vaše promene `character` na početku parametra `name`.

# --hints--

Tvoja funkcija od `padRow` treba da poveže `character` na početku `name`.

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

Vaša funkcija za ``padRow`` treba da vrati rezultate od ``character + name``.

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
