---
id: 6610c21b3ef82015573ffbbe
title: Korak 51
challengeType: 1
dashedName: step-51
---

# --description--

Tvoj parametar za `call` ima vrednost `undefined`, iako si ga spomenuo! Ovo je zato što tvoja funkcija `padRow` trenutno ne vraća nikakvu vrednost. Podrazumevano, funkcije vraćaju `undefined` kao svoju vrednost.

Da biste vratili nešto drugo, morate koristiti glavnu ključnu reč `return`. Evo primera funkcije koja vraća niz karaktera `"Functions are cool!"`:

```js
function demo() {
  return "Functions are cool!";
}
```

Koristi glavnu reč `return` ili svoj funkcioner da vrati niz slova `"Hello!"`.

# --hints--

Tvoja funkcija za `padRow` treba da koristi glavnu ključnu reč od `return`.

```js
assert.match(padRow.toString(), /return/);
```

Tvoja funkcija za `padRow` treba da vrati niz karaktera `"Hello!"`.

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
