---
id: 6610c87eac0f0b256d7b037e
title: Korak 62
challengeType: 1
dashedName: step-62
---

# --description--

Zato što tvoja funkcija više ne koristi parametar, promena argumenta nije uticala.

Uklonite deklaraciju `test` na mestu `return` iz vaše funkcije `padRow`, kako bi funkcija ponovo bila prazna.

# --hints--

Tvoja funkcija koja je `padRow` ne sme da ima parametar od `test`.

```js
assert.notMatch(padRow.toString(), /test/);
```

Vaša funkcija za `padRow` ne bi trebalo da vraća vrednost.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
