---
id: 6610b8017d1671f2814e8c77
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Kada deklariš promenu bez inicijalizovanja njene vrednosti, smatra se da je <dfn>haijaanzishwa</dfn>. Trenutno, vaše promene od `secondCharacter` nisu inicijalizovane.

Prikažite `console.log()` ili da vidite koja vrednost vaše izmene od `secondCharacter` ima.

# --hints--

Treba da dodaš `console.log` drugog.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Trebalo bi da napišete vaše promene za `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
