---
id: 6610c16c4fa0df12c0e30675
title: Korak 49
challengeType: 1
dashedName: step-49
---

# --description--

Pozivate svoju funkciju od `padRow`, ali ništa ne radite sa tim pozivom funkcije. Sve funkcije u JavaScript <dfn>hurejesha</dfn> vraćaju vrednost, što znači da vraćaju deklarisane rezultate pozivanja kako biste mogli koristiti negde drugde.

Da biste videli rezultate pozivanja vaše funkcije ``padRow``, deklarišite parametar ``call`` i postavite vrednost vašeg poziva ``padRow`` u taj parametar.

# --hints--

Trebalo bi da objavite parametar `call`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

Trebalo bi da koristiš `const` za deklarisanje svog parametra `call`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

Trebalo bi da postaviš vrednost `call` kao rezultat svog poziva `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
