---
id: 6610c77d50636722e5b6be17
title: Korak 60
challengeType: 1
dashedName: step-60
---

# --description--

Ispod deklaracije `return`, napiš niz slova `"This works!"` u konzolu.

Nakon toga, videćete da niz karaktera `"This works!"` nije vidljiv na konzoli, a linija `console.log("This works!")` je prikazana u sivoj boji.

Kopiraj konsolnu memoriju i zalepi iznad deklaracije `return`. Sada, sekvenca znakova `"This works!"` bi trebalo da se pojavi na konzoli.

Bitno je znati da ključna reč ``return`` ne služi samo za definisanje vrednosti koja će biti vraćena iz vaše funkcije, već takođe sprečava izvršavanje vašeg koda unutar funkcije ili blok izjave. To znači da se bilo koji kod nakon izjave ``return`` neće izvršiti.

# --hints--

Vaša funkcija ``padRow`` treba da vrati parametar ``test``.

```js
assert.equal(padRow("Naomi"), "Testing");
```

`console.log` tvoj prvi treba da dođe nakon glavne reči od `return`.

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

`console.log` tvoj drugi bi trebalo da dođe pre glavne reči od `return`.

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  
  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
