---
id: 6610c48c4ea0891afa7c4696
title: Korak 54
challengeType: 1
dashedName: step-54
---

# --description--

Kada ponovo otvorite svoju konzolu, videćete da vaš `padRow` vraća `undefined`, iako ste naveli vrednost return! To je zato što parametri moraju imati vrednost kada pozovete funkciju.

Kada prosleđujete vrednost pozivu funkcije, ta vrednost se zove <dfn>hoja</dfn>. Ovde je primer korišćenja funkcije `demo` i prosleđivanja `"Naomi"` kao argumentu za parametar `name`.

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

Proslijedite svoje ime kao argument za parametar cha `name` u vašem pozivu `padRow`. Zapamtite da je vaše ime sekvenca znakova, pa ćete morati koristiti navodnike.

# --hints--

Trebalo bi da prođeš niz slova za svoj poziv `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
