---
id: 6610bf6fa14d700beed1b109
title: Korak 87
challengeType: 1
dashedName: step-87
---

# --description--

Opereta za <dfn>usawa</dfn> `==` se koristi za proveru da li su dve vrednosti jednake. Za poređenje dve vrednosti, koristićete izraz kao što je `value == 8`.

Ispod `done++` unutar vašeg petlje, dodajte izraz `if`. Taj izraz mora da proveri da li je `done` jednak `count` koristeći operator jednakosti.

# --hints--

Treba da koristiš izraz `if` unutar svog kruga. Treba da se doda nakon `done++`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

Tvoj izraz od `if` treba da koristi operator jednakosti za poređenje `done` sa `count` u uslovu.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
