---
id: 66458f0a05df478aa627629e
title: Korak 82
challengeType: 1
dashedName: step-82
---

# --description--

Tekst je ponovo nestao! Sekvenca praznih karaktera se računa kao `false`, što ga čini nevažećom vrednošću <dfn></dfn>. Biće vam naučeno više o pozitivnim i nevažećim vrednostima u budućim projektima.

Dalje od deklaracija `if`, JavaScript takođe ima deklaracije za <dfn>else if</dfn>. Deklaracije za `else if` omogućavaju pregled mnogo uslova unutar jednog kodnog fragmenta.

Ovo je sintaksa deklaracije za `else if`:

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

Ako je prvi uslov `false`, JavaScript će proveriti sledeći uslov u lancu. Ako je drugi uslov `false`, JavaScript će proveriti treći uslov, i tako dalje.

Ispod vaše deklaracije `if`, dodajte deklaraciju `else if` koja proverava da li je `5` manje od `10`. Zatim, unutar tela deklaracije `else if`, napišite sekvencu karaktera `"5 is less than 10"` u konzoli.

Proverite konzolu da biste videli rezultate.

# --hints--

Trebalo bi da imate deklaraciju `else if`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

Vaša izjava za `else if` treba da proveri da li je `5` manje od `10`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

Trebalo bi da ukucate niz znakova `"5 is less than 10"` u konzolu.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
