---
id: 664599653fcd6e97104f9261
title: Korak 83
challengeType: 1
dashedName: step-83
---

# --description--

Ponekad ćete želeti da pokrenete drugačiji kod kada su svi uslovi za `if...else if` ni `false`. Možete to učiniti dodavanjem komada koda za `else`.

Kodni fragment `else` će se proveriti samo ako uslovi u kodnim fragmentima `if` i `else if` nisu navedeni.

Ovde deo koda za `else` je dodat u deo koda za `else if`.

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

Dodaj kodni deo ``else`` u kodni deo ``else if``. Unutar kodnog dela ``else``, napiši niz karaktera ``"This is the else block"`` na konzoli.

Da biste videli rezultate u konzoli, možete ručno promeniti `<` u deklaraciji `else if` u `>`. Ovo će stvoriti stanje `false` i deo koda sa `else` će biti izvršen.

# --hints--

Trebalo bi da imaš deo koda `else`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

Vaš kodni isječak za `else` bi trebalo da ispisa sekvencu karaktera `"This is the else block"` na konzoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
