---
id: 664ed9344ba3495e1f67f8ea
title: Korak 14
challengeType: 1
dashedName: step-14
---

# --description--

Sada, nakon pregleda deklaracije i inicijalizacije, uklonite kod koji ste napisali za svoj pregled.

Ne brišite vašu varijablu za `character`.

# --hints--

Nemoguće je da postoji varijabla `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

Ne bi trebalo da imaš varijablu za `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

Ne treba vam dubina od `console.log` za `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

Potrebna je dubina od `console.log` za `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

Još moraš imati svoju varijablu `character`.

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
