---
id: 664ed97c55a99e5ffab759dc
title: Korak 13
challengeType: 1
dashedName: step-13
---

# --description--

Pre nego što nastavite, trebalo biste da odvojite malo vremena za pregled koncepata koje ste naučili.

Koristi glavni izraz `let` da deklariraš varijablu `profession` u varijablu `age`. Inicijalizuj `profession` sa nizom karaktera `"teacher"`, ali ne inicijalizuj `age` sa bilo kojom vrednošću.

Napišite sve vaše varijable u konzolu da biste videli rezultate.

# --hints--

Treba da deklariraš varijablu `profession`.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

Trebalo bi da koristiš `let` za deklarisanje varijable `profession`.

```js
assert.match(code, /let\s+profession/);
```

Trebalo bi da postaviš vrednost niza znakova `"teacher"` za varijablu `profession`.

```js
assert.strictEqual(profession, "teacher");
```

Trebalo bi da deklariš varijablu `age`.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

Trebalo bi da koristiš `let` za deklarisanje varijable `age`.

```js
assert.match(code, /let\s+age/);
```

Ne smete postaviti vrednost promenljivoj `age`.

```js
assert.isUndefined(age);
```

Trebalo bi ti da napišeš `age` u konzoli.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

Trebalo bi da ukucate `profession` u konzolu.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
