---
id: 665732da4815b70bb083915e
title: Korak 24
challengeType: 1
dashedName: step-24
---


# --description--

U proteklih nekoliko koraka, naučili ste sve o radu sa nizovima podataka. Odvojite malo vremena da pregledate ono što ste naučili.

Počnite deklaracijom varijable `cities` i inicijalizujte je kao niz podataka sa sekvencama karaktera `"London"`, `"New York"`, i `"Mumbai"`. Zatim ispišite tu varijablu u konzolu.

Nakon pisanja, promenite poslednji element od `cities` u niz slova `"Mexico City"`, zatim ponovo napišite varijablu `cities`.

Ako ga pokrenete ispravno, trebalo bi da vidite ove rezultate u konzoli.

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

Trebalo bi da koristiš `let` za deklarisanje varijable `cities`.

```js
assert.match(code, /let\s+cities/);
```

Trebao bi da postaviš niz podataka sekvence karaktera `"London"`, `"New York"`, na `"Mumbai"` u varijablu `cities`.

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*,?\s*\]/);
```

Trebalo bi da koristite `console.log()` za pisanje celog reda podataka `cities` u konzolu.

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

Treba da ispravite poslednji element reda podataka za `cities` u niz znakova `"Mexico City"`. Zapamtite da možete dobiti poslednji element reda podataka korišćenjem `array[array.length - 1]`.

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

Trebalo bi da imaš dve deklaracije od `console.log(cities)` u svom kodu.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
