---
id: 665732e62eb8b30c249b5c2c
title: Korak 25
challengeType: 1
dashedName: step-25
---

# --description--

Sada ste spremni da nastavite sa sledećim setom lekcija o redovima podataka.

Uklonite sve vaše kodove iz prethodnog koraka.

# --hints--

Treba da uklonite sve reference vaše varijable `cities`.

```js
assert.notMatch(__helpers.removeJSComments(code), /cities/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let cities = ["London", "New York", "Mumbai"];
console.log(cities);
cities[cities.length - 1] = "Mexico City";
console.log(cities);
--fcc-editable-region--
console.log(rows);
```
