---
id: 66643c93e05093c728abdbe9
title: Korak 55
challengeType: 1
dashedName: step-55
---

# --description--

Pre nego što nastaviš, odvoji malo vremena da proveri kako funkcionišu radnici/komponente.

Definisati funkciju nazvanu `addTwoNumbers`. Ova funkcija treba da prihvati dva argumenta i vrati njihov zbir.

Vaša funkcija ne bi trebalo da koristi hardkodovane vrednosti. Primer funkcije sa hardkodovanom vrednošću može biti:

```js
function sayName(firstName, lastName) {
  return "John Doe";
}

sayName("Camper", "Cat");
```

Ova funkcija će vratiti `"John Doe"` bez obzira na argumente koje se proslede za parametre `firstName` i `lastName`, stoga je `"John Doe"` tretiran kao hardkodirana vrednost.

Deklariši parametar `sum` i postavi mu vrednost za poziv tvoje funkcije `addTwoNumbers`, koristeći `5` i `10` kao argumente. Ispiši parametar `sum` u konoli.

# --hints--

Trebalo bi da imaš funkciju nazvanu `addTwoNumbers`.

```js
assert.isFunction(addTwoNumbers);
```

Vaša funkcija za `addTwoNumbers` mora imati dva parametra.

```js
assert.lengthOf(addTwoNumbers, 2);
```

Tvoja funkcija mora da vrati zbir od dva parametra.

```js
assert.strictEqual(addTwoNumbers(5,10), 15);
```

Tvoja funkcija ne bi trebalo da vraća direktno postavljenu vrednost. To jest, treba da radi sa bilo kojim dva numerička argumenta.

```js
assert.strictEqual(addTwoNumbers(3, 5), 8);
```

Trebalo bi da objavite parametar `sum`.

```js
assert.isDefined(sum);
```

Vaš parametar za `sum` bi trebalo da ima vrednost `15`.

```js
assert.strictEqual(sum, 15);
```

Trebalo bi postaviti `sum` na vrednost koja proističe iz poziva funkcije `addTwoNumbers`, sa `5` i `10` kao argumentima.

```js
assert.match(code, /sum\s*=\s*addTwoNumbers\s*\(/);
```

Treba da napišete svoj parametar `sum`.

```js
assert.match(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--

--fcc-editable-region--

const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
