---
id: 60eebd07ea685b0e777b5583
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dobrodošli u projekat obrasca za registraciju! Počnite dodavanjem deklaracije `!DOCTYPE html` iznad dokumenta kako bi pretraživač znao koji je tip dokumenta koji čita.

# --hints--

Vaš kod treba da ima referencu na `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Trebalo bi da uključite razmak nakon reference na `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Trebalo bi da definišete tip dokumenta kao `html`.

```js
assert(code.match(/html/gi));
```

Morate zatvoriti deklaraciju za `DOCTYPE` za `>` nakon tipa.

```js
assert(code.match(/html\s*>/gi));
```

Tvoje saopštenje za `DOCTYPE` trebalo bi biti na početku tvog HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
