---
id: 60f030d388cb74067cf291c3
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Ispod naslova sekcije, koristite sledeći tekst unutar polja paragrafa kako biste podstakli korisnike da se registrušu:

```md
Please fill out this form with the required information
```

# --hints--

Treba da dodaš element `p` unutar `body`.

```js
assert.exists(document.querySelector('body > p'));
```

Trebalo bi da dodaš element `p` ispod `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

Trebalo bi mu dati atribut za tekstualni sadržaj `p` i `Please fill out this form with the required information`.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>

  </body>
--fcc-editable-region--
</html>
```
