---
id: 60f0286404aefb0562a4fdf9
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Dodaj element za `title` i `meta` unutar elementa `head`. Postavi naslov za svoj projekat koji je `Registration Form`, i dodaj atribut elementa `charset` sa vrednošću `utf-8` za svoj element `meta`.

# --hints--

Vaš kod mora imati karakteristiku `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tvoj element za `title` treba da bude unutar elementa `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Vaš projekat treba da ima naslov `Registration Form`. Imajte na umu, velika i mala slova zajedno sa pravopisom su važni za naslov.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Treba da kreiraš element `meta` unutar elementa `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Trebalo bi da uklonite atribut elementa `meta` sa svojstva elementa `charset` čija je vrednost `utf-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Vaš element za `meta` treba da bude neprekidni element, nema zatvarajući tag `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
