---
id: 60f02e7361b68405e27b62a5
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Unutar `body`, pružite kontekst zaglavlja sekcije za sadržaj, dodajući `h1` sa tekstom `Registration Form`.

# --hints--

Treba da ubaciš `h1` unutar `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Trebalo bi mu dati tekst `h1` `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
