---
id: 60f1a9cbd23023082e149fee
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Ovo je bolje. Sada, učinite temu prijatnijom za oči, menjanjem `body` `background-color` u `#1b1b32`. Zatim, da biste videli tekst, promenite `color` u `#f5f6f7`.

# --hints--

Trebalo bi da dodaš `background-color` unutar elementnog birača `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

Trebalo bi mu dati vrednost `background-color` od `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

Trebalo bi da dodaš `color` unutar elementarnog birača `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

Treba da mu date vrednost `#f5f6f7` za `color`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
