---
id: 60f1922fcbd2410527b3bd89
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--
Jedinica `vh` predstavlja visinu vidljivog područja, i jednako je 1% od `height` vidljivog područja. Ovo ga čini vezanim za visinu vidljivog područja.

Sada je vreme za stilizovanje projekta korišćenjem CSS. Počnite sa dodavanjem `body` `width` od `100%`, na `height` od `100vh`.

# --hints--

Trebalo bi da koristiš ekstrakter atributa `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Trebalo bi da mu/jajš (ili: trebao bi da mu/jajš) `body` `width` od `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Trebalo bi mu dati `body` `height` od `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
