---
id: 60f80e0081e0f2052ae5b505
title: Korak 19
challengeType: 0
dashedName: step-19
---

# --description--

Postavljanje atributa elementa `type` za element `input` je važno kako bi pretraživač znao kakvu vrstu podataka da očekuje. Ako `type` nije definisano, pretraživač će izabrati `text` kao podrazumevano.

Svojstvo elementa `input` sa vrednošću `type` za prva dva elementa `text`, svojstvo elementa `type` sa vrednošću `email` za treći element, i svojstvo elementa `type` sa vrednošću `password` za četvrti element.

Tipa `email` dozvoljava samo e-mailove sa `@` i `.` u domenu.
Tipa `password` skriva unos, i izdaje upozorenje ako sajt ne koristi HTTPS.

# --hints--

Trebalo bi da dodelite atribut elementu `input` sa vrednošću `type` za prvi element `text`.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

Morate dodeliti atribut elementu `input` sa vrednošću `type` za drugi element od `text`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

Trebalo bi da dodelite atribut elementu `input` sa vrednošću `type` za treći element `email`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

Trebalo bi da dodate atribut sa vrednošću `type` za četvrti element `password` na `input`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
        <label for="email">Enter Your Email: <input id="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
