---
id: 60f803d5241e6a0433a523a1
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--
Jedinica `rem` predstavlja koren `em`, i povezana je sa veličinom fonta elementa `html`.

Pošto su elementi `label` podrazumevano u liniji (inline), svi se prikazuju jedan pored drugog na istoj liniji, što čini njihov tekst teškim za čitanje. Da bi se pojavili na različitim linijama, dodajte `display: block` u element `label` i dodajte atribut `margin` sa vrednošću `0.5rem 0`, kako biste ih razdvojili.

# --hints--

Trebalo bi da koristiš alat `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

Trebalo bi dodati atribut `display` sa vrednošću `block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

Trebalo bi dodati atribut sa vrednošću `0.5rem 0` na `margin`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
