---
id: 60f81167d0d4910809f88945
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Prvi element `input` sa `type` od `submit` se direktno postavlja da prikaže element `form` najbližeg roditeljskog elemenata.

Da bi se obradio slanje forme, nakon elementa `fieldset` dodajte element `input` sa svojstvom `type` postavljeno na `submit` i sa svojstvom `value` postavljeno na `Submit`.

# --hints--

Trebalo bi da dodate element `input` nakon poslednjeg elementa `fieldset`.

```js
assert.exists(document.querySelectorAll('fieldset')?.[2]?.nextElementSibling?.tagName, 'input');
```

Trebalo bi mu dati atribut `input` svojstvo `type` sa vrednošću `submit`.

```js
assert.exists(document.querySelector('fieldset + input[type="submit"]'));
```

Trebalo bi mu dati atribut `input` svojstvo `value` sa vrednošću `Submit`.

```js
assert.exists(document.querySelector('fieldset + input[value="Submit"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>

--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
