---
id: 60f805f813eaf2049bc2ceea
title: Korak 17
challengeType: 0
dashedName: step-17
---

# --description--

Umestite element za `input` unutar svakog `label`. Proverite da dodate svaki `input` nakon teksta od `label`, i uključite razmak nakon dve tačke.

# --hints--

Treba da dodaš elemente iz `input` u element `fieldset`.

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

Trebalo bi da poređate sadržaj unutar elemenata `input` i unutar elemenata `label`.

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

Treba da dodaš prvi kod `input` nakon teksta `label` `Enter Your First Name:`, uključujući razmak nakon tačke.

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

Treba da dodaš `input` drugi nakon teksta `label` `Enter Your Last Name:`, uključujući razmak posle dve tačke.

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

Trebalo bi da dodate treći `input` nakon teksta `label` `Enter Your Email:`, uključujući razmak posle dve tačke.

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

Trebalo bi da dodate četvrti `input` nakon teksta `label` `Create a New Password:`, uključujući razmak nakon tačke.

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
  --fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
