---
id: 60fab8367d35de04e5cb7929
title: Korak 31
challengeType: 0
dashedName: step-31
---

# --description--

Umetnite tekst `I accept the terms and conditions` odmah nakon elementa `input` u `label` koji je nedavno dodan. Zatim spojite tekst `terms and conditions` sa sledećom lokacijom:

```md
https://www.freecodecamp.org/news/terms-of-service/
```

# --hints--

Trebalo bi da dodaš tekst `I accept the terms and conditions` nakon `label`, posle trećeg `fieldset`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label')?.innerText.trim(), 'I accept the terms and conditions');
```

Trebalo bi da koristiš element `a` za povezivanje sa uslovima i parametrima.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label a'));
```

Treba da umetnete novi tekst odmah nakon elementa za `input` u `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input + a'));
```

Trebalo bi mu dati atribut `a` `href` od `https://www.freecodecamp.org/news/terms-of-service/`.

```js
assert.match(document.querySelector('fieldset:nth-child(3) + label > input + a')?.href, /https:\/\/www\.freecodecamp\.org\/news\/terms-of-service\/?/);
```

Treba da kružiš samo oko elementa `a`, a ne oko teksta `terms and conditions`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent.trim(), 'terms and conditions');
```

Tekst unutar vašeg elementa sidrografa ima višak razmaka na početku ili na kraju. Razmak je dozvoljen samo u tekstu sidrografa između reči `terms`, `and`, i `conditions`.

```js
const nestedAnchor = document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent;
const innerContent = nestedAnchor?.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions"><input id="terms-and-conditions" type="checkbox" required /></label>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
