---
id: 60fad1cafcde010995e15306
title: Korak 42
challengeType: 0
dashedName: step-42
---

# --description--

Za slanje forme, važno je i dobra praksa da se navede atribut svakog elementa koji može biti poslat, atribut elementa `name`. Ovaj atribut elementa se koristi za identifikovanje elementa u slanju forme.

Osim dva ulaza za `radio` (koje ste već naveli), dostavite svaki element koji predstavlja jedinstvenu karakteristiku koju želite iz `name`.

# --hints--

Trebalo bi mu dati `input` očekivano ime prvog atributa `name`. _PS Izabrao bih `first-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

Trebalo bi mu dati `input` očekivano zadnje ime atributa elementa `name`. _PS Izabrao bih `last-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

Trebalo bi mu dati osobinu elementa `email` sa svojstvom `name`. _PS Izabrao bih `email`_

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

Trebalo bi mu dati atribut elementa `password` sa svojstvom `name`. _PS Izabrao bih `password`_

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

Trebalo bi mu dati atribut elementa `checkbox` sa svojstvom `name`. _PS Izabrao bih `terms`_

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

Trebalo bi mu dati atribut elementa `file` sa svojstvom `name`. _PS Izabrao bih `file`_

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

Trebalo bi mu dati osobinu elementa `number` od `name`. _PS Izabrao bih `age`_

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

Trebaš mu dati atribut elementu `select` sa atributom elementa `name`. _PS Izabrao bih `referrer`_

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

Trebalo bi mu dati atribut elementa `textarea` sa atributom elementa `name`. _PS Izabrao bih `bio`_

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

Nemoj dati bilo koji element od `option` atribut elemenata `name`.

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

Ne smete dodeliti bilo koji element od `label` svojstvo elementa `name`.

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

Ne smete da dodelite bilo koji element od `fieldset` sa svojstvom elementa `name`.

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
        <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset>
        <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file"/></label>
        <label for="age">Input your age (years): <input id="age" type="number" min="13" max="120" /></label>
        <label for="referrer">How did you hear about us?
          <select id="referrer">
            <option value="">(select one)</option>
            <option value="1">freeCodeCamp News</option>
            <option value="2">freeCodeCamp YouTube Channel</option>
            <option value="3">freeCodeCamp Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label for="bio">Provide a bio:
          <textarea id ="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
        </label>
      </fieldset>
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
      </label>
      --fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
