---
id: 60fabf0dd4959805dbae09e6
title: Korak 33
challengeType: 0
dashedName: step-33
---

# --description--

Dodajte još jedan `label` nakon prvog, sa tekstom `Input your age (years): `. Zatim, postavite unutar `input` sa `type` od `number`.

Sledeće, dodajte atribut za element `min` za `input` sa vrednošću `13` jer korisnici mlađi od 13 godina ne bi trebalo da se registruju. Takođe, korisnici možda neće preći starosnu granicu od 120 godina; dodajte atribut za element `max` sa vrednošću `120`.

Sada, ako korisnik pokuša da pošalje formular sa vrednošću izvan raspona, pojaviće se upozorenje i formular neće biti poslat. Isprobajte.

# --hints--

Treba dodati `label` na treći `fieldset`, nakon postojećeg `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

Trebalo bi mu dati tekst `label` `Input your age (years): `.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

Trebalo bi mu dodeliti `label` svojstvo elementa `input` sa `type` od `number`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

Elementu `input` treba da dodelite atribut `min` sa vrednošću `13`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

Elementu `input` treba da dodelite atribut `max` sa vrednošću `120`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        
      </fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
      </label>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
