---
id: 62cc5b1779e4d313466f73c5
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Postavi unutar elementa `link` unutar elementa `head`. Dodeli mu atribut/svojstvo `rel` sa vrednošću `stylesheet` i atribut/svojstvo `href` sa vrednošću `styles.css`.

# --hints--

Vaš kod mora imati element `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Ne smete da menjate vaše postojeće oznake za `head`. Proverite da niste izbrisali oznaku za zatvaranje.

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

Trebalo bi da imate jedan element `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Tvoje polje za `link` treba da bude prazno polje, ne bi trebalo da ima zatvarajuću oznaku `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Tvoj element `link` treba da bude unutar tvog elementa `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Vaš element od `link` treba da ima svojstvo atributa `rel` sa vrednošću `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Vaš element za `link` mora imati atribut elementa `href` sa vrednošću `styles.css`.

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
