---
id: 62ff8b9dab5ac88e4d3d43a3
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Prateći najbolje prakse pristupačnosti, povežite elemente od `input` i elemente od `label` koristeći svojstvo elementa `for`.

Koristi `first-name`, `last-name`, `email`, na `new-password` kao vrednosti za atribute `id`.

# --hints--

Prvi element `input` bi trebalo da ima `id` od `first-name`.

```js
assert(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

Drugi element od `input` treba da ima `id` od `last-name`.

```js
assert(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

Treći element `input` treba da ima `id` od `email`.

```js
assert(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

Četvrti element od `input` mora imati `id` iz `new-password`.

```js
assert(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

Prvi element `label` treba da ima svojstvo elementa `for` sa vrednošću `first-name`.

```js
assert(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

Drugi element `label` treba da ima svojstvo elementa `for` sa vrednošću `last-name`.

```js
assert(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

Treće polje `label` treba da ima svojstvo polja `for` sa vrednošću `email`.

```js
assert(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

Četvrti element od `label` mora imati poljsko svojstvo `for` sa vrednošću `new-password`.

```js
assert(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input /></label>
        <label>Enter Your Last Name: <input /></label>
        <label>Enter Your Email: <input /></label>
        <label>Create a New Password: <input /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
