---
id: 62ff8e998d3e7eae14d6ae3b
title: Korak 30
challengeType: 0
dashedName: step-30
---

# --description--

Morate potvrditi da je korisnik pročitao uslove i odredbe.

Dodaj element za `label`. Unutar novog elementa `label`, dodaj element `input` i postavi svojstvo elementa `type` na `checkbox`. Napravi ovaj element `input` da bude `required` kako bi korisnici ne mogli da se registruju bez pristanka uslovima i pravilima.

Dodaj `id` na svojstvo elementa `for` sa vrednošću `terms-and-conditions` za komponente/elemente za pristup.

# --hints--

Treba da dodaš `label` nakon trećeg elementa od `fieldset`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

Treba da dodaš `input` u element `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

Trebalo bi dodati atribut elementu `type` sa vrednošću `checkbox` na element `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

Trebalo bi da dodaš atribut elementa `required` na element `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

Element `input` treba da ima atribut elementa `id` sa vrednošću `terms-and-conditions`.

```js
assert(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

Element `label` mora imati svojstvo elementa `for` sa vrednošću `terms-and-conditions`.

```js
assert(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
