---
id: 63541ef4f96cd82e8e6c788a
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Atribut elementa `method` opisuje kako poslati podatke forme na URL, definisano u atributu elementa `action`. Podaci forme mogu se poslati putem zahteva `GET` kao parametri za URL (za `method="get"`) ili preko zahteva `POST` kao podaci u telu zahteva (za `method="post"`).

Postavite atribut elementa `method` ili pošaljite svoje podatke iz obrasca putem zahteva za `POST`.

# --hints--

Ne smete dodati novu funkciju za `form`.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

Tvoj element `form` bi trebalo da ima svojstvo elementa `method`.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

Tvoja svojstvo atributa `method` mora se postaviti na `post`.

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
