---
id: 7a97b79c096ecb000d03bd36
title: Implementacija kalkulatora za Nth Fibonacci broj
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Mora se kreirati funkcija nazvana `fibonacci`.
2. Treba definisati niz podataka nazvan `sequence` unutar funkcije `fibonacci`, i treba ga inicijalizovati sa vrednostima `[0, 1]`.
3. Funkcija `fibonacci` mora primiti jedan parametar, celobrojni neujem broj `n`.
4. Poziv `fibonacci(n)` treba da koristi metodu dinamičkog programiranja (dynamic programming) za izračunavanje i vraćanje `n`-ti broja iz Fibonacci sekvence, gde je svaki broj zbir prethodnih dve broja.
5. Svaki izračunati Fibonacci broj treba dodati u niz podataka `sequence`.

# --hints--

Trebalo bi da imaš funkciju koja se zove `fibonacci`.

```js
assert.isFunction(fibonacci);
```

Tvoja funkcija `fibonacci` treba da primi jedan parametar.

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

Trebalo bi da imaš niz podataka pod nazivom `sequence` unutar funkcije `fibonacci` definisane od strane `[0, 1]`.

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` trebalo bi da vrati `0`.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` treba vratiti `1`.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` treba da vrati `1`.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` treba da vrati `2`.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` treba da vrati `5`.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` treba vratiti `55`.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` treba da vrati `610`.

```js
assert.equal(fibonacci(15), 610);
```

Nije dozvoljeno koristiti recursion u vašem kodu.

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
