---
id: 686d749d6f32f708fdf2b7ad
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U prethodnoj lekciji, naučili ste o komponentama `svg` i kako se često koriste kao tagovi u projektima. U stvarnom kodu sveta, obično se oslanjate na biblioteke tagova, pa nema potrebe da kreirate komponente za `svg` od nule. Međutim, u ovoj radionici, kreiraćete srce oznaku kako biste naučili o osnovnim svojstvima korišćenim unutar komponente `svg`.

Počnite kreiranjem elementa `svg` na stranici.

# --hints--

Trebalo bi da imaš element `svg` na stranici.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
