---
id: 6551f6f9ec6db270eb83f3db
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Može se videti `11` ispisano na terminalu jer `'Hello World'` ima 11 karaktera.

Još jedna važna funkcija unutar je `type()`, koja vraća tip podataka parametra. Izmenite vaš poziv `print()` da ispišete tip podataka `text`.

# --hints--

Trebalo bi da pozoveš funkciju `type()`.

```js
assert.match(code, /type\s*\(.*\)/)
```

Treba da prođeš `text` kroz funkciju `type()` kako bi ga stavio između običnih zagrada.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

Treba da prođeš `type(text)` za funkciju `print()` kako bi ga stavio između običnih zagrada. Obrati pažnju na to da poziv funkcije stavi na početak linije.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Nemoj imati `print(len(text))` u svom kodu.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

Tvoj kod ima netačan sintaksa i/ili neispravan unos linija.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
