---
id: 6551fe3b1df7c9740f13f270
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Glavne tačke vezane za imenovanje parametara unutar Python su:

- Određene reči su rezervovane ključne reči (npr. `for`, `while`, `True`). One imaju posebna značenja u Python, pa ih ne možete koristiti kao imena parametara.
- Imena parametara ne mogu početi brojem i mogu sadržati samo slova, brojeve ili donje crte (underscores).
- Imena parametara su osetljiva na veličinu slova, tj. `my_var` je drugačije od `my_Var` i `MY_VAR`.
- Konačno, uobičajena praksa je pisanje imena parametara korišćenjem `snake_case`, gde se svaki razmak zamenjuje donjom crtom, a reči pišu se malim slovima.

Ukloni oba poziva za `print()` na atributu nazvanom `alphabet`. Postavite vrednost sekvence karaktera `'abcdefghijklmnopqrstuvwxyz'` za ovaj atribut.

# --hints--

Ne bi trebalo da imaš `print(type(text))` u svom kodu.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

Ne bi trebalo da imaš `print(type(shift))` u svom kodu.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/)
```

Trebalo bi da deklariraš parametar koji se zove `alphabet`. Obavezno stavi ime parametra na početak linije.

```js
assert.match(code, /^alphabet\s*=/m)
```

Trebalo bi da postavite vrednost sekvence karaktera `'abcdefghijklmnopqrstuvwxyz'` za vaš parametar `alphabet`. Zapamtite da koristite pojedinačne ili dvostruke navodnike za zatvaranje sekvence karaktera.

```js
assert.match(code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*(#.*)?$/m)
```

Vaš kod ima netačan sintaksa i/ili pogrešan ukazivač.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
