---
id: 6551ecf597038c36637f3759
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Parametri su važan deo u Python i bilo kom programskom jeziku. Parametar je ime koje se odnosi ili ukazuje na nešto. Možete deklarisati parametar tako što ćete napisati ime parametra na levoj strani operatera dodeljivanja `=` i navesti vrednost za taj parametar na desnoj strani operatera dodeljivanja:

```py
variable_name = value
```

Kreiraj parametar nazvan `number` i postavi vrednost `5` za tvoj novi parametar.

# --hints--

Trebalo bi da objavite parametar nazvan `number`. Obratite pažnju na postavljanje imena parametra na početak reda.

```js
assert.match(code, /^number\s*=/m)
```

Trebalo bi da postaviš vrednost `5` za svoj parametar `number`.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

Vaš kod ima neispravnu sintaksu i/ili neispravan unos reda.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
