---
id: 6551f5756c20146fc94f8675
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Možete dobiti broj karaktera u nizu slova koristeći funkciju koja se nalazi unutar `len()`.

Promeni svoj poziv od `print()` u prolazak `len(text)` umesto `text[-1]`.

# --hints--


Trebalo bi da pozoveš funkciju `len()`.

```js
assert.match(code, /len\s*\(.*\)/)
```

Trebalo bi da prosledite parametar `text` funkciji `len()` unutar običnih zagrada.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

Treba da prosledite `len(text)` funkciji `print()` za ubacivanje između običnih zagrada. Obratite pažnju da postavite poziv funkcije na početak linije.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

Ne smeš imati `print(text[-1])` u svom kodu.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

Tvoj kod ima neispravnu sintaksu i/ili neispravno poklapanje (matching).

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
