---
id: 6551f22653e6ce6e4f1d1370
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Promenite svoj poziv `print(shift)` ili unesite tip podataka vaše promene `shift`.

# --hints--

Trebalo bi da nastavite koristiti svoj postojeći poziv `print(type(text))`.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Trebalo bi proslediti `shift` funkciji `type()`.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

Trebalo bi da prođeš `type(shift)` kroz svoju funkciju `print()`. Obavezno pozovi funkciju na početku reda.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

Nemoj da imaš `print(shift)` u svom kodu.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

Tvoj kod ima neispravnu sintaksu i/ili pogrešno formatiranje.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
