---
id: 6551f3bfed58796ebb1268dc
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Svaka slovo sekvence slova može se referisati korišćenjem numeričkog indeksa. Izračunavanje indeksa počinje od nule. Stoga prvo slovo sekvence slova ima indeks `0`. Na primer, u sekvenci slova `'Hello World'`, `'H'` je na indeksu `0`, `'e'` je na indeksu `1`, i tako dalje.

Svaka slovo sekvence slova može se dobiti korišćenjem vitičastih zagrada. Morate napisati ime promenljive praćeno uglastim zagradama i dodati karakter slova unutar tih zagrada:

```py
text = 'Hello World'
r = text[8]
```

Sada, umesto štampanja `text`, samo štampaj slovo koje je u primeru `6`.

# --hints--

Trebalo bi da pozoveš funkciju `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Trebalo bi da proslediš `text[6]` funkciji `print()` kako bi ga stavio unutar običnih zagrada. Pažljivo pogledaj pozivanje funkcije na početku linije.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

Vaš kod ima neispravnu sintaksu i/ili greške u pisanju.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
