---
id: 6551fb531ce7f773071a0b84
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Kao što vidite, rezultati štampanja `type(text)` su `<class 'str'>`, što znači da je vaš parametar niz znakova, prikazan kao `str`.

Sada idite na novu liniju i kreirajte još jedan parametar nazvan `shift` i postavite vrednost `3` za ovaj parametar.

# --hints--

Treba da deklariš parametar koji se zove `shift`. Obrati pažnju na postavljanje imena parametra na početak linije.

```js
assert.match(code, /^shift\s*=/m)
```

Trebalo bi da postaviš vrednost `3` za svoj parametar `shift`.

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

Tvoj kod ima netačan sintaksa i/ili neispravan smernik.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
