---
id: 6551f4259d66026ee28cb205
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Takođe možete dobiti slova sekvence počevši od kraja te sekvence. Poslednje slovo ima uzorak za `-1`, drugo od kraja `-2` i tako dalje.

Sada promenite svoj postojeći poziv za `print()` da prikaže poslednji karakter u vašoj sekvenci karaktera.

# --hints--

Trebalo bi da pozoveš funkciju `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Treba da prosledite `text[-1]` funkciji `print()` kako biste ga stavili između običnih zagrada. Obratite pažnju da poziv funkcije postavite na početku linije.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

Vaš kod ima netačnu sintaksu i/ili nevažeću referencu.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
