---
id: 6551ef504d91936d2d4e54f8
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Argument je vrednost ili matematički izraz koji se prosleđuje funkciji — stavlja se između standardnih otvarajućih i zatvarajućih zagrada — kada se funkcija pozove:

```py
greet = 'Hello!'
print(greet)
```

Kod u gornjem primeru bi ispisao niz znakova `'Hello!'`, koji je vrednost parametra `greet` prosleđena za `print()` kao argument.

Prikažite svoj parametar od `text` na ekranu prosleđivanjem parametra `text` kao argument funkciji `print()`.

# --hints--

Treba da prosledite `text` funkciji `print()` kako biste dodali ime ovog parametra unutar standardnih otvarajućih i zatvarajućih zagrada. Zapamtite da poziv funkcije postavite na početak linije.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

Vaš kod ima netačnu sintaksu i/ili neispravnu referencu.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
