---
id: 6551f1e8737dbe6e0cb08ea4
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Sada objavi svoje nove promene.

# --hints--

Ne smete da uklonite svoj postojeći poziv od `print(type(text))`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

Trebalo bi da pozoveš funkciju `print()` prosleđujući argument `shift`. Obrati pažnju na to da poziv funkcije stavi na početak reda.

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

Vaš kod ima neispravnu sintaksu i/ili neispravan pokazivač/referencu.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
