---
id: 655208d59b131e7816f18c96
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Funkcija `print()` daje vam rezultate samo u konzoli, ali funkcije i metode mogu imati vrednost return koju možete koristiti u vašem kodu.

Sada postavi vrednost `alphabet.find(text[0])` u parametar nazvan `index`. Na ovaj način, `index` će sačuvati vrednost vraćenu od `alphabet.find(text[0])`.

# --hints--

Treba da deklariš parametar nazvan `index`. Pažljivo obratite pažnju na postavljanje imena parametra na početak linije.

```js
assert.match(code, /^index\s*=/m)
```

Trebalo bi da postaviš vrednost return koju pruža `alphabet.find(text[0])` u svoj parametar `index`.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

Tvoj kod ima netačne sintaksa i/ili neispravan ulaz linija.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
