---
id: 6552487e689f6e895f658717
title: Korak 33
challengeType: 20
dashedName: step-33
---

# --description--

Sada treba da kreirate parametar `new_char` na kraju vašeg ciklusa. Postavite mu vrednost `alphabet[new_index]`.

# --hints--

Treba da kreiraš parametar `new_char` unutar ciklusa `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_char\s*=/));
```

Trebalo bi da postaviš svoj parametar `new_char` na `alphabet[new_index]` na kraju svog ciklusa.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_char\s*=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
--fcc-editable-region--
```
