---
id: 65524d2a1a253b8bb5197ae2
title: Korak 28
challengeType: 20
dashedName: step-28
---

# --description--

`find` vraća ponovo `-1` sa velikim slovima, i takođe sa prostorom za pisanje. To mesto ćete obraditi kasnije.

Trenutno, umesto ponavljanja na `text`, promeni ciklus `for` da se ponavlja na `text.lower()`.

# --hints--

Trebalo bi da promeniš svoj ciklus `for` na onaj koji se ponavlja preko `text.lower()` umesto `text`.

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
