---
id: 6552008c0d9d9075cbec9772
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Koristi putanju `.find()` da pronađeš mesto u alfabetu svake slova u svojoj poruci. Putanja je kao funkcija, ali je svojstvo objekta.

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

Gore, metoda `.find()` se izvršava *na* `sentence` (sekvenca pretražujućih karaktera), a `'r'` (karakter za pretragu) se prosleđuje kao argument. Poziv `sentence.find('r')` vraća `4`, koji je primer prvog događaja/incidenata u `'r'` unutar `sentence`.

Nakon vašeg koda, pozovite `.find()` na `alphabet` i prosledite `'z'` kao argument na taj način.

# --hints--

Trebalo bi da pozoveš putanju `.find()`.

```js
assert.match(code, /find\s*\(.*\)/)
```

Treba da pozoveš metodu `.find()` na parametru `alphabet`.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```

Treba da izvršiš poziv metode `.find()` u parametar `alphabet` i prosledi `'z'` tom metodom. Zapamti da poziv metode stavi na početak linije.

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```

Tvoj kod ima netačan sintaksa i/ili neispravnu referencu.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
