---
id: 655261b2e1f2c197093f3993
title: Korak 41
challengeType: 20
dashedName: step-41
---

# --description--

Trenutno, prostor se piše kao `'c'`. Da bi se očuvao originalni razmak u običnoj poruci, bi Vam bio potreban izraz `if` uslova. Ovo je sastavljeno od glavne reči `if`, uslova i dve tačke `:`.

```py
if x != 0:
    print(x)
```

U gornjem primeru, uslov izjave za `if` je `x != 0`. Kod `print(x)`, unutar tela izjave za `if`, dešava se samo kada je uslov potvrđen da je `True` (u ovom primeru, njegovo značenje `x` je različito od nule).

Na samom mjestu u vašem ciklusu for, promijenite `print(char == ' ')` u izjavu `if`. Uslov za ovu izjavu `if` mora biti potvrđen da je `True` ako je `char` prazno i `False` inače. Unutar tijela `if`, ispišite sekvencu karaktera `'space!'`. Zapamtite da umetanje ovu liniju na tačnim mestom.

# --hints--

Ne smete imati `print(char == ' ')` u vašem kodu.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

Moraš da zameniš `print(char == ' ')` sa izrazom `if` koji se aktivira kada je `char == ' '`. Nemoj koristiti obične zagrade za obuhvat uslova `if` i zapamti da dodaš dve tačke na kraju.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

Treba da ispuniš niz znakova `'space!'` unutar svog novog izraza `if`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

Vaš kod ima neispravnu sintaksu i/ili unos neispravnih linija.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
