---
id: 655220953ba90d80514d7ee2
title: Korak 27
challengeType: 20
dashedName: step-27
---

# --description--

Trenutno, funkcija `print()` prima jedan argument `char`, ali može primiti više argumenata, razdvojena zarezima.

Dodaj drugi znak za `print(char)`, ili ću ispisati slovo i njegov simbol unutar alfabeta.

# --hints--

Moraš dodati ``index`` kao drugi argument za vaš postojeći poziv ``print(char)``. Ne zaboravi da razdvojiš argumente zarezom.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
