---
id: 65521fc818947e800bffe48a
title: Korak 26
challengeType: 20
dashedName: step-26
---

# --description--

Unutar petlje `for`, pre štampanja trenutnog karaktera, deklariš parametar nazvan `index` i postavi mu vrednost koja se vraća sa `alphabet.find(char)` u ovaj parametar.

# --hints--

Treba da deklariš novo polje pod nazivom `index` na početku svog ciklusa `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

Trebalo bi da staviš `alphabet.find(char)` u svoj novi parametar `index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

Vaš kod ima netačnu sintaksu i/ili nevažeću referencu.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
