---
id: 65520f8282faf57a0db4f7fe
title: Korak 21
challengeType: 20
dashedName: step-21
---

# --description--

Kao što vidite iz rezultata, `'h'` je u obrascu `7` u sekvenci znakova `alphabet`. Sada je potrebno dobiti karakter koji se nalazi u obrascu `7` zajedno sa vrednošću `shift`. Za to, možete koristiti operater za sabiranje, `+`, na isti način kako ga koristite za dodavanje broja.

Promeni svoju razliku od `shifted` ili sačuvaj vrednost `alphabet` u polje `index + shift`.

# --hints--

Trebalo bi da postaviš vrednost `alphabet[index + shift]` u svoju varijablu `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

Tvoj kod ima netačan sintaksa i/ili pogrešan format.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
