---
id: 65522255d5b9cd80f335c6f2
title: Korak 29
challengeType: 20
dashedName: step-29
---

# --description--

Na kraju vašeg ciklusa, deklariš parametar nazvan `new_index` i postavi vrednost `index + shift` za ovaj parametar.

# --hints--

Treba da deklarišete parametar nazvan `new_index` unutar vaše funkcije od `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

Treba da staviš `index + shift` kao parametar za svoju novu petlju na kraju tela svog ciklusa od `for`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
