---
id: 65525e359ca28d938baa82c5
title: Korak 39
challengeType: 20
dashedName: step-39
---

# --description--

Možete dobiti efekat isti kao `a = a + b` korišćenjem operatera sabiranja:

```py
a += b
```

Operater za povećanje omogućava vam da povećate vrednost varijable, nakon čega možete dodeliti inkrement toj varijabli.

Koristite operaciju `+=` da povećate vrednost i postavite procenat vrednosti istovremeno za `encrypted_text`.

# --hints--

Trebalo bi da koristiš operater za dodavanje (`+`) na vrednost ``alphabet[new_index]`` sa trenutnom vrednošću ``encrypted_text`` unutar tela svog ciklusa.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
