---
id: 655251308f31958d06cdf267
title: Korak 36
challengeType: 20
dashedName: step-36
---

# --description--

Trenutno, skriveni karakter se ažurira svaki put u svakom ciklusu. Bolje je sačuvati skriveni niz u novi parametar. Pre vašeg ciklusa `for`, deklarišite parametar nazvan `encrypted_text` i postavite vrednost praznog niza (`''`) za ovaj parametar.

# --hints--

Trebalo bi da objaviš parametar nazvan `encrypted_text` pre svog ciklusa `for`.

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

Treba da postaviš praznu vrednost niza karaktera za svoj parametar `encrypted_text`.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
