---
id: 655209a4a27dd37873c4cac3
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

`.find()` vraća obrazac odgovarajućeg slova unutar niza karaktera. Ako slovo nije pronađeno, vraća `-1`. Kao što vidite, prvo slovo u `text`, veliko slovo `'H'`, nije pronađeno, jer `alphabet` ima samo mala slova.

Možete promeniti sekvencu karaktera tako da odgovara malim slovima koristeći `.lower()`. Dodajte još jedan poziv `print()` da bi ispisao `text.lower()` i videli rezultate.

# --hints--

Trebalo bi da još uvek imaš `print(index)` u svom kodu. Obrati pažnju na pozivanje funkcije na početku linije.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Treba da odštampate `text.lower()`. Obratite pažnju na pozivanje funkcije na početku reda.

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

Vaš kod ima netačnu sintaksu i/ili neispravnu referencu.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
