---
id: 65520e6f2b9678799977f24d
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Ukloni poslednji poziv `print()`. Zatim, umesto `text[0]`, prosledite `text[0].lower()` kao argument svom pozivu `.find()` i vidite rezultate.

# --hints--

Ne bi trebalo da imaš `print(text.lower())` u svom kodu.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

Još moraš objaviti svoju razliku od `index`.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

Treba da ispraviš svoj poziv ``alphabet.find(text[0])`` ili koristi ``text[0].lower()`` kao argument. Obrati pažnju da staviš poziv metode na početak linije.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

Tvoj kod ima netačan sintaksa i/ili neispravan ulaz linija.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
