---
id: 65521203d7165c7b84b22ad4
title: Korak 23
challengeType: 20
dashedName: step-23
---

# --description--

Ciklus vam omogućava da prođete kroz uređeni niz elemenata i izvršite akcije za svaki od njih.

U ovom slučaju, koristićeš ciklus od `for`. Evo kako možeš da ponoviš na `text`:

```py
for i in text:
```

`for` je glavna reč koja ukazuje na tip kola. `i` je parametar koji uzima vrednost komponenti u `text` serijski. Rečenica se završava dve tačke, `:`.

Ispod linije gde ste objavili `alphabet`, napišite petlju `for` koja ponavlja preko `text`. Koristite `i` kao parametar petlje.

Stoga, postoji greška u terminalu. Naučićete o tome u sledećem koraku.

# --hints--

Moraš koristiti ključnu reč `for` da kreiraš ciklus. Osiguraj da staviš ključnu reč `for` na početak linije i ostavi prazan razmak nakon ključne reči.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

Trebalo bi da napišete parametar `i` nakon glavne reči `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

Treba da napišeš glavnu reč `in` nakon `for i `. Proveri da ostavi razmak blizu glavne reči `in`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

Trebaće ti da napišeš `text` nakon `for i in `. Nemoj zaboraviti dodati poslednji `:`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

Vaš ciklus za `for` treba da bude postavljen ispod linije koda `alphabet = 'abcdefghijklmnopqrstuvwxyz'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
