---
id: 6553a572f7a65718f1e42e18
title: Korak 49
challengeType: 20
dashedName: step-49
---

# --description--

Funkcija je deo koda koji se može ponovo koristiti. Već ste naučili neke ugrađene funkcije, kao što su `print()`, `find()` i `len()`. Ali možete deklarisati i prilagođene funkcije kao što je ovo:

```py
def function_name():
    <code>
```

Deklaracija funkcije počinje sa ključnom rečju `def` praćenu imenom funkcije — važeće ime parametra — i običnim zagradama. Deklaracija se završava dvotačkom tačkom.

Nakon vašeg šablona od `shift`, deklarišite funkciju nazvanu `caesar` i napišite sve sledeće linije da mu date telo vaše funkcije.

# --hints--

Trebalo bi da koristiš ključnu reč `def` za deklarisanje nove funkcije.

```js
assert.match(code, /^def\s+/m)
```

Treba da napišete `caesar` kao ime funkcije nakon glavne reči `def`. Zapamtite da dodate razmak nakon `def`.

```js
assert.match(code, /^def\s+caesar/m)
```

Treba da dodate standardne zagrade nakon imena funkcije. Ne zaboravite na zadnje dve tačke.

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

Trebalo bi da uklonite sve linije nakon `shift = 3` kako bi to bilo telo vaše nove funkcije.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
