---
id: 6553ed69ece88d29594748aa
title: Korak 52
challengeType: 20
dashedName: step-52
---

# --description--

Da bi se izvršila, funkcija mora da bude pozvana (ili aktivirana) dodavanjem para standardnih zagrada nakon njenog imena, kao što je ovo:

```py
function_name()
```

Kraj vašeg koda, ovo je poziv vaše funkcije za `caesar`. Obratite pažnju na poravnanje linija.

# --hints--

Trebalo bi da pozoveš svoju funkciju `caesar`. Proveri da li je poziv funkcije na početku reda.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
