---
id: 6553a7d8d05cbb1ae335a665
title: Korak 45
challengeType: 20
dashedName: step-45
---

# --description--

Kada ciklus dosegne karakter `Z`, ukupna vrednost `index + shift` prelazi poslednji model/primer sekvence karaktera `alphabet`. Stoga, `alphabet[new_index]` pokušava da koristi nevažeći model, što uzrokuje da se `IndexError` izbaci.

Možete videti da rezultati na terminalu se prikazuju u prostoru odmah pre `Z`, `print`, poslednji pre bacanja greške.

U ovom slučaju, modulo operator (`%`) može se koristiti za vraćanje ostatka deljenja između dva broja. Na primer: `5 % 2` je jednak `1`, jer 5 podeljeno sa 2 ima kvocijent od 2 i ostatak 1.

Obuhvati `index + shift` običnim zagradama, zatim primeni modulo na matematički izraz sa `26`, koji je dužina `alphabet`.

# --hints--

Treba da imaš `new_index = (index + shift) % 26` u svom `else`.

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

Vaš kod ima nevažeću sintaksu ili netačan pokazivač.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
