---
id: 6553980e0527fa115c705646
title: Korak 42
challengeType: 20
dashedName: step-42
---

# --description--

Sada, umesto štampanja `'space!'`, koristi operator dodavanja za povećanje (koji je trenutno sačuvan u `char`) na trenutnu vrednost `encrypted_text`.

# --hints--

Ne smete imati `print('space!')` u vašem kodu.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

Trebalo bi da koristiš operater `+=` da dodaš `char` na trenutnu vrednost `encrypted_text` unutar tvog novog deklarisanja `if`.

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

Tvoj kod ima netačan sintaksa i/ili neispravan unos linija.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
