---
id: 6553f8c570f9982e013a8886
title: Korak 57
challengeType: 20
dashedName: step-57
---

# --description--

Trenutno, svako slovo se zadržava sa istim slovom u svakoj prilici, u skladu sa deklarisanom udaljenošću. Šta ako bi taj pomeraj bio drugačiji za svako slovo? To bi bilo mnogo teže dešifrovati. Ovaj algoritam se naziva Vigenere šifra, gde se pomeraj za svako slovo određuje drugim tekstom, poznatim kao *ključ*.

Počni da konvertuješ svoj Caesar šifar u Vigenère šifar uklanjanjem poziva dve funkcije.

# --hints--

Treba da uklonite vaš poziv dve funkcije za `caesar()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
