---
id: 6553efd6ada3f42aa2d75448
title: Korak 53
challengeType: 20
dashedName: step-53
---

# --description--

Sada bi trebalo da vidite rezultate ponovo. Iako ova metoda radi, ne optimizuje korišćenje koda više i više. Pozivanje vaše funkcije često će dati isti rezultat. Međutim, funkciju se može deklarisati sa *parametarima* kako bi se postigla raznolikost i varijabilnost:

```py
def function_name(param_1, param_2):
    <code>
```

Parametri su varijabilni parametri koje možete koristiti unutar vaše funkcije. Funkcija može biti deklarisana sa različitim brojem parametara. U prethodnom primeru, `param_1` i `param_2` su parametri.

Izmeni deklaraciju svoje funkcije da prihvati dva parametra nazvana `message` i `offset`.

Nakon toga, videće grešku koja se pojavljuje u terminalu. Videćete kako da je rešite u sledećim koracima.

# --hints--

Funkcija vašeg `caesar` treba da primi `message` i `offset` kao parametre. Zapamtite da razdvojite parametre zarezom.

```js
assert.match(code, /^def\s+caesar\s*\(\s*message\s*,\s*offset\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar():
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
```
