---
id: 6553f4f66099802c6ae94613
title: Korak 55
challengeType: 20
dashedName: step-55
---

# --description--

Trenutno, vaš kod vraća `TypeError`, jer je funkcija `caesar` definisana sa dve parametre (`message` i `offset`), stoga očekuje da bude pozvana sa dva argumenta.

Pozvan sa `caesar()` bez potrebnih parametara za zaustavljanje izvršavanja koda.

Pružiti vrednosti za `message` i `offset` pri prenosu `text` i `shift` kao argumente u pozivu funkcije `caesar`.

# --hints--

Trebalo bi da proslediš `text` i `shift` kao argumente u tvom pozivu funkcije, stavivši ih unutar običnih zagrada. Ne zaboravi da razdvojiš argumente zarezom.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
