---
id: 6553995f412dd8122ed38e4a
title: Korak 43
challengeType: 20
dashedName: step-43
---

# --description--

Uslovna rečenica takođe može imati deo ``else``. Ovaj deo se može dodati na kraj rečenice ``if`` kako bi se izvršio alternativni kod ako je uslov rečenice ``if`` netačan:

```py
if x != 0:
    print(x)
else:
    print('x = 0')
```

Kao što vidite u svojim rezultatima, kada ciklus petlje dostigne poziciju, dodaje se pozicija na kodirani niz karaktera. Zatim se kod iz bloka `if` izvršava i `c` dodaje se na kodirani niz karaktera.

Za izmenu, dodajte segment `else` nakon `encrypted_text += char` i postavite sav sledeći kod u interni redosled, osim poziva `print()`.

# --hints--

Treba da kreiraš deo `else`. Ne zaboravi da uključiš dve tačke na kraju.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/else\s*:/));
```

Treba da postavite kod u interni redosled nakon vašeg odeljka `else`, osim poziva na `print()`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*new_index\s*=\s*index\s*\+\s*shift\s*encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
```

Tvoj kod ima neispravnu sintaksu i/ili pogređo interno formatiranje.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
