---
id: 6553a755879b131a445e664c
title: Korak 44
challengeType: 20
dashedName: step-44
---

# --description--

Pokušajte da postavite vrednost niza slova `'Hello Zaira'` u vašoj varijabli `text` i vidite šta se dešava na terminalu.

Videćeš grešku za `string index out of range`. Ne brini, uskoro ćeš saznati kako da je popraviš!

# --hints--

Trebalo bi da imaš razliku od `text`.

```js
assert.match(code, /^text\s*=/m)
```

Treba da unesete vrednost niza karaktera `'Hello Zaira'` u vašu varijablu `text`.

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
