---
id: 6554ad2463b8892748f8efdd
title: Korak 67
challengeType: 20
dashedName: step-67
---

# --description--

Trenutno, tvoja funkcija ispisuje neke nizove znakova, ali ove vrednosti ne mogu se koristiti od strane drugih delova koda za izvršavanje bilo kakvih radnji.

Za tu svrhu, morate koristiti deklaraciju `return`:

```python
def foo():
    return 'spam'
```

Potrebno je napisati ``return`` praćen razmakom i vrednošću koju funkcija mora da vrati. Kada se pronađe deklaracija ``return``, ta vrednost se vraća i izvršavanje funkcije se zaustavlja, nastavljajući sa sledećim redom koda nakon poziva funkcije. U gornjem primeru, funkcija ``foo`` vraća niz karaktera ``'spam'``.

Uklonite dva poziva od `print()` iz vaše funkcije i vratite `encrypted_text`.

# --hints--

Trebalo bi da uklonite dva poziva `print()` iz vaše funkcije.

```js
({ 
  test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/print\s*\(\s*("|')plain\stext:\1\s*,\s*message\s*\)/.test(commentless_code))
    assert.isFalse(/print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/.test(commentless_code))
  }
})
```

Tvoja funkcija mora da vrati `encrypted_text`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+message\.lower\s*\(\s*\)\s*/);
    const regex = /return\s+encrypted_text/
    assert(function_body.match(regex));
    assert.notMatch(block_body, regex);
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
