---
id: 6554a334a40edb1fb4eff827
title: Korak 62
challengeType: 20
dashedName: step-62
---

# --description--

Sledeće, unutar bloka `else`, deklaruj parametar nazvan `key_char` i postavi mu vrednost na `key` za izraz `key_index` mod(`%`) dužine `key`.

# --hints--

Treba da deklariš parametar nazvan `key_char` na početku tela tvog bloka `else`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");
assert(block_body.match(/^\s*key_char\s*=/));
```

Trebalo bi da staviš `key[key_index % len(key)]` za svoj parametar `key_char`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");    
assert(block_body.match(/key_char\s*=\s*key\s*\[\s*key_index\s*%\s*len\s*\(\s*key\s*\)\s*\]/));
```

Vaš kod ima netačnu sintaksu i/ili nevažeću referencu.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--        
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
--fcc-editable-region--        
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
