---
id: 6554e0adc7bb193cbfdb36d5
title: Korak 73
challengeType: 20
dashedName: step-73
---

# --description--

Proverite da li funkcija može konvertovati sekvencu karaktera nazad u običan tekst.

Deklariši još jedan parametar nazvan `decryption` i postavi vrednost za `vigenere(encryption, custom_key, -1)`.

# --hints--

Treba da pozoveš `vigenere` kada prosleđuješ `encryption`, `custom_key` i `-1` kao argument.

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/) })
```

Trebalo bi da deklariš parametar `decryption`.

```js
({ test: () => assert(__userGlobals.has("decryption")) })
```

Vaš parametar za `decryption` mora imati vrednost `vigenere(encryption, custom_key, -1)`.

```js
({ test: () => assert.match(code, /^decryption\s*=\s*vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--    
encryption = vigenere(text, custom_key, 1)
print(encryption)

--fcc-editable-region--
```
