---
id: 6554a88d5af937226f4a9121
title: Korak 65
challengeType: 20
dashedName: step-65
---

# --description--

Putanja za `.index()` je potpuno identična putanji `.find()`, ali izbacuje izuzetak `ValueError` ako ne može pronaći podnizu slova.

Greška za `ValueError` je interna greška koja se baca kada argument sa ispravnim tipom, ali neispravnom vrednošću, prođe kroz funkciju.

Nakon dodavanja `key_index`, deklaruj parametar nazvan `offset`. Pronađi model koji `key_char` ima u `alphabet` i postavi njegovu vrednost u `offset`. Koristi `.index()` da dobiješ taj model.

# --hints--

Trebalo bi da deklariš parametar nazvan `offset`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=/));
```

Tvoj parametar za `offset` treba da skladišti vrednost od `alphabet.index(key_char)`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=\s*alphabet\.index\s*\(\s*key_char\s*\)/));
```

Vaš kod ima netačnu sintaksu i/ili nevažeću referencu.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
            
--fcc-editable-region--            
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
