---
id: 655491bd5b98b813fa5bedca
title: Korak 58
challengeType: 20
dashedName: step-58
---

# --description--

Sada promeni deklaraciju svoje funkcije: promeni ime funkcije u `vigenere` i drugi parametar na `key`.

# --hints--

Trebalo bi da promeniš ime svoje funkcije u `vigenere`.

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

Vaša funkcija od `vigenere` treba da uzme `message` i `key` kao parametre.

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
