---
id: 65549561463f0016876e852c
title: Korak 60
challengeType: 20
dashedName: step-60
---

# --description--

Pošto je tvoj ključ kraći od teksta koji moraš sakriti, morać ćeš ga ponoviti kako bi generisao ceo skriveni tekst. Na početku tela svoje funkcije, deklariraj varijablu `key_index` i postavi joj vrednost nula.

# --hints--

Trebalo bi da deklarišete parametar nazvan `key_index` na početku tela vaše funkcije.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
assert(function_body.match(/^\s*key_index\s*=/));
```

Trebalo bi da postaviš vrednost `0` za svoj parametar `key_index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");   
assert(function_body.match(/key_index\s*=\s*0/));
```

Vaš kod ima netačnu sintaksu i/ili nevažeću referencu.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
