---
id: 6554e343caea913ffba7bec6
title: Korak 75
challengeType: 20
dashedName: step-75
---

# --description--

Sada, tvoja funkcija se može koristiti za šifrovanje i dešifrovanje poruka. Očisti svoj kod koristeći bolje imena parametara.

Promeni svaku instancu `encrypted_text` u `final_message`.

# --hints--

Treba da zameniš svaku instancu `encrypted_text` sa `final_message`.

```js
({
  test: () => {
    const l = code.match(/final_message/g).length
    assert.equal(l, 4)
  }
})
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'
--fcc-editable-region--
def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
