---
id: 6554a49a4f782f208abcc87e
title: Korak 63
challengeType: 20
dashedName: step-63
---

# --description--

Biće vam potrebno povećati broj `key_index` za sledeći ciklus. Da biste to uradili, nakon linije koju ste dodali i unutar istog dela koda, koristite inkrementalni operater da povećate `key_index` za jedan.

# --hints--

Trebalo bi da koristite operator `+=` da dodate `1` na `key_index` unutar tela `else` bloka.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");
assert(block_body.match(/key_char\s*=\s*key\s*\[\s*key_index\s*%\s*len\s*\(\s*key\s*\)\s*\]\s*key_index\s*\+=\s*1/));
```

Tvoj kod ima neispravnu sintaksu i/ili neispravno poklapanje (matching).

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
            key_char = key[key_index % len(key)]
--fcc-editable-region--
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
