---
id: 65549f90cf78131c96ebcf28
title: Korak 61
challengeType: 20
dashedName: step-61
---

# --description--

Kada pišete kod, čitljivost je veoma važna. Komentari se koriste kao efikasne napomene koje objašnjavaju logiku iza vašeg koda. Ovo je važno kada se vratite na projekat nakon dužeg vremena i takođe pomaže vašim kolegama da razumeju taj kod.

U Python, možete napisati komentare koristeći `#`. Sve što sledi `#` nije izvršeno.

Pre vašeg izjave o `if`, dodajte komentar koji kaže `Append space to the message`.

# --hints--

Treba da dodaš komentar koji kaže `Append space to the message`. Ne zaboravi da napišeš znak `#` na početku.

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
