---
id: 6554930320d70414e7b6acc6
title: Korak 59
challengeType: 20
dashedName: step-59
---

# --description--

Izbriši svoju varijablu `shift`. Zatim, deklariši novu varijablu pod nazivom `custom_key` i postavi vrednost niza karaktera `'python'` za ovu varijablu.

# --hints--

Trebalo bi da obrišeš varijablu `shift` i njenu vrednost.

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

Treba da deklarišete varijablu pod nazivom `custom_key`.

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

Treba da postaviš vrednost niza karaktera `'python'` za svoju varijablu `custom_key`.

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
