---
id: 6555e1bfcd374e18c6be8e58
title: Korak 91
challengeType: 20
dashedName: step-91
---

# --description--

U Python, postoji jednostavan način za formatiranje stringa. *f-strings* vam omogućavaju da ubacite vrednosti unutar vašeg niza karaktera.

Umetanje vrednosti znači pisanje zamenuvarača koji će biti zamenjeni navedenim vrednostima kada se program pokrene. Na primer, možete dobiti isti rezultat za `'Encrypted text: ' + text` i `f'Encrypted text: {text}'`. Morate postaviti `f` pre navodnika kako biste kreirali f-string i napisali parametre ili matematički izraz koji želite da umetnete unutar zagrada.

Promeni prvi poziv za `print()` ili koristi f-string.

# --hints--

Morate promeniti poziv `print('Encrypted text: ' + text)` ili ispisati isti rezultat koristeći f-string.

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Encrypted\stext:\s\{\s*text\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
