---
id: 6555231eac4d9f19bd3d44b5
title: Korak 83
challengeType: 20
dashedName: step-83
---

# --description--

Glavna reč `pass` može se koristiti kao predstavnik za budući kod. Nema nikakvog efekta u tvom kodu, ali te može spasiti od grešaka koje dobijaš tamo gde kod nije završen:

```py
def foo():
    pass
```

Pozivanje `vigenere` i `1` za kodiranje i `-1` za dekodiranje je u redu, ali može biti malo teško razumeti. Kreiraj novu funkciju nazvanu `encrypt` koja prima parametre `message` i `key`, i koristi `pass` da popuniš telo te funkcije.


# --hints--

Trebaće vam da deklarišete novu funkciju koja se zove `encrypt`. Zapamti da koristiš ključnu reč `pass` unutar njenog tela.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(encrypt)
  `))
})
```

Tvoja funkcija `encrypt` mora da primi parametre `message` i `key`. Red je važan.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(encrypt)
    str(sig) == "(message, key)"
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--

encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
