---
id: 655619327c7b364166f8dd6f
title: Korak 46
challengeType: 20
dashedName: step-46
---

# --description--

Ako želite da unesete više karaktera u niz karaktera ``alphabet``, kao što su brojevi ili specijalni znakovi, otvorićete da je neophodno ručno promeniti desni operand modulo operacije.

Promeni `26` u `len(alphabet)` ili izbegni ovaj problem.

# --hints--

Treba da promeniš vrednost `new_index` menjanjem `26` i `len(alphabet)`.

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
  assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/));
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % 26
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
