---
id: 655a2a7210094920069b117c
title: Korak 40
challengeType: 20
dashedName: step-40
---

# --description--

Operatori poređenja omogućavaju vam da uporedite dve stvari prema njihovim vrednostima. Možete koristiti operator poređenja postavljanjem između stvari koje želite da uporedite.
Oni vraćaju vrednost *tačno ili netačno* — tj. `True` ili `False` — u zavisnosti od tačnosti matematičkog izraza.

Python ima sledeće operacije poređenja:

<table>
  <thead>
    <tr>
      <th>Opereta</th>
      <th>Maana</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>Usawa</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>Si sawa</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>Kubwa kuliko</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>Ndogo kuliko</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>Kubwa kuliko au sawa na</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>Ndogo kuliko au sawa na</td>
    </tr>
  </tbody>
</table>

Na početku vašeg kruga, ispišite rezultate poređenja `char` i poziciju (`' '`). Koristite operator jednakosti `==` za to.

# --hints--

Trebalo bi da uporediš `char` sa pozicijom koristeći operator jednakosti unutar svog ciklusa `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

Trebalo bi da ispisaš rezultate poređenja `char` i polje unutar tvoje petlje `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

Trebalo bi da ispišete rezultate poređenja `char` na početku vašeg ciklusa.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
