---
id: 656886ec4e4d9b96ded3c565
title: Korak 25
challengeType: 20
dashedName: step-25
---

# --description--

Varijabla za petlju može imati bilo koje važeće ime, ali je lakše koristiti smisleno ime.

Promeni ime svoje parametre `i` u `char`.

# --hints--

Trebalo bi da zamenite `i` sa `char`.

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

Vaš kod ima nevažeću sintaksu ili netačan pokazivač.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
