---
id: 65688a50e6c998a21d8e41d3
title: Korak 30
challengeType: 20
dashedName: step-30
---

# --description--

Niz znakova ne može se promeniti, što znači da se ne može modifikovati nakon kreiranja. Na primer, možete misliti da sledeći kod menja vrednost `my_string` u niz znakova `'train'`, ali ovo nije važeće:

```py
my_string = 'brain'
my_string[0] = 't'
```

Potvrdite ovo koristeći zapisivanje u uglastim zagradama da biste dobili prvo slovo u `text` i pokušajte ga promeniti u slovo koje birate. Videćete da su rezultati nestali i pojaviće se greška.

# --hints--

Trebalo bi da još uvek imaš `text = 'Hello World'` u svom kodu.

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

Trebalo bi da dobiješ prvi slovo u `text` koristeći `text[0]`.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

Trebalo bi da koristiš atribut ``=`` da postavi vrednost slova koje biraš za ``text[0]``. Nemoj zaboraviti da taj slovo obuhvati pojedinačnim ili dvostrukim navodnicima.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
