---
id: 65688b5a1655a7a6caede847
title: Korak 31
challengeType: 20
dashedName: step-31
---

# --description--

Kada pokušate da promenite pojedinačna slova niza znakova kao što ste uradili u prethodnom koraku, dobijate `TypeError`, što se dešava kada se neadekvatan tip koristi u vašem kodu.

Kao što vidite iz poruke o grešci, niz znakova ne može dobiti indeks elementa jer je statičan. Međutim, parametar može dobiti drugi niz znakova:

```py
message = 'Hello World'
message = 'Hello there!'
```

Obriši red od `text[0]` na postavljanje vrednosti `text` sekvencu slova `'Albatross'`.

# --hints--

Ne bi trebalo da imaš `text[0] = 'C'` u svom kodu.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

Treba da postaviš vrednost `text` sledeva slova `'Albatross'`. Ne menjaj `text = 'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
