---
id: 65687d2f8c7ee27b0446cef3
title: Korak 24
challengeType: 20
dashedName: step-24
---

# --description--

Kod koji se izvršava svaki put — postavljen nakon `:` — jeste telo petlje. Ovaj kod mora biti formatiran razmakom. U Python, savetuje se koristiti 4 razmaka za svaki nivo uvlačenja. Ovaj nivo uvlačenja je deo koda.

```py
for i in text:
    print(i)
```

Python zavisi od prostora za prikaz delova koda. Vertikalna tačka na kraju linije je znak da će uslediti novi deo koda ako je napisan sa razmakom.

Stoga, kada nema definisanog kodnog segmenta za razmak nakon poslednje tačke, izvršavanje koda se zaustavlja i `IndentationError` je bačeno. Ovaj kod neće prikazati rezultate, već će umesto toga baciti `IndentationError`:

```py
for i in text:
print(i)
```

Dostavite telo za vaš ciklus od `for` da dodate poziv `print(i)`. Zapamtite da napišete telo ciklusa na mestu.

# --hints--

Trebalo bi dodati `print(i)` u telo svog kruga od `for`. Dajte prioritet korišćenom prostoru.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

Tvoj kod ima netačan sintaksa i/ili neispravan razmak.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
