---
id: 6568b7a40c7ba0ccbc4b4425
title: Korak 32
challengeType: 20
dashedName: step-32
---

# --description--

Kao što vidite, svaki karakter niza koji je ponovo dodeljen vrednost ispisuje se unutar petlje.

Vratite se početnoj sekvenci karaktera da biste izbrisali opsež vrednosti `text`.

# --hints--

Ne smije biti `text = 'Albatross'` u vašem kodu.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*=\s*("|')Albatross\1/.test(commentless_code))
```

I dalje je potrebno imati parametar za `text` sa vrednošću `'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /text\s*=\s*("|')Hello World\1/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text = 'Albatross'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
