---
id: 659526f3e5424b3eac4b6aa7
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Prikažite vrednost promene za `index` u konzoli.

# --hints--

Treba da objaviš svoje promene za `index`. Obrati pažnju na to da poziv funkcije stavi na početak linije.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Tvoj kod ima netačan sintaksa i/ili neispravan unos linija.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
