---
id: 65951b637ddec63611c482c0
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Možete koristiti funkciju koja se nalazi unutar `print()` da odštapiše rezultate vašeg koda na terminalu.

Funkcije su delovi koda koji se mogu ponovo koristiti koje možete pozvati ili pokrenuti kako biste izvršili svoj kod kada je potrebno. Da biste pozvali funkciju, samo morate napisati par standardnih zagrada pored njenog imena. Vi ćete saznati više o ovim funkcijama uskoro.

Trenutno, idite na novu liniju i dodajte prazan poziv za funkciju `print()`. Ne bi trebalo da vidite nikakve rezultate još uvek.

# --hints--

Trebalo bi da imaš `print()` u svom kodu. Obrati pažnju na pozivanje funkcije na početku linije.

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

Tvoj kod ima neispravnu sintaksu i/ili neispravno razmakovanje.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
