---
id: 65952b70924edb449faf1950
title: Korak 19
challengeType: 20
dashedName: step-19
---

# --description--

Deklariši novi parametar nazvan `shifted`. Upotrebite zagrade da dobijete vrednost `alphabet` za primer `index` i stavi tu vrednost u svoj novi parametar.

# --hints--

Treba da objaviš parametar koji se zove `shifted`. Obrati pažnju na to da postaviš ime parametra na početak reda.

```js
assert.match(code, /^shifted\s*=/m)
```

Treba da postaviš vrednost `alphabet[index]` u svoj parametar `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

Tvoj kod ima netačne sintakse i/ili pogrešno razmakovanje.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
