---
id: 659530ca823d6c47f61daa8a
title: Korak 20
challengeType: 20
dashedName: step-20
---

# --description--

Izdaj svoju razliku od `shifted`.

# --hints--

Trebalo bi da ispišete vašu razliku od `shifted`. Pažljivo pogledajte kako postaviti poziv funkcije na početku reda.

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

Tvoj kod ima netačan sintaksa i/ili neispravan unos linija.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
