---
id: 659db3e3670d3e712be82593
title: Korak 22
challengeType: 20
dashedName: step-22
---

# --description--

Ponavljanje procesa pronalaženja slova unutar alfabeta i identifikovanje slova koje je pomereno za svaki karakter u `text` može biti iscrpljujući zadatak. Srećom, možete ga učiniti lakšim korišćenjem petlje.

Trenutno, izbriši sve linije koda ispod deklaracije parametra `alphabet`.

# --hints--

Trebalo bi da još uvek imaš `text = 'Hello World'` u svom kodu.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

Trebalo bi da i dalje imaš `shift = 3` u svom kodu.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

Trebalo bi da još uvek imaš `alphabet = 'abcdefghijklmnopqrstuvwxyz'` u svom kodu.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

Trebalo bi da obrišete parametar `index` sa njegovom vrednošću.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

Trebalo bi da imaš `print(index)` u svom kodu.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

Trebalo bi da obrišete parametar `shifted` sa njegovom vrednošću.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

Trebalo bi da imaš `print(shifted)` u svoj kod.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

Vaš kod ima neispravnu sintaksu i/ili greške u pisanju.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
