---
id: 6624e0ec0bf47017eed84b1f
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Prvi tip šifrovanja (cipher) koji napraviš zove se *Caesar* cipher. Konkretno, uzmećeš svaki slovo u svojoj poruci, pronaći njegovo mesto u alfabetu, uzeti slovo koje je na mestu 3 posle u alfabetu, i zameniti originalno slovo sa novim slovom.

Da biste ovo izvršili, koristićete metodu ``.find()`` koja je diskutovana u prethodnom koraku. Izmenite vaš postojeći poziv od ``.find()`` dajući mu ``text[0]`` kao argument umesto ``'z'``.

# --hints--

Trebaće ti da promeniš svoj postojeći poziv `alphabet.find('z')` i zameniš ga sa `text[0]`.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
