---
id: 66605a64a9822f01ec5ae8ce
title: Korak 50
challengeType: 20
dashedName: step-50
---

# --description--

U Python, opseg parametra određuje gde taj parametar može biti dostupan:

- Parametri definisani van funkcije imaju *globalni opseg* i mogu se pristupiti iz bilo kog dela vašeg koda.

- Parametri definisani unutar funkcije su *interni* za tu funkciju i ne mogu se pristupiti izvan nje.

Da biste videli ovo, pokušajte da ispišete parametar ``alphabet`` na kraju vašeg koda. Ovo će izazvati grešku od ``NameError``.

Trebalo bi da vidite poruku o grešci koja pokazuje da `alphabet` nije definisan. To je zato što je `alphabet` definisan unutar funkcije `caesar` i nedostupan izvan nje.

# --hints--

Trebate pokušati da izpišete parametar `alphabet` van Caesar funkcije.

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
