---
id: a5229172f011153519423690
title: Izračunavanje zbora neobičnih Fibonaccijevih brojeva
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

U ovom laboratoriju ćete napraviti kalkulator za zbir Fibonaccijevih brojeva koji računa zbir svih Fibonaccijevih brojeva manjih ili jednakih datom pozitivnom broju.

**Cilj**: Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da postoji funkcija `sumFibs` koja prima broj kao argument.
1. Funkcija `sumFibs` treba da vrati sumu svih neparnih brojeva Fibonaccija koji su manji ili jednaki datoj broju.
1. Fibonacci niz počinje sa `0` i `1`, a svaki sledeći broj je zbir dva prethodna broja.
1. Samo neparni brojevi Fibonaccija će biti dodati u sumu.

# --hints--

Trebalo bi da imaš funkciju `sumFibs`.

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` trebalo bi da vrati broj.

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` treba da vrati `1785`.

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` treba da vrati `4613732`.

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` trebalo bi da vrati `5`.

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` trebalo bi da vrati `60696`.

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` treba da vrati `135721`.

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
