---
id: 5688e62ea601b2482ff8422b
title: Kreiranje pretrage profila
challengeType:  26
dashedName: lab-profile-lookup
---

# --description--

U ovoj laboratoriji, kreiraćeš pretragu profila koja traži informacije o osobama u listi kontakata.

Za ovaj primer, zamisli da postoji osoba za komunikacije po imenu Akira Laine, `lookUpProfile("Akira", "lastName")`, koja treba da vrati/pošalje `Laine`.

**Cilj:** Popunite priče korisnika ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Treba kreirati funkciju nazvanu `lookUpProfile` koja prima ime i atribut kao argumente.
2. Treba dobiti kontakt informacije iz niza podataka `contacts` koji je dostavljen.
3. Ako funkcija primi ime kontakta i atribut postoji za tu osobu, onda se vrednost tog atributa mora vratiti.
4. Ako ime poslato funkciji ne odgovara nikome u nizu podataka kontakata, funkcija mora vratiti `"No such contact"`.
5. Ako atribut ne postoji za pronađenog kontakta, funkcija mora vratiti `"No such property"`.

# --before-each--

```js
const _contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

contacts = structuredClone(_contacts); 
```

# --hints--

Red podataka za `contacts` bi trebalo da postoji u globalnom opsegu. Vrati funkciju kako biste ga ponovo dobili ako ste ga obrisali.

```js
assert.isArray(contacts);
```

Trebalo bi da imaš funkciju nazvanu `lookUpProfile`.

```js
assert.isFunction(lookUpProfile);
```

`lookUpProfile("Kristian", "lastName")` Trebalo bi da vrati niz znakova.

```js
assert.isString(lookUpProfile("Kristian", "lastName"));
```

`lookUpProfile("Kristian", "lastName")` treba da vrati sekvencu karaktera `Vos`

```js
assert.strictEqual(lookUpProfile('Kristian', 'lastName'), 'Vos');
```

`lookUpProfile("Sherlock", "likes")` treba da vrati `["Intriguing Cases", "Violin"]`

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")` trebalo bi vratiti red podataka

```js
assert.isArray(lookUpProfile('Harry', 'likes'));
```

`lookUpProfile("Bob", "number")` treba da vrati sekvencu karaktera `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'number'),'No such contact');
```

`lookUpProfile("Bob", "potato")` treba vratiti niz karaktera `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'potato'), 'No such contact');
```

`lookUpProfile("Akira", "address")` treba da vrati niz karaktera `No such property`

```js
assert.strictEqual(lookUpProfile('Akira', 'address'), 'No such property');
```

Podatni red za `contacts` treba da ostane nepromenjen nakon pokretanja `lookUpProfile`

```js
lookUpProfile("Sherlock", "likes");
assert.deepEqual(contacts,_contacts);
```

# --seed--

## --seed-contents--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


```

# --solutions--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
