---
id: 612e8279827a28352ce83a72
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Sada kopiraj set od sedam elemenata iz `.key`, zatim zalijepi još dva seta unutar div-a `.keys`.

# --hints--

Trebalo bi da imaš 21 element od `.key`.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 21);
```

Trebalo bi da imaš 15 elemenata za `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 15);
```

Svi elementi iz `.key` moraju biti unutar tvog elementa `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.key')?.length === 21);
```

Svi elementi iz `.black--key` moraju biti unutar tvog elementa `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.black--key')?.length === 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
