---
id: 612e89d254fe5d3df7d6693d
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Sada ste ponovo inicijalizovali model kontejnera za `html`, takođe morate da ga prosledite komponentama unutar njega. Da biste to uradili, možete postaviti svojstvo `box-sizing` na `inherit`, što će ciljanim komponentama reći da koriste istu vrednost kao roditeljsko svojstvo.

Takođe ćete morati da ciljate na fiktivne elemente, koji su specifične ključne reči koje prate izablon. Dva fiktivna elementa koja ćete koristiti su `::before` i `::after`.

Selektor `::before` kreira pseudo-element koji je prvo dete izabranog elementa, dok selektor `::after` kreira pseudo-element koji je poslednje dete izabranog elementa. Ovi pseudo-elementi se često koriste za kreiranje estetskog sadržaja, koji ćete kasnije videti u ovom projektu.

Trenutno, kreiraj selektor za CSS koji cilja sve elemente sa `*`, i uključi lažne elemente od `::before` i `::after`. Postavi atribut za `box-sizing` na `inherit`.

# --hints--

Trebalo bi da imate birač za `*, ::before, ::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Opcija vašeg `*, ::before, ::after` mora imati atribut `box-sizing` ako je postavljeno na `inherit`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing === 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
