---
id: 612e83ec2eca1e370f830511
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Dodaj element `link` unutar tvog elementa `head`. Dodeli mu atribut elementa `rel` postavljen na `stylesheet` i atribut elementa `href` postavljen na `styles.css`.

# --hints--

Vaš kod mora imati element `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Ne smete da menjate vaše postojeće oznake za `head`. Proverite da niste izbrisali oznaku za zatvaranje.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Trebalo bi da imate jedan element `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Tvoj element `link` treba da bude unutar tvog elementa `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Vaš element od `link` treba da ima svojstvo atributa `rel` sa vrednošću `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Vaš element za `link` mora imati atribut elementa `href` sa vrednošću `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
