---
id: 612e78af05201622d4bab8aa
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Vreme je da počnete sa radom na pijanou. Kreiraj element `div` unutar svog elementa `body`, sa `id` podešenim na `piano`.

# --hints--

Trebalo bi da kreirate novu komponentu za `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

`div` treba da bude unutar vašeg elementa od `body`.

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

`div` bi trebalo da ima `id` postavljeno na `piano`.

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
