---
id: 612ea4c4993aba52ab4aa32e
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Sada je vreme za korišćenje selektora veštačkog elementa koje ste prethodno pripremili. Da biste kreirali ključeve za crnu boju, dodajte novi selektor `.key.black--key::after`. Ovo će ciljati elemente sa klasom `key black--key`, i izabrati veštački element koji sledi ove elemente u HTML.

U novom izgrađivaču, postavite atribut `background-color` na `#1d1e22`. Takođe postavite atribut `content` na `""`. Ovo će učiniti veštačke elemente praznim.

Atribut ``content`` se koristi za postavljanje ili brisanje sadržaja elementa. Podrazumevano, lažni elementi kreirani pomoću selektora za lažne elemente ``::before`` i ``::after`` su prazni, pa se ništa neće prikazati na stranici. Postavljanje atributa ``content`` na prazan niz znakova ``""`` osiguraće da je element prikazan na stranici, a da i dalje bude prazan.

Ako želite da pokušate, pokušajte da uklonite atribut `background-color` i postavite različitu vrednost za atribut `content`, kao što je `"♥"`. Zapamtite da vratite ove promene kada završite kako bi testovi prošli.

# --hints--

Treba ti da imaš birač za `.key.black--key::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Tvoj selektor za `.key.black--key::after` mora imati svojstvo od `background-color` postavljeno na `#1d1e22`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

Vaš selektor za `.key.black--key::after` mora imati svojstvo `content` postavljeno na `""`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
