---
id: 612e813b3ba67633222cbe54
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Zapamtite da atribut elementa `class` može imati više vrednosti. Da biste razdvojili vaše bele i crne tastere, dodaćete drugu vrednost `class` od `black--key`. Dodajte ovo na vaše elemente drugi, treći, peti, šesti i sedmi od `.key`.

# --hints--

Tvoj drugi element za `.key` takođe bi trebalo da ima `class` od `black--key`.

```js
const key = document.querySelectorAll('.key')?.[1];
assert(key?.className?.includes('black--key'));
```

Treći vaš element od `.key` mora da ima `black--key` u `class`.

```js
const third = document.querySelectorAll('.key')?.[2];
assert(third?.classList?.contains('black--key'));
```

Tvoj peti element od `.key` mora imati `black--key` u `class`.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert(fifth?.classList?.contains('black--key'));
```

Šesta karakteristika vašeg `.key` mora biti `black--key` u `class`.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert(sixth?.classList?.contains('black--key'));
```

Tvoj sedmi element od `.key` treba da ima `black--key` u `class`.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert(seventh?.classList?.contains('black--key'));
```

Treba ti pet komponenti od `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 5);
```

Trebalo bi da ima sedam komponenti `.key`.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
