---
id: 612e9d142affc44a453655db
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Vreme je da se ukrasi sami ključevi. Kreiraj selektor za `class` za elemente `.key`. Postavi skup `background-color` na vrednost `#ffffff`, svojstvo `position` za `relative`, svojstvo `width` za `41px`, i svojstvo `height` za `175px`.

# --hints--

Trebalo bi da imaš birač za `.key`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key'));
```

Tvoj selektor za `.key` mora imati svojstvo od `background-color` postavljeno na `#ffffff`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor === 'rgb(255, 255, 255)');
```

Vaš selektor za `.key` mora imati svojstvo `position` postavljeno na `relative`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.position === 'relative');
```

Tvoj selektor za `.key` mora imati svojstvo `width` postavljeno na `41px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.width === '41px');
```

Vaš [selektor] za `.key` mora imati svojstvo `height` postavljeno na `175px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.height === '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
