---
id: 612e804c54d5e7308d7ebe56
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Unutar tvog elementa od `.keys`, dodaj sedam komponenti od `div`. Dodeli sve u klasu objekata od `key`.

# --hints--

Treba ti kreirati sedam novih elemenata za `div`.

```js
const divDivDiv = document.querySelectorAll('div');
assert(divDivDiv?.length === 9);
```

Sedam novih elemenata vašeg `div` treba da budu unutar vašeg elementa `.keys`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].length === 7);
assert([...keys?.children].every(child => child?.tagName === 'DIV'));
```

Sedam tvojih novih elemenata za `div` treba da imaju `class` postavljeno na `key`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].every(child => child?.classList?.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
