---
id: 612e7d1c29fb872d6384379c
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Postavi unutar `div`, drugi deo unutar `div` koji imaš, i postavi `class` da bude `keys`.

# --hints--

Treba da kreiraš drugi deo `div`.

```js
const divDiv = document.querySelectorAll('div');
assert(divDiv?.length === 2);
```

Tvoja nova komponenta za `div` treba da bude unutar tvog postojećeg elementa koji je definisan sa `div`.

```js
const div = document.querySelector('div');
assert(div?.children?.length === 1);
assert(div?.children?.[0]?.localName === 'div');
```

Vaša nova funkcija od `div` mora imati `class` postavljeno na `keys`.

```js
const div = document.querySelector('div');
assert(div?.children?.[0]?.className === 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
