---
id: 612e96fc87fe8e44f69f7ec5
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Vreme je za konfigurisanje ključeva. Pod pravilom `#piano`, postavite element za `.keys` za selektor `class`. Dodajte novo pravilo sa svojstvom `background-color` sa vrednošću `#040404`, svojstvom `width` sa vrednošću `949px` i svojstvom `height` sa vrednošću `180px`.

# --hints--

Trebalo bi da imaš birač za `.keys`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys'));
```

Vaš selektor za `.keys` mora imati svojstvo od `background-color` postavljeno na `#040404`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.backgroundColor === 'rgb(4, 4, 4)');
```

Vaš selektor za `.keys` mora imati svojstvo `width` postavljeno na `949px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.width === '949px');
```

Tvoj selektor za `.keys` mora imati svojstvo od `height` postavljeno na `180px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.height === '180px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
