---
id: 612e6afc009b450a437940a1
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Počni sa osnovnim dizajnom HTML. Dodaj deklaraciju `DOCTYPE` na elemente `html`, `head`, `body`, i `title`.

Postavite jezik ove stranice na engleski. Zamenite `title` sa `Piano`.

# --hints--

Vaš kod treba da ima referencu na `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Trebalo bi da uključite razmak nakon reference na `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Trebalo bi da definišete tip dokumenta kao `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Trebalo bi da zatvorite deklaraciju `DOCTYPE` sa etiketom za zatvaranje `>` nakon tipa.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Tvoj element `html` treba da ima labelu za otvaranje sa svojstvom elementa `lang` od `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Element vašeg `html` treba da ima zatvarajući tag.

```js
assert(code.match(/<\/html\s*>/));
```

Tvoje saopštenje za `DOCTYPE` trebalo bi biti na početku tvog HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Trebalo bi da imaš etiketu za otvaranje `head`.

```js
assert(code.match(/<head\s*>/i));
```

Trebalo bi da ima oznaku za zatvaranje `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Trebalo bi da imaš etiketu za otvaranje `body`.

```js
assert(code.match(/<body\s*>/i));
```

Trebalo bi da ima oznaku za zatvaranje `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Elementi za `head` i `body` bi trebalo da budu sestrinski elementi.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Element `head` treba da bude unutar elementa `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Element za `body` treba da bude unutar elementa `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Vaš kod mora imati karakteristiku `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Vaš projekat treba da ima zaglavlje `Piano`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

Zapamtite, velika i mala slova zajedno sa pravopisom su važne za naslov.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
