---
id: 612e89562043183c86df287c
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Pretraživači mogu koristiti ugrađene vrednosti ivice i unutrašnji prostor za specifične elemente. Da biste osigurali da vaša stranica izgleda konzistentno, morate ispraviti kutni model.

Dodaj selektor za `html` u svoj fajl od CSS, i postavi atribut `box-sizing` na `border-box`.

# --hints--

Trebalo bi da imaš birač za `html`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

Vaš selektor za `html` mora imati svojstvo `box-sizing` ako je podešen na `border-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
