---
id: 612e77aba7ca691f598feb02
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Dodajte dve oznake za `meta`, jedna da poboljša vašu stranicu za mobilne uređaje, i jedna da identifikuje `charset` koja je prihvatljiva za tu stranicu.

# --hints--

Treba da dodaš dva elementa `meta` na svoju stranicu.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Trebalo bi da ima dva elementa za `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Jedan element od `meta` mora imati svojstvo elementa `name` koje je postavljeno na `viewport`, i `content` ako je postavljeno na `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Drugi element od `meta` mora imati svojstvo elementa `charset` postavljeno na `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
