---
id: 687d8c6f066877768f7ee441
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

U vašoj klasi `LinkedList`, zamenite ključnu reč `pass` unutrašnjom klasom nazvanom `Node`. Ovo će predstavljati pojedinačne čvorove koji skladište podatke u listi.

# --hints--

You should create a `Node` class within the `LinkedList` class. Remember to add the `pass` keyword inside the `Node` class.

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
