---
id: 687d7deea6a2a22fa8816a73
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U prethodnim lekcijama, naučili ste o nekim uobičajenim strukturalima podataka. U ovoj radionici, implementiraćete spajanu listu.

Prvo, kreirajte klasu pod nazivom `LinkedList` da bi predstavljala strukturu koja će držati čvorove vaše liste. Dodajte ključnu reč `pass` unutar klase kako bi test mogao da prođe.

# --hints--

You should create a class named `LinkedList`. Remember to add the `pass` keyword.

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
