---
id: 687d8e05a529a4838f1c961c
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Zatim, uklonite ključnu reč `pass` unutar klase `Node`, zatim dodelite klasi `Node` metod `__init__` sa dva parametra, `self` i `element`.

Unutar metode, dodelite `element` na `self.element`. Ovo skladišti vrednost koja se prosleđuje prilikom kreiranja čvora.

# --hints--

You should define an `__init__` method inside your `Node` class.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

Your `__init__` method should have two parameters, `self` and `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

You should assign `element` to `self.element` inside the `__init__` method.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
