---
id: 687e366b3004a9831e4cbd41
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Dodajte metod nazvan `is_empty` koji vraća `self.length == 0`.

Koristićete ovo da pratite koliko je čvorova u listi. Dakle, sa `self.length == 0`, vraća `True`, što znači da lista nema čvorove.

Ako je `self.length` veće od `0`, vraća `False`, što znači da lista ima najmanje jedan čvor.

# --hints--

You should create an `is_empty` method that returns `self.length == 0`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    
--fcc-editable-region--
```
