---
id: 688003f589b875af60273dd9
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

I dalje unutar metode `__init__`, dodeli `None` na `self.next`. Ovo postavlja pokazivač `next` na `None`, što znači da novi čvor još uvek ne pokazuje na bilo koji drugi čvor u listi.

Kasnije, `next` se može ažurirati da pokaže na drugi čvor i stvori lanac povezanog spiska.

# --hints--

Inside the `__init__` method, you should initialize the `next` attribute to `None`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
