---
id: 68800a64b06e80d9a31c131b
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Još u metodi `__init__` klase `LinkedList`, postavić `self.head` na `None`. Ovo služi kao referenca do prvog čvora u listi. Lista je trenutno prazna, pa moraš da ga postaviš na `None`.

# --hints--

You should set `self.head` to `None` inside the `__init__` method of the `LinkedList` class.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    def __init__(self):
        self.length = 0
--fcc-editable-region--
```
