---
id: 68857e0a49baf4751f737a2a
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Sada, nakon metode `is_empty`, kreiraj metodu nazvanu `add` unutar klase `LinkedList`. Nova metoda treba da ima dva parametra, `self` i `element`.

Unutar vaše nove metode, kreiraj varijablu nazvanu `node` i dodeli joj novu instancu od `Node` koristeći `self.Node(element)`. Ovo će kreirati novi čvor koji treba da se doda u spajanu listu.

# --hints--

You should create an `add` method inside the `LinkedList` class.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

Your `add` method should have two parameters, `self` and `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

You should assign `self.Node(element)` to `node` in the `add` method.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
