---
id: 688649c8673f842f49228a5e
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Ako lista nije prazna, metoda bi trebalo da pronađe poslednji čvor i poveže novi čvor sa njegovim pokazivačem `next`.

Da biste to uradili, kreirajte blok `else` nakon bloka `if`. Unutar bloka `else`, kreirajte varijablu `current_node` i dodelite joj `self.head`.

# --hints--

You should assign `self.head` to `current_node` inside the `else` block.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
