---
id: 68b84df2648b4c96b25022f3
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Vreme je da testirate svoju do sada napravljenu spajanu listu! Na dnu vašeg koda, kreirajte instancu `LinkedList` nazvanu `my_list`. Zatim, koristeći metodu `is_empty` koju ste kreirali u poslednjem koraku, ispišite `my_list.is_empty()`.

# --hints--

You should create an instance of `LinkedList` called `my_list`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

You should print `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
