---
id: 68b9e951e4892721632acc26
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Odlično! Sada je vreme da testirate dodavanje elemenata u vašu spajanu listu. Na dnu vašeg koda, nakon kreiranja instance `LinkedList` nazvane `my_list` i štampanja `my_list.is_empty()`, dodajte `1` i `2` u listu koristeći metod `add`.

Zatim, ponovo štampajte `my_list.is_empty()` da biste potvrdili da lista više nije prazna. Na kraju, štampajte `my_list.length` da vidite trenutnu dužinu spajane liste.

# --hints--

You should add `1` to the list using the `add` method.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

You should add `2` to the list using the `add` method.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

You should print `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

You should print `my_list.length`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
