---
id: 66e2d680e129e1423116a541
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, napravićete analizator rečenica koji uzima rečenicu i dobija broj riječi, glasokorisnih slova, suglasnika i interpunkcijskih znakova.

Počni, kreiraj funkciju za `getVowelCount` koja ima parametar nazvan `sentence`. Tvoja funkcija treba da vrati ukupan broj samoglasnika u rečenici.

# --hints--

Moraš da kreiraš funkciju za `getVowelCount`.

```js
assert.isFunction(getVowelCount);
```

Vaša funkcija za `getVowelCount` treba da ima parametar `sentence`.

```js
assert.match(getVowelCount.toString(), /sentence/);
```

Tvoja funkcija za `getVowelCount` treba da vrati broj.

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

Kada je rečenica `"Apples are tasty fruits"`, funkcija `getVowelCount` bi trebalo da vrati `7`.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

Kada je rečenica `"Hello, World!"`, funkcija bi trebalo da vrati `3`.

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

Brojanje vaših samoglasnika treba da bude neosvisljivo na velika ili mala slova.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

Tvoja funkcija `getVowelCount` treba da vrati tačan broj vokala za bilo koju rečenicu.

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
