---
id: 66e2e5d890ac8a54c9d703bd
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Nije vreme za brojanje konsonanti. Kreiraj funkciju `getConsonantCount` sa parametrom `sentence`.

Unutar funkcije, koristite petlju da izračunate broj suglasnika u `sentence` koji će biti proslijeđen funkciji kada se pozove. Suglasnik je bilo koji slovo koje nije među sljedećim slovima: `"aeiou"`.

Tvoja funkcija za `getConsonantCount` mora da vrati broj.

# --hints--

Trebalo bi da kreirate funkciju za `getConsonantCount`.

```js
assert.isFunction(getConsonantCount);
```

Tvoja funkcija za `getConsonantCount` mora imati parametar od `sentence`.

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

Tvoja funkcija za `getConsonantCount` treba da vrati broj.

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

Kada je rečenica `"Coding is fun"`, funkcija `getConsonantCount` treba da vrati `7`.

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

Kada je rečenica `"Hello, World!"`, funkcija `getConsonantCount` bi trebalo da vrati `7`.

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

Brojanje suglasnika bi trebalo da bude nezavisno od veličine slova.

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

Tvoja funkcija ``getConsonantCount`` treba da vrati tačan broj suglasnika za bilo koju rečenicu.

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
