---
id: 66e2edb23013ca596f8d9568
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Sada je vreme da testirate vašu funkciju `getPunctuationCount`.

Kreiraj parametar za `punctuationCount` i postavi mu vrednost kao rezultat pozivanja funkcije `getPunctuationCount` sa argumentom `"WHAT?!?!?!?!?"`

Nakon toga, napišite sledeće u konzolu: `"Punctuation Count: [Punctuation count goes here]"`. Promenite `[Punctuation count goes here]` u stvarni naziv parametra. Možete izabrati da koristite tekstualni niz šablona ili spojiti tekstualne nizove koristeći operator `+` ovde.

# --hints--

Trebalo bi da kreirate parametar `punctuationCount`.

```js
assert.isNotNull(punctuationCount)
```

Vaš parametar za `punctuationCount` treba biti postavljen na rezultat `getPunctuationCount("WHAT?!?!?!?!?")`.

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

Trebalo bi da napišete sledeće u konzolu: `"Punctuation Count: [Punctuation count goes here]"`. Zamenite `[Punctuation count goes here]` sa stvarnim imenom parametra. Uverite se da ovde koristite ispravnu sintaksu za konkatenaciju niza karaktera.

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
