---
id: 66e2f06b191e305c00574e4d
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Na kraju, izračunaj broj reči za kreiranje funkcije `getWordCount` koja ima parametar `sentence`. Ta funkcija treba da vrati ukupan broj reči u poslatu rečenicu kada se pozove.

# --hints--

Treba da kreiraš funkciju za `getWordCount`.

```js
assert.isFunction(getWordCount);
```

Tvoja funkcija za `getWordCount` mora da ima parametar od `sentence`.

```js
assert.match(getWordCount.toString(), /sentence/);
```

Tvoja funkcija `getWordCount` treba da vrati broj.

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

Kada je rečenica `"When are you gonna start learning to code?"`, funkcija za `getWordCount` bi trebalo da vrati `8`.

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

Kada je rečenica `"What's going on?"`, funkcija za `getWordCount` bi trebalo da vrati `3`.

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

Brojanje vaših reči trebalo bi da bude nezavisno od velikih ili malih slova.

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

Vaša funkcija za `getWordCount` treba da vrati tačan broj reči za bilo koju rečenicu.

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both front-end and back-end development"), 7);
```

Tvoja funkcija za `getWordCount` treba da vrati tačan broj reči za prazan niz znakova, ili niz koji sadrži samo razmake.

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
