---
id: 66e2eab8a5638f57b637b7cc
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Sada moraš da izbrojiš broj znakova.

Kreiraj funkciju za `getPunctuationCount` sa parametrom `sentence`.

Unutar te funkcije, napravite petlju za brojanje broja znakova za pisanje u `sentence` koji će biti prosleđeni funkciji kada se funkcija pozove. Za naše svrhe, znak za pisanje je bilo koji karakter koji nije razmak (`" "`) ili obična slova.

Tvoja funkcija za `getPunctuationCount` mora da vrati broj.

# --hints--

Trebalo bi da kreirate funkciju za `getPunctuationCount`.

```js
assert.isFunction(getPunctuationCount);
```

Vaša funkcija za `getPunctuationCount` bi trebalo da ima parametar `sentence`.

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

Vaša funkcija `getPunctuationCount` treba da vrati broj.

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

Kada je rečenica `"What's going on here?"`, funkcija koja `getPunctuationCount` treba vratiti `2`.

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

Kada je rečenica `"What????!"`, funkcija `getPunctuationCount` treba da vrati `5`.

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

Vaša funkcija `getPunctuationCount` treba da vrati tačan broj interpunkcijskih znakova za bilo koju rečenicu.

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
