---
id: 66e2df8c5676fe51cead8ec8
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Sada je vreme da testiraš svoju funkciju `getVowelCount`.

Kreiraj parametar za `vowelCount` i postavi njegovu vrednost na rezultat pozivanja funkcije `getVowelCount` sa argumentom `"Apples are tasty fruits"`

Nakon toga, napišite sledeće u konzolu: `"Vowel Count: [vowel count goes here]"`. Zamenite `[vowel count goes here]` stvarnim imenom parametra. Možete izabrati da koristite sekvencu slova šablona ili spojiti sekvencu slova koristeći parametar za `+` ovde.

# --hints--

Trebalo bi da kreirate parametar `vowelCount`.

```js
assert.isNotNull(vowelCount)
```

Vaš parametar za `vowelCount` treba da bude postavljen na rezultat od `getVowelCount("Apples are tasty fruits")`.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

Trebalo bi da unesete parametar `vowelCount` u konzolu, uz parametar `"Vowel Count: ${vowelCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
