---
id: 66e2e99caa0cb3570f4822b4
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada je vreme da testiraš svoju funkciju `getConsonantCount`.

Kreiraj parametar za `consonantCount` i postavi njegovu vrednost na rezultat pozivanja funkcije `getConsonantCount` sa argumentom `"Coding is fun"`

Nakon toga, napišite sledeće u konzolu: `"Consonant Count: [Consonant count goes here]"`. Zamenite `[Consonant count goes here]` sa stvarnim imenom parametra. Možete izabrati da koristite nizu karaktera iz šablona ili spojiti niz karaktera koristeći parametar `+` ovde.

# --hints--

Treba da kreiraš parametar `consonantCount`.

```js
assert.isNotNull(consonantCount)
```

Vaš parametar za `consonantCount` bi trebalo da bude postavljen na rezultat `getConsonantCount("Coding is fun")`.

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

Treba da napišete parametar za `consonantCount` u konzoli za parametar `"Consonant Count: ${consonantCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
