---
id: a0b5010f579e69b815e7c5d6
title: Napravi funkciju za pretvaranje reči sa inteligencijom
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

U ovoj laboratoriji, kreiraćete funkciju koja obavlja pretragu i modifikaciju u rečenicama koristeći zadate argumente, a zatim vraća novu rečenicu.

**Napomena:** Sačuvajte prvo slovo originalne reči kada je menjate. Na primer, ako nameravate da promenite reč `Book` sa reč `dog`, trebalo bi da se promeni kao `Dog`

**Cilj**: Popunite scenarije korisnika ispod i prođite sve testove kako biste završili laboratorijski rad.

Priče o korisniku:

1. Morate kreirati funkciju nazvanu `myReplace`.
1. Funkcija `myReplace` mora da primi tri argumenta: niz znakova, reč za zamenu i zamenjujuću reč.
1. Funkcija `myReplace` mora da vrati novi niz znakova sa izmenama.
1. Morate sačuvati prvo slovo originalne reči kada je menjate.

# --hints--

Treba da napraviš funkciju koja se zove `myReplace`.

```js
assert.isFunction(myReplace);
```

`myReplace` treba da uzme u obzir tri tačke.

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")` trebalo bi da vrati sekvencu znakova `Let us go to the mall`.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` treba da vrati sekvencu karaktera `He is Sitting on the couch`.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` treba da vrati niz znakova `I think we should look down there`.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` treba da vrati niz znakova `This has a spelling error`.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` trebalo bi da vrati sekvencu karaktera `His name is John`.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` trebalo bi da vrati niz znakova `Let us get back to more Algorithms`.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
