---
id: 684cd8bf5445de089af9c837
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Sada je vreme da primenite boju ivice. Da biste postavili svojstvo `border-color` koristeći Tailwind CSS, možete koristiti utilitarni klas sa sledećim obrascem: `border-color-number`, gde `number` predstavlja nijansu za boju ivice.

Evo primera primene ljubičaste boje ivice na element:

```html
<p class="border-purple-400">This is border-color.</p>
```

Unutar vašeg otvarajućeg taga `div`, primenite crvenu boju ivice sa nijansom `300`.

# --hints--

Vaš otvarajući `div` tag bi trebalo da ima utilitet `border-red-300` class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-red-300"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
