---
id: 684cd027744c2a0520ba4d58
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS koristi utilitarne klase koje se primenjuju direktno na oznake (markup). Ove klase uključuju stilizovanje za fontove, boje, Flexbox rasporede, Grid i više.

Da biste primenili pozadinsku boju u Tailwind CSS-u, možete koristiti prefiks `bg`, praćen imenom boje i brojem. Broj predstavlja nijansu boje.

Evo primera primene tamno plave pozadinske boje na paragraf element.

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

U vašem otvarajućem tagu `div`, dodajte klasu sa utilitarnom klasom `bg-red-100` koja predstavlja svetlo crvenu pozadinsku boju.

# --hints--

Vaš otvarajući `div` tag bi trebalo da ima class atribut.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

Vaš class atribut bi trebalo da ima korisnost class `bg-red-100`.

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
