---
id: 684cff0bf0a4ca0ad7e42990
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Sledeći korak bi bio za zaokruživanje uglova ivice/granice. Da biste postavili svojstvo `border-radius` koristeći Tailwind CSS, možete koristiti utilitarni klas sa sledećim obrascem: `rounded-size`, gde `size` predstavlja radijus.

Evo primera primene ekstra velikog radijusa ivice na element:

```html
<p class="rounded-xl">This is border radius.</p>
```

Unutar vašeg otvarajućeg taga `div`, primenite radijus ivice veličine `md`.

# --hints--

Vaš početni `div` tag bi trebalo da ima `rounded-md` utility class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
