---
id: 684cca09aac0fc035db4099f
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Unutar vašeg `p` elementa, dodajte tekst `Error! Something went wrong. Please try again.` gde bi `Error!` trebalo da bude umotan unutar `strong` elementa.

# --hints--

Trebalo bi da imate element `strong` unutar vašeg elementa `p`.

```js
assert.exists(document.querySelector("p strong"));
```

Vaš element `strong` bi trebalo da sadrži tekst `Error!`.

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

Vaš element `p` treba da ima tačan tekst koji se završava sa `Something went wrong. Please try again.`

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      --fcc-editable-region--
      <p>

      </p>
      --fcc-editable-region--
    </div>
  </body>
</html>
```
