---
id: 684cd59ec94b5e063bf776f0
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Biće lepo da imaš ivicu oko te svetlocrvene pozadine.

Da biste postavili `border-width` u Tailwind CSS-u, možete koristiti utilitarnu klasu `border-number` gde će `number` predstavljati svetlinu ili debljinu vaše ivice.

Evo primera primene `border-width` od `4px` na svim stranama:

```html
<p class="border-4">This is how to set the border-width.</p>
```

Unutar vašeg otvarajućeg taga `div`, dodajte odgovarajuću utilitarnu klasu za `border-width` podešen na `2px` na svim stranama.

# --hints--

Vaš otvarajući `div` tag treba da ima `border-2` class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
