---
id: 684ccbdc1ef89e042b744e09
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sledeći deo radionice je dodavanje dugmeta za zatvaranje.

Ispod vašeg elementa `p`, dodajte element `button` sa atributom `aria-label` podešenim na `"Dismiss"`. Unutar tog elementa `button`, dodajte element `span` sa atributom `aria-hidden`. Tekst elementa `span` treba da bude `&times;`.

*NAPOMENA*: Pošto će se ova radionica fokusirati samo na rad sa Tailwind CSS stilovima, element `button` neće imati dodatu nikakvu interaktivnost.

# --hints--

Trebalo bi da imate element `button`.

```js
assert.exists(document.querySelector("button"));
```Vaš `button` element treba da ima `aria-label` atribut.

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

Vaš `button` element treba da ima atribut `aria-label` set do `"Dismiss"`.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

Vaš element `button` bi trebalo da sadrži element `span` ugnježdjen unutar.

```js
assert.exists(document.querySelector("button span"));
```

Vaš `span` element bi trebalo da ima `aria-hidden` atribut.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

Vaš element `span` bi trebalo da sadrži tekst `&times;`.

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```
