---
id: 684d07cbfa43a20d6a002ed7
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Stilizacije za element `div` počinju da se slažu. Ali bilo bi lepo da postoji nešto razmaka oko teksta.

Da biste primenili `padding` na element koristeći Tailwind CSS, možete koristiti sledeći obrazac: `p-number`, gde `number` predstavlja piksele koje želite da primenite sa svih četiri strane.

Evo primera primene `padding` od `8px` na element:

```html
<p class="p-8">This is padding.</p>
```

Unutar vašeg otvarajućeg taga `div`, koristite odgovarajuću utilitarnu klasu da primenite `padding` od `4px`.

# --hints--

Vaš otvarajući `div` tag trebao bi imati utilitet `p-4` class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
