---
id: 684d160a2370b01310ae0fda
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Nova crvena boja izgleda lepo, ali bi bilo bolje da je tekst veći. Da biste postavili `font-size` za element koristeći Tailwind CSS, možete koristiti sledeći obrazac: `text-value`.

Evo primera postavljanja `font-size` za element na `0.875rem` ili (`14px`):

```html
<p class="text-sm">This is for font-size.</p>
```

Unutar vašeg otvarajućeg `p` elementa, koristite odgovarajuću utilitarnu klasu da postavite `font-size` na `xl`, što predstavlja `1.25rem` ili (`20px`).

# --hints--

Vaš otvarajući `p` tag treba da ima utilitet `text-xl` class.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
