---
id: 684d03df2a0ddf0c735865f0
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Sada vidite da je širina elementa pola veličine od `body`. Ovo dobro funkcioniše za veće uređaje, ali ne tako dobro za manje uređaje.

Tailwind CSS prati pristup "mobile-first" u dizajnu. Dakle, ako želite da kreirate responsivne dizajne za druge uređaje, možete koristiti Tailwindove responsivne varijante kao što su `md` i `lg` na ovaj način:

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

Sledeći primer će postaviti širinu na `66%` roditeljskog elementa za veličine uređaja od `1024px` i više.

Unutar vašeg otvarajućeg taga `div`, modifikujte svoj `w-1/2` tako da se primenjuje samo na veličine uređaja od `768px` i više. Možete dodati prefiks ove klase sa varijantom `md`.

# --hints--

Vaš početni `div` tag bi trebalo da ima `md:w-1/2` utilitet class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

Ne bi trebalo da imate `w-1/2` class u vašem otvarajućem `div` tag-u.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
