---
id: 684d17a3b350bb13d3c19c0f
title: Korak 19
challengeType: 0
dashedName: step-19
---

# --description--

Sada je vreme da primenite neke stilove na `button` element.

Počnite dodavanjem atributa klase vašem `button` elementu. Zatim, primenite odgovarajuću utilitarnu klasu za postavljanje crvene boje teksta sa nijansom postavljenu na `700`.

Zatim, unutar vašeg otvarajućeg taga `button`, koristite odgovarajuću utilitarnu klasu da postavite `font-size` na `2xl` što predstavlja `1.5rem` ili (`24px`).

# --hints--

Vaš `button` element treba da ima class atribut.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

Vaš otvorajući `button` tag bi trebalo da ima utilitet `text-red-700` class.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

Vaš otvarajući `button` tag bi trebalo da ima utilitet `text-2xl` class.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
