---
id: 684d0eeb9c81e5109e52fe8b
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Elemenat `div` počinje da izgleda bolje. Međutim, sada je element za zatvaranje (dismiss) `button` malo preblizak paragrafskom elementu.

Da biste stvorili razmak između fleks elemenata, možete koristiti sledeći obrazac: `gap-number`. `number` predstavlja fiksni razmak između redova i kolona u fleks rasporedu (flex layout).

Evo primera primene `gap` na element:

```html
<p class="flex gap-8">This is gap.</p>
```

Unutar vašeg otvarajućeg taga `div`, koristite odgovarajuću utilitarnu klasu za postavljanje `gap` na `4`.

# --hints--

Vaš početni `div` tag treba da ima `gap-4` utilitet class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
