---
id: 684d1942a1904a14a3709c6c
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Stilovanje za vaš komponentu poruke o grešci je skoro gotovo, ali postoji jedna stvar koja treba popraviti. Pošto ste postavili različite veličine fonta za tekst paragrafa i dugmeta, vertikalno poravnanje između tih fleks elemenata je pogrešno.

Da biste pozicionirali fleks elemente duž poprečne ose kontejnera koristeći Tailwind CSS, možete koristiti sledeći obrazac: `items-value`. `value` može biti bilo šta od `start`, `end`, `center`, `baseline` i više.

Unutar vašeg otvarajućeg taga `div`, koristite odgovarajuću utilitarnu klasu za postavljanje svojstva `align-items` na `center`.

I sa tom poslednjom promenom, vaša komponenta poruke o grešci je gotova!

# --hints--

Vaš otvarajući `div` tag bi trebalo da ima `items-center` utilitet class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
    --fcc-editable-region--
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto flex gap-4 justify-center items-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
