---
id: 684d11ffd15a531225d081e9
title: Korak 17
challengeType: 0
dashedName: step-17
---

# --description--

Sada kada su stilovi primenjeni na element ``div``, vreme je da primenimo neke stilove i na element ``p``.

Trenutno, boja teksta za paragrafski element je crna. Ali bilo bi lepo da je boja umesto toga tamno crvena.

U ranijem koraku, naučili ste kako da radite sa bojama pozadine kao što je ovo: ``bg-red-100``. Sličnu stvar možete uraditi i sa bojom teksta prateći ovaj obrazac: ``text-color-shade``.

Evo primera kako primeniti plavu boju teksta na element:

```html
<p class="text-blue-500">This has blue text.</p>
```

Počnite dodavanjem atributa klase na vaš ``p`` element. Zatim primenite odgovarajuću utilitarnu klasu za postavljanje crvene boje teksta sa nijansom postavljenoj na ``700``.

# --hints--

Vaš `p` element bi trebalo da ima atribut class.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

Vaš otvarajući `p` tag bi trebalo da ima `text-red-700` alat class.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
