---
id: 684d105e488dbd1166a717ae
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Ako povećate veličinu prozora za pregled, primetićete da su svi elementi i dalje poravnati levo u kontejneru. Bilo bi lepo da fleks elementi budu centrirani.

Da biste pozicionirali fleks elemente duž glavne ose kontejnera sa Tailwind CSS-om, možete koristiti sledeći obrazac: `justify-value`. `value` može biti `start`, `end`, `center`, `between` i više. Biloja važeća vrednost korišćena u čistom (vanilla) CSS-u sa `justify-content` može se ovde koristiti i sa Tailwind CSS-om.

Unutar vašeg otvarajućeg taga `div`, koristite odgovarajuću utilitarnu klasu za postavljanje svojstva `justify-content` na `center`.

# --hints--

Vaš otvarajući `div` tag treba da ima `justify-center` utilitet class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
