---
id: 684d009994a3840b98c9aa3a
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Trenutno, širina elementa `div` pokriva celu širinu stranice. Da biste primenili širine u Tailwind CSS-u, imate nekoliko opcija, kao što je korišćenje obrasca `w-number` za fiksne širine ili obrasca `w-fraction` za procente.

Evo primera primene širine postavljene na `66%` njegovog roditeljskog kontejnera:

```html
<p class="w-2/3">This is a percentage width.</p>
```

Unutar vašeg otvarajućeg taga `div`, primenite širinu postavljenu na `50%` njegovog roditeljskog kontejnera.

# --hints--

Vaš otvarajući `div` tag bi trebalo da ima `w-1/2` utilitet class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
