---
id: 684d09cad4ec900e41db2e74
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Trenutno je element `div` postavljen u gornji levi ugao ekrana za pregled. Ali bilo bi lepo dodati malo prostora na vrh elementa. Da biste primenili `margin`, na element koji koristi Tailwind CSS, možete koristiti sledeći obrazac: `m-number`, gde `number` predstavlja fiksne piksele za primenu na sve četiri strane.

Ako želite da primenite margina samo na jednu stranu, možete koristiti bilo koji od sledećih obrazaca: `mt-number` (margin top), `mb-number` (margin bottom), `mr-number` (margin right) i `ml-number` (margin left).

Evo primera primene `margin-bottom` na element:

```html
<p class="mb-8">This is margin.</p>
```

Unutar vašeg otvarajućeg taga `div`, koristite odgovarajuću utilitarnu klasu za primenu `margin-top` od `4px`.

# --hints--

Vaš otvarajući `div` tag treba da ima utiliteti `mt-4` za class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
