---
id: 68501d32b69f9d7756334ffc
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Ako povećate prozor za pregled, trebalo bi da vidite element `div` centriran horizontalno na stranici. Ali za manje uređaje, širina za element `div` pokriva celu širinu stranice.

Ažurirajte klasu `mx-auto` tako da se primenjuje samo za uređaje `md` i veće.

# --hints--

Vaš otvarajući `div` tag bi trebalo da ima `md:mx-auto` utilitet class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

Ne bi trebalo da imate `mx-auto` class u vašem početnom `div` tagu.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
