---
id: 6690e10ebe2181212abc9652
title: Napravi poslovnu vizitkartu
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba da imaš element `div` sa atributom `class` čija je vrednost `business-card`, koji će sadržati sve ostale elemente.
1. Unutar elementa `.business-card` treba da postoji element `img` sa atributom `class` čija je vrednost `profile-image`. Ako želiš, možeš postaviti izvor slike na `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg`. Element treba da ima smislen atribut `alt`.
1. Unutar elementa `.business-card`, nakon elementa `img`, treba da postoje tri elementa `p` sa atributom `class` čije su vrednosti `full-name`, `designation` i `company`, tim redosledom.
1. Prvi element `p` treba da sadrži tvoje ime.
1. Drugi element `p` treba da sadrži tvoju poziciju.
1. Treći element `p` treba da sadrži naziv tvoje kompanije.
1. Ispod trećeg elementa `p` treba da se nalazi element `hr`.
1. Nakon elementa `hr` treba da postoje dva elementa `p`. Prvi element `p` treba da sadrži adresu e-pošte kao tekst. Drugi element `p` treba da sadrži broj telefona kao tekst.
1. Nakon ta dva elementa `p` treba da postoji element `a` sa klasom `portfolio-link`, tekstom `Portfolio` i važećim URL-om.
1. Nakon elementa `a` koji vodi do portfolija treba da postoji element `hr`.
1. Treba da imaš još jedan element `div` sa atributom `class` čija je vrednost `social-media`. Unutar njega treba da postoji element `h2` sa tekstom `Connect with me`.
1. Unutar elementa `.social-media` treba da postoje tri elementa `a`. Oni treba da imaju tekst `Twitter`, `LinkedIn` i `GitHub`, tim redosledom, i da vode ka važećim URL-ovima.
1. Treba da povežeš datoteku stilova `styles.css` pomoću elementa `link` u delu `head` HTML datoteke.
1. Pozadina stranice treba da bude `rosybrown`. Glavni font treba da bude `Arial`, a rezervni font `sans-serif`.
1. Selektor `.business-card` treba da ima širinu `300px` i boju pozadine po tvom izboru. Takođe treba da ima `padding` od `20px` sa svih strana i `margin-top` od `100px`. Tekst treba da bude centriran, veličina fonta `16px`, a leva i desna margina postavljene na `auto`.
1. Selektor `.profile-image` treba da ima svojstvo `max-width` postavljeno na `100%`.
1. Svi elementi `p` treba da imaju gornju i donju marginu od `5px`.
1. Nijedan link na stranici ne treba da bude podvučen.

# --hints--

Trebalo bi da imaš glavni `div` sa svojstvom elementa `class` sa vrednošću `business-card`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

Unutar elementa `.business-card`, mora da postoji element `img` sa atributom polja `class` koji ima vrednost `profile-image`.

```js
assert.exists(document.querySelector('img.profile-image'));
```

Izvor slike mora biti postavljen na validnu sliku.

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

Atribut elementa `alt` slike treba da se postavi na smislen opis.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

Unutar polja `.business-card`, nakon polja `img`, trebalo bi da postoji polje `p` koje ima svojstvo polja `class` sa vrednošću `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

Nakon elementa `.full-name`, trebalo bi da postoji element/polje `p` sa svojstvom polja `class` i vrednošću `designation`.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

Nakon elementa `.designation`, trebalo bi da postoji element `p` sa atributom elementa `class` sa vrednošću od `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

Prvi element `p`, onaj sa klasom objekta `full-name`, treba da nosi tvoje ime.

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

Drugi element od `p`, onaj koji ima svojstvo elementa `class` sa vrednošću `designation`, treba da ima svoju klasu.

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

Treće polje od `p`, onaj koji ima svojstvo polja `class` sa vrednošću `company`, treba da nosi naziv vaše kompanije.

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

Nakon trećeg elementa od `p`, onaj koji ima atribut polja `class` sa vrednošću `company`, mora postojati element `hr`.

```js
assert.exists(document.querySelector('p.company + hr'));
```

Nakon prvog elementa od `hr`, treba da postoji element `p` sa adresom e-pošte kao vaš tekst.

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

Nakon polja za e-poštu sa `p`, treba da postoji drugo polje sa `p` koje ima broj telefona kao vaš tekst.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

Nakon polja broja telefona za `p`, treba da postoji polje `a` sa klasom objekata `portfolio-link`.

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

Element `.portfolio-link` mora imati tekst od `Portfolio`.

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

Vaš element od `.portfolio-link` treba da ima važeći `href`.

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

Nakon komponente za `a` sa vašim portfolio-om, treba da postoji još jedna komponenta od `hr` kao separator.

```js
assert.exists(document.querySelector('a + hr'));
```

Nakon drugog elementa od `hr`, treba da postoji element `div` sa karakteristikom elementa `class` sa vrednošću `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

Unutar elementa `.social-media`, trebalo bi da postoji element `h2` sa tekstom `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

Unutar elementa za `.social-media`, treba da postoji tri komponente od `a`.

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

Prvi element od `a` bi trebalo da ima tekst `Twitter`.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

Atribut elementa ``href`` prvog elementa ``a`` mora biti povezan sa validnom vezom.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

Drugi element od `a` treba da sadrži tekst `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

Atribut elementa `href` drugog elementa `a` mora biti povezan sa validnim linkom.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

Treći element od `a` treba da ima tekst `GitHub`.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

Atribut elementa `href`, trećeg elementa `a`, treba da se poveže sa validnom vezom.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

Tvoj element `link` treba da bude unutar tvog elementa `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Tvoj element za `link` treba da ima svojstvo elemenata `rel` sa vrednošću `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Vaše polje `link` treba da ima atribut polja `href` sa vrednošću `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

Trebalo bi da postavite boju teme stranice u izbornik ``body`` na ``rosybrown``.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

Trebalo bi da postavite font stranice na `Arial` u polje `body` i alternativno na `sans-serif`.

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

Tvoj objekat `.business-card` mora imati svojstvo `width` sa vrednošću `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

Vaš selektor za `.business-card` mora imati svojstvo `background-color` sa validnom vrednošću boje.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

Tvoj selektor za `.business-card` mora imati svojstvo `padding` sa vrednošću `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

Tvoj selektor za `.business-card` mora imati maksimalnu granicu od `100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

Tekst unutar polja `.business-card` treba da bude centriran.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

Veličina fonta teksta unutar elementa `.business-card` treba da bude `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

Leva i desna ivice elementa ``.business-card`` treba da se podele na ``auto``.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

Tvoj selektor za `.profile-image` mora imati svojstvo `max-width` sa vrednošću `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

Vaši elementi za `p` treba da imaju gornju i donju ivicu od `margin` od `5px`.

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

Elementi stranice ne bi trebalo da imaju podvlačenje.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
