---
id: ac6993d51946422351508a41
title: Algoritam za sečenje sekvence slova
challengeType: 26
dashedName: truncate-a-string
---

# --description--

U ovoj laboratoriji, vežbaćeš da sečeš niz znakova do određene dužine.

**Cilj:** Popunite korisničke priče ispod i prođite kroz sve testove kako biste završili ovaj laboratorijski rad.

Priče o korisniku:

1. Treba da postoji funkcija `truncateString` koja prima dva argumenta: prvi je niz znakova, a drugi je broj.
1. Ako je dužina niza znakova veća od dato broja, taj niz mora biti skraćen da odgovara toj dužini, a `...` treba dodati na kraj skraćenog niza.
1. Ako je dužina niza znakova jednaka ili manja od dato broja, niz mora biti vraćen nepromenjen.

# --hints--

`truncateString("A-tisket a-tasket A green and yellow basket", 8)` treba da vrati sekvencu `A-tisket...`.

```js
assert.strictEqual(
  truncateString('A-tisket a-tasket A green and yellow basket', 8),
  'A-tisket...'
);
```

`truncateString("Peter Piper picked a peck of pickled peppers", 11)` treba da vrati sekvencu `Peter Piper...`.

```js
assert.strictEqual(
  truncateString('Peter Piper picked a peck of pickled peppers', 11),
  'Peter Piper...'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length)` trebalo bi da vrati sekvencu `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length + 2)` treba da vrati sekvencu `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length + 2
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-", 1)` treba vratiti sekvencu `A...`.

```js
assert.strictEqual(truncateString('A-', 1), 'A...');
```

`truncateString("Absolutely Longer", 2)` treba da vrati sekvencu `Ab...`.

```js
assert.strictEqual(truncateString('Absolutely Longer', 2), 'Ab...');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function truncateString(str, num) {
  if (num >= str.length) {
    return str;
  }

  return str.slice(0, num) + '...';
}

```
