---
id: 681cb1b0dab50c87ddb2e518
title: "Zoezi la 7: Jumla iliyolengwa"
challengeType: 28
dashedName: challenge-7
---

# --description--
Kada imate niz brojeva i ciljnu brojčanu vrednost, pronađite dva različita broja u nizu podataka čija je suma jednaka ciljnoj vrednosti. Vratite niz podataka sa indeksima tih dva broja, ili `"Target not found"` ako ne postoji par brojeva čija je suma jednaka cilju.

- Red podataka koji se vraća treba da ima indekse brojeve u rastućem redosledu.

# --hints--

`findTarget([2, 7, 11, 15], 9)` inapaswa kurudisha `[0, 1]`.

```js
assert.deepEqual(findTarget([2, 7, 11, 15], 9), [0, 1]);
```

`findTarget([3, 2, 4, 5], 6)` inapaswa kurudisha `[1, 2]`.

```js
assert.deepEqual(findTarget([3, 2, 4, 5], 6), [1, 2]);
```

`findTarget([1, 3, 5, 6, 7, 8], 15)` inapaswa kurudisha `[4, 5]`.

```js
assert.deepEqual(findTarget([1, 3, 5, 6, 7, 8], 15), [4, 5]);
```

`findTarget([1, 3, 5, 7], 14)` treba da vrati `"Target not found"`.

```js
assert.equal(findTarget([1, 3, 5, 7], 14), "Target not found");
```

# --seed--

## --seed-contents--

```js
function findTarget(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findTarget(arr, target) {
  for (let i = 0; i < arr.length; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] + arr[j] === target) {
        return [i, j];
      }
    }
  }
  return 'Target not found'
}
```
