---
id: 681cb1afdab50c87ddb2e517
title: "Zoezi la 6: Kagua anagramu"
challengeType: 28
dashedName: challenge-6
---

# --description--
Razmatrajući sekvencu od dva slova, odredite da li su anagrami jedno drugog (imaju iste slova u bilo kom redosledu).

- Ignorišite razlike između velikih i malih slova, kao i prazne prostore.

# --hints--

`areAnagrams("listen", "silent")` inapaswa kurudisha `true`.

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")` inapaswa kurudisha `true`.

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")` inapaswa kurudisha `true`.

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")` inapaswa kurudisha `false`.

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")` inapaswa kurudisha `false`.

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")` inapaswa kurudisha `false`.

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```
