---
id: 681cb1b0dab50c87ddb2e51a
title: "Zoezi la 9: Jumla ya miraba"
challengeType: 28
dashedName: challenge-9
---

# --description--

Ako je dat pozitivan ceo broj do 1.000, vrati zbir svih celih brojeva koji su kvadrirani od 1 do tog broja.

# --hints--

`sumOfSquares(5)` inapaswa kurudisha `55`.

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)` inapaswa kurudisha `385`.

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)` inapaswa kurudisha `5525`.

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)` inapaswa kurudisha `41791750`.

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)` inapaswa kurudisha `333833500`.

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```
