---
id: 681cb1a2dab50c87ddb2e514
title: "Zoezi la 3: Mfululizo wa Fibonacci"
challengeType: 28
dashedName: challenge-3
---

# --description--
Niz Fibonaccija je sekvenca brojeva gde je svaki broj zbir prethodnih dva broja. Počevši od `0` i `1`, prvih deset brojeva u sekvenci su `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

Uz zadati niz podataka koji sadrži dve prva broja Fibonaccijevog niza i ceo broj koji predstavlja dužinu sekvence, vratite niz podataka sa sekvencijom zadate dužine.

- Vaša funkcija bi trebalo da obrađuje sekvencu bilo koje dužine veće ili jednake nuli.
- Ako je dužina nula, vratite prazan niz podataka.
- Napomena da su početne brojeve deo sekvence.

# --hints--

`fibonacciSequence([0, 1], 20)` inapaswa kurudisha `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]);
```

`fibonacciSequence([21, 32], 1)` inapaswa kurudisha `[21]`.

```js
assert.deepEqual(fibonacciSequence([21, 32], 1), [21]);
```

`fibonacciSequence([0, 1], 0)` inapaswa kurudisha `[]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 0), []);
```

`fibonacciSequence([10, 20], 2)` inapaswa kurudisha `[10, 20]`.

```js
assert.deepEqual(fibonacciSequence([10, 20], 2), [10, 20]);
```

`fibonacciSequence([123456789, 987654321], 5)` inapaswa kurudisha `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
assert.deepEqual(fibonacciSequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541]);
```

# --seed--

## --seed-contents--

```js
function fibonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function fibonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2];
    sequence.push(nextValue);
  }
  return sequence;
}
```
