---
id: 681cb1b0dab50c87ddb2e519
title: "Zoezi la 8: Faktoria"
challengeType: 28
dashedName: challenge-8
---

# --description--

Ako je dat ceo broj od nule do dvadeset, vrati faktorijel tog broja. Faktorijel broja je proizvod svih brojeva od 1 do datog broja.

- Faktoria ya sifuri ni 1.

# --hints--

`factorial(0)` inapaswa kurudisha `1`.

```js
assert.equal(factorial(0), 1);
```

`factorial(5)` inapaswa kurudisha `120`.

```js
assert.equal(factorial(5), 120);
```

`factorial(20)` inapaswa kurudisha `2432902008176640000`.

```js
assert.equal(factorial(20), 2432902008176640000);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

  return n;
}
```

# --solutions--

```js
function factorial(n) {
  return n == 0 ? 1 : n * factorial(n - 1);
}
```
