---
id: 681cb1afdab50c87ddb2e515
title: "Zoezi la 4: S  P  A  C  E  J  A  M"
challengeType: 28
dashedName: challenge-4
---

# --description--
Za dati niz slova, uklonite sve razmake iz tog niza, stavite po dva razmaka između svakog slova, pretvorite sva slova abecede u velika slova i vratite rezultat.

- Znakovi koji nisu slova treba da ostanu nepromenjeni (osim razmaka).

# --hints--

`spaceJam("freeCodeCamp")` inapaswa kurudisha `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("   free   Code   Camp   ")` inapaswa kurudisha `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("Hello World?!")` inapaswa kurudisha `"H  E  L  L  O  W  O  R  L  D  ?  !"`.

```js
assert.equal(spaceJam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !");
```

`spaceJam("C@t$ & D0g$")` inapaswa kurudisha `"C  @  T  $  &  D  0  G  $"`.

```js
assert.equal(spaceJam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $");
```

`spaceJam("allyourbase")` inapaswa kurudisha `"A  L  L  Y  O  U  R  B  A  S  E"`.

```js
assert.equal(spaceJam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E");
```

# --seed--

## --seed-contents--

```js
function spaceJam(s) {

  return s;
}
```

# --solutions--

```js
function spaceJam(s) {
  return s.toUpperCase().replace(/\s+/g, '').split('').join('  ');
}
```
