---
id: 681cb05adab50c87ddb2e513
title: "Zoezi la 2: Ukaguzi wa Msingi"
challengeType: 28
dashedName: challenge-2
---

# --description--

Ako je dat niz znakova koji predstavlja broj i ceo broj osnovu sistema od 2 do 36, odredi da li je taj broj validan u toj osnovi.

- Niz znakova može sadržati cijeli broj, sa velikim ili malim slovima.
- Validacija treba da bude neosjetljiva na razliku između velikih i malih slova.
- Osnova može biti bilo koji broj između 2 i 36.
- Broj je validan ako je svaki znak važeći cifarica u datoj osnovi.
- Primjer validnih cifara za osnove:
  - Osnova 2: 0-1
  - Osnova 8: 0-7
  - Osnova 10: 0-9
  - Osnova 16: 0-9 i A-F
  - Osnova 36: 0-9 i A-Z

# --hints--

`isValidNumber("10101", 2)` should return `true`.

```js
assert.isTrue(isValidNumber("10101", 2))
```

`isValidNumber("10201", 2)` should return `false`.

```js
assert.isFalse(isValidNumber("10201", 2))
```

`isValidNumber("76543210", 8)` should return `true`.

```js
assert.isTrue(isValidNumber("76543210", 8))
```

`isValidNumber("9876543210", 8)` should return `false`.

```js
assert.isFalse(isValidNumber("9876543210", 8))
```

`isValidNumber("9876543210", 10)` should return `true`.

```js
assert.isTrue(isValidNumber("9876543210", 10))
```

`isValidNumber("ABC", 10)` should return `false`.

```js
assert.isFalse(isValidNumber("ABC", 10))
```

`isValidNumber("ABC", 16)` should return `true`.

```js
assert.isTrue(isValidNumber("ABC", 16))
```

`isValidNumber("Z", 36)` should return `true`.

```js
assert.isTrue(isValidNumber("Z", 36))
```

`isValidNumber("ABC", 20)` should return `true`.

```js
assert.isTrue(isValidNumber("ABC", 20))
```

`isValidNumber("4B4BA9", 16)` should return `true`.

```js
assert.isTrue(isValidNumber("4B4BA9", 16))
```

`isValidNumber("5G3F8F", 16)` should return `false`.

```js
assert.isFalse(isValidNumber("5G3F8F", 16))
```

`isValidNumber("5G3F8F", 17)` should return `true`.

```js
assert.isTrue(isValidNumber("5G3F8F", 17))
```

`isValidNumber("abc", 10)` should return `false`.

```js
assert.isFalse(isValidNumber("abc", 10))
```

`isValidNumber("abc", 16)` should return `true`.

```js
assert.isTrue(isValidNumber("abc", 16))
```

`isValidNumber("AbC", 16)` should return `true`.

```js
assert.isTrue(isValidNumber("AbC", 16))
```

`isValidNumber("z", 36)` should return `true`.

```js
assert.isTrue(isValidNumber("z", 36))
```

# --seed--

## --seed-contents--

```js
function isValidNumber(n, base) {

  return n;
}
```

# --solutions--

```js
function isValidNumber(n, base) {
  return new RegExp(`^[${'0123456789abcdefghijklmnopqrstuvwxyz'.substring(0, base)}]+\$`, "i").test(n);
}
```
