---
id: 68216ef80f957572e7c340c5
title: "Zoezi la 12: Mtafiti wa ujumbe"
challengeType: 28
dashedName: challenge-12
---

# --description--

Ako je dat niz znakova tajne poruke i ceo broj koji predstavlja broj pomeranja slova korišćen za sakrivanje poruke, vrati dešifrovani niz znakova.

- Pozitivan broj znači poruku pomerenu napred u alfabetu.
- Negativan broj znači poruku pomerenu nazad u alfabetu.
- Razlika između velikih i malih slova je važna; dekodovana slova moraju sačuvati razliku velikog ili malog slova kao što su bila u prethodnoj poruci.
- Slova koja nisu deo alfabeta ne bi trebalo da se dekodiraju.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` inapaswa kurudisha `"This is a secret message."`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` inapaswa kurudisha `"Hello World!"`

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` inapaswa kurudisha `"Are you okay?"`

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` inapaswa kurudisha `"freeCodeCamp"`

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```
