---
id: 6821ebee237de8297eaee796
title: "Vežba 20: Kopije reda podataka
challengeType: 28
dashedName: challenge-20
---

# --description--
Ako je dat niz celih brojeva, vrati niz celih brojeva koji se ponavljaju više od jednom u originalnom nizu, sortiran po rastućem redosledu. Ako ne postoji vrednost koja se ponavlja više od jednom, vrati prazan niz podataka.

- Uključite samo jedan primer svake vrednosti u vraćeni niz podataka.

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` inapaswa kurudisha `[]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` inapaswa kurudisha `[1, 2]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` inapaswa kurudisha `[-6, 0, 2, 4, 5, 23]`.

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
