---
id: 6821ebd4237de8297eaee791
title: "Zoezi la 15: camelCase"
challengeType: 28
dashedName: challenge-15
---

# --description--
Kada uneseš niz znakova, vrati mu verziju u camelcase formatu koristeći sledeća pravila:

- Reči u sekvenci slova su razdvojene jednim ili više znakova iz sledećeg skupa: razmak (` `), crtica (`-`) ili donja crta (`_`). Svaku sekvencu ovoga tretirati kao razdelnik reči.
- Prva reč mora biti isključivo mala slova.
- Svaka sledeća reč mora početi velikim slovom, a ostatak mora biti malih.
- Svi razmaci i separatori moraju se ukloniti.

# --hints--

`toCamelCase("hello world")` inapaswa kurudisha `"helloWorld"`.

```js
assert.equal(toCamelCase("hello world"), "helloWorld");
```

`toCamelCase("HELLO WORLD")` inapaswa kurudisha `"helloWorld"`.

```js
assert.equal(toCamelCase("HELLO WORLD"), "helloWorld");
```

`toCamelCase("secret agent-X")` inapaswa kurudisha `"secretAgentX"`.

```js
assert.equal(toCamelCase("secret agent-X"), "secretAgentX");
```

`toCamelCase("FREE cODE cAMP")` inapaswa kurudisha `"freeCodeCamp"`.

```js
assert.equal(toCamelCase("FREE cODE cAMP"), "freeCodeCamp");
```

`toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` inapaswa kurudisha `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
assert.equal(toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk");
```

# --seed--

## --seed-contents--

```js
function toCamelCase(s) {

  return s;
}
```

# --solutions--

```js
function toCamelCase(s) {
  const words = s.replace(/[_\- ]+/g, ' ').split(' ');

  return words.map((word, i) => {
    if (i === 0) {
      return word.toLowerCase();
    } else {
      const tempWord = word.split('');
      return tempWord.shift().toUpperCase() + tempWord.join('').toLowerCase();
    }
  }).join('')
}
```
