---
id: 6821ebea237de8297eaee795
title: "Zoezi la 19: Mwanga wa Mishumaa"
challengeType: 28
dashedName: challenge-19
---

# --description--
Ako unesete kompletan broj koji predstavlja početnu količinu sveća, i kompletan broj koji predstavlja koliko je spaljenih sveća potrebno za napraviti novu sveću, vratite broj sveća koje ćete koristiti nakon što napravite i sagorite sve moguće sveće.

Na primer, ako vam se da 7 sveća i potrebno je 2 spaljene sveće za napraviti novu sveću:

1. Choma mishumaa 7 kupata mabaki 7,
2. Rudisha mabaki 6 kuwa mishumaa 3 mipya (mshumaa 1 unabaki),
3. Choma mishumaa 3 kupata mabaki 3 zaidi (jumla 4),
4. Rudisha mabaki 4 kuwa mishumaa 2 mipya,
5. Choma mishumaa 2 kupata mabaki 2,
6. Rudisha mabaki 2 kuwa mshumaa 1 mpya,
7. Choma mshumaa 1.

U ovom primeru će biti upaljeno ukupno 13 sveća.

# --hints--

`burnCandles(7, 2)` inapaswa kurudisha `13`

```js
assert.equal(burnCandles(7, 2), 13);
```

`burnCandles(10, 5)` inapaswa kurudisha `12`

```js
assert.equal(burnCandles(10, 5), 12);
```

`burnCandles(20, 3)` inapaswa kurudisha `29`

```js
assert.equal(burnCandles(20, 3), 29);
```

`burnCandles(17, 4)` inapaswa kurudisha `22`

```js
assert.equal(burnCandles(17, 4), 22);
```

`burnCandles(2345, 3)` inapaswa kurudisha `3517`

```js
assert.equal(burnCandles(2345, 3), 3517);
```

# --seed--

## --seed-contents--

```js
function burnCandles(candles, leftoversNeeded) {

  return candles;
}
```

# --solutions--

```js
function burnCandles(candles, leftoversNeeded) {
  let totalBurned = 0;
  let unusedLeftovers = 0;

  while (candles > 0) {
    totalBurned += candles;
    const leftovers = candles + unusedLeftovers;
    candles = Math.floor(leftovers / leftoversNeeded);
    unusedLeftovers = leftovers % leftoversNeeded;
  }

  return totalBurned;
}
```
