---
id: 6821ebce237de8297eaee790
title: "Zoezi la 14: Mapigano ya herufi"
challengeType: 28
dashedName: challenge-14
---

# --description--
Daje se sekvenca od dva slova koja predstavljaju vašu vojsku i vojsku protivnika; svako slovo iz vaše vojske bori se sa slovom koje je na istoj poziciji u vojskoj protivnika koristeći sledeća pravila:

- Slovo `a-z` ima snagu `1-26`, po odelima.
- Slovo `A-Z` ima snagu `27-52`, po odelima.
- Broj `0-9` ima snagu jednaku njegovoj nominalnoj vrednosti.
- Sva ostala slova imaju nulu vrednosti.
- Svako slovo može se boriti samo u jednoj bitci.

U svakoj bitci pobeđuje najjače slovo. Vojska sa najviše pobeda osvaja rat. Vratite sledeće vrednosti:

- `"Opponent retreated"` ako tvoja vojska ima više vojnika od protivničke.
- `"We retreated"` ako protivnička vojska ima više vojnika od tvoje.
- `"We won"` ako je tvoja vojska pobedila u više bitaka.
- `"We lost"` ako je protivnička vojska pobedila u više bitaka.
- `"It was a tie"` ikiwa majeshi yote mawili yameshinda vita idadi sawa.

# --hints--

`battle("Hello", "World")` inapaswa kurudisha `"We lost"`.

```js
assert.equal(battle("Hello", "World"), "We lost");
```

`battle("pizza", "salad")` inapaswa kurudisha `"We won"`.

```js
assert.equal(battle("pizza", "salad"), "We won");
```

`battle("C@T5", "D0G$")` inapaswa kurudisha `"We won"`.

```js
assert.equal(battle("C@T5", "D0G$"), "We won");
```

`battle("kn!ght", "orc")` inapaswa kurudisha `"Opponent retreated"`.

```js
assert.equal(battle("kn!ght", "orc"), "Opponent retreated");
```

`battle("PC", "Mac")` inapaswa kurudisha `"We retreated"`.

```js
assert.equal(battle("PC", "Mac"), "We retreated");
```

`battle("Wizards", "Dragons")` inapaswa kurudisha `"It was a tie"`.

```js
assert.equal(battle("Wizards", "Dragons"), "It was a tie");
```

`battle("Mr. Smith", "Dr. Jones")` inapaswa kurudisha `"It was a tie"`.

```js
assert.equal(battle("Mr. Smith", "Dr. Jones"), "It was a tie");
```

# --seed--

## --seed-contents--

```js
function battle(myArmy, opposingArmy) {

  return myArmy;
}
```

# --solutions--

```js
function getStrength(soldier) {
  const soldiers = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ';
  let strength = 0;

  if (/\d/.test(soldier)) {
    strength = parseInt(soldier);
  } else if (soldiers.includes(soldier)) {
    strength = soldiers.indexOf(soldier) + 1;
  }

  return strength;
}

function battle(myArmy, opposingArmy) {
  if (myArmy.length > opposingArmy.length) return 'Opponent retreated';
  if (opposingArmy.length > myArmy.length) return 'We retreated';

  let myWins = 0;
  let theirWins = 0;

  for (let i = 0; i < myArmy.length; i++) {
    const mySoldier = myArmy[i];
    const theirSoldier = opposingArmy[i];

    const myStrength = getStrength(mySoldier);
    const theirStrength = getStrength(theirSoldier);

    if (myStrength > theirStrength) myWins++;
    if (theirStrength > myStrength) theirWins++;
  }

  return myWins > theirWins ? 'We won' : theirWins > myWins ? 'We lost' : 'It was a tie';
}
```
