---
id: 6821ebda237de8297eaee792
title: "Zoezi la 16: Rudisha Mabano ya Kawaida"
challengeType: 28
dashedName: challenge-16
---

# --description--
Ako naiđete na sekvencu slova koja sadrži pravilno postavljene obične zagrade, vratite interpretiranu verziju te sekvence koristeći sledeće pravilo:

- Svi znakovi unutar svake pare običnih zagradi treba se vratiti (backtrack).
- Obične zagrade treba ukloniti u konačnom rezultatu.
- Ako postoje ugnježđene obične zagrade, prvo se mora izvršiti povratak najunutrašnje pare, a zatim njen rezultat uključiti u povratak spoljašnje pare.
- Pretpostaviti da su sve obične zagrade pravilno i tačno ugnježđene.

# --hints--

`decode("(f(b(dc)e)a)")` inapaswa kurudisha `"abcdef"`.

```js
assert.equal(decode("(f(b(dc)e)a)"), "abcdef");
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` inapaswa kurudisha `"Can you read this?"`.

```js
assert.equal(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?");
```

`decode("f(Ce(re))o((e(aC)m)d)p")` inapaswa kurudisha `"freeCodeCamp"`.

```js
assert.equal(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(s) {

  return s;
}
```

# --solutions--

```js
function decode(s) {
  while (s.includes(')')) {
    const closeIndex = s.indexOf(')');
    const openIndex = s.lastIndexOf('(', closeIndex);
    const before = s.slice(0, openIndex);
    const group = s.slice(openIndex + 1, closeIndex).split('').reverse().join('');
    const after = s.slice(closeIndex + 1);
    s = before + group + after;
  }

  return s;
}
```
