---
id: 68adce01c0e1144d0a902958
title: "Zoezi la 26: Kagua anwani halali ya IPv4"
challengeType: 28
dashedName: challenge-26
---

# --description--

Ako je data niska, odredi da li predstavlja ispravnu IPv4 adresu. Ispravna IPv4 adresa sastoji se od četiri cela broja odvojena tačkama (`.`). Svaki broj mora da ispunjava sledeće uslove:

- Između 0 i 255, uključujući.
- Nema vodećih nula (primer 0 je dozvoljeno, 01 nije dozvoljeno).
- Dozvoljeni karakteri su samo brojevi.

# --hints--

`isValidIPv4("192.168.1.1")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

`isValidIPv4("0.0.0.0")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

`isValidIPv4("255.01.50.111")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

`isValidIPv4("255.00.50.111")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

`isValidIPv4("256.101.50.115")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

`isValidIPv4("192.168.101.")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

`isValidIPv4("192168145213")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidIPv4("192168145213"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```
