---
id: 68adce01c0e1144d0a902956
title: "Zoezi la 25: Kurudia vokali"
challengeType: 28
dashedName: challenge-25
---

# --description--
Ako izvučete sekvencu slova, vratite novu verziju te sekvence gde se svaka samoglasanka ponavlja za jedan put više nego prethodna samoglasanka koju ste naišli. Na primer, prva samoglasanka u rečenici ostaje nepromenjena. Druga samoglasanka pojavljuje se dvaput uzastopno. Treća samoglasanka pojavljuje se tri puta uzastopno, i tako dalje.

- Slova `a`, `e`, `i`, `o` i `u`, bilo velika ili mala, smatraju se samoglasnicima.
- Originalni samoglasnik treba da zadrži svoj oblik slova.
- Ponavljajući samoglasnici treba da budu mala slova.
- Svi znakovi koji nisu samoglasnici treba da zadrže svoj originalni oblik slova.

# --hints--

`repeatVowels("hello world")` inapaswa kurudisha `"helloo wooorld"`.

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` inapaswa kurudisha `"freeeCooodeeeeCaaaaamp"`.

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` inapaswa kurudisha `"AEeIiiOoooUuuuu"`.

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("I like eating ice cream in Iceland")` inapaswa kurudisha `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```
