---
id: 68adce01c0e1144d0a90295e
title: "Vežba 29: Generator početnih slova"
challengeType: 28
dashedName: challenge-29
---

# --description--
Ako vam je dat niz slova koji sadrži jedno ili više reči, vratite početna slova tih reči koristeći sledeće uslove:

- Početna slova treba da budu velika slova, osim ako nije drugačije navedeno.
- Početna slova sledećih reči ne treba da budu velika, osim ako nisu prva reč u datoj sekvenci: `a`, `for`, `an`, `and`, `by`, i `of`.
- Početna slova treba da se vrate u redosledu kojim su data.
- Početna slova ne smeju imati razmak između njih.

# --hints--

`buildAcronym("Search Engine Optimization")` inapaswa kurudisha `"SEO"`.

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Frequently Asked Questions")` inapaswa kurudisha `"FAQ"`.

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("National Aeronautics and Space Administration")` inapaswa kurudisha `"NASA"`.

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")` inapaswa kurudisha `"FBI"`.

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("For your information")` inapaswa kurudisha `"FYI"`.

```js
assert.equal(buildAcronym("Light Amplification by Stimulated Emission of Radiation"), "LASER");
```

`buildAcronym("By the way")` inapaswa kurudisha `"BTW"`.

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` inapaswa kurudisha `"AUHWPOTIMSH"`.

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```
