---
id: 68af0687ef34c76c28ffa54f
title: "Zoezi la 34: Nambari Zinazokosekana"
challengeType: 28
dashedName: challenge-34
---

# --description--
Ako vam je dat opseg celih brojeva od 1 do `n`, uključujući, vratite opseg svih nedostajućih celih brojeva između 1 i `n` (gde je `n` najveći broj u datom opsegu podataka).

- Izvučeni red podataka može biti neuređen i može sadržavati duplikate.
- Red podataka koji će se obrađivati treba da bude u rastućem redosledu.
- Ako ne nedostaju potpuni brojevi, vratite prazan red podataka.

# --hints--

`findMissingNumbers([1, 3, 5])` inapaswa kurudisha `[2, 4]`.

```js
assert.deepEqual(findMissingNumbers([1, 3, 5]), [2, 4]);
```

`findMissingNumbers([1, 2, 3, 4, 5])` inapaswa kurudisha `[]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 10])` inapaswa kurudisha `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
assert.deepEqual(findMissingNumbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([10, 1, 10, 1, 10, 1])` inapaswa kurudisha `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
assert.deepEqual(findMissingNumbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([3, 1, 4, 1, 5, 9])` inapaswa kurudisha `[2, 6, 7, 8]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])` inapaswa kurudisha `[11]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11]);
```

# --seed--

## --seed-contents--

```js
function findMissingNumbers(arr) {

  return arr;
}
```

# --solutions--

```js
function findMissingNumbers(arr) {
  if (arr.length === 0) return [];
  
  const n = Math.max(...arr);
  const set = new Set(arr);
  const result = [];

  for (let i = 1; i <= n; i++) {
    if (!set.has(i)) {
      result.push(i);
    }
  }
  return result;
}
```
