---
id: 68af0687ef34c76c28ffa547
title: "Zoezi la 30: Herufi za kipekee"
challengeType: 28
dashedName: challenge-30
---

# --description--

Ako je dat niz znakova, odredi da li su sva slova u tom nizu jedinstvena.

- Velika i mala slova moraju se tretirati kao različita slova.

# --hints--

`allUnique("abc")` inapaswa kurudisha `true`.

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` inapaswa kurudisha `true`.

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` inapaswa kurudisha `true`.

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` inapaswa kurudisha `true`.

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` inapaswa kurudisha `false`.

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` inapaswa kurudisha `false`.

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` inapaswa kurudisha `false`.

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
