---
id: 68af0687ef34c76c28ffa54b
title: "Zoezi la 32: Geuza Sentensi"
challengeType: 28
dashedName: challenge-32
---

# --description--
Dato je sekvenca reči, vrati novu sekvencu reči u obrnutom redosledu. Na primer, prva reč bi trebalo da bude na kraju vraćene sekvence, a poslednja reč bi trebalo da bude na početku vraćene sekvence.

- U datom nizu reči mogu biti razdvojene jednim ili više razmaka.
- Vraćeni niz treba da ima samo jedan razmak između reči.

# --hints--

`reverseSentence("world hello")` inapaswa kurudisha `"hello world"`.

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` treba da vrati `"git commit push"`.

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` inapaswa kurudisha `"sudo install npm"`.

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` inapaswa kurudisha `"export function default import"`.

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
