---
id: 68b06e589bf227324381476f
title: "Vežba 35: Pojavljivanje reči
challengeType: 28
dashedName: challenge-35
---

# --description--
Ako uzmete paragraf, vratite red podataka sa tri najčešće pojavljene reči.

- Reči u rečenici će biti razdvojene razmakom.
- Ignorišite razlike između malih i velikih slova u datoj rečenici. Na primer, tretirajte `Hello` i `hello` kao jednu reč.
- Ignorišite interpunkcijske znake u datoj rečenici. Interpunkcijski znaci su kao što su zarez (`,`), tačka (`.`) i uzvičnik (`!`).
- Povraćeni red podataka treba da sadrži sve reči u malim slovima.
- Povraćeni red podataka treba da bude u silaznom redosledu po učestalosti, sa najčešće korišćenom reči na prvom mestu.

# --hints--

`getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` inapaswa kurudisha `["coding", "python", "in"]`.

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")` inapaswa kurudisha `["i", "like", "coding"]`.

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` inapaswa kurudisha `["debug", "test", "deploy"]`.

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```
