---
id: 68b06e589bf2273243814773
title: "Vežba 37: Stavljanje velikog slova na početak rečenice
challengeType: 28
dashedName: challenge-37
---

# --description--
Ako dobiješ pasus, vrati novi pasus u kojem je prvo slovo svake rečenice napisano velikim slovom.

- Sve ostale slova treba da ostanu kako su.
- Rečenice mogu se završiti tačkom (`.`), jednim ili više znakova pitanja (`?`), ili jednim ili više znakova uzvika (`!`).

# --hints--

`capitalize("this is a simple sentence.")` treba da vrati `"This is a simple sentence."`.

```js
assert.equal(capitalize("this is a simple sentence."), "This is a simple sentence.");
```

`capitalize("hello world. how are you?")` inapaswa kurudisha `"Hello world. How are you?"`.

```js
assert.equal(capitalize("hello world. how are you?"), "Hello world. How are you?");
```

`capitalize("i did today's coding challenge... it was fun!!")` inapaswa kurudisha `"I did today's coding challenge... It was fun!!"`.

```js
assert.equal(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!");
```

`capitalize("crazy!!!strange???unconventional...sentences.")` inapaswa kurudisha `"Crazy!!!Strange???Unconventional...Sentences."`.

```js
assert.equal(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.");
```

`capitalize("there's a space before this period . why is there a space before that period ?")` inapaswa kurudisha `"There's a space before this period . Why is there a space before that period ?"`.

```js
assert.equal(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?");
```

# --seed--

## --seed-contents--

```js
function capitalize(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function capitalize(paragraph) {
  return paragraph.replace(/([.!?]+)(\s*)([a-z])/g, (match, punc, spaces, char) => {
    return punc + spaces + char.toUpperCase();
  })
  .replace(/^([a-z])/, (match, char) => char.toUpperCase());
}
```
