---
id: 68b06e589bf2273243814771
title: "Zoezi la 36: Mrekebishaji wa Thermostat"
challengeType: 28
dashedName: challenge-36
---

# --description--
Izbaci trenutnu temperaturu u sobi i željenu temperaturu, vrati niz znakova koji pokazuje kako podesiti temperaturu u sobi prema ovim pragovima:

- Vrati `"heat"` ako je trenutna temperatura niža od ciljne temperature.
- Vrati `"cool"` ako je trenutna temperatura viša od ciljne temperature.
- Vrati `"hold"` ako je trenutna temperatura jednaka ciljnoj temperaturi.

# --hints--

`adjustThermostat(68, 72)` inapaswa kurudisha `"heat"`.

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)` inapaswa kurudisha `"cool"`.

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)` inapaswa kurudisha `"hold"`.

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)` inapaswa kurudisha `"heat"`.

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)` inapaswa kurudisha `"cool"`.

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)` inapaswa kurudisha `"hold"`.

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```
