---
id: 68b06e589bf2273243814775
title: "Zoezi la 38: Kizalishaji cha slug"
challengeType: 28
dashedName: challenge-38
---

# --description--
Ako naiđete na sekvencu karaktera, vratite odgovarajuću verziju za URL te sekvence koristeći sledeće uslove:

- Sve slovo moraju biti mala slova.
- Slova, brojevi ili nealfabetički razmaci, brojevi ili razmaci treba da se obrišu.
- Svi razmaci treba da se zamene sa kodom razmaka dodeljenim za URL `%20`.
- Ponovljeni razmaci treba da se izmene u samo `%20`.
- Vraćeni niz ne sme da počinje ili završi sa `%20`.

# --hints--

`generateSlug("helloWorld")` inapaswa kurudisha `"helloworld"`.

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")` inapaswa kurudisha `"hello%20world"`.

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")` inapaswa kurudisha `"helloworld"`.

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")` inapaswa kurudisha `"hello%20world"`.

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")` inapaswa kurudisha `"h3110%20w0r1d"`.

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
