---
id: 68b1f72371a5ac895ac70a06
title: "Zoezi la 42: Hifadhi ya video"
challengeType: 28
dashedName: challenge-42
---

# --description--
Uzimajući u obzir veličinu videa, jedinicu veličine videa, kapacitet hard diska i jedinicu hard diska, vratite broj videa koje hard disk može da snimi koristeći sljedeće kriterijume:

- Jedinica veličine videa može biti bajt(s) (`"B"`), kilobajt(s) (`"KB"`), megabajt(s) (`"MB"`) ili gigabajt(s) (`"GB"`).
- Ako nije navedena nijedna od gore pomenutih jedinica za video, vrati `"Invalid video unit"`.
- Jedinica kapaciteta tvrdog diska može biti gigabajt(s) (`"GB"`) ili terabajt(s) (`"TB"`).
- Ako nije navedena nijedna od gore pomenutih jedinica za tvrdog disk, vrati `"Invalid drive unit"`.
- Vrati broj punih videa koje disk može da snimi.
- Koristite sledeće parametre:

| Kitengo | Sawia na |
|:-------:|:--------:|
| 1 B     |   1 B    |
| 1 KB    |  1000 B  |
| 1 MB    |  1000 KB |
| 1 GB    |  1000 MB |
| 1 TB    |  1000 GB |

Na primer, ako primi `500`, `"MB"`, `100`, na `"GB"` kao argumente, odredi koliko je videa od 500 MB može se sačuvati na glavnom disku od 100 GB.

# --hints--

`numberOfVideos(500, "MB", 100, "GB")` inapaswa kurudisha `200`.

```js
assert.equal(numberOfVideos(500, "MB", 100, "GB"), 200);
```

`numberOfVideos(1, "TB", 10, "TB")` inapaswa kurudisha `"Invalid video unit"`.

```js
assert.equal(numberOfVideos(1, "TB", 10, "TB"), "Invalid video unit");
```

`numberOfVideos(2000, "MB", 100000, "MB")` inapaswa kurudisha `"Invalid drive unit"`.

```js
assert.equal(numberOfVideos(2000, "MB", 100000, "MB"), "Invalid drive unit");
```

`numberOfVideos(500000, "KB", 2, "TB")` inapaswa kurudisha `4000`.

```js
assert.equal(numberOfVideos(500000, "KB", 2, "TB"), 4000);
```

`numberOfVideos(1.5, "GB", 2.2, "TB")` inapaswa kurudisha `1466`.

```js
assert.equal(numberOfVideos(1.5, "GB", 2.2, "TB"), 1466);
```

# --seed--

## --seed-contents--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {

  return videoSize;
}
```

# --solutions--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {
  const videoUnits = { "KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000 };
  const driveUnits = { "GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000 };

  if (!videoUnits[videoUnit]) return "Invalid video unit";
  if (!driveUnits[driveUnit]) return "Invalid drive unit";

  const videoBytes = videoSize * videoUnits[videoUnit];
  const driveBytes = driveSize * driveUnits[driveUnit];

  return Math.floor(driveBytes / videoBytes);
}
```
