---
id: 68b1f72371a5ac895ac70a08
title: "Vežba 43: Brojevi protiv slova"
challengeType: 28
dashedName: challenge-43
---

# --description--

Ako je dat niz znakova, vrati `"digits"` ako niz sadrži više brojeva nego slova, `"letters"` ako sadrži više slova nego brojeva, i `"tie"` ako ima isti broj brojeva i slova.

- Nambari zinajumuisha `0-9`.
- Slova uključuju `a-z` u velikim ili malim slovima.
- Puuza herufi nyingine yoyote.

# --hints--

`digitsOrLetters("abc123")` inapaswa kurudisha `"tie"`.

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")` inapaswa kurudisha `"letters"`.

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")` inapaswa kurudisha `"digits"`.

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")` inapaswa kurudisha `"letters"`.

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")` inapaswa kurudisha `"digits"`.

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")` inapaswa kurudisha `"tie"`.

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```
