---
id: 68b1f72371a5ac895ac70a04
title: "Zoezi la 41: Hifadhi ya faili"
challengeType: 28
dashedName: challenge-41
---

# --description--
S obzirom na veličinu fajla, jedinice veličine fajlova i kapacitet diska u gigabajtima (GB), vratite broj fajlova koje hard disk može da sačuva koristeći sledeće kriterijume:

- Jedinice veličine fajlova mogu biti bait (`"B"`), kilobajt (`"KB"`), ili megabajt (`"MB"`).
- Vraća broj kompletnih fajlova koje disk može da primi.
- Koristite sledeće promene:

| Kitengo | Sawia na |
|:-------:|:--------:|
| 1 B     |   1 B    |
| 1 KB    |  1000 B  |
| 1 MB    |  1000 KB |
| 1 GB    |  1000 MB |

Na primer, ako imate `500`, `"KB"`, na `1` kao pretpostavku, odredite koliko fajlova od 500 KB može biti sačuvano na hard disku od 1 GB.

# --hints--

`numberOfFiles(500, "KB", 1)` inapaswa kurudisha `2000`.

```js
assert.equal(numberOfFiles(500, "KB", 1), 2000);
```

`numberOfFiles(50000, "B", 1)` inapaswa kurudisha `20000`.

```js
assert.equal(numberOfFiles(50000, "B", 1), 20000);
```

`numberOfFiles(5, "MB", 1)` inapaswa kurudisha `200`.

```js
assert.equal(numberOfFiles(5, "MB", 1), 200);
```

`numberOfFiles(4096, "B", 1.5)` inapaswa kurudisha `366210`.

```js
assert.equal(numberOfFiles(4096, "B", 1.5), 366210);
```

`numberOfFiles(220.5, "KB", 100)` inapaswa kurudisha `453514`.

```js
assert.equal(numberOfFiles(220.5, "KB", 100), 453514);
```

`numberOfFiles(4.5, "MB", 750)` inapaswa kurudisha `166666`.

```js
assert.equal(numberOfFiles(4.5, "MB", 750), 166666);
```

# --seed--

## --seed-contents--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {

  return fileSize;
}
```

# --solutions--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {
  const driveSizeBytes = driveSizeGb * 1000 * 1000 * 1000;

  let fileSizeBytes;
  if (fileUnit === "B") {
    fileSizeBytes = fileSize;
  } else if (fileUnit === "KB") {
    fileSizeBytes = fileSize * 1000;
  } else {
    fileSizeBytes = fileSize * 1000 * 1000;
  }

  return Math.floor(driveSizeBytes / fileSizeBytes);
}
```
