---
id: 68b1f72371a5ac895ac70a0a
title: "Vežba 44: Ogledalo niske"
challengeType: 28
dashedName: challenge-44
---

# --description--
Uzimajući u obzir sekvencu od dva slova, utvrdite da li je druga sekvenca ogledalo prve.

- Sekvenca slova se smatra palindromom ako ima ista slova u obrnutom redosledu.
- Razmatrati velika i mala slova kao različita.
- Ignorišite slova koja nisu deo celog alfabeta.

# --hints--

`isMirror("helloworld", "helloworld")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
