---
id: 68b7687dded630607aceccab
title: "Zoezi la 45: Mraba kamili"
challengeType: 28
dashedName: challenge-45
---

# --description--

Ako je dat ceo broj, odredi da li je savršen kvadrat.

- Broj je savršeno kvadratan ako ga možeš dobiti množenjem celog broja sa njim samim. Na primer, 9 je savršeno kvadratan broj zato što možeš pomnožiti 3 sa njim samim i dobiti 9.

# --hints--

`isPerfectSquare(9)` inapaswa kurudisha `true`.

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)` inapaswa kurudisha `true`.

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)` inapaswa kurudisha `true`.

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)` inapaswa kurudisha `false`.

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)` inapaswa kurudisha `false`.

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)` inapaswa kurudisha `false`.

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)` inapaswa kurudisha `true`.

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)` inapaswa kurudisha `true`.

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```
