---
id: 68b7687dded630607aceccb3
title: "Zoezi la 49: Mchambuzi wa kichwa cha CSV"
challengeType: 28
dashedName: challenge-49
---

# --description--
Kada uzmete prvi red iz fajla sa vrednostima razdvojenim zarezom (CSV), vratite red podataka koji sadrži zaglavlja kolona.

- Prva linija fajla CSV ima zaglavlja polja razdvojena zarezima.
- Uklonite sve razmake na početku ili kraju iz svakog zaglavlja polja.

# --hints--

`getHeadings("name,age,city")` treba da vrati `["name", "age", "city"]`.

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")` treba da vrati `["first name", "last name", "phone"]`.

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")` inapaswa kurudisha `["username", "email", "signup date"]`.

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```
