---
id: 68b7687dded630607aceccb1
title: "Zoezi la 48: Kigeuzi cha spam"
challengeType: 28
dashedName: challenge-48
---

# --description--

Ako je broj telefona u formatu `"+A (BBB) CCC-DDDD"`, gde svaki znak predstavlja cifru na sledeći način:

- `A` predstavlja kod zemlje i može imati bilo koji broj cifara.
- `BBB` predstavlja regionalni kod i uvek će imati tri cifre.
- `CCC` i `DDDD` predstavljaju lokalni broj, a uvek će imati po tri i četiri cifre, respektivno.

Odredite da li je reč o spam broju prema sledećim kriterijumima:

- Kod države mora imati više od 2 cifre ili ne početi sa nultom (`0`).
- Kod područja mora biti veći od 900 ili manji od 200.
- Prvih tri cifre lokalnog broja moraju biti prisutne unutar zadnje četiri cifre lokalnog broja.
- Broj mora imati istu cifru četiri ili više puta zaredom (bez obzira na strukturalne slova).

# --hints--

`isSpam("+0 (200) 234-0182")` inapaswa kurudisha `false`.

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` inapaswa kurudisha `true`.

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` inapaswa kurudisha `false`.

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
