---
id: 68b7cadffed0e75a517da675
title: "Vežba 53: Korišćenje pozitivnog celog broja za konverziju iz decimalnog u binarni sistem"
challengeType: 28
dashedName: challenge-53
---

# --description--

Ako je dat nenegativan ceo broj, vrati njegov binarni prikaz kao niz znakova.

Binarni broj koristi samo cifre `0` i `1` za predstavljanje bilo kog broja. Da bi se decimalni broj pretvorio u binarni, deli broj uzastopno sa `2` i beleži ostatke. Ponavljaj dok broj ne postane nula. Pročitaj zabeležene ostatke od poslednjeg do prvog. Na primer, pretvaranje `12` u binarni oblik:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` katika binary ni `1100`.

# --hints--

`toBinary(5)` inapaswa kurudisha `"101"`.

```js
assert.equal(toBinary(5), "101");
```

`toBinary(12)` inapaswa kurudisha `"1100"`.

```js
assert.equal(toBinary(12), "1100");
```

`toBinary(50)` inapaswa kurudisha `"110010"`.

```js
assert.equal(toBinary(50), "110010");
```

`toBinary(99)` inapaswa kurudisha `"1100011"`.

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```
