---
id: 68c1a929005bf54d342aa8d3
title: "Zoezi la 56: Wiki ya Anga la 2: Utafutaji wa Sayari Nje ya Mfumo"
challengeType: 28
dashedName: challenge-56
---

# --description--
Na drugi dan Nedelje svemira/astronomije, dobijaš niz slova gde svako slovo predstavlja merenje sjajnosti zvezde. Odredi da li su merenja otkrila egzoplanetu korišćenjem tranzitne metode (transit method). Tranzitna metoda je kada planeta prođe ispred zvezde, smanjujući vidljivu sjajnost.

- Mjerenja osvetljenosti uključuju samo karaktere `0-9` i `A-Z`, gde se svako merenje odnosi na sledeće numeričke vrednosti:
- Karakteri `0-9` predstavljaju nivoe osvetljenosti `0-9`.
- Karakteri `A-Z` predstavljaju nivoe osvetljenosti `10-35`.

Zvezda se smatra da ima eksoplanetu ako je bilo koja jedna merač manja ili jednaka 80 procenata proseka svih merenja. Na primer, ako je prosečna svetlina zvezde 10, smatraće se da ima eksoplanetu ako je bilo koja jedna merač 8 ili niža.

# --hints--

`hasExoplanet("665544554")` inapaswa kurudisha `false`.

```js
assert.isFalse(hasExoplanet("665544554"));
```

`hasExoplanet("FGFFCFFGG")` inapaswa kurudisha `true`.

```js
assert.isTrue(hasExoplanet("FGFFCFFGG"));
```

`hasExoplanet("MONOPLONOMONPLNOMPNOMP")` inapaswa kurudisha `false`.

```js
assert.isFalse(hasExoplanet("MONOPLONOMONPLNOMPNOMP"));
```

`hasExoplanet("FREECODECAMP")` inapaswa kurudisha `true`.

```js
assert.isTrue(hasExoplanet("FREECODECAMP"));
```

`hasExoplanet("9AB98AB9BC98A")` inapaswa kurudisha `false`.

```js
assert.isFalse(hasExoplanet("9AB98AB9BC98A"));
```

`hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` inapaswa kurudisha `true`.

```js
assert.isTrue(hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"));
```

# --seed--

## --seed-contents--

```js
function hasExoplanet(readings) {

  return readings;
}
```

# --solutions--

```js
function hasExoplanet(readings) {
  let total = 0;
  const values = readings.split('').map(c => parseInt(c, 36));
  
  total = values.reduce((sum, v) => sum + v, 0);
  const average = total / values.length;
  const threshold = average * 0.8;

  return values.some(v => v <= threshold);
}
```
