---
id: 68c1a929005bf54d342aa8d5
title: "Vežba 58: Nedelja u svemiru Dan 4: Sekcija sletanja
challengeType: 28
dashedName: challenge-58
---

# --description--
Na četvrti dan Svemirske nedelje, dobijate matricu brojeva (mrežu podataka redova podataka), koja predstavlja zone gde vaš rover može da sleti. Pronađite najbezbednije mesto za slanje prema sledećim pravilima:

- Svaka ćelija u matrici će imati broj od `0-9`, uključujući.
- Svaki `0` predstavlja moguće mesto za sletanje.
- Svaki broj, osim `0`, je veoma opasan za sletanje. Kako se broj povećava, tako raste i opasnost.
- Najbezbednije mesto definisano je kao ćelija `0` čije susedne ćelije (do 4 suseda, bez uzimanja uglova u obzir) imaju najniži ukupni nivo opasnosti.
- Ignorišite susede van granica (uglovi i ivice imaju samo malo suseda).
- Vratite indikatore najbezbednijeg mesta za sletanje. Uvek će postojati samo jedno najbezbednije mesto.

Na primer, ako dobijete:

```js
[
  [1, 0],
  [2, 0]
]
```

Vrati `[0, 1]`, indikatore za `0` u prvom redu.

# --hints--

`findLandingSpot([[1, 0], [2, 0]])` inapaswa kurudisha `[0, 1]`.

```js
assert.deepEqual(findLandingSpot([[1, 0], [2, 0]]), [0, 1]);
```

`findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` inapaswa kurudisha `[1, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1]);
```

`findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` inapaswa kurudisha `[2, 2]`.

```js
assert.deepEqual(findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2]);
```

`findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` inapaswa kurudisha `[2, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1]);
```

# --seed--

## --seed-contents--

```js
function findLandingSpot(matrix) {

  return matrix;
}
```

# --solutions--

```js
function findLandingSpot(matrix) {
  let bestSpot = null;
  let lowestNeighborSum = Infinity;

  for (let i = 0; i < matrix.length; i++) {
    for (let j = 0; j < matrix[i].length; j++) {

      if (matrix[i][j] === 0) {
        let currentNeighborSum = 0;

        if (i > 0) currentNeighborSum += matrix[i - 1][j];
        if (j < matrix[i].length - 1) currentNeighborSum += matrix[i][j + 1];
        if (i < matrix.length - 1) currentNeighborSum += matrix[i + 1][j];
        if (j > 0) currentNeighborSum += matrix[i][j - 1];

        if (currentNeighborSum < lowestNeighborSum) {
          lowestNeighborSum = currentNeighborSum;
          bestSpot = [i, j];
        }
      }
    }
  }

  return bestSpot;
}
```
