---
id: 68c1a929005bf54d342aa8d6
title: "Zoezi la 59: Wiki ya Anga Siku ya 5: Eneo la Goldilocks"
challengeType: 28
dashedName: challenge-59
---

# --description--
Na petom danu Svemirske nedelje, ćete izračunati "Goldiloks zonu" zvezde – oblast koja okružuje zvezdu gde su uslovi "savršeni" za postojanje tečne vode.

Uzimajući u obzir masu zvezde, vrati red podataka sa početnom i krajnjom udaljenošću njegove Zone Goldiloks za astrološke mere.

Ili kuhesabu Eneo la Goldilocks:

1. Izračunajte luminostet zvezde podizanjem njene intenzivnosti na stepen 3,5.
2. Početak površine je 0,95 puta kvadratnog korena njenog sjaja/luminositeta.
3. Kraj površine je 1,37 puta kvadratnog korena njenog sjaja/luminositeta.

- Zaokruži vrednost na dve decimale.

Na primer, ako dobijete `1` kao težinu, vratite `[0.95, 1.37]`.

# --hints--

`goldilocksZone(1)` inapaswa kurudisha `[0.95, 1.37]`.

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)` inapaswa kurudisha `[0.28, 0.41]`.

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)` inapaswa kurudisha `[21.85, 31.51]`.

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)` inapaswa kurudisha `[9.38, 13.52]`.

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)` inapaswa kurudisha `[179.69, 259.13]`.

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```
