---
id: 68cae5b538ff798bbd4da005
title: "Zoezi la 67: Kagua barua pepe"
challengeType: 28
dashedName: challenge-67
---

# --description--
Uz dati niz znakova, odredite da li je važeća e-poštanska adresa koristeći sledeće kriterijume:

- Mora da ima samo jedan karakter `@`.
- Interni deo (pre `@`):
  - Može sadržati samo slova (`a-z`, `A-Z`), brojeve (`0-9`), tačke (`.`), donji crte (`_`) ili tirete (`-`).
  - Ne sme počinjati niti završavati sa tačkom.
- Domen deo (nakon `@`):
  - Mora imati najmanje jednu tačku.
  - Mora se završavati tačkom praćenom najmanje dva slova.
- Interni deo niti domen deo ne sme imati dve uzastopne tačke.

# --hints--

`validate("a@b.cd")` inapaswa kurudisha `true`.

```js
assert.isTrue(validate("a@b.cd"));
```

`validate("hell.-w.rld@example.com")` inapaswa kurudisha `true`.

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

`validate(".b@sh.rc")` inapaswa kurudisha `false`.

```js
assert.isFalse(validate(".b@sh.rc"));
```

`validate("example@test.c0")` inapaswa kurudisha `false`.

```js
assert.isFalse(validate("example@test.c0"));
```

`validate("freecodecamp.org")` inapaswa kurudisha `false`.

```js
assert.isFalse(validate("freecodecamp.org"));
```

`validate("develop.ment_user@c0D!NG.R.CKS")` inapaswa kurudisha `true`.

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

`validate("hello.@wo.rld")` inapaswa kurudisha `false`.

```js
assert.isFalse(validate("hello.@wo.rld"));
```

`validate("hello@world..com")` inapaswa kurudisha `false`.

```js
assert.isFalse(validate("hello@world..com"));
```

`validate("develop..ment_user@c0D!NG.R.CKS")` inapaswa kurudisha `false`.

```js
assert.isFalse(validate("develop..ment_user@c0D!NG.R.CKS"));
```

`validate("git@commit@push.io")` treba da vrati `false`.

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```
