---
id: 68cae5b538ff798bbd4da001
title: "Zoezi la 63: Mapambano ya Maneno"
challengeType: 28
dashedName: challenge-63
---

# --description--
Uzimajući u obzir dve rečenice koje predstavljaju vaš tim i protivnički tim, gde svaka reč iz vašeg tima se takmiči sa odgovarajućom reči protivničkog tima, utvrdite koji tim pobeđuje koristeći sledeća pravila:

- Rečenice koje se daju će uvek imati isti broj reči.
- Reči su razdvojene samo razmakom i sadržiće samo slova.
- Vrednost svake reči je zbir vrednosti njenih slova.
- Slova `a` do `z` odgovaraju vrednosti `1` do `26`. Na primer, `a` je `1`, a `z` je `26`.
- Velika slova udvostručuju vrednost tog slova. Na primer, `A` je `2`, a `Z` je `52`.
- Reči se takmiče po redu: prva reč vašeg tima takmiči se sa prvom rečju protivničkog tima, i tako dalje.
- Reč pobeđuje ako je njena vrednost veća od vrednosti reči protivnika.
- Pobednički tim je onaj koji ima najviše pobedničkih reči.

Vrati `"We win"` ako je tvoj tim pobednik, `"We lose"` ako tvoj tim gubi, i `"Draw"` ako oba tima imaju isti broj pobeda.

# --hints--

`battle("hello world", "hello word")` inapaswa kurudisha `"We win"`.

```js
assert.equal(battle("hello world", "hello word"), "We win");
```

`battle("Hello world", "hello world")` inapaswa kurudisha `"We win"`.

```js
assert.equal(battle("Hello world", "hello world"), "We win");
```

`battle("lorem ipsum", "kitty ipsum")` inapaswa kurudisha `"We lose"`.

```js
assert.equal(battle("lorem ipsum", "kitty ipsum"), "We lose");
```

`battle("hello world", "world hello")` inapaswa kurudisha `"Draw"`.

```js
assert.equal(battle("hello world", "world hello"), "Draw");
```

`battle("git checkout", "git switch")` inapaswa kurudisha `"We win"`.

```js
assert.equal(battle("git checkout", "git switch"), "We win");
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` inapaswa kurudisha `"We lose"`.

```js
assert.equal(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose");
```

`battle("We must never surrender", "Our team must win")` inapaswa kurudisha `"Draw"`.

```js
assert.equal(battle("We must never surrender", "Our team must win"), "Draw");
```

# --seed--

## --seed-contents--

```js
function battle(ourTeam, opponent) {

  return ourTeam;
}
```

# --solutions--

```js
function getCharacterValue(char) {
  const lowerCaseValue = char.toLowerCase().charCodeAt(0) - 96;
  return char >= 'A' && char <= 'Z' ? lowerCaseValue * 2 : lowerCaseValue;
}

function getWordValue(word) {
  let wordValue = 0;
  for (let i=0; i<word.length; i++) {
    wordValue += getCharacterValue(word[i])
  }

  return wordValue;
}

function battle(ourTeam, opponent) {
  const myWords = ourTeam.split(' ');
  const opponentWords = opponent.split(' ');

  let myWins = 0, opponentWins = 0;

  for (let i=0; i<myWords.length; i++) {
    const myWordValue = getWordValue(myWords[i]);
    const opponentWordValue = getWordValue(opponentWords[i])

    if (myWordValue > opponentWordValue) myWins++;
    if (opponentWordValue > myWordValue) opponentWins++;
  }

  if (myWins > opponentWins) return 'We win';
  if (opponentWins > myWins) return 'We lose';
  return 'Draw';
}
```
