---
id: 68cae5b538ff798bbd4da004
title: "Vežba 66: Uklanjanje HTML oznaka"
challengeType: 28
dashedName: challenge-66
---

# --description--

Dat je niz znakova HTML koda, ukloni oznake i vrati običan tekstualni sadržaj.

- Sekvenca znakova unosa će imati samo validan HTML.
- Etikete za HTML mogu biti interno konfigurisane.
- Uklonite sve oznake i atribute elemenata.

Na primer, `'<a href="#">Click here</a>'` treba da vrati `"Click here"`.

# --hints--

`stripTags('<a href="#">Click here</a>')` inapaswa kurudisha `"Click here"`.

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

`stripTags('<p class="center">Hello <b>World</b>!</p>')` inapaswa kurudisha `"Hello World!"`.

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

`stripTags('<img src="cat.jpg" alt="Cat">')` treba da vrati prazan niz znakova (`""`).

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

`stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` inapaswa kurudisha `sectionsection`.

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```
