---
id: 68cae5b538ff798bbd4da003
title: "Vežba 65: Brojanje niza znakova"
challengeType: 28
dashedName: challenge-65
---

# --description--
Uzimajući u obzir sekvencu od dve slova, utvrdite koliko puta se sekvenca drugog slova pojavljuje u prvom.

- Sekvenca uzorka može se preklapati u sekvencu prvog slova/karaktera. Na primer, `"aaa"` ima `"aa"` dva puta. `a` dve prve i `a` dve druge.

# --hints--

`count('abcdefg', 'def')` inapaswa kurudisha `1`.

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')` inapaswa kurudisha `0`.

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')` inapaswa kurudisha `2`.

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')` inapaswa kurudisha `3`.

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')` inapaswa kurudisha `10`.

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
