---
id: 68cae5b538ff798bbd4da007
title: "Zoezi la 69: Soksi Zilizokosekana"
challengeType: 28
dashedName: challenge-69
---

# --description--
Uzimajući u obzir ukupan broj koji predstavlja početni broj para čarapa, i drugi broj koji predstavlja koliko ste vešeračkih ciklusa uradili, vratite trenutni broj kompletnih para čarapa koristeći sledeće uslove:

- Svakih 2 ciklusa pranja, gubiš jednu čarapeu.
- Svakih 3 ciklusa pranja, nalaziš jednu izgubljenu čarapeu.
- Svakih 5 ciklusa pranja, jedna čarapa se ošteti i mora se baciti.
- Svakih 10 ciklusa pranja, kupuješ par novih čarapa.
- Ne možeš imati negativan broj čarapa.
- Pravila se mogu kombinovati. Na primer, u ciklusu pranja od 10, gubiš jednu čarapeu, bacaš jednu čarapeu i kupuješ par novih čarapa.
- Vrati broj kompletnih para čarapa.

# --hints--

`sockPairs(2, 5)` inapaswa kurudisha `1`.

```js
assert.equal(sockPairs(2, 5), 1);
```

`sockPairs(1, 2)` inapaswa kurudisha `0`.

```js
assert.equal(sockPairs(1, 2), 0);
```

`sockPairs(5, 11)` inapaswa kurudisha `4`.

```js
assert.equal(sockPairs(5, 11), 4);
```

`sockPairs(6, 25)` inapaswa kurudisha `3`.

```js
assert.equal(sockPairs(6, 25), 3);
```

`sockPairs(1, 8)` inapaswa kurudisha `0`.

```js
assert.equal(sockPairs(1, 8), 0);
```

# --seed--

## --seed-contents--

```js
function sockPairs(pairs, cycles) {

  return pairs;
}
```

# --solutions--

```js
function sockPairs(pairs, cycles) {
  let socks = pairs * 2;

  for (let i = 1; i <= cycles; i++) {
    if (i % 2 === 0) socks -= 1;
    if (i % 3 === 0) socks += 1;
    if (i % 5 === 0) socks -= 1;
    if (i % 10 === 0) socks += 2;

    if (socks < 0) socks = 0;
  }

  return Math.floor(socks / 2);
}
```
