---
id: 68cae5b538ff798bbd4da009
title: "Zoezi la 71: Kihesabu cha Tipi"
challengeType: 28
dashedName: challenge-71
---

# --description--
Uzimajući u obzir cenu vaše hrane i željeni procenat naplate, vratite red podataka sa tri vrednosti naplate: 15%, 20% i željeni iznos naplate.

- Cena će biti data u formatu: `"$N.NN"`.
- Procenti standardnih tipova će biti dati ovim formatom: `"25%"`.
- Zaokružite iznos na isti format kao i `"$N.NN"`, zaokružen do dve decimalne mesta.

Na primer, uzimajući u obzir cenu hrane za `"$10.00"`, i vrednost tipične vrste od `"25%"`, vratite `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculateTips("$10.00", "25%")` inapaswa kurudisha `["$1.50", "$2.00", "$2.50"]`.

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")` inapaswa kurudisha `["$13.45", "$17.93", "$23.31"]`.

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")` inapaswa kurudisha `["$2.98", "$3.97", "$1.79"]`.

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```
