---
id: 68cae5b538ff798bbd4da00a
title: "Zoezi la 72: Kirekebishaji cha Thermostat 2"
challengeType: 28
dashedName: challenge-72
---

# --description--
Uzimajući u obzir trenutnu temperaturu sobe u Farenhajtu i željenu temperaturu u Celzijusima, vratite niz znakova koji opisuje kako podesiti temperaturu u sobi prema ovim uslovima:

- Vrati `"Heat: X degrees Fahrenheit"` ako je trenutna temperatura ispod cilja. Gde su `X` Fahrenheit stepeni za zagrevanje prostorije do cilja, zaokruženi na jednu decimalnu tačku.
- Vrati `"Cool: X degrees Fahrenheit"` ako je trenutna temperatura iznad cilja. Gde su `X` Fahrenheit stepeni za hlađenje prostorije do cilja, zaokruženi na jednu decimalnu tačku.
- Vrati `"Hold"` ako je trenutna temperatura jednaka cilju.

Da biste pretvorili Celzijus u Farenhajt, pomnožite Celzijus temperaturu sa 1.8, a zatim dodajte 32 na rezultat (`F = (C * 1.8) + 32`).

# --hints--

`adjustThermostat(32, 0)` inapaswa kurudisha `"Hold"`.

```js
assert.equal(adjustThermostat(32, 0), "Hold");
```

`adjustThermostat(70, 25)` inapaswa kurudisha `"Heat: 7.0 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(70, 25), "Heat: 7.0 degrees Fahrenheit");
```

`adjustThermostat(72, 18)` inapaswa kurudisha `"Cool: 7.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(72, 18), "Cool: 7.6 degrees Fahrenheit");
```

`adjustThermostat(212, 100)` inapaswa kurudisha `"Hold"`.

```js
assert.equal(adjustThermostat(212, 100), "Hold");
```

`adjustThermostat(59, 22)` inapaswa kurudisha `"Heat: 12.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(59, 22), "Heat: 12.6 degrees Fahrenheit");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(currentF, targetC) {

  return currentF;
}
```

# --solutions--

```js
function adjustThermostat(currentF, targetC) {
  const targetF = targetC * 1.8 + 32;
  const diff = Math.abs(targetF - currentF).toFixed(1);

  if (currentF < targetF) {
    return `Heat: ${diff} degrees Fahrenheit`;
  } else if (currentF > targetF) {
    return `Cool: ${diff} degrees Fahrenheit`;
  } else {
    return "Hold"
  }
}
```
