---
id: 68cae5b538ff798bbd4da008
title: "Vežba 70: Filter za atribute HTML elemenata
challengeType: 28
dashedName: challenge-70
---

# --description--
Ako dobijete niz znakova za važeći atribut od HTML, vratite atribute tog elementa koristeći sledeće parametre:

- Biće vam dat samo jedan element.
- Atributi će biti u formatu: `attribute="value"`.
- Vratite red podataka koji sadrži svaki atribut elemenata i njegovu vrednost, razdvojene zarezima, u ovom formatu: `["attribute1, value1", "attribute2, value2"]`.
- Vratite atribute u redosledu kojim su dati.
- Ako nema dostupnih atributa, vratite prazan red podataka.

# --hints--

`extractAttributes('<span class="red"></span>')` inapaswa kurudisha `["class, red"]`.

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` inapaswa kurudisha `["charset, UTF-8"]`.

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` inapaswa kurudisha `[]`.

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` treba da vrati `["name, email", "type, email", "required, true"]`.

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` inapaswa kurudisha `["id, submit", "class, btn btn-primary"]`.

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
