---
id: 68cae5b538ff798bbd4da006
title: "Zoezi la 68: Kifichuzi cha Kadi ya Mikopo"
challengeType: 28
dashedName: challenge-68
---

# --description--
Izdavanje sekvence brojeva kreditne kartice, vratite skrivenu verziju koristeći sledeće uslove:

- Sekvenca će imati četiri seta od po četiri broja svaki (`0-9`), i svi setovi će biti razdvojeni jednim razmakom ili jednom oznakom križa (`-`).
- Promenite sve brojeve osim poslednje četiri u zvezdice (`*`).
- Ostavite preostala slova nepromenjena.

Na primer, ako se ukloni `"4012-8888-8888-1881"`, vratiti `"****-****-****-1881"`.

# --hints--

`mask("4012-8888-8888-1881")` inapaswa kurudisha `"****-****-****-1881"`.

```js
assert.equal(mask("4012-8888-8888-1881"), "****-****-****-1881");
```

`mask("5105 1051 0510 5100")` inapaswa kurudisha `"**** **** **** 5100"`.

```js
assert.equal(mask("5105 1051 0510 5100"), "**** **** **** 5100");
```

`mask("6011 1111 1111 1117")` inapaswa kurudisha `"**** **** **** 1117"`.

```js
assert.equal(mask("6011 1111 1111 1117"), "**** **** **** 1117");
```

`mask("2223-0000-4845-0010")` inapaswa kurudisha `"****-****-****-0010"`.

```js
assert.equal(mask("2223-0000-4845-0010"), "****-****-****-0010");
```

# --seed--

## --seed-contents--

```js
function mask(card) {

  return card;
}
```

# --solutions--

```js
function mask(card) {
  const split = card.split(card.includes('-') ? '-' : ' ');
  
  return card.includes('-') ?
    `****-****-****-${split[3]}` :
    `**** **** **** ${split[3]}`;
}
```
