---
id: 68d2ba1468508398389487d0
title: "Zoezi la 75: Hazina iliyofichwa"
challengeType: 28
dashedName: challenge-75
---

# --description--
Daje se 2D niz podataka koji predstavlja mapu dna mora sa skrivenim blagom, i niz podataka sa koordinatama (`[row, column]`) za plivanje do vaše blago-traženja, vratite `"Empty"`, `"Found"`, ili `"Recovered"` koristeći sledeće pravilo:

- Generisani 2D niz podataka će sadržati samo jednu neispravljenu vrednost koja će zauzeti mnogo ćelija.
- Svaka ćelija u 2D nizu podataka sadrži jednu od sledećih vrednosti:
  - `"-"`: Nema vrednosti.
  - `"O"`: Deo vrednosti koji nije dostupan.
  - `"X"`: Deo vrednosti koji je već dostupan.

- Ako zona nema blaga, vrati `"Empty"`.
- Ako zona pronalazi blago, ali najmanje deo tog blaga još uvek nije dostupan, vrati `"Found"`.
- Ako zona pronalazi poslednji deo blaga koji nije dostupan, vrati `"Recovered"`.

Na primer, ako se ukloni:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

Na `[2, 1]` za koordinaciju površine za plivanje, vrati `"Recovered"` jer je plivanje pronašlo poslednji deo blaga koji nije bio dostupan.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` inapaswa kurudisha `"Recovered"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered");
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` inapaswa kurudisha `"Empty"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty");
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` treba da vrati `"Found"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found");
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` treba da vrati `"Found"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` inapaswa kurudisha `"Recovered"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` inapaswa kurudisha `"Empty"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty");
```

# --seed--

## --seed-contents--

```js
function dive(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function dive(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Empty";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Found";
        }
      }
    }
    return "Recovered";
  }

  return "Found";
}
```
