---
id: 68d30fc57588d97fd3027b30
title: "Zoezi la 79: Navigator"
challengeType: 28
dashedName: challenge-79
---

# --description--

Mnamo Oktoba 28, 1994, Netscape Navigator ilizinduliwa, ikisaidia mamilioni ya watu kuchunguza mtandao wa mwanzo.

S obzirom na niz podataka o komandama pretraživača koje ste izveli na Netscape Navigator-u, obnovite trenutnu stranicu koju imate nakon izvršavanja svih komandi koristeći sledeće pravilo:

- Uvek počinjete na stranici `"Home"`, koja nije uključena u listu komandi.
- Važeće komande su:
  - `"Visit Page"`: Gde je `"Page"` naziv stranice koju posetite. Na primer, `"Visit About"` vodi vas na stranicu `"About"`. Kada posetite novu stranicu, uverite se da obrišete bilo koju istoriju napredovanja koju imate.
  - `"Back"`: Vodi vas na prethodno posetu stranicu u vašoj istoriji ili ostaje na trenutnoj stranici ako ne postoji.
  - `"Forward"`: Vodi vas napred u istoriji do stranice sa koje ste otišli ili ostaje na trenutnoj stranici ako ne postoji.

Na primer, ako dobijete `["Visit About Us", "Back", "Forward"]`, vratite `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` inapaswa kurudisha `"About Us"`.

```js
assert.equal(navigate(["Visit About Us", "Back", "Forward"]), "About Us");
```

`navigate(["Forward"])` inapaswa kurudisha `"Home"`.

```js
assert.equal(navigate(["Forward"]), "Home");
```

`navigate(["Back"])` inapaswa kurudisha `"Home"`.

```js
assert.equal(navigate(["Back"]), "Home");
```

`navigate(["Visit About Us", "Visit Gallery"])` inapaswa kurudisha `"Gallery"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery"]), "Gallery");
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` inapaswa kurudisha `"Home"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home");
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` inapaswa kurudisha `"Contact"`.

```js
assert.equal(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact");
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` inapaswa kurudisha `"Visit Us"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us");
```

# --seed--

## --seed-contents--

```js
function navigate(commands) {

  return commands;
}
```

# --solutions--

```js
function navigate(commands) {
  const history = ["Home"];
  let currentPageIndex = 0;

  for (const command of commands) {
    if (command.startsWith("Visit")) {
      history.splice(currentPageIndex + 1);
      history.push(command.slice(6));
      currentPageIndex++;
    } else if(command === "Back" && currentPageIndex > 0) {
      currentPageIndex--;
    } else if (command === "Forward" && currentPageIndex < history.length - 1) {
      currentPageIndex++;
    }
  }

  return history[currentPageIndex];
}
```
