---
id: 68d30845cc08266018fc46bc
title: "Zoezi la 76: DNA ya nyongeza"
challengeType: 28
dashedName: challenge-76
---

# --description--

Data je niska koja predstavlja DNK sekvencu. Vrati njen komplementarni niz koristeći sledeća pravila:

- DNA uključuje slova `"A"`, `"C"`, `"G"`, na `"T"`.
- Slova `"A"` i `"T"` su opciona za svako.
- Slova `"C"` i `"G"` su opciona za svako.

Na primer, ako dobijete `"ACGT"`, vratite `"TGCA"`.

# --hints--

`complementaryDNA("ACGT")` inapaswa kurudisha `"TGCA"`.

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")` inapaswa kurudisha `"TACGCATGCAATCG"`.

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")` inapaswa kurudisha `"CCGAATGCTAGCTTC"`.

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")` inapaswa kurudisha `"CTAGATCGATCCGATCGATC"`.

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
