---
id: 68d30845cc08266018fc46be
title: "Vežba 78: Niz celih brojeva"
challengeType: 28
dashedName: challenge-78
---

# --description--
Ako pronađeš pozitivan pun broj, vrati niz karaktera koji prikazuje sve pune brojeve počevši od `1` do, uključujući taj broj, u numeričkom redosledu.

Na primer, ako dobiješ `5`, vrati `"12345"`.

# --hints--

`sequence(5)` inapaswa kurudisha `"12345"`.

```js
assert.equal(sequence(5), "12345");
```

`sequence(10)` inapaswa kurudisha `"12345678910"`.

```js
assert.equal(sequence(10), "12345678910");
```

`sequence(1)` inapaswa kurudisha `"1"`.

```js
assert.strictEqual(sequence(1), "1");
```

`sequence(27)` inapaswa kurudisha `"123456789101112131415161718192021222324252627"`.

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```
