---
id: 68d30845cc08266018fc46bd
title: "Zoezi la 77: Mtafiti wa muda"
challengeType: 28
dashedName: challenge-77
---

# --description--
Kada dobijete pun broj sekundi, vratite niz znakova koji predstavlja isto vreme u formatu `"H:MM:SS"`, gde je `"H"` broj sati, `"MM"` je broj minuta, a `"SS"` je broj sekundi. Vratite vreme koristeći sledeće pravila:

- Sekunde: Uvek moraju imati dve cifre.
- Minuti: Treba ukloniti vodeće nule kada nisu potrebne. Koristi `"0"` ako je vreme kraće od jednog minuta.
- Sati: Treba ih uključiti samo ako su veći od nule.

# --hints--

`format(500)` inapaswa kurudisha `"8:20"`.

```js
assert.equal(format(500), "8:20");
```

`format(4000)` inapaswa kurudisha `"1:06:40"`.

```js
assert.equal(format(4000), "1:06:40");
```

`format(1)` inapaswa kurudisha `"0:01"`.

```js
assert.equal(format(1), "0:01");
```

`format(5555)` inapaswa kurudisha `"1:32:35"`.

```js
assert.equal(format(5555), "1:32:35");
```

`format(99999)` inapaswa kurudisha `"27:46:39"`.

```js
assert.equal(format(99999), "27:46:39");
```

# --seed--

## --seed-contents--

```js
function format(seconds) {

  return seconds;
}
```

# --solutions--

```js
function format(seconds) {
  const h = Math.floor(seconds / 3600);
  const m = Math.floor((seconds % 3600) / 60);
  const s = seconds % 60;

  const secondsStr = s.toString().padStart(2, "0");
  const minutesStr = m.toString();

  if (h > 0) {
    const hoursStr = h.toString();
    return `${hoursStr}:${minutesStr.padStart(2, "0")}:${secondsStr}`;
  } else {
    return `${minutesStr}:${secondsStr}`;
  }
}
```
