---
id: 68e39ed6106dac2f0a98fd65
title: "Zoezi la 83: Uthibitisho wa Sahihi"
challengeType: 28
dashedName: challenge-83
---

# --description--
Dato je sekvenca znakova poruke, sekvenca znakova tajnog ključa i broj potpisa, određuje da li je potpis validan koristeći ovu metodu kodiranja:

- Karakteri u poruci i tajni ključeve imaju ove vrednosti:
  - `a` do `z` imaju vrednosti `1` do `26` respektivno.
  - `A` do `Z` imaju vrednosti `27` do `52` respektivno.
- Svi ostali karakteri nemaju vrednost.
- Izračunavanje je uzimanjem zbi poruke zajedno sa zbirom tajnih ključeva.

Na primer, ako se pošalje poruka `"foo"` i tajni ključ `"bar"`, potpis će biti `57`:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

Proveri da li se izračunati kontrolni broj poklapa sa datim kontrolnim brojem.

# --hints--

`verify("foo", "bar", 57)` inapaswa kurudisha `true`.

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)` inapaswa kurudisha `false`.

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)` inapaswa kurudisha `true`.

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)` inapaswa kurudisha `false`.

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)` inapaswa kurudisha `true`.

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` inapaswa kurudisha `true`.

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
