---
id: 68e39ed6106dac2f0a98fd66
title: "Zoezi la 84: Kuambukizwa"
challengeType: 28
dashedName: challenge-84
---

# --description--
2. novembra 1988. godine, oslobođen je prvi veliki internet vir, ostavljajući približno 10 procenata računara povezanih na internet da budu zaraženi nakon samo jednog dana.

U ovom izazovu, dobijate broj dana koji su prošli od oslobađanja internet buga, i morate da utvrdite koliko računara je zaraženo koristeći sledeća pravila:

- Na dan 0, prvi računar je zaražen.
- Svaki sledeći dan, broj zaraženih računara se udvostručuje.
- Svaki treći dan, koristi se vakcina nakon što virus proširi i smanji broj zaraženih računara za 20%. Zaokružite broj vakcinisanih računara na najbliži ceo broj.

Na primer, dan:

- Dan 0: računara 1 ukupno zaraženo.
- Dan 1: računara 2 ukupno zaraženo.
- Dan 2: računara 4 ukupno zaraženo.
- Dan 3: računara 8 ukupno zaraženo. Zatim, primeni imunitet: računara 8 zaraženih * procenat 20 = 1.6 dobili imunitet. Zaokruži 1.6 na 2. Računara 8 zaraženih - 2 dobili imunitet = računara 6 ukupno zaraženo nakon dana 3.

Rudisha idadi ya kompyuta jumla zilizambukizwa baada ya idadi ya siku zilizotolewa kupita.

# --hints--

`infected(1)` inapaswa kurudisha `2`.

```js
assert.equal(infected(1), 2);
```

`infected(3)` inapaswa kurudisha `6`.

```js
assert.equal(infected(3), 6);
```

`infected(8)` inapaswa kurudisha `152`.

```js
assert.equal(infected(8), 152);
```

`infected(17)` inapaswa kurudisha `39808`.

```js
assert.equal(infected(17), 39808);
```

`infected(25)` inapaswa kurudisha `5217638`.

```js
assert.equal(infected(25), 5217638);
```

# --seed--

## --seed-contents--

```js
function infected(days) {

  return days;
}
```

# --solutions--

```js
function infected(days) {
  let infected = 1;

  for (let day = 1; day <= days; day++) {
    infected *= 2;

    if (day % 3 === 0) {
      let patched = Math.ceil(infected * 0.2);
      infected -= patched;
    }
  }

  return infected;
}
```
