---
id: 68e39ed6106dac2f0a98fd64
title: "Zoezi la 82: SpOoKy~CaSe"
challengeType: 28
dashedName: challenge-82
---

# --description--
Ako vam je dat niz slova koji predstavlja ime parametra, promenite ga u "spooky case" koristeći sledeće uslove:

- Zameni sve podvučene znakove (`_`) i znake križa (`-`) sa tildom (`~`).
- Napišite prvi slovo sekvence velikim slovom, i svako drugo sledeće naizmenično slovo takođe velikim slovom. Ignorišite tilde znak prilikom brojanja. Sve ostale slova napišite malim slovima.

Na primer, ako primiš `hello_world`, vrati `HeLlO~wOrLd`.

# --hints--

`spookify("hello_world")` inapaswa kurudisha `"HeLlO~wOrLd"`.

```js
assert.equal(spookify("hello_world"), "HeLlO~wOrLd");
```

`spookify("Spooky_Case")` inapaswa kurudisha `"SpOoKy~CaSe"`.

```js
assert.equal(spookify("Spooky_Case"), "SpOoKy~CaSe");
```

`spookify("TRICK-or-TREAT")` inapaswa kurudisha `"TrIcK~oR~tReAt"`.

```js
assert.equal(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt");
```

`spookify("c_a-n_d-y_-b-o_w_l")` inapaswa kurudisha `"C~a~N~d~Y~~b~O~w~L"`.

```js
assert.equal(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L");
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` inapaswa kurudisha `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`.

```js
assert.equal(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS");
```

# --seed--

## --seed-contents--

```js
function spookify(boo) {

  return boo;
}
```

# --solutions--

```js
function spookify(boo) {
  const replaced = boo.replace(/[_-]/g, "~");
  
  let result = "";
  let capitalize = true;

  for (let char of replaced) {
    if (char === "~") {
      result += char;
    } else {
      result += capitalize ? char.toUpperCase() : char.toLowerCase();
      capitalize = !capitalize;
    }
  }

  return result;
}
```
