---
id: 68e39ed6106dac2f0a98fd62
title: "Zoezi la 80: Mpangaji wa barua pepe"
challengeType: 28
dashedName: challenge-80
---

# --description--
Na 29. oktobra 1971., poslata je prva e-pošta, uspostavljajući format `username@domain` koji i dalje koristimo. Sada postoji milijarde adresa za e-poštu.

Za ovim izazovom, dobijate listu e-mail adresa i morate ih sortirati po abecedi prvo po domenu (deo nakon `@`), a korisničko ime drugo (deo pre `@`).

- Formatiranje treba da bude nezavisno od velikih i malih slova.
- Ako više e-poštama imaju isti naziv domena, formatirati ih po korisničkom imenu.
- Vratiti spisak formatiranih adresa.
- Vraćene adrese treba da sačuvaju svoje izvorne veličine slova.

Na primer, ako dobijete `["jill@mail.com", "john@example.com", "jane@example.com"]`, vratite `["jane@example.com", "john@example.com", "jill@mail.com"]`.

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` inapaswa kurudisha `["jane@example.com", "john@example.com", "jill@mail.com"]`.

```js
assert.deepEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"]);
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` inapaswa kurudisha `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`.

```js
assert.deepEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"]);
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` inapaswa kurudisha `["user@x.com", "user@y.com", "user@z.com"]`.

```js
assert.deepEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"]);
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` inapaswa kurudisha `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`.

```js
assert.deepEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]);
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` treba da vrati `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`.

```js
assert.deepEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]);
```

# --seed--

## --seed-contents--

```js
function sort(emails) {

  return emails;
}
```

# --solutions--

```js
function sort(emails) {
  return emails.slice().sort((a, b) => {
    const [userA, domainA] = a.split('@');
    const [userB, domainB] = b.split('@');

    const domainCompare = domainA.toLowerCase().localeCompare(domainB.toLowerCase());
    if (domainCompare !== 0) return domainCompare;

    return userA.toLowerCase().localeCompare(userB.toLowerCase());
  });
}
```
