---
id: 68ee9e3066cfd4eb2328e8a8
title: "Zoezi la 89: Kuhesabu Kadi"
challengeType: 28
dashedName: challenge-89
---

# --description--
Standardni komplet igračkih karata ima 13 jedinstvenih karata u svakoj mašti. Ako dobijete kompletan broj koji predstavlja broj karata za izbor iz paketa, vratite broj jedinstvenih kombinacija karata koje možete odabrati.

- Redos nije važan. Izbor kartice A, a zatim kartice B, isto je kao izbor kartice B, a zatim kartice A.

Na primer, ako dobijete `52`, vratite `1`. Postoji samo jedna kombinacija od 52 karte koju možete izabrati iz paketa od 52 karata. I ako dobijete `2`, vratite `1326`. Postoji 1326 kombinacija karata koje možete dobiti kada izaberete 2 karte iz paketa.

# --hints--

`combinations(52)` inapaswa kurudisha `1`.

```js
assert.equal(combinations(52), 1);
```

`combinations(1)` inapaswa kurudisha `52`.

```js
assert.equal(combinations(1), 52);
```

`combinations(2)` inapaswa kurudisha `1326`.

```js
assert.equal(combinations(2), 1326);
```

`combinations(5)` inapaswa kurudisha `2598960`.

```js
assert.equal(combinations(5), 2598960);
```

`combinations(10)` inapaswa kurudisha `15820024220`.

```js
assert.equal(combinations(10), 15820024220);
```

`combinations(50)` inapaswa kurudisha `1326`.

```js
assert.equal(combinations(50), 1326);
```

# --seed--

## --seed-contents--

```js
function combinations(cards) {

  return cards;
}
```

# --solutions--

```js
function combinations(cards) {
  const n = 52;

  function factorial(x) {
    let result = 1;
    for (let i = 2; i <= x; i++) {
      result *= i;
    }
    return result;
  }

  return factorial(n) / (factorial(cards) * factorial(n - cards));
}
```
