---
id: 68ee9e3066cfd4eb2328e8a5
title: "Zoezi la 86: Utafutaji wa Picha"
challengeType: 28
dashedName: challenge-86
---

# --description--
Dana 4. novembra 2001. godine, Google je lansirao svoj pretraga slika, omogućavajući ljudima da pronađu slike koristeći ključne reči za pretragu. U ovom izazovu, ćete simulirati pretraživanje slika.

Ako pružite red podataka sa imenima slika i traženi termin, vratite red podataka sa imenima slika koja sadrže traženi termin.

- Ignorišite razlike između velikih i malih slova prilikom poređenja pretraga reči.
- Vratite slike u istom redosledu kakvim se pojavljuju u unosnim podacima.

# --hints--

`imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` inapaswa kurudisha `["dog.png"]`.

```js
assert.deepEqual(imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"]);
```

`imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` inapaswa kurudisha `["Sunset.jpg", "sunflower.jpeg"]`.

```js
assert.deepEqual(imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"]);
```

`imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` inapaswa kurudisha `["Moon.png", "stars.png"]`.

```js
assert.deepEqual(imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"]);
```

`imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` inapaswa kurudisha `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]`.

```js
assert.deepEqual(imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]);
```

# --seed--

## --seed-contents--

```js
function imageSearch(images, term) {

  return images;
}
```

# --solutions--

```js
function imageSearch(images, term) {
  const lowerTerm = term.toLowerCase();
  return images.filter(img => img.toLowerCase().includes(lowerTerm));
}
```
