---
id: 68ee9e3066cfd4eb2328e8a4
title: "Zoezi la 85: Hesabu ya maneno"
challengeType: 28
dashedName: challenge-85
---

# --description--

Ako je data rečenica kao niska, vrati broj reči u toj rečenici.

- Reč je bilo koja sekvenca znakova koji nisu razmaci i odvojena je samo jednim razmakom.

# --hints--

`countWords("Hello world")` inapaswa kurudisha `2`.

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")` inapaswa kurudisha `9`.

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")` inapaswa kurudisha `4`.

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")` treba da vrati `7`.

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")` inapaswa kurudisha `7`.

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```
