---
id: 68f6587287ad1f4ad39b0c84
title: "Zoezi la 98: Hesabu ya mstatili"
challengeType: 28
dashedName: challenge-98
---

# --description--
Ako su date dva pozitivna celog broja koja predstavljaju širinu i dužinu pravougaonika, odrediti koliko linija tog pravougaonika se može smestiti unutar datog.

- Računaj samo pravougaonike čija su širina i visina celi brojevi.

Na primer, uklanjanjem `1` i `3`, vratite `6`. Tri linije od 1x1, dve linije od 1x2, i jedna linija od 1x3.

# --hints--

`countRectangles(1, 3)` inapaswa kurudisha `6`.

```js
assert.equal(countRectangles(1, 3), 6);
```

`countRectangles(3, 2)` inapaswa kurudisha `18`.

```js
assert.equal(countRectangles(3, 2), 18);
```

`countRectangles(1, 2)` inapaswa kurudisha `3`.

```js
assert.equal(countRectangles(1, 2), 3);
```

`countRectangles(5, 4)` inapaswa kurudisha `150`.

```js
assert.equal(countRectangles(5, 4), 150);
```

`countRectangles(11, 19)` inapaswa kurudisha `12540`.

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```
