---
id: 68f6587287ad1f4ad39b0c7f
title: "Zoezi la 93: Vowels na Consonants"
challengeType: 28
dashedName: challenge-93
---

# --description--
Ako dobiješ niz slova, vrati red podataka koji sadrži broj samoglasnika i broj suglasnika u tom nizu.

- Vowels ni `a`, `e`, `i`, `o`, `u` katika hali yoyote.
- Suglasnici su sva ostala slova, bez obzira na veličinu.
- Zanemari sve znakove koji nisu slova.

Na primer, ako dobijete `"Hello World"`, vratite `[3, 7]`.

# --hints--

`count("Hello World")` inapaswa kurudisha `[3, 7]`.

```js
assert.deepEqual(count("Hello World"), [3, 7]);
```

`count("JavaScript")` treba da vrati `[3, 7]`.

```js
assert.deepEqual(count("JavaScript"), [3, 7]);
```

`count("Python")` inapaswa kurudisha `[1, 5]`.

```js
assert.deepEqual(count("Python"), [1, 5]);
```

`count("freeCodeCamp")` inapaswa kurudisha `[5, 7]`.

```js
assert.deepEqual(count("freeCodeCamp"), [5, 7]);
```

`count("Hello, World!")` inapaswa kurudisha `[3, 7]`.

```js
assert.deepEqual(count("Hello, World!"), [3, 7]);
```

`count("The quick brown fox jumps over the lazy dog.")` inapaswa kurudisha `[11, 24]`.

```js
assert.deepEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24]);
```

# --seed--

## --seed-contents--

```js
function count(str) {

  return str;
}
```

# --solutions--

```js
function count(str) {
  const vowels = 'aeiou';
  const consonants = 'bcdfghjklmnpqrstvwxyz';
  let v = 0, c = 0;

  for (let i=0; i<str.length; i++) {
    if (vowels.includes(str[i].toLowerCase())) {
      v++;
    }
    if (consonants.includes(str[i].toLowerCase())) {
      c++;
    }
  }

  return [v, c];
}
```
