---
id: 68f6587287ad1f4ad39b0c83
title: "Zoezi la 97: GCD"
challengeType: 28
dashedName: challenge-97
---

# --description--

Data su dva pozitivna cela broja. Vrati njihov najveći zajednički delilac (GCD).

- Najveći zajednički delilac dve celobrojke je najveći broj koji deli oba broja po celom bez ostataka.

Primer, delitelji `4` su `1`, `2`, i `4`. Delitelji `6` su `1`, `2`, `3`, i `6`. Stoga, uzimajući u obzir `4` i `6`, vrati `2`, najveći broj koji se pojavljuje u oba skupa delitelja.

# --hints--

`gcd(4, 6)` inapaswa kurudisha `2`.

```js
assert.equal(gcd(4, 6), 2);
```

`gcd(20, 15)` inapaswa kurudisha `5`.

```js
assert.equal(gcd(20, 15), 5);
```

`gcd(13, 17)` inapaswa kurudisha `1`.

```js
assert.equal(gcd(13, 17), 1);
```

`gcd(654, 456)` inapaswa kurudisha `6`.

```js
assert.equal(gcd(654, 456), 6);
```

`gcd(3456, 4320)` inapaswa kurudisha `864`.

```js
assert.equal(gcd(3456, 4320), 864);
```

# --seed--

## --seed-contents--

```js
function gcd(x, y) {

  return x;
}
```

# --solutions--

```js
function gcd(x, y) {
  while (y !== 0) {
    const temp = y;
    y = x % y;
    x = temp;
  }
  return x;
}
```
