---
id: 68f6587287ad1f4ad39b0c82
title: "Vežba 96: Da li je kraj nedelje?
challengeType: 28
dashedName: challenge-96
---

# --description--
Kada primiš datum u formatu `"YYYY-MM-DD"`, vrati broj dana preostalih do kraja nedelje.

- Kraj nedelje počinje u subotu.
- Ako je dato datum subota ili nedelja, vrati `"It's the weekend!"`.
- Inače, vrati `"X days until the weekend."`, gde je `X` broj dana do subote.
- Ako je `X` jednako `1`, koristi `"day"` (jednina) umesto `"days"` (množina).
- Proveri da izračunavanje ne uzima u obzir lokalno vreme.

# --hints--

`daysUntilWeekend("2025-11-14")` inapaswa kurudisha `"1 day until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")` inapaswa kurudisha `"3 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")` inapaswa kurudisha `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")` inapaswa kurudisha `"4 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")` inapaswa kurudisha `"5 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")` inapaswa kurudisha `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```
