---
id: 68f6587287ad1f4ad39b0c7e
title: "Zoezi la 92: Kiondoaji cha Kiendelezi"
challengeType: 28
dashedName: challenge-92
---

# --description--
Uklanjanje sekvence znakova koja predstavlja ime fajla, obnavljanjem ekstenzije tog fajla.

- Ekstenzija je deo imena fajla koji sledi nakon poslednje tačke (`.`).
- Ako ime fajla nema tačku ili se završava tačkom, vrati `"none"`.
- Ekstenzija treba da bude vraćena onakva kakva jeste, sačuvanjem velikih i malih slova.

# --hints--

`getExtension("document.txt")` inapaswa kurudisha `"txt"`.

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")` inapaswa kurudisha `"none"`.

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")` inapaswa kurudisha `"PNG"`.

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")` inapaswa kurudisha `"gitignore"`.

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")` inapaswa kurudisha `"gz"`.

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")` inapaswa kurudisha `"none"`.

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```
