---
id: 68f6587287ad1f4ad39b0c80
title: "Zoezi la 94: Kizalishaji cha Sahihi ya Barua Pepe"
challengeType: 28
dashedName: challenge-94
---

# --description--
Izvući sekvencu slova iz imena osobe, titule i kompanije, vratiti e-mail adresu kao jednu sekvencu slova koristeći sledeća pravila:

- Ime mora biti prikazano prvo, praćeno prefiksom koji zavisi od prvog slova imena. Za imena koja počinju sa (bez obzira na veliko ili malo slovo):
  - `A-I`: Koristiti `>>` kao prefiks.
  - `J-R`: Koristiti `--` kao prefiks.
  - `S-Z`: Koristiti `::` kao prefiks.
- Zarez i razmak (`, `) treba da prate ime.
- Titula i kompanija treba da prate zarez i razmak, odvojeni sa `" at "` (sa razmacima sa obe strane).

Na primer, ako se ukloni `"Quinn Waverly"`, `"Founder and CEO"`, i `"TechCo"`, vrati `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")` inapaswa kurudisha `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")` inapaswa kurudisha `">>Alice Reed, Engineer at TechCo"`.

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")` inapaswa kurudisha `"::Tina Vaughn, Developer at example.com"`.

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")` inapaswa kurudisha `">>B. B., Product Tester at AcmeCorp"`.

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")` inapaswa kurudisha `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```
