---
id: 68f6587287ad1f4ad39b0c85
title: "Zoezi la 99: Jaribio la Alama za Vidole"
challengeType: 28
dashedName: challenge-99
---

# --description--
Dati niz od dva slova koja predstavljaju otiske prstiju, utvrđuje da li je podudarnost koristeći sledeća pravila:

- Svaki otisak prsta će sadržati samo mala slova (`a-z`).
- Dva otiska prsta se smatraju podudarnim ako:
  - Zina urefu sawa.
- Broj različitih znakova nije veći od 10% dužine otiska prsta.

# --hints--

`isMatch("helloworld", "helloworld")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMatch("helloworld", "helloworld"));
```

`isMatch("helloworld", "helloworlds")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMatch("helloworld", "helloworlds"));
```

`isMatch("helloworld", "jelloworld")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMatch("helloworld", "jelloworld"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog"));
```

`isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog")` inapaswa kurudisha `true`.

```js
assert.isTrue(isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat")` inapaswa kurudisha `false`.

```js
assert.isFalse(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat"));
```

# --seed--

## --seed-contents--

```js
function isMatch(fingerprintA, fingerprintB) {

  return fingerprintA;
}
```

# --solutions--

```js
function isMatch(fingerprintA, fingerprintB) {
  if (fingerprintA.length !== fingerprintB.length) return false;

  const length = fingerprintA.length;
  let mismatches = 0;

  for (let i = 0; i < length; i++) {
    if (fingerprintA[i] !== fingerprintB[i]) {
      mismatches++;
      if (mismatches > length * 0.1) return false;
    }
  }

  return true;
}
```
