---
id: 68f6587287ad1f4ad39b0c7c
title: "Zoezi la 90: Kiwango cha herufi"
challengeType: 28
dashedName: challenge-90
---

# --description--
U ovom izazovu, dobijate niz karaktera i morate utvrditi da li odgovara objavi na društvenim mrežama. Vratite sledeći niz karaktera prema datim pravilima:

- `"short post"` ikiwa unaendana na kiwango cha herufi 40.
- `"long post"` ako ima više od 40 znakova i odgovara ograničenju od 80 znakova.
- `"invalid post"` ako je predugačak da bi odgovarao bilo kom od dva ograničenja.

# --hints--

`canPost("Hello world")` inapaswa kurudisha `"short post"`.

```js
assert.equal(canPost("Hello world"), "short post");
```

`canPost("This is a longer message but still under eighty characters.")` inapaswa kurudisha `"long post"`.

```js
assert.equal(canPost("This is a longer message but still under eighty characters."), "long post");
```

`canPost("This message is too long to fit into either of the character limits for a social media post.")` inapaswa kurudisha `"invalid post"`.

```js
assert.equal(canPost("This message is too long to fit into either of the character limits for a social media post."), "invalid post");
```

# --seed--

## --seed-contents--

```js
function canPost(message) {

  return message;
}
```

# --solutions--

```js
function canPost(message) {
  if (message.length <= 40) {
    return "short post"
  } else if (message.length <= 80) {
    return "long post"
  } else {
    return "invalid post"
  }
}
```
