---
id: 68ffb91507a5b645769328cc
title: "Vežba 109: Koliko godina imam ponovo?"
challengeType: 28
dashedName: challenge-109
---

# --description--
Uzimajući u obzir datum rođenja osobe u formatu `YYYY-MM-DD`, vratite godine te osobe do datuma 27. novembra 2025.

- Pretpostavi da su svi rođendani validni datumi pre 27. novembra 2025.
- Vrati godine kao ceo broj.
- Vodite računa da proverite da li je ta osoba već proslavila svoj rođendan u 2025. godini.

# --hints--

`calculateAge("2000-11-20")` inapaswa kurudisha `25`.

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` inapaswa kurudisha `24`.

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` inapaswa kurudisha `11`.

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` inapaswa kurudisha `31`.

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` inapaswa kurudisha `30`.

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
