---
id: 68ffb91507a5b645769328c9
title: "Zoezi la 106: Mthibitishaji wa ujumbe"
challengeType: 28
dashedName: challenge-106
---

# --description--

Ako su dati niska poruke i niska za potvrdu, odredi da li je poruka važeća.

- Poruka je važeća ako svaka reč u poruci počinje slovom koje odgovara sekvenci karaktera u kodu za potvrdu, po redu.
- Slova ne razmatraju razliku između velikih i malih slova.
- Reči u poruci su razdvojene pojedinačnim razmakom.

# --hints--

`isValidMessage("hello world", "hw")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
