---
id: 68ffb91507a5b645769328c5
title: "Vežba 102: Najduža reč"
challengeType: 28
dashedName: challenge-102
---

# --description--

Ako je dat niz znakova rečenice, vrati najdužu reč u toj rečenici.

- Reči su razdvojene samo jednim razmakom.
- Samo slova (`a-z`, nezavisno od velikih ili malih slova) se računaju za dužinu reči.
- Ako postoji više reči iste dužine, vrati prvu pronađenu.
- Vrati reč kako je navedena u datoj sekvenci, bez interpunkcijskih znakova.

# --hints--

`longestWord("The quick red fox")` inapaswa kurudisha `"quick"`.

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` inapaswa kurudisha `"challenge"`.

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` inapaswa kurudisha `"This"`.

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` treba da vrati `"exclamation"`.

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` inapaswa kurudisha `"tie"`.

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` inapaswa kurudisha `"Wouldnt"`.

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
