---
id: 68ffb91507a5b645769328c8
title: "Zoezi la 105: Hesabu ya herufi"
challengeType: 28
dashedName: challenge-105
---

# --description--
Ako izvučete sekvencu slova rečenice, vratite niz podataka koji sadrži brojač svakog slova, sortiran po abecednom redu.

- Računajte velika i mala slova kao jedno slovo tokom brojanja.
- Ignorišite brojeve, razmake, interpunkciju, itd.
- Vratite broj i slova u formatu `"letter count"`. Primer je `"a 3"`.
- Sva vraćena slova moraju biti mala slova.
- Ne vratite broj slova koja nisu deo navedenog niza slova.

# --hints--

`countCharacters("hello world")` inapaswa kurudisha `["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]`.

```js
assert.deepEqual(countCharacters("hello world"), ["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]);
```

`countCharacters("I love coding challenges!")` inapaswa kurudisha `["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]`.

```js
assert.deepEqual(countCharacters("I love coding challenges!"), ["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]);
```

`countCharacters("// TODO: Complete this challenge ASAP!")` inapaswa kurudisha `["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]`.

```js
assert.deepEqual(countCharacters("// TODO: Complete this challenge ASAP!"), ["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]);
```

# --seed--

## --seed-contents--

```js
function countCharacters(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countCharacters(sentence) {
  const counts = {};

  for (let char of sentence.toLowerCase()) {
    if (char >= "a" && char <= "z") {
      counts[char] = (counts[char] || 0) + 1;
    }
  }

  return Object.keys(counts)
               .sort()
               .map(letter => `${letter} ${counts[letter]}`);
}
```
