---
id: 68ffb91507a5b645769328c4
title: "Učenje 101: Konvertor zaglavlja sekcije za Markdown
challengeType: 28
dashedName: challenge-101
---

# --description--
Ako vam je dat niz slova koji predstavlja zaglavlje sekcije u Markdown formatu, vratite odgovarajuće zaglavlje sekcije za HTML.

Naslov važeće sekcije u Markdown-u je obavezan:

- Mora početi sa nula ili više razmaka, praćenih
- Slova broja 1 do 6 iz heš-a (`#`) uzastopno, zatim
- Najmanje jedan razmak. I na kraju,
- Tekst naslova sekcije.

Broj znaka haš određuje nivo zaglavlja sekcije. Na primer, jedan znak haš odgovara etiketi `h1`, a šest znaka haš odgovara etiketi `h6`.

Ako dati niz slova ne odgovara tačno prethodno opisanom formatu, vrati `"Invalid format"`.

Na primer, ako primiš `"# My level 1 heading"`, vrati `"<h1>My level 1 heading</h1>"`.

Napomena: Konsola možda neće prikazati oznake za HTML unutar sekvence karaktera prilikom snimanja poruka. Proverite konzolu pregledača da vidite zapise sa označenjima koji su uneseni.

# --hints--

`convert("# My level 1 heading")` inapaswa kurudisha `"<h1>My level 1 heading</h1>"`.

```js
assert.equal(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>");
```

`convert("My heading")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convert("My heading"), "Invalid format");
```

`convert("##### My level 5 heading")` inapaswa kurudisha `"<h5>My level 5 heading</h5>"`.

```js
assert.equal(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>");
```

`convert("#My heading")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convert("#My heading"), "Invalid format");
```

`convert("  ###  My level 3 heading")` inapaswa kurudisha `"<h3>My level 3 heading</h3>"`.

```js
assert.equal(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>");
```

`convert("####### My level 7 heading")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convert("####### My level 7 heading"), "Invalid format");
```

`convert("## My #2 heading")` inapaswa kurudisha `"<h2>My #2 heading</h2>"`.

```js
assert.equal(convert("## My #2 heading"), "<h2>My #2 heading</h2>");
```

# --seed--

## --seed-contents--

```js
function convert(heading) {

  return heading;
}
```

# --solutions--

```js
function convert(heading) {
  const trimmed = heading.trim()

  if (!/^#{1,6}\s+\S+/.test(trimmed)) return "Invalid format"

  const hashCount = trimmed.indexOf(' ');
  const text = trimmed.slice(hashCount).trimStart();

  return `<h${hashCount}>${text}</h${hashCount}>`;
}
```
