---
id: 68ffb91507a5b645769328c6
title: "Zoezi la 103: LCM"
challengeType: 28
dashedName: challenge-103
---

# --description--

Ako su data dva cela broja, vrati njihov najmanji zajednički sadržalac (LCM).

Najmanji zajednički sadržalnik dve brojeve je najmanji pozitivan ceo broj koji je višekratnik oba broja. Na primer, uzimajući u obzir `4` i `6`, vrati `12` jer:

- Mzalishaji wa `4` ni `4`, `8`, `12` na kadhalika.
- Mzalishaji wa `6` ni `6`, `12`, `18` na kadhalika.
- `12` je najmanji broj koji je deljiv sa oba broja.

# --hints--

`lcm(4, 6)` inapaswa kurudisha `12`.

```js
assert.equal(lcm(4, 6), 12);
```

`lcm(9, 6)` inapaswa kurudisha `18`.

```js
assert.equal(lcm(9, 6), 18);
```

`lcm(10, 100)` inapaswa kurudisha `100`.

```js
assert.equal(lcm(10, 100), 100);
```

`lcm(13, 17)` inapaswa kurudisha `221`.

```js
assert.equal(lcm(13, 17), 221);
```

`lcm(45, 70)` inapaswa kurudisha `630`.

```js
assert.equal(lcm(45, 70), 630);
```

# --seed--

## --seed-contents--

```js
function lcm(a, b) {

  return a;
}
```

# --solutions--

```js
function lcm(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  return Math.abs(a * b) / gcd(a, b);
}
```
