---
id: 69162d64f96574d9bb629f00
title: "Vežba 115: Zamena stavke iz Markdown liste
challengeType: 28
dashedName: challenge-115
---

# --description--
Ako vam je dat niz slova koji predstavlja stavku u listi formatiranoj u Markdown-u, vratite odgovarajući niz slova za HTML.

Važeći element liste formatiran u Markdown-u mora:

- Počinje sa nultom ili više razmaka, praćenim
- Broj (jedan ili više) i tačka (`.`), praćenim
- Najmanje jedan razmak, zatim
- Tekst stavke liste.

Ako niz znakova u potpunosti ne odgovara strukturi iznad, vratite `"Invalid format"`. Inače, obuhvatite tekst stavke liste unutar oznaka `li` i vratite niz znakova.

Na primer, ako primite `"1. My item"`, vratite `"<li>My item</li>"`.

Napomena: Konzola možda neće prikazati HTML oznake u nizu znakova prilikom beleženja poruka. Pogledaj konzolu pregledača da vidiš zapise koji sadrže oznake.

# --hints--

`convertListItem("1. My item")` inapaswa kurudisha `"<li>My item</li>"`.

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")` inapaswa kurudisha `"<li>Another item</li>"`.

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")` inapaswa kurudisha `"<li>list item text</li>"`.

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")` inapaswa kurudisha `"Invalid format"`.

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
