---
id: 69162d64f96574d9bb629efb
title: "Izazov 110: Prediktor reči
challengeType: 28
dashedName: challenge-110
---

# --description--
S obzirom na par od dva karaktera iste dužine, tajnu reč i pogađanu reč, uporedite pogađanu reč sa tajnom rečju koristeći sledeća pravila:

- Tajni reč i pokušaj će sadržati samo velika slova (`"A"` do `"Z"`);
- Za svako slovo u reči za pogađanje, zamenite ga brojem u skladu sa načinom na koji odgovara tajnoj reči:
  - `"2"` ako je slovo u tajnoj reči i na pravom mestu.
  - `"1"` ako je slovo u tajnoj reči, ali na pogrešnom mestu.
  - `"0"` ako slovo nije u tajnoj reči.
- Svako slovo u tajnoj reči može se koristiti samo jednom.
- Potpuni podatak (`"2"`) se predaje prvi, zatim delimični podatak (`"1"`) se predaje od leve ka desnoj za preostala slova.
- Ako slovo često występuje u reči za pogađanje, može se pogoditi samo onoliko puta koliko występuje u tajnoj reči.

Na primer, uzimajući u obzir tajnu reč `"APPLE"` i nasumičnu reč `"POPPA"`, vratite `"10201"`:

Prvo slovo od `"P"` nije na pravom mestu (`"1"`), `"O"` nije u tajnoj reči (`"0"`), drugo slovo od `"P"` je na pravom mestu (`"2"`), treće slovo od `"P"` je nula (`"0"`) jer su dva slova od `"P"` u tajnoj reči korišćena, i `"A"` nije na pravom mestu (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` inapaswa kurudisha `"10201"`.

```js
assert.equal(compare("APPLE", "POPPA"), "10201");
```

`compare("REACT", "TRACE")` inapaswa kurudisha `"11221"`.

```js
assert.equal(compare("REACT", "TRACE"), "11221");
```

`compare("DEBUGS", "PYTHON")` inapaswa kurudisha `"000000"`.

```js
assert.equal(compare("DEBUGS", "PYTHON"), "000000");
```

`compare("JAVASCRIPT", "TYPESCRIPT")` treba da vrati `"0000222222"`.

```js
assert.equal(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222");
```

`compare("ORANGE", "ROUNDS")` inapaswa kurudisha `"110200"`.

```js
assert.equal(compare("ORANGE", "ROUNDS"), "110200");
```

`compare("WIRELESS", "ETHERNET")` inapaswa kurudisha `"10021000"`.

```js
assert.equal(compare("WIRELESS", "ETHERNET"), "10021000");
```

# --seed--

## --seed-contents--

```js
function compare(word, guess) {

  return word;
}
```

# --solutions--

```js
function compare(word, guess) {
  const n = word.length;
  let result = Array(n).fill("0");
  let used = Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    if (guess[i] === word[i]) {
      result[i] = "2";
      used[i] = true;
    }
  }

  for (let i = 0; i < n; i++) {
    if (result[i] === "0") {
      for (let j = 0; j < n; j++) {
        if (!used[j] && guess[i] === word[j]) {
          result[i] = "1";
          used[j] = true;
          break;
        }
      }
    }
  }

  return result.join("");
}
```
