---
id: 69162d64f96574d9bb629efd
title: "Zoezi la 112: Kigeuzi cha AI"
challengeType: 28
dashedName: challenge-112
---

# --description--
Danasnji izazov proizašao je iz lansiranja ChatGPT-a dana 30. novembra 2022.

Uzevši u obzir sekvencu karaktera jedne ili više rečenica, utvrdite da li je možda generisano od strane AI koristeći sledeća pravila:

- Imaš dve ili više dash-eva (`-`).
- Imaš dva ili više skupova zagrada (`()`). Tekst može biti unutar zagrada.
- Imaš tri ili više reči koja imaju 7 ili više slova.

- Reči su razdvojene samo razmakom i sastoje se isključivo od slova (`A-Z`). Ne uključujte interpunkciju ili ne-slovočne znakove kao deo reči.

Ako data sentensi inakidhi yoyote kati ya sheria hizi, rudisha `"AI"`, vinginevyo rudisha `"Human"`.

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")` inapaswa kurudisha `"Human"`.

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")` inapaswa kurudisha `"Human"`.

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")` inapaswa kurudisha `"AI"`.

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")` inapaswa kurudisha `"AI"`.

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")` inapaswa kurudisha `"AI"`.

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
