---
id: 69162d64f96574d9bb629f04
title: "Vežba 119: Kompresija niza znakova"
challengeType: 28
dashedName: challenge-119
---

# --description--

Ako je data rečenica kao niz znakova, vrati kompresovanu verziju rečenice gde se uzastopno ponovljene reči zamenjuju sledećim formatom: reč i broj ponavljanja u običnim zagradama.

- Samo su kompresovane samo reči koje se ponavljaju uzastopno.
- Reči su razdvojene pojedinačnim razmakom.

Na primer, ako dobijete `"yes yes yes please"`, vratite `"yes(3) please"`.

# --hints--

`compressString("yes yes yes please")` inapaswa kurudisha `"yes(3) please"`.

```js
assert.equal(compressString("yes yes yes please"), "yes(3) please");
```

`compressString("I have have have apples")` inapaswa kurudisha `"I have(3) apples"`.

```js
assert.equal(compressString("I have have have apples"), "I have(3) apples");
```

`compressString("one one three and to the the the the")` inapaswa kurudisha `"one(2) three and to the(4)"`.

```js
assert.equal(compressString("one one three and to the the the the"), "one(2) three and to the(4)");
```

`compressString("route route route route route route tee tee tee tee tee tee")` inapaswa kurudisha `"route(6) tee(6)"`.

```js
assert.equal(compressString("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)");
```

# --seed--

## --seed-contents--

```js
function compressString(sentence) {

  return sentence;
}
```

# --solutions--

```js
function compressString(sentence) {
  const words = sentence.split(" ");
  const result = [];
  let count = 1;

  for (let i = 0; i < words.length; i++) {
    if (words[i] === words[i + 1]) {
      count++;
    } else {
      if (count > 1) {
        result.push(`${words[i]}(${count})`);
      } else {
        result.push(words[i]);
      }
      count = 1;
    }
  }

  return result.join(" ");
}
```
