---
id: 69162d64f96574d9bb629f02
title: "Vežba 117: Razlika jednakosti
challengeType: 28
dashedName: challenge-117
---

# --description--
Kada uzmete dva reda podataka, vratite novi red koji sadrži razliku između njih.

- Simetrična razlika između dva skupa je skup vrednosti koje se pojavljuju u jednom ili drugom skupu, ali ne i u oba.
- Vrati vrednosti u redosledu njihovog prvog pojavljivanja u redovima ulaznih podataka.

# --hints--

`difference([1, 2, 3], [3, 4, 5])` inapaswa kurudisha `[1, 2, 4, 5]`.

```js
assert.deepEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5]);
```

`difference(["a", "b"], ["c", "b"])` inapaswa kurudisha `["a", "c"]`.

```js
assert.deepEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"]);
```

`difference([1, "a", 2], [2, "b", "a"])` inapaswa kurudisha `[1, "b"]`.

```js
assert.deepEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"]);
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])` inapaswa kurudisha `[2, 4, 6, 8]`.

```js
assert.deepEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8]);
```

# --seed--

## --seed-contents--

```js
function difference(arr1, arr2) {

  return arr1;
}
```

# --solutions--

```js
function difference(arr1, arr2) {
  const diff = [];

  for (const v of arr1) {

    if (!diff.includes(v) && !arr2.includes(v)) diff.push(v)
  }
  
  for (const v of arr2) {
    if (!diff.includes(v) && !arr1.includes(v)) diff.push(v)
  }

  return diff;
}
```
