---
id: 69162d64f96574d9bb629eff
title: "Vežba 114: Iz camel u snake"
challengeType: 28
dashedName: challenge-114
---

# --description--
Ako vam je dat niz znakova u stilu camel case, vratite verziju u snake stilu case tog niza koristeći sledeće pravilo:

- Ulazni niz znakova će sadržati samo slova (`A-Z` i `a-z`) i uvek mora počinjati malim slovom.
- Svako veliko slovo u camel case započinje novu reč.
- Konvertujte sva slova u mala.
- Razdvojite reči donjom crtom (underscore) (`_`).

# --hints--

`toSnake("helloWorld")` inapaswa kurudisha `"hello_world"`.

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` treba da vrati `"my_variable_name"`.

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` inapaswa kurudisha `"freecodecamp_daily_challenges"`.

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
