---
id: 69162d64f96574d9bb629f03
title: "Zoezi la 118: Mbadalishaji wa tarehe"
challengeType: 28
dashedName: challenge-118
---

# --description--
Ako je datum formatiran kao `"Month day, year"`, vratite taj datum na format `"YYYY-MM-DD"`.

- Izvučeni mesec će biti pun naziv meseca na engleskom jeziku. Primer: `"January"`, `"February"`, n.k.
- U vrednosti return, dodajte vodeću nulu mesecu i dan ako je potrebno kako biste osigurali dve cifre.

Na primer, ako je `"December 6, 2025"` uklonjeno, vrati `"2025-12-06"`.

# --hints--

`formatDate("December 6, 2025")` inapaswa kurudisha `"2025-12-06"`.

```js
assert.equal(formatDate("December 6, 2025"), "2025-12-06");
```

`formatDate("January 1, 2000")` inapaswa kurudisha `"2000-01-01"`.

```js
assert.equal(formatDate("January 1, 2000"), "2000-01-01");
```

`formatDate("November 11, 1111")` inapaswa kurudisha `"1111-11-11"`.

```js
assert.equal(formatDate("November 11, 1111"), "1111-11-11");
```

`formatDate("September 7, 512")` inapaswa kurudisha `"512-09-07"`.

```js
assert.equal(formatDate("September 7, 512"), "512-09-07");
```

`formatDate("May 4, 1950")` inapaswa kurudisha `"1950-05-04"`.

```js
assert.equal(formatDate("May 4, 1950"), "1950-05-04");
```

`formatDate("February 29, 1992")` inapaswa kurudisha `"1992-02-29"`.

```js
assert.equal(formatDate("February 29, 1992"), "1992-02-29");
```

# --seed--

## --seed-contents--

```js
function formatDate(dateString) {

  return dateString;
}
```

# --solutions--

```js
function formatDate(dateString) {
  const months = {
    January: "01",
    February: "02",
    March: "03",
    April: "04",
    May: "05",
    June: "06",
    July: "07",
    August: "08",
    September: "09",
    October: "10",
    November: "11",
    December: "12"
  };

  const [month, day, year] = dateString.replace(",", "").split(" ");

  return `${year}-${months[month]}-${day.padStart(2, "0")}`;
}
```
