---
id: 69162d64f96574d9bb629efe
title: "Zoezi la 113: Maili hadi kilomita"
challengeType: 28
dashedName: challenge-113
---

# --description--
Ako unesete udaljenost u miljama kao broj, konvertujte istu udaljenost u kilometre.

- Ulaz će uvek biti broj bez negativne vrednosti.
- 1 milja je jednaka 1,60934 kilometara.
- Zaokruži rezultat na dve decimale.
- Ukloni nepotrebne nule sa kraja zaokruženog rezultata.

# --hints--

`convertToKm(1)` inapaswa kurudisha `1.61`.

```js
assert.equal(convertToKm(1), 1.61);
```

`convertToKm(21)` inapaswa kurudisha `33.8`.

```js
assert.equal(convertToKm(21), 33.8);
```

`convertToKm(3.5)` inapaswa kurudisha `5.63`.

```js
assert.equal(convertToKm(3.5), 5.63);
```

`convertToKm(0)` inapaswa kurudisha `0`.

```js
assert.equal(convertToKm(0), 0);
```

`convertToKm(0.621371)` inapaswa kurudisha `1`.

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```
