---
id: 691b559495c5cb5a37b9b487
title: "Zoezi la 127: Kagua Kasi"
challengeType: 28
dashedName: challenge-127
---

# --description--
Na osnovu brzine koju vozite u miljama po satu (MPH) i ograničenja brzine u kilometrima po satu (KPH), utvrdite da li prelazite dozvoljenu brzinu i da li ćete dobiti upozorenje ili kaznu.

- Milja 1 iznosi 1,60934 kilometara.
- Ako putujete brzinom koja ne prelazi ili je jednaka ograničenju brzine, vratite `"Not Speeding"`.
- Ako putujete brzinom od 5 KPH ili manjom od ograničenja brzine, vratite `"Warning"`.
- Ako putujete brzinom većom od 5 KPH i veće od ograničenja brzine, vratite `"Ticket"`.

# --hints--

`speedCheck(30, 70)` inapaswa kurudisha `"Not Speeding"`.

```js
assert.equal(speedCheck(30, 70), "Not Speeding");
```

`speedCheck(40, 60)` inapaswa kurudisha `"Warning"`.

```js
assert.equal(speedCheck(40, 60), "Warning");
```

`speedCheck(40, 65)` inapaswa kurudisha `"Not Speeding"`.

```js
assert.equal(speedCheck(40, 65), "Not Speeding");
```

`speedCheck(60, 90)` inapaswa kurudisha `"Ticket"`.

```js
assert.equal(speedCheck(60, 90), "Ticket");
```

`speedCheck(65, 100)` inapaswa kurudisha `"Warning"`.

```js
assert.equal(speedCheck(65, 100), "Warning");
```

`speedCheck(88, 40)` inapaswa kurudisha `"Ticket"`.

```js
assert.equal(speedCheck(88, 40), "Ticket");
```

# --seed--

## --seed-contents--

```js
function speedCheck(speedMph, speedLimitKph) {
  return speedMph;
}
```

# --solutions--

```js
function speedCheck(speedMph, speedLimitKph) {
  const speedKph = speedMph * 1.60934;

  if (speedKph <= speedLimitKph) {
    return "Not Speeding";
  }

  const over = speedKph - speedLimitKph;

  if (over <= 5) {
    return "Warning";
  }

  return "Ticket";
}
```
