---
id: 691b559495c5cb5a37b9b489
title: "Zoezi la 129: Parser ya blockquote ya Markdown"
challengeType: 28
dashedName: challenge-129
---

# --description--

Ako je dat niz znakova koji sadrži Markdown blockquote, vrati odgovarajući HTML niz znakova.

Blockquote katika Markdown ni mstari wowote unao:

- Počinje sa nula ili više razmaka
- Praćen znakom veće od (>)(`>`)
- Zatim, jedan ili više razmaka
- I na kraju, tekst u bloku citata (blockquote).

Vrati tekst blockquote-a okružen početnim i završnim HTML oznakama `blockquote`.

Na primer, ako primite `"> This is a quote"`, vratite `<blockquote>This is a quote</blockquote>`.

Napomena: Konsola možda neće prikazati oznake HTML u nizu karaktera tokom snimanja poruka. Proverite konzolu pretraživača da vidite zapise sa priloženim oznakama.

# --hints--

`parseBlockquote("> This is a quote")` inapaswa kurudisha `"<blockquote>This is a quote</blockquote>"`.

```js
assert.equal(parseBlockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>");
```

`parseBlockquote(" > This is also a quote")` inapaswa kurudisha `"<blockquote>This is also a quote</blockquote>"`.

```js
assert.equal(parseBlockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>");
```

`parseBlockquote("  >    So  Is  This")` inapaswa kurudisha `"<blockquote>So  Is  This</blockquote>"`.

```js
assert.equal(parseBlockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>");
```

# --seed--

## --seed-contents--

```js
function parseBlockquote(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBlockquote(markdown) {
  const blockquoteRegex = /^\s*>\s+(.+)$/;
  const match = markdown.match(blockquoteRegex);

  const text = match[1];
  return `<blockquote>${text}</blockquote>`;
}
```
