---
id: 691b559495c5cb5a37b9b485
title: "Izazov 125: Igra života
challengeType: 28
dashedName: challenge-125
---

# --description--
Ako vam je data matrica (redovi redova) koja predstavlja trenutno stanje u Konvejevoj igri života, vratite sledeće stanje matrice koristeći ova pravila:

- Svaka ćelija je `1` (život) ili `0` (smrt).
- Susedi ćelije su do osam ćelija koje okružuju (vertikalno, horizontalno i dijagonalno).
- Ćelije blizu granica imaju manje suseda od osam.

Pravila za ažuriranje svake ćelije:

- Svaka živa ćelija sa manje od dva živa suseda umire (usamljenost).
- Svaka živa ćelija sa dva ili tri živa suseda nastavlja da živi.
- Svaka živa ćelija sa više od tri živa suseda umire (prenatrpanost).
- Svaka mrtva ćelija sa tačno tri živa suseda oživljava (rađanje).

Na primer, ako je uklonjeno:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

rudisha:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Svaka ćelija se ažurira u zavisnosti od broja živih suseda. Na primer, `[0][0]` ostaje mrtva (2 živa suseda), `[0][1]` ostaje živa (2 živa suseda), `[0][2]` umire (3 živa suseda), itd.

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` inapaswa kurudisha `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` inapaswa kurudisha `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` inapaswa kurudisha `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` inapaswa kurudisha `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```
