---
id: 691b559495c5cb5a37b9b486
title: "Vežba 126: Napiši velikim slovima
challengeType: 28
dashedName: challenge-126
---

# --description--
Ako pronađeš niz slova koji se zove naslov (title), vrati novi niz formatiran velikim slovima koristeći sledeća pravila:

- Napiši prvo slovo svake reči velikim slovom.
- Sve ostale slova u svakoj reči učini malim slovom.
- Reči ostaju razdvojene samo jednim razmakom uvek.

# --hints--

`titleCase("hello world")` inapaswa kurudisha `"Hello World"`.

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` inapaswa kurudisha `"The Quick Brown Fox"`.

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` treba da vrati `"Javascript And Python"`.

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` inapaswa kurudisha `"Avocado Toast For Breakfast"`.

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
