---
id: 691b559495c5cb5a37b9b488
title: "Zoezi la 128: Hesabu ya Konsonanti"
challengeType: 28
dashedName: challenge-128
---

# --description--

Na osnovu niza znakova i ciljnog broja, odredi da li taj niz ima tačan broj suglasnika kao ciljni broj.

- Suglasnici su sva slova abecede osim `"a"`, `"e"`, `"i"`, `"o"` i `"u"` u bilo kom obliku slova.
- Zanemari brojeve, znakove interpunkcije, razmake i druge znakove koji nisu slova prilikom brojanja.

# --hints--

`hasConsonantCount("helloworld", 7)` inapaswa kurudisha `true`.

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)` inapaswa kurudisha `false`.

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)` inapaswa kurudisha `true`.

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` inapaswa kurudisha `false`.

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` inapaswa kurudisha `true`.

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
