---
id: 691b559495c5cb5a37b9b480
title: "Zoezi la 120: Pauni hadi Kilogramu"
challengeType: 28
dashedName: challenge-120
---

# --description--
Ako navedete težinu u funtama kao broj, vratite sekvencu slova `"(lbs) pounds equals (kgs) kilograms."`.

- Zameni `"(lbs)"` sa brojem unosa.
- Zameni `"(kgs)"` sa unosom konvertovanim u kilograme, zaokruži na dve decimale i uvek uključuj dve decimale u vrednost.
- 1 funta je jednako 0,453592 kilograma.
- Ako je unos `1`, koristi `"pound"` umesto `"pounds"`.
- Ako je konvertovana vrednost `1`, koristi `"kilogram"` umesto `"kilograms"`.

# --hints--

`convertToKgs(1)` inapaswa kurudisha `"1 pound equals 0.45 kilograms."`.

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)` inapaswa kurudisha `"0 pounds equals 0.00 kilograms."`.

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)` inapaswa kurudisha `"100 pounds equals 45.36 kilograms."`.

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)` inapaswa kurudisha `"2.5 pounds equals 1.13 kilograms."`.

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)` inapaswa kurudisha `"2.20462 pounds equals 1.00 kilogram."`.

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```
