---
id: 691b559495c5cb5a37b9b483
title: "Zoezi la 123: Mtengenezaji wa Nambari za Kirumi"
challengeType: 28
dashedName: challenge-123
---

# --description--
Ako dobiješ kompletan broj, vrati mu ekvivalentnu vrednost u rimskim brojevima.

Rimski brojevi koriste sledeće simbole:

| Simbol | Vrednost |
| ------ | -------- |
| I      | 1        |
| V      | 5        |
| X      | 10       |
| L      | 50       |
| C      | 100      |
| D      | 500      |
| M      | 1000     |

Brojevi za Kirimi se pišu od velikog ka malom, od leve ka desnoj koristeći sledeće pravilo:

- Dodavanje se koristi kada je znak praćen drugim znakiem jednakog ili manjeg vrednosti. Na primer, `18` piše se kao `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- Oduzimanje se koristi kada manji znak stoji pre većeg, predstavljajući 4 ili 9 u bilo kojoj pozicionoj vrednosti. Na primer, 19 piše se kao `XIX` (10 + (10 - 1)).
- Nijedan znak se ne može ponoviti više od tri puta zaredom. Oduzimanje se koristi kada je potrebno napisati znak više od tri puta zaredom. Stoga, najveći broj koji možete napisati je 3999.

Ovo je još jedan primer: ukoliko dobijete `1464`, vratite `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`toRoman(18)` inapaswa kurudisha `"XVIII"`.

```js
assert.equal(toRoman(18), "XVIII");
```

`toRoman(19)` inapaswa kurudisha `"XIX"`.

```js
assert.equal(toRoman(19), "XIX");
```

`toRoman(1464)` inapaswa kurudisha `"MCDLXIV"`.

```js
assert.equal(toRoman(1464), "MCDLXIV");
```

`toRoman(2025)` inapaswa kurudisha `"MMXXV"`.

```js
assert.equal(toRoman(2025), "MMXXV");
```

`toRoman(3999)` inapaswa kurudisha `"MMMCMXCIX"`.

```js
assert.equal(toRoman(3999), "MMMCMXCIX");
```

# --seed--

## --seed-contents--

```js
function toRoman(num) {
  return num;
}
```

# --solutions--

```js
function toRoman(num) {
  const symbols = [
    ["M", 1000],
    ["CM", 900],
    ["D", 500],
    ["CD", 400],
    ["C", 100],
    ["XC", 90],
    ["L", 50],
    ["XL", 40],
    ["X", 10],
    ["IX", 9],
    ["V", 5],
    ["IV", 4],
    ["I", 1]
  ];

  let result = "";
  for (const [sym, val] of symbols) {
    while (num >= val) {
      result += sym;
      num -= val;
    }
  }
  return result;
}
```
