---
id: 6925e2068081f40f549ced1a
title: "Zoezi la 136: Mchambuzi wa picha wa Markdown"
challengeType: 28
dashedName: challenge-136
---

# --description--

Ako je dat niz znakova slike u Markdown formatu, vrati odgovarajući HTML niz znakova.

Markdown slika ima sledeći format: `"![alt text](image_url)"`. Gde:

- `alt text` je opis slike (vrednost svojstva elementa za `alt`).
- `image_url` je URL izvor slike (vrednost svojstva elementa za `src`).

Vrati HTML oznaku `img` gde je `src` postavljen na URL slike, a `alt` postavljen na alternativni tekst.

Na primer, ako dobijete `"![Cute cat](cat.png)"` vratite `'<img src="cat.png" alt="Cute cat">'`;

- Osigurajte da su oznake, raspored atributa elementa, razmak i korišćenje navodnika u skladu sa gore.

Napomena: Konzola možda neće prikazati HTML oznake u nizu znakova prilikom beleženja poruka — proveri konzolu pregledača da vidiš zapise koji sadrže oznake.

# --hints--

`parseImage("![Cute cat](cat.png)")` inapaswa kurudisha `'<img src="cat.png" alt="Cute cat">'`.

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` inapaswa kurudisha `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` inapaswa kurudisha `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```
