---
id: 69272dcf1c24b44fd79137c6
title: "Vežba 143: Markdown parser za naglašavanje/kurziv
challengeType: 28
dashedName: challenge-143
---

# --description--

Ako je dat niz znakova koji može sadržati tekst u Markdown italik stilu, vrati odgovarajući HTML niz znakova.

- Italik tekst u Markdown-u je svaki tekst koji počinje i završava jednom zvezdicom (`*`) ili jednom donjom crtom (`_`).
- Nisu dozvoljeni razmaci između teksta i zvezdice ili donje crte, ali razmaci mogu postojati unutar samog teksta.
- Pretvori sve italik delove u HTML `i` oznake i vrati niz znakova.

Na primer, ako primite `"*This is italic*"`, vratite `"<i>This is italic</i>"`.

Napomena: Konsola možda neće prikazati tagove HTML unutar niza karaktera prilikom snimanja poruka. Proverite konzolu pretraživača da biste videli rekorde sa tim oznakama.

# --hints--

`parseItalics("*This is italic*")` inapaswa kurudisha `"<i>This is italic</i>"`.

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")` inapaswa kurudisha `"<i>This is also italic</i>"`.

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")` inapaswa kurudisha `"*This is not italic *"`.

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")` inapaswa kurudisha `"_ This is also not italic_"`.

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")` inapaswa kurudisha `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
