---
id: 69306364df283fcaff2e1ada
title: "Zoezi la 149: vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

Ako je dat niz znakova, vrati novi niz gde su svi samoglasnici pretvoreni u velika slova, a svi ostali alfabetni znakovi u mala slova.

- Samoglasnici su `"a"`, `"e"`, `"i"`, `"o"` i `"u"` u bilo kom obliku slova.
- Herufi zisizo za alfabeti hazibadilishwi.

# --hints--

`vowelCase("vowelcase")` inapaswa kurudisha `"vOwElcAsE"`.

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")` inapaswa kurudisha `"cOdIng Is fUn"`.

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")` inapaswa kurudisha `"hEllO, wOrld!"`.

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")` inapaswa kurudisha `"gIt chErry-pIck"`.

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")` inapaswa kurudisha `"hEAd~1"`.

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```
