---
id: 69306364df283fcaff2e1ad5
title: "Zoezi la 144: Mfululizo wa azimio"
challengeType: 28
dashedName: challenge-144
---

# --description--
Ako dobijete red podataka sa podredovima, gde svaki podred predstavlja dan vaše aktivnosti u vezi sa ciljem i ima tri potpuna broja: broj koraka koje ste tog dana prešli, minute provedene na ekranu tog dana i broj pročitanih stranica tog dana; odlučite da li pratite svoju rezoluciju.

- Prvi mali red je dan 1, a drugi je dan 2, i tako dalje.

Dan se smatra uspešnim ako su ispunjena sledeća tri cilja:

- Prošli ste najmanje 10.000 koraka.
- Prošlo je više od 120 minuta korišćenja ekrana.
- Pročitali ste najmanje pet stranica.

Ako su svi dati dani uspešni, vrati `"Resolution on track: N day streak."` gde je `N` broj uspešnih dana.

Ako je jedan ili više dana neuspešno, vrati `"Resolution failed on day X: N day streak."`. Gde `X` je broj prvog neuspešnog dana, i `N` je broj uspešnih dana pre prvog neuspešnog dana.

# --hints--

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])` inapaswa kurudisha `"Resolution on track: 3 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.");
```

`resolutionStreak([[10000, 120, 5], [10950, 121, 11]])` inapaswa kurudisha `"Resolution failed on day 2: 1 day streak."`

```js
assert.equal(resolutionStreak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.");
```

`resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])` inapaswa kurudisha `"Resolution failed on day 3: 2 day streak."`

```js
assert.equal(resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.");
```

`resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])` inapaswa kurudisha `"Resolution failed on day 4: 3 day streak."`

```js
assert.equal(resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.");
```

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])` inapaswa kurudisha `"Resolution on track: 10 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.");
```

# --seed--

## --seed-contents--

```js
function resolutionStreak(days) {

  return days;
}
```

# --solutions--

```js
function resolutionStreak(days) {
  let streak = 0;

  for (let i = 0; i < days.length; i++) {
    const [steps, screen, pages] = days[i];
    if (steps >= 10000 && screen <= 120 && pages >= 5) {
      streak++;
    } else {
      return `Resolution failed on day ${i + 1}: ${streak} day streak.`;
    }
  }

  return `Resolution on track: ${streak} day streak.`;
}
```
