---
id: 69306364df283fcaff2e1ad7
title: "Vežba 146: Stolica za osobu koja piše levom rukom za sto
challengeType: 28
dashedName: challenge-146
---

# --description--
Dato je tabela dimenzija 4x2 (red po red) koja predstavlja raspored stolica za jelo/obrok; odredite koliko stolica osoba koja piše levom rukom može zauzeti.

- Osoba koja piše levom rukom ne može da sedi na mestu gde bi osoba koja piše desnom rukom sedela na stolicu sa leve strane te osobe.

Katika jedwali lililotolewa:

- `"R"` je stolica zauzeta od strane osobe koja piše desnom rukom.
- `"L"` je stolica zauzeta od strane osobe koja piše levom rukom.
- `"U"` je prazna stolica.
- To su samo prazne stolice koje se mogu koristiti za sedenje.
- Stolice u gornjem redu okrenute su "dole", a stolice u donjem redu okrenute su "gore" (kao sto), tako da zavisi šta je levo, a šta desno od orijentacije stolice.
- Uglove stolice imaju samo jednu stolicu pored sebe.

Na primer, u datoj tabeli:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

Ne može se sedeti na gornjo-levu stolicu jer neko piše desnom rukom sa leve strane. Ostale dve slobodne stolice mogu se zauzeti jer niko ne piše desnom rukom sa leve strane.

# --hints--

`findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` inapaswa kurudisha `2`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2);
```

`findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` inapaswa kurudisha `8`.

```js
assert.equal(findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8);
```

`findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` inapaswa kurudisha `0`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0);
```

`findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` inapaswa kurudisha `1`.

```js
assert.equal(findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1);
```

`findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` inapaswa kurudisha `5`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5);
```

# --seed--

## --seed-contents--

```js
function findLeftHandedSeats(table) {

  return table;
}
```

# --solutions--

```js
function findLeftHandedSeats(table) {
  let availableSeats = 0;

  const [topRow, bottomRow] = table;

  for (let i=0; i<topRow.length; i++) {
    if (topRow[i] === 'U') {
      if (i === 3) {
        availableSeats++;
      } else if (topRow[i+1] !== 'R') {
        availableSeats++;
      }
    }
  }

  for (let i=0; i<bottomRow.length; i++) {
    if (bottomRow[i] === 'U') {
      if (i === 0) {
        availableSeats++;
      } else if (bottomRow[i-1] !== 'R') {
        availableSeats++;
      }
    }
  }

  return availableSeats;
}
```
