---
id: 69373793f5a867f769cde135
title: "Vežba 150: Analiziranje neuređene liste u Markdown formatu
challengeType: 28
dashedName: challenge-150
---

# --description--
Kada je dato niz znakova iz validne liste koja nije formatirana u Markdown-u, vrati odgovarajući niz znakova za HTML.

Neprestrukturisana lista ima jedan ili više stavki liste. Važeća stavka liste je vidljiva na svom redu i:

- Počinje sa crticom (`"-"`), praćeno sa
- Bar jedan razmak, zatim
- Tekst stavke liste.

Orodha se daje kao niz znakova sa novim redovima, razdvojenih znakom nova linija (`"\n"`). Ne uključujte znake nove linije u tekst polja.

Obroti svaki stav liste unutar oznaka HTML `li`, i celu listu stavki unutar oznaka `ul`.

Na primer, ako primiš `"- Item A\n- Item B"`, vrati `"<ul><li>Item A</li><li>Item B</li></ul>"`.

Napomena: Konzola možda neće prikazati HTML oznake u nizu znakova prilikom beleženja poruka. Pogledaj konzolu pregledača da vidiš zapise sa ubačenim oznakama.

# --hints--

`parseUnorderedList("- Item A\n- Item B")` inapaswa kurudisha `"<ul><li>Item A</li><li>Item B</li></ul>"`.

```js
assert.equal(parseUnorderedList("- Item A\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>");
```

`parseUnorderedList("- JavaScript\n- Python")` treba da vrati `"<ul><li>JavaScript</li><li>Python</li></ul>"`.

```js
assert.equal(parseUnorderedList("-  JavaScript\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>");
```

`parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` inapaswa kurudisha `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`.

```js
assert.equal(parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>");
```

`parseUnorderedList("- A-1\n- A-2\n- B-1")` inapaswa kurudisha `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`.

```js
assert.equal(parseUnorderedList("- A-1\n- A-2\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>");
```

# --seed--

## --seed-contents--

```js
function parseUnorderedList(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseUnorderedList(markdown) {
  const lines = markdown.split("\n");
  const listItems = [];
  
  for (let line of lines) {
    listItems.push(`<li>${line.slice(1).trim()}</li>`);
  }
  
  return `<ul>${listItems.join("")}</ul>`;
}
```
