---
id: 69373793f5a867f769cde13a
title: "Zoezi la 155: Panda Mazao"
challengeType: 28
dashedName: challenge-155
---

# --description--

Na osnovu celog broja koji predstavlja veličinu tvoje farme, i `"acres"` ili `"hectares"` koji predstavljaju jedinicu površine farme, zajedno sa vrstom useva, odredi koliko sadnica te vrste možeš posaditi na svojoj farmi.

- Ekari 1 ni sawa na mita za mraba 4046.86.
- Hekta 1 ni sawa na mita za mraba 10,000.

Ovo je lista useva koji pružaju unos, zajedno sa površinom koju jedna biljka zauzima:

|Površina|Razmak po biljci|
|-|-|
|`"corn"`|jedan kvadrat metar|
|`"wheat"`|0,1 kvadrat metra|
|`"soybeans"`|0,5 kvadrat metra|
|`"tomatoes"`|0,25 kvadrat metra|
|`"lettuce"`|0,2 kvadrat metra|

Vrati broj biljaka koje mogu stati na farmu, zaokružujući naniže na najbliži ceo broj biljaka.

# --hints--

`getNumberOfPlants(1, "acres", "corn")` inapaswa kurudisha `4046`.

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")` inapaswa kurudisha `100000`.

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")` inapaswa kurudisha `161874`.

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")` inapaswa kurudisha `150000`.

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")` inapaswa kurudisha `271139`.

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```
