---
id: 69373793f5a867f769cde136
title: "Vežba 151: Redak organizovanih podataka?
challengeType: 28
dashedName: challenge-151
---

# --description--
Ako dobijete red numeričkih podataka, potvrdite da li su brojevi poređani u rastućem, opadajućem ili da nisu poređani po bilo kom redosledu.

Ako je red podataka koji je izvučen:

- U rastućem redosledu (od malog do velikog), vratite `"Ascending"`.
- U opadajućem redosledu (od velikog do malog), vratite `"Descending"`.
- Nije sortirano u rastućem ili opadajućem redosledu, vratite `"Not sorted"`.

# --hints--

`isSorted([1, 2, 3, 4, 5])` inapaswa kurudisha `"Ascending"`.

```js
assert.equal(isSorted([1, 2, 3, 4, 5]), "Ascending");
```

`isSorted([10, 8, 6, 4, 2])` inapaswa kurudisha `"Descending"`.

```js
assert.equal(isSorted([10, 8, 6, 4, 2]), "Descending");
```

`isSorted([1, 3, 2, 4, 5])` inapaswa kurudisha `"Not sorted"`.

```js
assert.equal(isSorted([1, 3, 2, 4, 5]), "Not sorted");
```

`isSorted([3.14, 2.71, 1.61, 0.57])` inapaswa kurudisha `"Descending"`.

```js
assert.equal(isSorted([3.14, 2.71, 1.61, 0.57]), "Descending");
```

`isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9])` inapaswa kurudisha `"Ascending"`.

```js
assert.equal(isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9]), "Ascending");
```

`isSorted([0.4, 0.5, 0.3])` inapaswa kurudisha `"Not sorted"`.

```js
assert.equal(isSorted([0.4, 0.5, 0.3]), "Not sorted");
```

# --seed--

## --seed-contents--

```js
function isSorted(arr) {

  return arr;
}
```

# --solutions--

```js
function isSorted(arr) {
  let ascending = true;
  let descending = true;

  for (let i = 1; i < arr.length; i++) {
    if (arr[i] < arr[i - 1]) ascending = false;
    if (arr[i] > arr[i - 1]) descending = false;
  }

  if (ascending) return "Ascending";
  if (descending) return "Descending";
  return "Not sorted";
}
```
