---
id: 6939b873185d8e00d453563b
title: "Vežba 156: Neparan ili paran broj?"
challengeType: 28
dashedName: challenge-156
---

# --description--

Ako je dat pozitivan ceo broj, vrati `"Odd"` ako je neparan, a `"Even"` ako je paran.

# --hints--

`oddOrEven(1)` inapaswa kurudisha `"Odd"`.

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` inapaswa kurudisha `"Even"`.

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` inapaswa kurudisha `"Odd"`.

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` inapaswa kurudisha `"Even"`.

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` inapaswa kurudisha `"Odd"`.

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
