---
id: 6939b873185d8e00d453563c
title: "Zoezi la 157: Kuchambua kiungo cha Markdown"
challengeType: 28
dashedName: challenge-157
---

# --description--

Ako je data Markdown veza, vrati odgovarajući HTML.

Markdown veza ima sledeći format: `"[link_text](link_url)"`. Vrati HTML `a` element sa atributom `href` postavljenim na `link_url`, a `link_text` kao sadržajem elementa.

Na primer, ako dobijete `"[freeCodeCamp](https://freecodecamp.org/)"` vratite `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`;

Napomena: Konzola možda neće prikazati oznake HTML u nizu karaktera tokom snimanja poruke — proverite konzolu pregledača da biste videli zapise sa priloženim oznakama.

# --hints--

`parseLink("[freeCodeCamp](https://freecodecamp.org/)")` inapaswa kurudisha `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
assert.equal(parseLink("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>');
```

`parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` inapaswa kurudisha `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
assert.equal(parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>');
```

`parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` inapaswa kurudisha `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
assert.equal(parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>');
```

# --seed--

## --seed-contents--

```js
function parseLink(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseLink(markdown) {
  const match = markdown.match(/\[(.*?)\]\((.*?)\)/);

  const [, text, url] = match;
  return `<a href="${url}">${text}</a>`;
}
```
