---
id: 694596b0585c11170ac7c7fc
title: "Vežba 164: Parsiranje koda unutar Markdown linije"
challengeType: 28
dashedName: challenge-164
---

# --description--
Ako vam je dat Markdown niz znakova koji sadrži jednu ili više kodnih sekcija unutar linije, vratite niz znakova HTML koji odgovara.

Inline kodni dijelovi u liniji u Markdownu koriste jedan znak backtick-a (``` ` ```) Početkom i krajem teksta kod sekcije.

Vratite sekvencu sa svim kodovima zamijenjenim oznakama HTML `code`.

Na primer, ako dobiješ sekvencu ```"Use `let` to declare the variable."```, vrati `"Koristi <code>let</code> za deklarisanje promenljive."`.

Napomena: Konzola možda neće prikazati HTML oznake unutar niza znakova prilikom beleženja poruka. Pogledaj konzolu pregledača da vidiš zapise sa ubačenim oznakama.

# --hints--

```parseInlineCode("Use `let` to declare the variable.")``` should return `"Use <code>let</code> to declare the variable."`.

```js
`assert.equal(parseInlineCode("Koristi `let` za deklarisanje promenljive."), "Koristi <code>let</code> za deklarisanje promenljive.");`
```

```parseInlineCode("Tumia `let` au `const` kutangaza kigezo.")``` should return `"Tumia <code>let</code> au <code>const</code> kutangaza kigezo."`.

```js
assert.equal(parseInlineCode("Use `let` or `const` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.");
```

```parseInlineCode("Run `npm install` then `npm start`.")``` should return `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
`assert.equal(parseInlineCode("Pokreni `npm install` zatim `npm start`."), "Pokreni <code>npm install</code> zatim <code>npm start</code>.");`
```

# --seed--

## --seed-contents--

```js
function parseInlineCode(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseInlineCode(markdown) {
  return markdown.replace(/`([^`]+)`/g, "<code>$1</code>");
}
```
