---
id: 694596b0585c11170ac7c7fa
title: "Zoezi la 162: Matumizi ya nishati"
challengeType: 28
dashedName: challenge-162
---

# --description--

Na osnovu broja kalorija sagorelih tokom vežbanja i broja vat-sati koje su vaši elektronski uređaji potrošili tokom tog vežbanja, odredite šta je potrošilo više energije.

Da biste ih uporedili, konvertujte obe vrednosti u džoule koristeći sledeće transformacije:

- 1 Kalori ni sawa na jouli 4184.
- 1 watt-saa ni sawa na jouli 3600.

Rudisha:

- `"Workout"` ako je vežbanje potrošilo više energije.
- `"Devices"` ako je uređaj potrošio više energije.
- `"Equal"` ako su sve koristile sličnu količinu energije.

# --hints--

`compareEnergy(250, 50)` inapaswa kurudisha `"Workout"`.

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)` inapaswa kurudisha `"Devices"`.

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)` inapaswa kurudisha `"Equal"`.

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)` inapaswa kurudisha `"Workout"`.

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)` inapaswa kurudisha `"Devices"`.

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)` inapaswa kurudisha `"Equal"`.

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```
