---
id: 694596b0585c11170ac7c7fb
title: "Vežba 163: Mala slova suglasnika"
challengeType: 28
dashedName: challenge-163
---

# --description--
Za niz slova koji predstavlja naziv parametra, pretvorite ga u male konsonante koristeći sledeća pravila:

- Sve suglasnike treba da budu u velikim slovima.
- Vowels zote (`a`, `e`, `i`, `o`, `u` kwa hali yoyote) treba da budu u malim slovima.
- Reči separatori (hyphens) (`-`) treba da se izmene u donje crte (underscores) (`_`).

# --hints--

`toConsonantCase("helloworld")` inapaswa kurudisha `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")` inapaswa kurudisha `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-")` inapaswa kurudisha `"_HeLLo_WoRLD_"`.

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")` treba da vrati `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```
