---
id: 696655d24b614176d4c9b78d
title: "Vežba 170: Neparan ili paran dan"
challengeType: 28
dashedName: challenge-170
---

# --description--

Ukitoa timestamp (idadi ya millisekunde tangu mwanzo wa Unix epoch), rudisha:

- `"odd"` ako je broj dana u mesecu za taj vremenski pečat pogrešan.
- `"even"` ako je broj dana u mesecu za taj vremenski pečat tačan.

Na primer, kada izvučete `1769472000000`, timestamp za datum 27. januara 2026., vratite `"odd"` jer je broj dana (`27`) pogrešan.

Napomena: Timestamp je u milisekundama i treba da koristiš datum u UTC vremenskoj zoni, a ne lokalno vreme.

# --hints--

`oddOrEvenDay(1769472000000)` inapaswa kurudisha `"odd"`.

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)` inapaswa kurudisha `"even"`.

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)` inapaswa kurudisha `"odd"`.

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)` inapaswa kurudisha `"odd"`.

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)` inapaswa kurudisha `"even"`.

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```
