---
id: 696655d24b614176d4c9b78b
title: "Zoezi la 168: Picha iliyopimwa upya"
challengeType: 28
dashedName: challenge-168
---

# --description--

Ako je dat niz znakova koji predstavlja širinu i visinu slike i broj za promenu veličine slike, vrati novu širinu i visinu nakon promene veličine.

- Sekvenca ulaznih karakteristika je u formatu `"WxH"`. Na primer, `"800x600"`.
- Merevanje je broj za množenje širine i visine.

Vrati ponovna mjerenja sa istom strukturom kao kod `"WxH"`.

# --hints--

`scaleImage("800x600", 2)` inapaswa kurudisha `"1600x1200"`.

```js
assert.equal(scaleImage("800x600", 2), "1600x1200");
```

`scaleImage("100x100", 10)` inapaswa kurudisha `"1000x1000"`.

```js
assert.equal(scaleImage("100x100", 10), "1000x1000");
```

`scaleImage("1024x768", 0.5)` inapaswa kurudisha `"512x384"`.

```js
assert.equal(scaleImage("1024x768", 0.5), "512x384");
```

`scaleImage("300x200", 1.5)` inapaswa kurudisha `"450x300"`.

```js
assert.equal(scaleImage("300x200", 1.5), "450x300");
```

# --seed--

## --seed-contents--

```js
function scaleImage(size, scale) {

  return size;
}
```

# --solutions--

```js
function scaleImage(size, scale) {
  const [width, height] = size.split("x").map(Number);

  const newWidth = width * scale;
  const newHeight = height * scale;

  return `${newWidth}x${newHeight}`;
}
```
