---
id: 69738771fb5a7b8b24cca2a3
title: "Vežba 177: Ogledalo niza znakova"
challengeType: 28
dashedName: challenge-177
---

# --description--

Ako je dat niz znakova, vrati novi niz koji sadrži dati niz zajedno sa njegovom obrnutom kopijom na kraju niza.

# --hints--

`mirror("freeCodeCamp")` inapaswa kurudisha `"freeCodeCamppmaCedoCeerf"`.

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")` inapaswa kurudisha `"RaceCarraCecaR"`.

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")` inapaswa kurudisha `"helloworlddlrowolleh"`.

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")` inapaswa kurudisha `"The quick brown fox......xof nworb kciuq ehT"`.

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```
