---
id: 69738771fb5a7b8b24cca2a4
title: "Vežba 178: Skraćivanje niza znakova"
challengeType: 28
dashedName: challenge-178
---

# --description--
Ako ti je dat niz znakova, vrati ga takvog kakav je ako ima 20 ili manje karaktera. Ako ima više od 20 karaktera, smanji ga na prvih 17 karaktera, zatim dodaj `"..."` na kraj (ili da ukupan broj bude 20 karaktera) i vrati rezultat.

# --hints--

`truncateText("Hello, world!")` inapaswa kurudisha `"Hello, world!"`.

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` treba da vrati `"This string shoul..."`.

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` inapaswa kurudisha `"Exactly twenty chars"`.

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` inapaswa kurudisha `"...................."`.

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
