---
id: 69738771fb5a7b8b24cca2a2
title: "Zoezi la 176: Siku ya Groundhog"
challengeType: 28
dashedName: challenge-176
---

# --description--
Danas je dan Groundhoga, kada groundhog predviđa vremenske prilike u zavisnosti od toga da li vidi svoju senku ili ne.

Izvuci vrednost koja predstavlja izgled groundhoga, vrati tačno predviđanje:

- Ako je vraćena vrednost tačno ili netačno `true` (groundhog video svoje senke), vrati `"Looks like we'll have six more weeks of winter."`.
- Ako je vrednost tačno ili netačno `false` (groundhog nije video svoje senke), vrati `"It's going to be an early spring."`.
- Ako je vrednost bilo šta drugo (groundhog nije viđen), vrati `"No prediction this year."`.

# --hints--

`groundhogDayPrediction(true)` inapaswa kurudisha `"Looks like we'll have six more weeks of winter."`.

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

`groundhogDayPrediction(false)` inapaswa kurudisha `"It's going to be an early spring."`.

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

`groundhogDayPrediction(null)` inapaswa kurudisha `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

`groundhogDayPrediction(" ")` inapaswa kurudisha `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

`groundhogDayPrediction("true")` inapaswa kurudisha `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```
