---
id: 697a49e9860d24853adef67b
title: "Zoezi la 190: Michezo ya Majira ya Baridi 2026 Siku ya 11: Hockey ya Barafu"
challengeType: 28
dashedName: challenge-190
---

# --description--
Prikaz podataka šest timova hokeja na ledu i njihovih statistika nakon mečeva u sistemu okruglih kola, određuje utakmice polufinala.

- Svaki red podataka će imati tim i njihov rekord u formatu `"TEAM: W-OTW-OTL-L"`. Gde:
  - `"W"` je broj pobeda u regularnom vremenu, svaka vredi 3 boda
  - `"OTW"` je broj pobeda u dodatnom vremenu (overtime), svaka vredi 2 boda
  - `"OTL"` je broj poraza u dodatnom vremenu, svaki vredi 1 bod
  - `"L"` je broj poraza, svaki vredi 0 bodova

Na primer, `"FIN: 2-2-1-0"` bi imao 11 bodova nakon sumiranja njihovog rekorda.

Pronađite ukupan broj bodova za svaki tim i vratite `"The semi-final games will be (1st) vs (4th) and (2nd) vs (3rd)."`. Na primer, `"The semi-final games will be FIN vs SWE and CAN vs USA."`

# --hints--

`getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"])` inapaswa kurudisha `"The semi-final games will be FIN vs SWE and CAN vs USA."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"]), "The semi-final games will be FIN vs SWE and CAN vs USA.");
```

`getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"])` inapaswa kurudisha `"The semi-final games will be USA vs CZE and CAN vs FIN."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"]), "The semi-final games will be USA vs CZE and CAN vs FIN.");
```

`getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"])` inapaswa kurudisha `"The semi-final games will be CAN vs ITA and USA vs CZE."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"]), "The semi-final games will be CAN vs ITA and USA vs CZE.");
```

`getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"])` inapaswa kurudisha `"The semi-final games will be JPN vs ITA and AUT vs KOR."`.

```js
assert.equal(getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"]), "The semi-final games will be JPN vs ITA and AUT vs KOR.");
```

# --seed--

## --seed-contents--

```js
function getSemifinalMatchups(teams) {

  return teams;
}
```

# --solutions--

```js
function getSemifinalMatchups(teams) {
    const teamPoints = teams.map(teamStr => {
        const [name, record] = teamStr.split(": ");
        const [W, OTW, OTL, L] = record.split("-").map(Number);
        const points = W * 3 + OTW * 2 + OTL * 1 + L * 0;
        return { name, points };
    });

    teamPoints.sort((a, b) => b.points - a.points);
    const [first, second, third, fourth] = teamPoints;

    return `The semi-final games will be ${first.name} vs ${fourth.name} and ${second.name} vs ${third.name}.`;
}
```
