---
id: 697a49e9860d24853adef67c
title: "Zoezi la 191: Michezo ya Majira ya Baridi 2026 Siku ya 12: Bobsled"
challengeType: 28
dashedName: challenge-191
---

# --description--
Pružen je red podataka koji predstavlja težinu sportista u bobsled timu i broj koji predstavlja težinu bobsleada; utvrdite da li je tim kvalifikovan za takmičenje.

- Dužina reda podataka određuje veličinu tima: timovi od 1, 2 ili 4 osobe.
- Sve vrednosti težine koje su date su u kilogramima (kg).

Bobsled (sled pekee) lazima iwe na kiwango cha chini cha uzito cha:

- 162 kg za tim od 1 osobe
- 170 kg za tim od 2 osobe
- 210 kg za tim od 4 osobe

Ukupna težina boba (sportisti zajedno sa sankama) ne sme da pređe:

- 247 kg za tim od 1 osobe
- 390 kg za tim od 2 osobe
- 630 kg za tim od 4 osobe

Vrati `"Eligible"` ako tim ispunjava sve uslove, ili `"Not Eligible"` ako ne ispunjava jedan ili više uslova.

# --hints--

`checkEligibility([78], 165)` inapaswa kurudisha `"Eligible"`.

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)` inapaswa kurudisha `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)` inapaswa kurudisha `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)` inapaswa kurudisha `"Eligible"`.

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)` inapaswa kurudisha `"Not Eligible"`.

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)` inapaswa kurudisha `"Not Eligible"`.

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)` inapaswa kurudisha `"Eligible"`.

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)` inapaswa kurudisha `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)` inapaswa kurudisha `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```
