---
id: 697a49e6ff50d756c9b6935e
title: "Zoezi la 181: Michezo ya Majira ya Baridi 2026 Siku ya 2: Snowboarding"
challengeType: 28
dashedName: challenge-181
---

# --description--
Položaj koji snowboarding sportista zauzme kada pokrene rotaciju po stepenima, određuje mu položaj prilikom pada.

- Mkao wa mwanariadha wa snowboarding ni `"Regular"` au `"Goofy"`.
- Rotacije metode su uglavnom za 90 stepeni. Pozitivan broj predstavlja rotaciju u smeru kazaljke na satu, a negativan u suprotnom smeru.
- Položaj pada menja se pri svakoj rotaciji od 180 stepeni.

Na primer, ako se ukloni `"Regular"` i `90`, vrati `"Regular"`. Ako se ukloni `"Regular"` i stepen `180`, vrati `"Goofy"`.

# --hints--

`getLandingStance("Regular", 90)` inapaswa kurudisha `"Regular"`.

```js
assert.equal(getLandingStance("Regular", 90), "Regular");
```

`getLandingStance("Regular", 180)` inapaswa kurudisha `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 180), "Goofy");
```

`getLandingStance("Goofy", -270)` inapaswa kurudisha `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -270), "Regular");
```

`getLandingStance("Regular", 2340)` inapaswa kurudisha `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 2340), "Goofy");
```

`getLandingStance("Goofy", 2160)` inapaswa kurudisha `"Goofy"`.

```js
assert.equal(getLandingStance("Goofy", 2160), "Goofy");
```

`getLandingStance("Goofy", -540)` inapaswa kurudisha `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -540), "Regular");
```

# --seed--

## --seed-contents--

```js
function getLandingStance(startStance, rotation) {

  return startStance;
}
```

# --solutions--

```js
function getLandingStance(startStance, rotation) {
  const flips = Math.floor(Math.abs(rotation) / 180);
  if (flips % 2 === 0) {
    return startStance;
  } else {
    return startStance === "Regular" ? "Goofy" : "Regular";
  }
}
```
