---
id: 697a49e6ff50d756c9b69362
title: "Izazov 185: Zimske Igre 2026 Dan 6: Klizanje po crtežu"
challengeType: 28
dashedName: challenge-185
---

# --description--
Smatrajući red podataka ocjena sudija i dobrovoljne kazne, izračunajte konačnu ocjenu pomoću procedure glatkanja dijagrama.

Prvi red je niz podataka sa ocjenima od 10 sudija, svaka je broj od 0 do 10. Uklonite najviše i najniže ocjene sudija i izračunajte sumu preostalih 8 ocjena kako biste dobili osnovnu ocjenu.

Bilo koji dodatni bodovi preneseni funkciji su kazna. Smanjite sve kazne od baznog rezultata da biste dobili konačni rezultat.

# --hints--

`computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1)` inapaswa kurudisha `64`.

```js
assert.equal(computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1), 64);
```

`computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10])` inapaswa kurudisha `80`.

```js
assert.equal(computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10]), 80);
```

`computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1)` inapaswa kurudisha `67`.

```js
assert.equal(computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1), 67);
```

`computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0)` inapaswa kurudisha `67.5`.

```js
assert.equal(computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0), 67.5);
```

`computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5)` inapaswa kurudisha `59`.

```js
assert.equal(computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5), 59);
```

# --seed--

## --seed-contents--

```js
function computeScore(judgeScores, ...penalties) {

  return judgeScores;
}
```

# --solutions--

```js
function computeScore(judgeScores, ...penalties) {
  const sortedScores = [...judgeScores].sort((a, b) => a - b);
  const baseScore = sortedScores.slice(1, 9).reduce((sum, s) => sum + s, 0);
  const totalPenalty = penalties.reduce((sum, p) => sum + p, 0);

  return baseScore - totalPenalty;
}
```
