---
id: 697a49e6ff50d756c9b69364
title: "Zoezi la 187: Michezo ya Majira ya Baridi 2026 Siku ya 8: Luge"
challengeType: 28
dashedName: challenge-187
---

# --description--
Dati je red podataka sa pet brojeva, svaki od kojih predstavlja vreme (u sekundama) koje je trebalo luđeru da završi deo staze. Utvrdite koji deo ima najveću brzinu i koja je ta brzina.

Putnja je podeljena na sledeće delove:

- Deo 1: metara 320
- Deo 2: metara 280
- Deo 3: metara 350
- Deo 4: metara 300
- Deo 5: metara 250

Prva vrednost u nizu podataka odgovara vremenskom delu 1, druga vrednost za deo 2, i tako dalje.

Da bi se izračunala brzina (u metrima po sekundi) za određeni deo, podeli rastojanje sa vremenom.

Vratite `"The luger's fastest speed was X m/s on segment Y."`. Gde je `X` maksimalna brzina, zaokružena na dve decimale, i gde je `Y` broj sekcije gde je maksimalna brzina nastala.

# --hints--

`getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128])` inapaswa kurudisha `"The luger's fastest speed was 35.07 m/s on segment 5."`

```js
assert.equal(getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.");
```

`getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284])` inapaswa kurudisha `"The luger's fastest speed was 37.75 m/s on segment 2."`

```js
assert.equal(getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.");
```

`getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912])` inapaswa kurudisha `"The luger's fastest speed was 38.49 m/s on segment 3."`

```js
assert.equal(getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.");
```

`getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840])` inapaswa kurudisha `"The luger's fastest speed was 37.69 m/s on segment 1."`

```js
assert.equal(getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.");
```

`getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508])` inapaswa kurudisha `"The luger's fastest speed was 39.24 m/s on segment 4."`

```js
assert.equal(getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.");
```

# --seed--

## --seed-contents--

```js
function getFastestSpeed(times) {

  return times;
}
```

# --solutions--

```js
function getFastestSpeed(times) {
  const distances = [320, 280, 350, 300, 250];
  let fastestSpeed = 0;
  let fastestSegment = 0;

  for (let i = 0; i < times.length; i++) {
    const speed = distances[i] / times[i];
    if (speed > fastestSpeed) {
      fastestSpeed = speed;
      fastestSegment = i + 1;
    }
  }

  return `The luger's fastest speed was ${fastestSpeed.toFixed(2)} m/s on segment ${fastestSegment}.`;
}
```
