---
id: 697a49e6ff50d756c9b6935d
title: "Vežba 180: Zimske olimpijske igre 2026, 1. dan: Dan otvaranja"
challengeType: 28
dashedName: challenge-180
---

# --description--
Danas je početak Zimske Igre 2026. Sledećih 17 dana donosiće vam izazove koda inspirisane tim igrama.

Za prvi izazov, dobijaš dvosločni kod zemlje i moraš vratiti emoji zastave te zemlje.

Koristi ovu listu:

| Nchi | Msimbo | Bendera |
| - | -: | :--: |
| Albania                    |            `"AL"` | `"🇦🇱"` |
| Andorra                    |            `"AD"` | `"🇦🇩"` |
| Argentina                  |            `"AR"` | `"🇦🇷"` |
| Armenia                    |            `"AM"` | `"🇦🇲"` |
| Australia                  |            `"AU"` | `"🇦🇺"` |
| Austria                    |            `"AT"` | `"🇦🇹"` |
| Azerbaijan                 |            `"AZ"` | `"🇦🇿"` |
| Belgium                    |            `"BE"` | `"🇧🇪"` |
| Benin                      |            `"BJ"` | `"🇧🇯"` |
| Bolivia                    |            `"BO"` | `"🇧🇴"` |
| Bosnia and Herzegovina     |            `"BA"` | `"🇧🇦"` |
| Brazil                     |            `"BR"` | `"🇧🇷"` |
| Bulgaria                   |            `"BG"` | `"🇧🇬"` |
| Canada                     |            `"CA"` | `"🇨🇦"` |
| Chile                      |            `"CL"` | `"🇨🇱"` |
| China                      |            `"CN"` | `"🇨🇳"` |
| Colombia                   |            `"CO"` | `"🇨🇴"` |
| Croatia                    |            `"HR"` | `"🇭🇷"` |
| Cyprus                     |            `"CY"` | `"🇨🇾"` |
| Czech Republic             |            `"CZ"` | `"🇨🇿"` |
| Denmark                    |            `"DK"` | `"🇩🇰"` |
| Ecuador                    |            `"EC"` | `"🇪🇨"` |
| Eritrea                    |            `"ER"` | `"🇪🇷"` |
| Estonia                    |            `"EE"` | `"🇪🇪"` |
| Finland                    |            `"FI"` | `"🇫🇮"` |
| France                     |            `"FR"` | `"🇫🇷"` |
| Georgia                    |            `"GE"` | `"🇬🇪"` |
| Germany                    |            `"DE"` | `"🇩🇪"` |
| Great Britain              |            `"GB"` | `"🇬🇧"` |
| Greece                     |            `"GR"` | `"🇬🇷"` |
| Guinea-Bissau              |            `"GW"` | `"🇬🇼"` |
| Haiti                      |            `"HT"` | `"🇭🇹"` |
| Hong Kong                  |            `"HK"` | `"🇭🇰"` |
| Hungary                    |            `"HU"` | `"🇭🇺"` |
| Iceland                    |            `"IS"` | `"🇮🇸"` |
| India                      |            `"IN"` | `"🇮🇳"` |
| Iran                       |            `"IR"` | `"🇮🇷"` |
| Ireland                    |            `"IE"` | `"🇮🇪"` |
| Israel                     |            `"IL"` | `"🇮🇱"` |
| Italy                      |            `"IT"` | `"🇮🇹"` |
| Jamaica                    |            `"JM"` | `"🇯🇲"` |
| Japan                      |            `"JP"` | `"🇯🇵"` |
| Kazakhstan                 |            `"KZ"` | `"🇰🇿"` |
| Kenya                      |            `"KE"` | `"🇰🇪"` |
| Kosovo                     |            `"XK"` | `"🇽🇰"` |
| Kyrgyzstan                 |            `"KG"` | `"🇰🇬"` |
| Latvia                     |            `"LV"` | `"🇱🇻"` |
| Lebanon                    |            `"LB"` | `"🇱🇧"` |
| Liechtenstein              |            `"LI"` | `"🇱🇮"` |
| Lithuania                  |            `"LT"` | `"🇱🇹"` |
| Luxembourg                 |            `"LU"` | `"🇱🇺"` |
| Madagascar                 |            `"MG"` | `"🇲🇬"` |
| Malaysia                   |            `"MY"` | `"🇲🇾"` |
| Malta                      |            `"MT"` | `"🇲🇹"` |
| Mexico                     |            `"MX"` | `"🇲🇽"` |
| Moldova                    |            `"MD"` | `"🇲🇩"` |
| Monaco                     |            `"MC"` | `"🇲🇨"` |
| Mongolia                   |            `"MN"` | `"🇲🇳"` |
| Montenegro                 |            `"ME"` | `"🇲🇪"` |
| Morocco                    |            `"MA"` | `"🇲🇦"` |
| Netherlands                |            `"NL"` | `"🇳🇱"` |
| New Zealand                |            `"NZ"` | `"🇳🇿"` |
| Nigeria                    |            `"NG"` | `"🇳🇬"` |
| North Macedonia            |            `"MK"` | `"🇲🇰"` |
| Norway                     |            `"NO"` | `"🇳🇴"` |
| Pakistan                   |            `"PK"` | `"🇵🇰"` |
| Philippines                |            `"PH"` | `"🇵🇭"` |
| Poland                     |            `"PL"` | `"🇵🇱"` |
| Portugal                   |            `"PT"` | `"🇵🇹"` |
| Puerto Rico                |            `"PR"` | `"🇵🇷"` |
| Romania                    |            `"RO"` | `"🇷🇴"` |
| San Marino                 |            `"SM"` | `"🇸🇲"` |
| Saudi Arabia               |            `"SA"` | `"🇸🇦"` |
| Serbia                     |            `"RS"` | `"🇷🇸"` |
| Singapore                  |            `"SG"` | `"🇸🇬"` |
| Slovakia                   |            `"SK"` | `"🇸🇰"` |
| Slovenia                   |            `"SI"` | `"🇸🇮"` |
| South Africa               |            `"ZA"` | `"🇿🇦"` |
| South Korea                |            `"KR"` | `"🇰🇷"` |
| Spain                      |            `"ES"` | `"🇪🇸"` |
| Sweden                     |            `"SE"` | `"🇸🇪"` |
| Switzerland                |            `"CH"` | `"🇨🇭"` |
| Thailand                   |            `"TH"` | `"🇹🇭"` |
| Trinidad & Tobago          |            `"TT"` | `"🇹🇹"` |
| Turkey                     |            `"TR"` | `"🇹🇷"` |
| Ukraine                    |            `"UA"` | `"🇺🇦"` |
| United Arab Emirates       |            `"AE"` | `"🇦🇪"` |
| United States              |            `"US"` | `"🇺🇸"` |
| Uruguay                    |            `"UY"` | `"🇺🇾"` |
| Uzbekistan                 |            `"UZ"` | `"🇺🇿"` |
| Venezuela                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`getFlag("AL")` inapaswa kurudisha `"🇦🇱"`.

```js
assert.equal(getFlag("AL"), "🇦🇱");
```

`getFlag("AD")` inapaswa kurudisha `"🇦🇩"`.

```js
assert.equal(getFlag("AD"), "🇦🇩");
```

`getFlag("AR")` inapaswa kurudisha `"🇦🇷"`.

```js
assert.equal(getFlag("AR"), "🇦🇷")
```

`getFlag("AM")` inapaswa kurudisha `"🇦🇲"`.

```js
assert.equal(getFlag("AM"), "🇦🇲")
```

`getFlag("AU")` inapaswa kurudisha `"🇦🇺"`.

```js
assert.equal(getFlag("AU"), "🇦🇺")
```

`getFlag("AT")` inapaswa kurudisha `"🇦🇹"`.

```js
assert.equal(getFlag("AT"), "🇦🇹")
```

`getFlag("AZ")` inapaswa kurudisha `"🇦🇿"`.

```js
assert.equal(getFlag("AZ"), "🇦🇿")
```

`getFlag("BE")` inapaswa kurudisha `"🇧🇪"`.

```js
assert.equal(getFlag("BE"), "🇧🇪")
```

`getFlag("BJ")` inapaswa kurudisha `"🇧🇯"`.

```js
assert.equal(getFlag("BJ"), "🇧🇯")
```

`getFlag("BO")` inapaswa kurudisha `"🇧🇴"`.

```js
assert.equal(getFlag("BO"), "🇧🇴")
```

`getFlag("BA")` inapaswa kurudisha `"🇧🇦"`.

```js
assert.equal(getFlag("BA"), "🇧🇦")
```

`getFlag("BR")` inapaswa kurudisha `"🇧🇷"`.

```js
assert.equal(getFlag("BR"), "🇧🇷")
```

`getFlag("BG")` inapaswa kurudisha `"🇧🇬"`.

```js
assert.equal(getFlag("BG"), "🇧🇬")
```

`getFlag("CA")` inapaswa kurudisha `"🇨🇦"`.

```js
assert.equal(getFlag("CA"), "🇨🇦")
```

`getFlag("CL")` inapaswa kurudisha `"🇨🇱"`.

```js
assert.equal(getFlag("CL"), "🇨🇱")
```

`getFlag("CN")` inapaswa kurudisha `"🇨🇳"`.

```js
assert.equal(getFlag("CN"), "🇨🇳")
```

`getFlag("CO")` inapaswa kurudisha `"🇨🇴"`.

```js
assert.equal(getFlag("CO"), "🇨🇴")
```

`getFlag("HR")` inapaswa kurudisha `"🇭🇷"`.

```js
assert.equal(getFlag("HR"), "🇭🇷")
```

`getFlag("CY")` inapaswa kurudisha `"🇨🇾"`.

```js
assert.equal(getFlag("CY"), "🇨🇾")
```

`getFlag("CZ")` inapaswa kurudisha `"🇨🇿"`.

```js
assert.equal(getFlag("CZ"), "🇨🇿")
```

`getFlag("DK")` inapaswa kurudisha `"🇩🇰"`.

```js
assert.equal(getFlag("DK"), "🇩🇰")
```

`getFlag("EC")` inapaswa kurudisha `"🇪🇨"`.

```js
assert.equal(getFlag("EC"), "🇪🇨")
```

`getFlag("ER")` inapaswa kurudisha `"🇪🇷"`.

```js
assert.equal(getFlag("ER"), "🇪🇷")
```

`getFlag("EE")` inapaswa kurudisha `"🇪🇪"`.

```js
assert.equal(getFlag("EE"), "🇪🇪")
```

`getFlag("FI")` inapaswa kurudisha `"🇫🇮"`.

```js
assert.equal(getFlag("FI"), "🇫🇮")
```

`getFlag("FR")` inapaswa kurudisha `"🇫🇷"`.

```js
assert.equal(getFlag("FR"), "🇫🇷")
```

`getFlag("GE")` inapaswa kurudisha `"🇬🇪"`.

```js
assert.equal(getFlag("GE"), "🇬🇪")
```

`getFlag("DE")` inapaswa kurudisha `"🇩🇪"`.

```js
assert.equal(getFlag("DE"), "🇩🇪")
```

`getFlag("GB")` inapaswa kurudisha `"🇬🇧"`.

```js
assert.equal(getFlag("GB"), "🇬🇧")
```

`getFlag("GR")` inapaswa kurudisha `"🇬🇷"`.

```js
assert.equal(getFlag("GR"), "🇬🇷")
```

`getFlag("GW")` inapaswa kurudisha `"🇬🇼"`.

```js
assert.equal(getFlag("GW"), "🇬🇼")
```

`getFlag("HT")` inapaswa kurudisha `"🇭🇹"`.

```js
assert.equal(getFlag("HT"), "🇭🇹")
```

`getFlag("HK")` inapaswa kurudisha `"🇭🇰"`.

```js
assert.equal(getFlag("HK"), "🇭🇰")
```

`getFlag("HU")` inapaswa kurudisha `"🇭🇺"`.

```js
assert.equal(getFlag("HU"), "🇭🇺")
```

`getFlag("IS")` inapaswa kurudisha `"🇮🇸"`.

```js
assert.equal(getFlag("IS"), "🇮🇸")
```

`getFlag("IN")` inapaswa kurudisha `"🇮🇳"`.

```js
assert.equal(getFlag("IN"), "🇮🇳")
```

`getFlag("IR")` inapaswa kurudisha `"🇮🇷"`.

```js
assert.equal(getFlag("IR"), "🇮🇷")
```

`getFlag("IE")` inapaswa kurudisha `"🇮🇪"`.

```js
assert.equal(getFlag("IE"), "🇮🇪")
```

`getFlag("IL")` inapaswa kurudisha `"🇮🇱"`.

```js
assert.equal(getFlag("IL"), "🇮🇱")
```

`getFlag("IT")` inapaswa kurudisha `"🇮🇹"`.

```js
assert.equal(getFlag("IT"), "🇮🇹")
```

`getFlag("JM")` inapaswa kurudisha `"🇯🇲"`.

```js
assert.equal(getFlag("JM"), "🇯🇲")
```

`getFlag("JP")` inapaswa kurudisha `"🇯🇵"`.

```js
assert.equal(getFlag("JP"), "🇯🇵")
```

`getFlag("KZ")` inapaswa kurudisha `"🇰🇿"`.

```js
assert.equal(getFlag("KZ"), "🇰🇿")
```

`getFlag("KE")` inapaswa kurudisha `"🇰🇪"`.

```js
assert.equal(getFlag("KE"), "🇰🇪")
```

`getFlag("XK")` inapaswa kurudisha `"🇽🇰"`.

```js
assert.equal(getFlag("XK"), "🇽🇰")
```

`getFlag("KG")` inapaswa kurudisha `"🇰🇬"`.

```js
assert.equal(getFlag("KG"), "🇰🇬")
```

`getFlag("LV")` inapaswa kurudisha `"🇱🇻"`.

```js
assert.equal(getFlag("LV"), "🇱🇻")
```

`getFlag("LB")` inapaswa kurudisha `"🇱🇧"`.

```js
assert.equal(getFlag("LB"), "🇱🇧")
```

`getFlag("LI")` inapaswa kurudisha `"🇱🇮"`.

```js
assert.equal(getFlag("LI"), "🇱🇮")
```

`getFlag("LT")` inapaswa kurudisha `"🇱🇹"`.

```js
assert.equal(getFlag("LT"), "🇱🇹")
```

`getFlag("LU")` inapaswa kurudisha `"🇱🇺"`.

```js
assert.equal(getFlag("LU"), "🇱🇺")
```

`getFlag("MG")` inapaswa kurudisha `"🇲🇬"`.

```js
assert.equal(getFlag("MG"), "🇲🇬")
```

`getFlag("MY")` inapaswa kurudisha `"🇲🇾"`.

```js
assert.equal(getFlag("MY"), "🇲🇾")
```

`getFlag("MT")` inapaswa kurudisha `"🇲🇹"`.

```js
assert.equal(getFlag("MT"), "🇲🇹")
```

`getFlag("MX")` inapaswa kurudisha `"🇲🇽"`.

```js
assert.equal(getFlag("MX"), "🇲🇽")
```

`getFlag("MD")` inapaswa kurudisha `"🇲🇩"`.

```js
assert.equal(getFlag("MD"), "🇲🇩")
```

`getFlag("MC")` inapaswa kurudisha `"🇲🇨"`.

```js
assert.equal(getFlag("MC"), "🇲🇨")
```

`getFlag("MN")` inapaswa kurudisha `"🇲🇳"`.

```js
assert.equal(getFlag("MN"), "🇲🇳")
```

`getFlag("ME")` inapaswa kurudisha `"🇲🇪"`.

```js
assert.equal(getFlag("ME"), "🇲🇪")
```

`getFlag("MA")` inapaswa kurudisha `"🇲🇦"`.

```js
assert.equal(getFlag("MA"), "🇲🇦")
```

`getFlag("NL")` inapaswa kurudisha `"🇳🇱"`.

```js
assert.equal(getFlag("NL"), "🇳🇱")
```

`getFlag("NZ")` inapaswa kurudisha `"🇳🇿"`.

```js
assert.equal(getFlag("NZ"), "🇳🇿")
```

`getFlag("NG")` inapaswa kurudisha `"🇳🇬"`.

```js
assert.equal(getFlag("NG"), "🇳🇬")
```

`getFlag("MK")` inapaswa kurudisha `"🇲🇰"`.

```js
assert.equal(getFlag("MK"), "🇲🇰")
```

`getFlag("NO")` inapaswa kurudisha `"🇳🇴"`.

```js
assert.equal(getFlag("NO"), "🇳🇴")
```

`getFlag("PK")` inapaswa kurudisha `"🇵🇰"`.

```js
assert.equal(getFlag("PK"), "🇵🇰")
```

`getFlag("PH")` inapaswa kurudisha `"🇵🇭"`.

```js
assert.equal(getFlag("PH"), "🇵🇭")
```

`getFlag("PL")` inapaswa kurudisha `"🇵🇱"`.

```js
assert.equal(getFlag("PL"), "🇵🇱")
```

`getFlag("PT")` inapaswa kurudisha `"🇵🇹"`.

```js
assert.equal(getFlag("PT"), "🇵🇹")
```

`getFlag("PR")` inapaswa kurudisha `"🇵🇷"`.

```js
assert.equal(getFlag("PR"), "🇵🇷")
```

`getFlag("RO")` inapaswa kurudisha `"🇷🇴"`.

```js
assert.equal(getFlag("RO"), "🇷🇴")
```

`getFlag("SM")` inapaswa kurudisha `"🇸🇲"`.

```js
assert.equal(getFlag("SM"), "🇸🇲")
```

`getFlag("SA")` inapaswa kurudisha `"🇸🇦"`.

```js
assert.equal(getFlag("SA"), "🇸🇦")
```

`getFlag("RS")` inapaswa kurudisha `"🇷🇸"`.

```js
assert.equal(getFlag("RS"), "🇷🇸")
```

`getFlag("SG")` inapaswa kurudisha `"🇸🇬"`.

```js
assert.equal(getFlag("SG"), "🇸🇬")
```

`getFlag("SK")` inapaswa kurudisha `"🇸🇰"`.

```js
assert.equal(getFlag("SK"), "🇸🇰")
```

`getFlag("SI")` inapaswa kurudisha `"🇸🇮"`.

```js
assert.equal(getFlag("SI"), "🇸🇮")
```

`getFlag("ZA")` inapaswa kurudisha `"🇿🇦"`.

```js
assert.equal(getFlag("ZA"), "🇿🇦")
```

`getFlag("KR")` inapaswa kurudisha `"🇰🇷"`.

```js
assert.equal(getFlag("KR"), "🇰🇷")
```

`getFlag("ES")` inapaswa kurudisha `"🇪🇸"`.

```js
assert.equal(getFlag("ES"), "🇪🇸")
```

`getFlag("SE")` inapaswa kurudisha `"🇸🇪"`.

```js
assert.equal(getFlag("SE"), "🇸🇪")
```

`getFlag("CH")` inapaswa kurudisha `"🇨🇭"`.

```js
assert.equal(getFlag("CH"), "🇨🇭")
```

`getFlag("TH")` inapaswa kurudisha `"🇹🇭"`.

```js
assert.equal(getFlag("TH"), "🇹🇭")
```

`getFlag("TT")` inapaswa kurudisha `"🇹🇹"`.

```js
assert.equal(getFlag("TT"), "🇹🇹")
```

`getFlag("TR")` inapaswa kurudisha `"🇹🇷"`.

```js
assert.equal(getFlag("TR"), "🇹🇷")
```

`getFlag("UA")` inapaswa kurudisha `"🇺🇦"`.

```js
assert.equal(getFlag("UA"), "🇺🇦")
```

`getFlag("AE")` inapaswa kurudisha `"🇦🇪"`.

```js
assert.equal(getFlag("AE"), "🇦🇪")
```

`getFlag("US")` inapaswa kurudisha `"🇺🇸"`.

```js
assert.equal(getFlag("US"), "🇺🇸")
```

`getFlag("UY")` inapaswa kurudisha `"🇺🇾"`.

```js
assert.equal(getFlag("UY"), "🇺🇾")
```

`getFlag("UZ")` inapaswa kurudisha `"🇺🇿"`.

```js
assert.equal(getFlag("UZ"), "🇺🇿")
```

`getFlag("VE")` inapaswa kurudisha `"🇻🇪"`.

```js
assert.equal(getFlag("VE"), "🇻🇪")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
