---
id: 697a49e9860d24853adef67e
title: "Vežba 193: Zimske igre 2026 Dan 14: Planinsko skijanje
challengeType: 28
dashedName: challenge-193
---

# --description--

Na osnovu dubine snega i nagiba planine, odredi da li postoji opasnost od snežne lavine.

- Vrednosti za dubinu snega su `"Shallow"`, `"Moderate"`, ili `"Deep"`.
- Vrednosti za strminu su `"Gentle"`, `"Steep"`, ili `"Very Steep"`.

Vratite `"Safe"` ili `"Risky"` prema ovoj tabeli:

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalancheRisk("Shallow", "Gentle")` inapaswa kurudisha `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Gentle"), "Safe");
```

`avalancheRisk("Shallow", "Steep")` inapaswa kurudisha `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Steep"), "Safe");
```

`avalancheRisk("Shallow", "Very Steep")` inapaswa kurudisha `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Very Steep"), "Safe");
```

`avalancheRisk("Moderate", "Gentle")` inapaswa kurudisha `"Safe"`.

```js
assert.equal(avalancheRisk("Moderate", "Gentle"), "Safe");
```

`avalancheRisk("Moderate", "Steep")` inapaswa kurudisha `"Risky"`.

```js
assert.equal(avalancheRisk("Moderate", "Steep"), "Risky");
```

`avalancheRisk("Moderate", "Very Steep")` inapaswa kurudisha `"Risky"`.

```js
assert.equal(avalancheRisk("Moderate", "Very Steep"), "Risky");
```

`avalancheRisk("Deep", "Gentle")` inapaswa kurudisha `"Safe"`.

```js
assert.equal(avalancheRisk("Deep", "Gentle"), "Safe");
```

`avalancheRisk("Deep", "Steep")` inapaswa kurudisha `"Risky"`.

```js
assert.equal(avalancheRisk("Deep", "Steep"), "Risky");
```

`avalancheRisk("Deep", "Very Steep")` inapaswa kurudisha `"Risky"`.

```js
assert.equal(avalancheRisk("Deep", "Very Steep"), "Risky");
```

# --seed--

## --seed-contents--

```js
function avalancheRisk(snowDepth, slope) {

  return snowDepth;
}
```

# --solutions--

```js
function avalancheRisk(snowDepth, slope) {
  const riskTable = {
    Gentle: {
      Shallow: "Safe",
      Moderate: "Safe",
      Deep: "Safe",
    },
    Steep: {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
    "Very Steep": {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
  };

  return riskTable[slope][snowDepth];
}
```
