---
id: 697a49e9860d24853adef680
title: "Zoezi la 195: Michezo ya Majira ya Baridi 2026 Siku ya 16: Curling"
challengeType: 28
dashedName: challenge-195
---

# --description--

Ako je data matrica 5x5 koja predstavlja "kuću" na kraju runde curlinga, odredi koji tim osvaja poene i koliko poena dobija.

Mpangilio:

- Centralni modul (ilustracija `[2, 2]`) je „taster“.
- Os modula koji okružuju taster direktno predstavljaju prsten 1.
- A šesnaest modula na spoljašnjem obodu strukture predstavljaju prsten 2.

U datoj matrici:

- `"."` predstavlja prazno mesto.
- `"R"` predstavlja mesto sa crvenim kamenom.
- `"Y"` predstavlja mesto sa žutim kamenom.

Sheria za kupata alama:

- Samo jedna ekipa može osvojiti bodove po rundi.
- Ekipa čiji kamen je najbliži dugmetu osvaja bodove.
- Ekipa koja osvoji bodove dobija 1 poen za svaki svoj kamen koji je bliži dugmetu od najbližeg kamena protivnika.
- Broj najmanjeg prstena znači da je kamen najbliži centru.
- Ako su najbliži kamenovi obe ekipe jednako udaljeni od centra, nijedna ekipa ne osvaja bodove.

Rudisha:

- Sekvenca slova u formatu `"team: number_of_points"`. Primer: `"R: 2"`.
- ili `"No points awarded"` ako nijedan tim nije osvojio poene.

Na primer, ako se ukloni:

```js
[
  [".", ".", "R", ".", "."],
  [".", "R", ".", ".", "."],
  ["Y", ".", ".", ".", "."],
  [".", "R", ".", ".", "."],
  [".", ".", ".", ".", "."]
]
```

Rudisha `"R: 2"`. Dva crvena dragulja u prstenu 1 pobedila su jer su bliži, i oni su jedini dva dragulja koja su bliža od najbližeg žutog dragulja.

# --hints--

`scoreCurling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]])` inapaswa kurudisha `"R: 2"`.

```js
assert.equal(scoreCurling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]]), "R: 2");
```

`scoreCurling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]])` inapaswa kurudisha `"Y: 3"`.

```js
assert.equal(scoreCurling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]]), "Y: 3");
```

`scoreCurling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]])` inapaswa kurudisha `"No points awarded"`.

```js
assert.equal(scoreCurling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]]), "No points awarded");
```

`scoreCurling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]])` inapaswa kurudisha `"R: 4"`.

```js
assert.equal(scoreCurling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]]), "R: 4");
```

`scoreCurling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]])` inapaswa kurudisha `"Y: 1"`.

```js
assert.equal(scoreCurling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]]), "Y: 1");
```

`scoreCurling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]])` inapaswa kurudisha `"No points awarded"`.

```js
assert.equal(scoreCurling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]]), "No points awarded");
```

# --seed--

## --seed-contents--

```js
function scoreCurling(house) {

  return house;
}
```

# --solutions--

```js
function scoreCurling(house) {
  const stones = { R: [], Y: [] };

  for (let i = 0; i < 5; i++) {
    for (let j = 0; j < 5; j++) {
      const cell = house[i][j];
      if (cell === 'R' || cell === 'Y') {
        const distance = Math.max(Math.abs(i - 2), Math.abs(j - 2));
        stones[cell].push(distance);
      }
    }
  }

  const redClosest = stones.R.length > 0 ? Math.min(...stones.R) : Infinity;
  const yellowClosest = stones.Y.length > 0 ? Math.min(...stones.Y) : Infinity;

  if (redClosest === Infinity && yellowClosest === Infinity) {
    return "No points awarded";
  }

  if (redClosest === yellowClosest) {
    return "No points awarded";
  }

let winner, opponentClosest, winningStones;
  if (redClosest < yellowClosest) {
    winner = 'R';
    opponentClosest = yellowClosest;
    winningStones = stones.R;
  } else {
    winner = 'Y';
    opponentClosest = redClosest;
    winningStones = stones.Y;
  }

  const points = winningStones.filter(d => d < opponentClosest).length;

  return points === 0 ? "No points awarded" : `${winner}: ${points}`;
}
```
