---
id: 697a49e9860d24853adef67d
title: "Zoezi la 192: Michezo ya Majira ya Baridi 2026 Siku ya 13: Mchanganyiko wa Nordic"
challengeType: 28
dashedName: challenge-192
---

# --description--
Generisanje reda podataka sa rezultatima skokova sportista i izračunavanje njihovog vremena kašnjenja za kros-kanjeći deo trka nordijskog kombinovanja.

Atlet sa najvišom visinom skoka počinje prvi (bez kašnjenja od 0 sekundi). Svi ostali sportisti počinju kasnije, u zavisnosti od udaljenosti njihovih skokova od najboljeg rezultata.

Da biste izračunali kašnjenje za svakog sportistu, oduzmite ocenu skoka sportiste od najbolje ukupne ocene skoka, zatim rezultat pomnožite sa 1,5. Zaokružite kašnjenje na sledeći ceo broj.

# --hints--

`calculateStartDelays([120, 110, 125])` inapaswa kurudisha `[8, 23, 0]`.

```js
assert.deepEqual(calculateStartDelays([120, 110, 125]), [8, 23, 0]);
```

`calculateStartDelays([118, 125, 122, 120])` inapaswa kurudisha `[11, 0, 5, 8]`.

```js
assert.deepEqual(calculateStartDelays([118, 125, 122, 120]), [11, 0, 5, 8]);
```

`calculateStartDelays([100, 105, 95, 110, 120, 115, 108])` inapaswa kurudisha `[30, 23, 38, 15, 0, 8, 18]`.

```js
assert.deepEqual(calculateStartDelays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18]);
```

`calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` inapaswa kurudisha `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`.

```js
assert.deepEqual(calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12]);
```

# --seed--

## --seed-contents--

```js
function calculateStartDelays(jumpScores) {

  return jumpScores
}
```

# --solutions--

```js
function calculateStartDelays(jumpScores) {
  const bestJump = Math.max(...jumpScores);
  return jumpScores.map(score => Math.ceil((bestJump - score) * 1.5));
}
```
