---
id: 698a1a863194f1f4e63f645e
title: "Vežba 207: Najmanji razmak"
challengeType: 28
dashedName: challenge-207
---

# --description--
Ako se primi sekvenca karaktera, vrati najmanji podniz između dva identična karaktera koje razdvaja razmak već veći od jednog karaktera (minimalni razmak).

- Uvek će postojati najmanje par poklapanja slova.
- Mali deo sekvence koji se vraća ne sme da sadrži poklapanja slova.
- Ako su dve ili više razmaka iste dužine, vrati slova iz prvog razmaka.

Na primer, ako se ukloni `"ABCDAC"`, vratite `"DA"` zato što:

- Ni `"A"` na `"C"` pekee zinazorudiwa katika mfuatano.
- Broj znakova između dva slova `"A"` je 3, a između dva slova `"C"` je 2.
- Zato vrati podniz između dva slova `"C"`.

# --hints--

`smallestGap("ABCDAC")` inapaswa kurudisha `"DA"`.

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")` inapaswa kurudisha `"e"`.

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` inapaswa kurudisha `"#q6e&rkf(p"`.

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")` inapaswa kurudisha `""`.

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")` inapaswa kurudisha `"fox"`.

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
