---
id: 698a1a863194f1f4e63f6460
title: "Zoezi la 209: Ukubwa wa kipengee"
challengeType: 28
dashedName: challenge-209
---

# --description--
Uzimajući u obzir veličinu prozora, širinu elementa u dimenzijama vidljivog područja za `"vw"` i visinu elementa u dimenzijama vidljivog područja za `"vh"`, odredite veličinu elementa u pikselima.

Veličina prozora i veličina vraćenog elementa su niz karaktera u formatu `"width x height"`, `"1200 x 800"` na primer.

- Dimenzije za `"vw"` su procenat širine prozora. `"50vw"` na primer, je 50% širine prozora.
- Dimenzije za `"vh"` su procenat visine prozora. `"50vh"` na primer, je 50% visine prozora.

# --hints--

`getElementSize("1200 x 800", "50vw", "50vh")` inapaswa kurudisha `"600 x 400"`.

```js
assert.equal(getElementSize("1200 x 800", "50vw", "50vh"), "600 x 400");
```

`getElementSize("320 x 480", "25vw", "50vh")` inapaswa kurudisha `"80 x 240"`.

```js
assert.equal(getElementSize("320 x 480", "25vw", "50vh"), "80 x 240");
```

`getElementSize("1000 x 500", "7vw", "3vh")` inapaswa kurudisha `"70 x 15"`.

```js
assert.equal(getElementSize("1000 x 500", "7vw", "3vh"), "70 x 15");
```

`getElementSize("1920 x 1080", "95vw", "100vh")` inapaswa kurudisha `"1824 x 1080"`.

```js
assert.equal(getElementSize("1920 x 1080", "95vw", "100vh"), "1824 x 1080");
```

`getElementSize("1200 x 800", "0vw", "0vh")` inapaswa kurudisha `"0 x 0"`.

```js
assert.equal(getElementSize("1200 x 800", "0vw", "0vh"), "0 x 0");
```

`getElementSize("1440 x 900", "100vw", "114vh")` inapaswa kurudisha `"1440 x 1026"`.

```js
assert.equal(getElementSize("1440 x 900", "100vw", "114vh"), "1440 x 1026");
```

# --seed--

## --seed-contents--

```js
function getElementSize(windowSize, elementVw, elementVh) {

  return windowSize;
}
```

# --solutions--

```js
function getElementSize(windowSize, elementVw, elementVh) {
  const [windowWidth, windowHeight] = windowSize
    .split(" x ")
    .map(Number);

  const vw = Number(elementVw.replace("vw", ""));
  const vh = Number(elementVh.replace("vh", ""));

  const widthPx = (vw / 100) * windowWidth;
  const heightPx = (vh / 100) * windowHeight;

  return `${widthPx} x ${heightPx}`;
}
```
