---
id: 698a1a73ade5ac0e19180fa1
title: "Zoezi la 198: Hesabu ya siku za biashara"
challengeType: 28
dashedName: challenge-198
---

# --description--

Ako su dati početni i završni datum, vrati broj radnih dana između njih.

- Navedeni datumi su u formatu `"YYYY-MM-DD"`.
- Radni dani su poslovni dani (Ponedeljak do Petka).
- Vikend nije radni dan (Subota i Nedelja).
- Uključite sve datume, od početnog do krajnjeg, prilikom računanja.

# --hints--

`countBusinessDays("2026-02-24", "2026-02-26")` inapaswa kurudisha `3`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

`countBusinessDays("2026-02-24", "2026-02-28")` inapaswa kurudisha `4`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

`countBusinessDays("2026-02-21", "2026-03-01")` inapaswa kurudisha `5`.

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

`countBusinessDays("2026-03-08", "2026-03-17")` inapaswa kurudisha `7`.

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

`countBusinessDays("2026-02-24", "2027-02-24")` inapaswa kurudisha `262`.

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```
