---
id: 698a1a73ade5ac0e19180fa0
title: "Vežba 197: Kompatibilnost krvnih grupa"
challengeType: 28
dashedName: challenge-197
---

# --description--
Uzevši u obzir tip krvi donatora i tip krvi primaoca, utvrdite da li donator može dati krv primaocu.

Kila aina ya damu ina:

- Slovi: `"A"`, `"B"`, `"AB"`, ili `"O"`
- I za element Rh: `"+"` ili `"-"`

Tipovi krvi će biti jedan od validnih slova praćen faktorom Rh. Na primer, `"AB+"` i `"O-"` su validni tipovi krvi.

Sheria za herufi:

- `"O"` može da donira krv za drugi tip slova.
- `"A"` može da donira krv za `"A"` i `"AB"`.
- `"B"` može da donira krv za `"B"` i `"AB"`.
- `"AB"` može da donira krv za `"AB"` samo.

Sheria za Rh:

- Hasi (`"-"`) može dati krv kom `"-"` i `"+"`.
- Chanya (`"+"`) može dati krv samo kom `"+"`.

Moraju biti ispunjeni svi uslovi kompatibilnosti za tip krvi i Rh faktor kako bi donator mogao dati krv primaocu.

# --hints--

`canDonate("B+", "B+")` inapaswa kurudisha `true`.

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")` inapaswa kurudisha `true`.

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")` inapaswa kurudisha `false`.

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")` inapaswa kurudisha `true`.

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")` inapaswa kurudisha `false`.

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")` inapaswa kurudisha `true`.

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")` inapaswa kurudisha `false`.

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")` inapaswa kurudisha `true`.

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")` inapaswa kurudisha `false`.

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")` inapaswa kurudisha `false`.

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```
