---
id: 698a1a73ade5ac0e19180fa9
title: "Vežba 206: Vrednosti karata igre
challengeType: 28
dashedName: challenge-206
---

# --description--
Kada izvučete red podataka o kartama igre, vratite novi red podataka sa vrednošću broja svake karte.

Vrednosti kartica:

- As (`"A"`) ima vrednost 1.
- Kartice sa brojevima (`"2"` - `"10"`) imaju vrednost jednaku brojevima ispred njih: 2 - 10, respektivno.
- Figurativne kartice: Džek (`"J"`), Kraljica (`"Q"`) i Kralj (`"K"`) svaka ima vrednost 10.

Aina za kadi:

- Kila kadi ina aina: Spades (`"S"`), Clubs (`"C"`), Diamonds (`"D"`), au Hearts (`"H"`).

Muundo wa kadi:

- Svaka karta je predstavljena kao niz slova: `"valueSuit"`. Na primer: `"AS"` je As Špila, `"10H"` je Deset Srca, i `"QC"` je Kraljica Špila.

# --hints--

`cardValues(["3H", "4D", "5S"])` inapaswa kurudisha `[3, 4, 5]`.

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` inapaswa kurudisha `[1, 10, 10, 6, 7]`.

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` inapaswa kurudisha `[8, 10, 2, 10, 9]`.

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` inapaswa kurudisha `[1, 10]`.

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` inapaswa kurudisha `[10, 10, 10, 10, 1]`.

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
