---
id: 6994cff2290543b3aec9f511
title: "Vežba 212: Unos u red podataka
challengeType: 28
dashedName: challenge-212
---

# --description--
Kada se pruži red podataka, vrednost za ubacivanje u taj red i indikator mesta gde treba ubaciti tu vrednost, vrati novi red podataka sa vrednošću umetnutom na navedeno mesto.

# --hints--

`insertIntoArray([2, 4, 8, 10], 6, 2)` inapaswa kurudisha `[2, 4, 6, 8, 10]`.

```js
assert.deepEqual(insertIntoArray([2, 4, 8, 10], 6, 2), [2, 4, 6, 8, 10]);
```

`insertIntoArray(["the", "quick", "fox"], "brown", 2)` inapaswa kurudisha `["the", "quick", "brown", "fox"]`.

```js
assert.deepEqual(insertIntoArray(["the", "quick", "fox"], "brown", 2), ["the", "quick", "brown", "fox"]);
```

`insertIntoArray([], 0, 0)` inapaswa kurudisha `[0]`.

```js
assert.deepEqual(insertIntoArray([], 0, 0), [0]);
```

`insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5)` inapaswa kurudisha `[0, 1, 1, 2, 3, 5, 8, 13]`.

```js
assert.deepEqual(insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5), [0, 1, 1, 2, 3, 5, 8, 13]);
```

# --seed--

## --seed-contents--

```js
function insertIntoArray(arr, value, index) {

  return arr;
}
```

# --solutions--

```js
function insertIntoArray(arr, value, index) {
  return [
    ...arr.slice(0, index),
    value,
    ...arr.slice(index)
  ];
}
```
