---
id: 699c8e045ee7cb94ed2322d6
title: "Zoezi la 215: Kihesabu cha Ada ya Maegesho"
challengeType: 28
dashedName: challenge-215
---

# --description--

Dat je niz od dva znaka koji predstavlja vreme kada si parkirao automobil i vreme kada si ga preuzeo, izračunaj cenu parkinga.

- Generisana sekvenca karaktera će biti u formatu `"HH:MM"` koristeći 24-satni sistem. Na primer, `"14:00"` je 2 ujutru.
- Prva sekvenca će biti vreme kada ste parkirali vozilo, a druga će biti vreme kada ste ga preuzeli.
- Ako je vreme preuzimanja ranije nego vreme dolaska, to znači da je parkiranje trajalo do nakon šest uveče do sledećeg dana.

Sheria za ada:

- Kila saa ya maegesho inagharimu $3.
- Saa zisizokamilika zinarundikwa hadi saa kamili inayofuata.
- Ako je parking trajao preko noći (posle ponoći), naplaćuje se dodatna noćna naknada od 10$.
- Postoji minimalna naknada od 5$ (primenjuje se samo ako bi ukupan iznos bio manji od 5$).

Vratite ukupne troškove u formatu `"$cost"`, `"$5"` na primer.

# --hints--

`calculateParkingFee("09:00", "11:00")` inapaswa kurudisha `"$6"`.

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")` inapaswa kurudisha `"$5"`.

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")` inapaswa kurudisha `"$9"`.

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")` inapaswa kurudisha `"$27"`.

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")` inapaswa kurudisha `"$34"`.

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")` inapaswa kurudisha `"$82"`.

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
