---
id: 699c8e045ee7cb94ed2322d5
title: "Zoezi la 214: Kagua mnyororo wa domino"
challengeType: 28
dashedName: challenge-214
---

# --description--
Ako vam je dat red podataka iz domino sekvence, utvrdite da li formira validnu lančanu vezu.

- Svaki element u nizu podataka predstavlja domino i sastoji se od para brojeva od 1 do 6, (uključujući).
- Da bi niz bio validan, drugi broj svake domino mora odgovarati prvom broju naredne domino.
- Prvi broj prve domino i poslednji broj zadnje domino ne moraju odgovarati ničemu.

# --hints--

`isValidDominoChain([[1, 3], [3, 6], [6, 5]])` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidDominoChain([[1, 3], [3, 6], [6, 5]]));
```

`isValidDominoChain([[6, 2], [3, 4], [4, 1]])` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidDominoChain([[6, 2], [3, 4], [4, 1]]));
```

`isValidDominoChain([[2, 5], [5, 6], [5, 1]])` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidDominoChain([[2, 5], [5, 6], [5, 1]]));
```

`isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]));
```

`isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]));
```

# --seed--

## --seed-contents--

```js
function isValidDominoChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function isValidDominoChain(dominoes) {
  for (let i = 0; i < dominoes.length - 1; i++) {
    if (dominoes[i][1] !== dominoes[i+1][0]) return false
  }

  return true;
}
```
