---
id: 699c8e045ee7cb94ed2322db
title: "Vežba 220: Najveći broj"
challengeType: 28
dashedName: challenge-220
---

# --description--

Na osnovu niza brojeva razdvojenih različitim znakovima interpunkcije, vrati najveći broj.

- Dati niz će sadržati samo brojeve i interpunkcijske znakove.
- Interpunkcijski znakovi mogu biti zarez (`","`), uzvičnik (`"!"`), upitnik (`"?"`), dvotačka (`":"`) ili tačka-zarez (`";"`).

# --hints--

`largestNumber("1,2")` inapaswa kurudisha `2`.

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")` inapaswa kurudisha `60`.

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` inapaswa kurudisha `-402`.

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")` inapaswa kurudisha `99.9`.

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```
