---
id: 699c8e045ee7cb94ed2322dd
title: "Zoezi la 222: Vivuli vya Equinox"
challengeType: 28
dashedName: challenge-222
---

# --description--
Danas je dan ekvinokcija, kada je sunce direktno iznad središnje linije Zemlje i najviše tačno podne. U nekom trenutku, izmerite senku koju baca uspravna šipka od 4 stopa.

- Vreme će biti sekvenca slova u 24-satnom formatu `"HH:MM"` (na primer, `"15:00"` je 3 ujutru).
- Biće vam dato vreme samo u intervalima od 30 minuta.

Sheria:

- Sunce izlazi tačno u 6 ujutro `"east"` i zalazi tačno u 6 popodne `"west"`.
- Senka je uvek usmerena suprotno od sunca.
- Dužina senke (u stopama) je broj sati od podne, podignut na treći stepen.
- Nema senke pre nego što sunce izađe (pre 6 ujutro), nakon zalaska sunca (od 6 popodne ili kasnije), niti podne.

Rudisha:

- Ako postoji senka, vrati `"(length)ft (direction)"`. Na primer, `"8ft west"`.
- Inače, vrati `"No shadow"`.

Na primer, ako dobijete `"10:00"`, vratite `"8ft west"` jer 10 ujutro je 2 sata od šest popodne, stoga 2<sup>3</sup> = futi 8, i senka pokazuje na zapad jer je sunce na istoku u 10 ujutro.

# --hints--

`getShadow("10:00")` inapaswa kurudisha `"8ft west"`.

```js
assert.equal(getShadow("10:00"), "8ft west");
```

`getShadow("15:00")` inapaswa kurudisha `"27ft east"`.

```js
assert.equal(getShadow("15:00"), "27ft east");
```

`getShadow("12:00")` inapaswa kurudisha `"No shadow"`.

```js
assert.equal(getShadow("12:00"), "No shadow");
```

`getShadow("17:30")` inapaswa kurudisha `"166.375ft east"`.

```js
assert.equal(getShadow("17:30"), "166.375ft east");
```

`getShadow("05:00")` inapaswa kurudisha `"No shadow"`.

```js
assert.equal(getShadow("05:00"), "No shadow");
```

`getShadow("06:00")` inapaswa kurudisha `"216ft west"`.

```js
assert.equal(getShadow("06:00"), "216ft west");
```

`getShadow("18:00")` inapaswa kurudisha `"No shadow"`.

```js
assert.equal(getShadow("18:00"), "No shadow");
```

`getShadow("07:30")` inapaswa kurudisha `"91.125ft west"`.

```js
assert.equal(getShadow("07:30"), "91.125ft west");
```

`getShadow("00:00")` inapaswa kurudisha `"No shadow"`.

```js
assert.equal(getShadow("00:00"), "No shadow");
```

# --seed--

## --seed-contents--

```js
function getShadow(time) {

  return time;
}
```

# --solutions--

```js
function getShadow(time) {
  const [hourStr, minuteStr] = time.split(":");
  const hours = Number(hourStr);
  const minutes = Number(minuteStr);
  const timeInHours = hours + minutes / 60;

  if (timeInHours < 6 || timeInHours >= 18 || timeInHours === 12) {
    return "No shadow";
  }

  const hoursFromNoon = Math.abs(12 - timeInHours);
  const length = Math.pow(hoursFromNoon, 3);
  const direction = timeInHours < 12 ? "west" : "east";

  return `${length}ft ${direction}`;
}
```
