---
id: 699c8e045ee7cb94ed2322d8
title: "Zoezi la 217: Vipande vya chess vilivyokamatwa"
challengeType: 28
dashedName: challenge-217
---

# --description--
Generisati red podataka sekvence karaktera koji predstavljaju šahovske figure koje još uvek imate na tabli, i izračunati vrednost figura koje je vaš protivnik uhvatio.

Katika chess, unaanza na vipande 16:

| Figura | Kifupisho | Broj | Vrednost |
| :---: | :---: | :---: | :---: |
| Pešak | `"P"` | 8 | 1 |
| Toranj | `"R"` | 2 | 5 |
| Konj | `"N"` | 2 | 3 |
| Slonac | `"B"` | 2 | 3 |
| Dama | `"Q"` | 1 | 9 |
| Kralj | `"K"` | 1 | 0 |

- Izvučeni set podataka će sadržati samo navedene skraćenice.
- Bilo koji delovi od 16 koji nisu u izvučenom setu podataka su uhvaćeni.
- Vratite ukupnu vrednost svih uhvaćenih delova, osim...
- Ako je Kralj uhvaćen, vratite `"Checkmate"`.

# --hints--

`getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` inapaswa kurudisha `8`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8);
```

`getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"])` inapaswa kurudisha `26`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"]), 26);
```

`getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` inapaswa kurudisha `16`.

```js
assert.equal(getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16);
```

`getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` inapaswa kurudisha `4`.

```js
assert.equal(getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4);
```

`getCapturedValue(["P", "K"])` inapaswa kurudisha `38`.

```js
assert.equal(getCapturedValue(["P", "K"]), 38);
```

`getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` inapaswa kurudisha `0`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0);
```

`getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` inapaswa kurudisha `"Checkmate"`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate");
```

# --seed--

## --seed-contents--

```js
function getCapturedValue(pieces) {

  return pieces;
}
```

# --solutions--

```js
function getCapturedValue(pieces) {
  if (!pieces.includes("K")) return "Checkmate";

  const values = { P:1, R:5, N:3, B:3, Q:9, K:0 };
  let remaining = "PPPPPPPPRRNNBBQK";

  for (const p of pieces) {
    remaining = remaining.replace(p, "");
  }

  return [...remaining].reduce((sum, p) => sum + values[p], 0);
}
```
