---
id: 699c8e045ee7cb94ed2322d9
title: "Jedinica 218: Ravnomerna podela
challengeType: 28
dashedName: challenge-218
---

# --description--
Kada uzmete dva celobrojna broja, utvrdite da li prvi broj može biti podeljen ravnomerno sa drugim.

# --hints--

`isEvenlyDivisible(4, 2)` inapaswa kurudisha `true`.

```js
assert.isTrue(isEvenlyDivisible(4, 2));
```

`isEvenlyDivisible(7, 3)` inapaswa kurudisha `false`.

```js
assert.isFalse(isEvenlyDivisible(7, 3));
```

`isEvenlyDivisible(5, 10)` inapaswa kurudisha `false`.

```js
assert.isFalse(isEvenlyDivisible(5, 10));
```

`isEvenlyDivisible(48, 6)` inapaswa kurudisha `true`.

```js
assert.isTrue(isEvenlyDivisible(48, 6));
```

`isEvenlyDivisible(3186, 9)` inapaswa kurudisha `true`.

```js
assert.isTrue(isEvenlyDivisible(3186, 9));
```

`isEvenlyDivisible(4192, 11)` inapaswa kurudisha `false`.

```js
assert.isFalse(isEvenlyDivisible(4192, 11));
```

# --seed--

## --seed-contents--

```js
function isEvenlyDivisible(a, b) {

  return a;
}
```

# --solutions--

```js
function isEvenlyDivisible(a, b) {
  return a % b === 0;
}
```
