---
id: 69a890af247de743333bd4d2
title: "Vežba 229: Skraćivanje niza znakova 2"
challengeType: 28
dashedName: challenge-229
---

# --description--

Ako je dat niz znakova, vrati novi skraćeni niz tako da ukupna širina znakova ne prelazi 50 jedinica.

Svaki znak ima određenu širinu:

| Herufi | Upana |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

Gornja tabela uključuje sva velika i mala slova. Pored toga:

- Nafasi (`" "`) zina upana wa 2
- Nukta (`"."`) zina upana wa 1

- Ako dati niz ima širinu od 50 jedinica ili manje, vrati ga nepromenjenog, u suprotnom
- Skrati niz i dodaj tri tačke na kraju (`"..."`) tako da ukupna širina, uključujući tri tačke, bude što je moguće bliža 50 jedinica bez prekoračenja.

# --hints--

`truncateText("The quick brown fox")` inapaswa kurudisha `"The quick brown f..."`.

```js
assert.equal(truncateText("The quick brown fox"), "The quick brown f...");
```

`truncateText("The silky smooth sloth")` inapaswa kurudisha `"The silky smooth s..."`.

```js
assert.equal(truncateText("The silky smooth sloth"), "The silky smooth s...");
```

`truncateText("THE LOUD BRIGHT BIRD")` inapaswa kurudisha `"THE LOUD BRIG..."`.

```js
assert.equal(truncateText("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...");
```

`truncateText("The fast striped zebra")` inapaswa kurudisha `"The fast striped z..."`.

```js
assert.equal(truncateText("The fast striped zebra"), "The fast striped z...");
```

`truncateText("The big black bear")` inapaswa kurudisha `"The big black bear"`.

```js
assert.equal(truncateText("The big black bear"), "The big black bear");
```

# --seed--

## --seed-contents--

```js
function truncateText(str) {

  return str;
}
```

# --solutions--

```js
function truncateText(str) {
  const MAX_WIDTH = 50;
  const ELLIPSIS = "...";
  const ELLIPSIS_WIDTH = 3;
  const TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH;

  const charGroups = {
    'ilI.': 1,
    'fjrt ': 2,
    'abcdeghkmnopqrstuvwxyzJL': 3,
    'ABCDEFGHKMNOPQRSTUVWXYZ': 4
  };

  function getCharWidth(char) {
    for (const key in charGroups) {
      if (key.includes(char)) return charGroups[key];
    }
    return 3;
  }

  function stringWidth(str) {
    let totalWidth = 0;
    for (const char of str) {
      totalWidth += getCharWidth(char);
    }
    return totalWidth;
  }

  if (stringWidth(str) <= MAX_WIDTH) return str;

  let result = "";
  let totalWidth = 0;

  for (const char of str) {
    const charWidth = getCharWidth(char);

    if (totalWidth + charWidth > TRUNCATE_LIMIT) break;

    result += char;
    totalWidth += charWidth;
  }

  return result + ELLIPSIS;
}
```
