---
id: 69a890af247de743333bd4ce
title: "Vežba 225: Nema uzastopnog ponavljanja niza znakova"
challengeType: 28
dashedName: challenge-225
---

# --description--

Ako je dat niz znakova, odredi da li nema ponovljenih znakova.

- Niz znakova nema ponavljanje ako se isti znak ne pojavljuje dva ili više puta uzastopno.

# --hints--

`hasNoRepeats("hi world")` inapaswa kurudisha `true`.

```js
assert.isTrue(hasNoRepeats("hi world"));
```

`hasNoRepeats("hello world")` inapaswa kurudisha `false`.

```js
assert.isFalse(hasNoRepeats("hello world"));
```

`hasNoRepeats("abcdefghijklmnopqrstuvwxyz")` inapaswa kurudisha `true`.

```js
assert.isTrue(hasNoRepeats("abcdefghijklmnopqrstuvwxyz"));
```

`hasNoRepeats("freeCodeCamp")` inapaswa kurudisha `false`.

```js
assert.isFalse(hasNoRepeats("freeCodeCamp"));
```

`hasNoRepeats("The quick brown fox jumped over the lazy dog.")` inapaswa kurudisha `true`.

```js
assert.isTrue(hasNoRepeats("The quick brown fox jumped over the lazy dog."));
```

`hasNoRepeats("Mississippi")` inapaswa kurudisha `false`.

```js
assert.isFalse(hasNoRepeats("Mississippi"));
```

# --seed--

## --seed-contents--

```js
function hasNoRepeats(str) {

  return str;
}
```

# --solutions--

```js
function hasNoRepeats(str) {
  for (let i = 1; i < str.length; i++) {
    if (str[i] === str[i - 1]) {
      return false;
    }
  }
  return true;
}
```
