---
id: 69a890af247de743333bd4d1
title: "Zoezi la 228: Usiku wa Filamu"
challengeType: 28
dashedName: challenge-228
---

# --description--
Ako dobijete niz slova koji ukazuje dan u nedelji, drugi niz slova koji ukazuje vreme projekcije filma i kompletan broj ulaznice, vratite ukupnu cenu film ulaznica za tu projekciju.

Siku iliyotolewa itakuwa moja kati ya:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

Saa za prikaz će biti dostupne u formatu `"H:MMam"` ili `"H:MMpm"`. Na primer `"10:00am"` ili `"10:00pm"`.

Vratite ukupnu cenu za strukturu `"$D.CC"` koristeći ove pravila:

- Vikend (Petak - Nedelja): $12.00 po ulazu.
- Radni dani (Ponedeljak - Četvrtak): $10.00 po ulazu.
- Popodnevni nastup (pre 23 časova): popust od $2.00 po ulazu (osim utorka).
- Sreda: Sve karte su $5.00 svaka.

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` inapaswa kurudisha `"$12.00"`.

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` inapaswa kurudisha `"$10.00"`.

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` inapaswa kurudisha `"$10.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` inapaswa kurudisha `"$30.00"`.

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` inapaswa kurudisha `"$32.00"`.

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` inapaswa kurudisha `"$50.00"`.

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` inapaswa kurudisha `"$5.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
