---
id: 69a890af247de743333bd4cd
title: "Zoezi la 224: Kigeuzi cha Kuangalia Kawa ya Kuoka"
challengeType: 28
dashedName: challenge-224
---

# --description--
Generisanje sekvence slova koja predstavlja zrna korišćena za pripremu šolje kafe, određuje nivo prženja te šolje.

- Data niska sadržaće sledeće znakove, od kojih svaki predstavlja određenu vrstu zrna kafe:

- Znak apostrofa (`'`) je karakter sagorevanja za svetlosni intenzitet sa vrednošću od 1 poen svaki.
- Znak razmaka (`-`) je karakter sagorevanja srednje snage sa vrednošću od 2 poena svaki.
- Znak tačke (`.`) je karakter sagorevanja u mraku sa vrednošću od 3 poena svaki.

- Nivo prženja određuje se prosečnom vrednošću svih zrna.

Rudisha:

- `"Light"` ako je prosek manji od 1.75.
- `"Medium"` ako je prosek između 1.75 i 2.5.
- `"Dark"` ako je prosek veći od 2.5.

# --hints--

`detectRoast("''-''''''-'-''--''''")` inapaswa kurudisha `"Light"`.

```js
assert.equal(detectRoast("''-''''''-'-''--''''"), "Light");
```

`detectRoast(".'-''-''..'''.-.-''-")` inapaswa kurudisha `"Medium"`.

```js
assert.equal(detectRoast(".'-''-''..'''.-.-''-"), "Medium");
```

`detectRoast("--.''--'-''.--..-.--")` inapaswa kurudisha `"Medium"`.

```js
assert.equal(detectRoast("--.''--'-''.--..-.--"), "Medium");
```

`detectRoast("-...'-......-..-...-")` inapaswa kurudisha `"Dark"`.

```js
assert.equal(detectRoast("-...'-......-..-...-"), "Dark");
```

`detectRoast(".--.-..-......----.'")` inapaswa kurudisha `"Medium"`.

```js
assert.equal(detectRoast(".--.-..-......----.'"), "Medium");
```

`detectRoast("..-..-..-..-....-.-.")` inapaswa kurudisha `"Dark"`.

```js
assert.equal(detectRoast("..-..-..-..-....-.-."), "Dark");
```

`detectRoast("-'-''''''..-'.''-'.'")` inapaswa kurudisha `"Light"`.

```js
assert.equal(detectRoast("-'-''''''..-'.''-'.'"), "Light");
```

# --seed--

## --seed-contents--

```js
function detectRoast(beans) {

  return beans;
}
```

# --solutions--

```js
function detectRoast(beans) {
  let total = 0;

  for (const bean of beans) {
    if (bean === "'") total += 1;
    else if (bean === "-") total += 2;
    else if (bean === ".") total += 3;
  }

  const avg = total / beans.length;

  if (avg < 1.75) return "Light";
  if (avg <= 2.5) return "Medium";
  return "Dark";
}
```
