---
id: 69b58ce40693f140c84c855b
title: "Zoezi la 242: Nambari inayofuata ya Bingo"
challengeType: 28
dashedName: challenge-242
---

# --description--
Ako dobijete broj za bingo, vratite sledeći broj za bingo u redosledu.

Bingo broj je jedno slovo, praćeno brojem u njegovom redu prema ovoj tabeli:

| Karakteri | Opseg brojeva |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

Na primer, ako dobiješ `"B10"`, vrati `"B11"`, sledeći broj za bingo. Ako dobiješ poslednji broj za bingo, vrati `"B1"`.

# --hints--

`getNextBingoNumber("B10")` inapaswa kurudisha `"B11"`.

```js
assert.equal(getNextBingoNumber("B10"), "B11");
```

`getNextBingoNumber("N33")` inapaswa kurudisha `"N34"`.

```js
assert.equal(getNextBingoNumber("N33"), "N34");
```

`getNextBingoNumber("I30")` inapaswa kurudisha `"N31"`.

```js
assert.equal(getNextBingoNumber("I30"), "N31");
```

`getNextBingoNumber("G60")` inapaswa kurudisha `"O61"`.

```js
assert.equal(getNextBingoNumber("G60"), "O61");
```

`getNextBingoNumber("O75")` inapaswa kurudisha `"B1"`.

```js
assert.equal(getNextBingoNumber("O75"), "B1");
```

# --seed--

## --seed-contents--

```js
function getNextBingoNumber(n) {

  return n;
}
```

# --solutions--

```js
function getNextBingoNumber(n) {
  const num = parseInt(n.slice(1));
  const next = num === 75 ? 1 : num + 1;

  if (next <= 15) return "B" + next;
  if (next <= 30) return "I" + next;
  if (next <= 45) return "N" + next;
  if (next <= 60) return "G" + next;
  return "O" + next;
}
```
