---
id: 69b559d2903b9e4afe9075f8
title: "Vežba 238: Izbegavanje petlje brojeva
challengeType: 28
dashedName: challenge-238
---

# --description--
Ako je dat pozitivan ceo broj, proveri da li taj broj, ili bilo koji ciklus tog broja, može biti savršeno podeljen sa brojem njegovih cifara.

Rotacija znači pomjeranje prve cifre na kraj. Na primer, nakon jedne rotacije, 123 postaje 231.

- Proverite broj kruga `0` (izdati broj) prvo.
- Izdati brojevi neće sadržavati nijedan nulti.
- Vratite prvi broj kruga ako postoji, ili `"none"` ako ne postoji.

# --hints--

`getRotation(123)` inapaswa kurudisha `0`.

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)` inapaswa kurudisha `3`.

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)` inapaswa kurudisha `"none"`.

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)` inapaswa kurudisha `6`.

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```
