---
id: 69b58ce40693f140c84c8559
title: "Zoezi la 240: Herufi za Palindrome"
challengeType: 28
dashedName: challenge-240
---

# --description--

Za zadatu nisku odredi da li je palindrom i vrati srednji znak (ako je dužina neparna) ili dva srednja znaka (ako je dužina parna).

- Palindrom je niska koja se isto čita unapred i unazad.
- Ikiwa si palindrome, rudisha `"none"`.

# --hints--

`palindromeLocator("racecar")` inapaswa kurudisha `"e"`.

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")` inapaswa kurudisha `"v"`.

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")` inapaswa kurudisha `"none"`.

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")` inapaswa kurudisha `"oo"`.

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")` inapaswa kurudisha `"00"`.

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```
