---
id: 69b58ce40693f140c84c855a
title: "Zoezi la 241: Kagua mfululizo wa FizzBuzz"
challengeType: 28
dashedName: challenge-241
---

# --description--
Dati je niz uzastopnih celih brojeva, gde su multiplikovi sa 3 i 5 transformisani; potvrdite da li je to validna FizzBuzz sekvenca.

Katika mfululizo halali wa FizzBuzz:

- Maradufu ya 3 yamebadilishwa na "Fizz".
- Maradufu ya 5 yamebadilishwa na "Buzz".
- Maradufu ya 3 na 5 yamebadilishwa na "FizzBuzz".
- Svi ostali brojevi ostaju celi brojevi.

# --hints--

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz"])` inapaswa kurudisha `true`.

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4, "Buzz"]));
```

`isFizzBuzz([13, 14, "FizzBuzz", 16, 17])` inapaswa kurudisha `true`.

```js
assert.isTrue(isFizzBuzz([13, 14, "FizzBuzz", 16, 17]));
```

`isFizzBuzz([1, 2, "Fizz", 4, 5])` inapaswa kurudisha `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, 5]));
```

`isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])` inapaswa kurudisha `true`.

```js
assert.isTrue(isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]));
```

`isFizzBuzz([1, 2, "Fizz", "Buzz", 5])` inapaswa kurudisha `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", "Buzz", 5]));
```

`isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103])` inapaswa kurudisha `false`.

```js
assert.isFalse(isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103]));
```

`isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"])` inapaswa kurudisha `true`.

```js
assert.isTrue(isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"]));
```

# --seed--

## --seed-contents--

```js
function isFizzBuzz(arr) {

  return arr;
}
```

# --solutions--

```js
function isFizzBuzz(arr) {
  let start = null;
  for (let i = 0; i < arr.length; i++) {
    if (typeof arr[i] === "number") {
      start = arr[i] - i;
      break;
    }
  }

  if (start === null) return false;

  for (let i = 0; i < arr.length; i++) {
    const n = start + i;
    const expected = n % 15 === 0 ? "FizzBuzz" : n % 3 === 0 ? "Fizz" : n % 5 === 0 ? "Buzz" : n;
    if (arr[i] !== expected) return false;
  }

  return true;
}
```
