---
id: 69bc6cb30c1d112a2e110a07
title: "Zoezi la 250: Ufunguo uliofichwa"
challengeType: 28
dashedName: challenge-250
---

# --description--
Dobrodošli na dnevnu vežbu od 250!

Ako imate kodirani niz slova, dekodirajte ga koristeći ključ za šifrovanje i vratite rezultat.

Ili kupata ufunguo:

- Pogledajte sve dnevne izazove do danas čiji je broj izazova deo 25 (uključujući i ovaj).
- Uzmite prvo slovo iz naslova svih tih izazova i spojite ih u niz slova. Ako naslov počinje sa karakterom koji nije slovo, pronađite njegovo prvo slovo.

Da bi dešifrovao poruku, prođi kroz svaki znak šifrovane poruke i:

1. Pronađite odgovarajuće slovo u ključu (ponovite ključ ako je poruka duža od ključa).
  2. Pretvorite slovo ključa u odgovarajući broj: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
  3. Pomjerite šifrovano slovo unazad u alfabet za taj broj.
  4. Ako pomjeranje dođe pre `"A"`, zaokružite nazad na `"Z"`.

Na primer, ako kodirani poruka počinje sa `"Y"` i prvi slovo ključa je `"V"` (22), gurnite `"Y"` nazad na poziciju 22 da biste dobili `"C"`. Ponovite ovaj proces za svako slovo kako biste dešifrovali celu poruku.

- Samo slova se pomeraju, razmaci ostaju nepromenjeni.
- Sva ulazna i izlazna slova su velika slova.

# --hints--

`decode("YAVJYNXE")` inapaswa kurudisha `"CONGRATS"`.

```js
assert.equal(decode("YAVJYNXE"), "CONGRATS");
```

`decode("YALLUT PQUMJP")` inapaswa kurudisha `"CODING LEGEND"`.

```js
assert.equal(decode("YALLUT PQUMJP"), "CODING LEGEND");
```

`decode("UAC DYR EISAKYM")` inapaswa kurudisha `"YOU ARE AWESOME"`.

```js
assert.equal(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME");
```

`decode("GQMS NBMZU")` inapaswa kurudisha `"KEEP GOING"`.

```js
assert.equal(decode("GQMS NBMZU"), "KEEP GOING");
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` inapaswa kurudisha `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`.

```js
assert.equal(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP");
```

# --seed--

## --seed-contents--

```js
function decode(message) {

  return message;
}
```

# --solutions--

```js
function decode(message) {
  const key = "VLHCGMDLNH";
  let keyIndex = 0;
  return message.split('').map(char => {
    if (char === ' ') return char;
    const shift = key[keyIndex % key.length].charCodeAt(0) - 64;
    keyIndex++;
    return String.fromCharCode(((char.charCodeAt(0) - 65 - shift + 26) % 26) + 65);
  }).join('');
}
```
