---
id: 69bc6cb30c1d112a2e110a03
title: "Vežba 246: Slova imena
challengeType: 28
dashedName: challenge-246
---

# --description--
Ako unesete puno ime kao niz slova, vratite mu početna slova.

- Imena za inicijalizaciju se razdvajaju razmakom.
- Početna slova bi trebalo da budu u velikim slovima.
- Početna slova bi trebalo da se razdvajaju tačkom.

Na primer, `"Tommy Millwood"` vraća `"T.M."`.

# --hints--

`getInitials("Tommy Millwood")` inapaswa kurudisha `"T.M."`.

```js
assert.equal(getInitials("Tommy Millwood"), "T.M.");
```

`getInitials("Savanna Puddlesplash")` inapaswa kurudisha `"S.P."`.

```js
assert.equal(getInitials("Savanna Puddlesplash"), "S.P.");
```

`getInitials("Frances Cowell Conrad")` inapaswa kurudisha `"F.C.C."`.

```js
assert.equal(getInitials("Frances Cowell Conrad"), "F.C.C.");
```

`getInitials("Dragon")` inapaswa kurudisha `"D."`.

```js
assert.equal(getInitials("Dragon"), "D.");
```

`getInitials("Dorothy Vera Clump Haverstock Norris")` inapaswa kurudisha `"D.V.C.H.N."`.

```js
assert.equal(getInitials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.");
```

# --seed--

## --seed-contents--

```js
function getInitials(name) {

  return name;
}
```

# --solutions--

```js
function getInitials(name) {
  return name.split(' ').map(word => word[0].toUpperCase() + '.').join('');
}
```
