---
id: 69bc6cb30c1d112a2e110a06
title: "Vežba 249: Matematika niza znakova"
challengeType: 28
dashedName: challenge-249
---

# --description--
Ako vam je dat niz slova koji sadrži brojeve i druga slova, izračunajte broj prema broju ne-numeričkih slova između tih brojeva.

- Ako je broj znakova koji razdvajaju dva broja parno, koristi sabiranje.
- Ako je neparno, koristi oduzimanje.
- Uzastopni brojevi čine jedan broj.
- Operacije se izvode od leve ka desnoj.
- Ignoriši početne i krajnje znakove koji nisu brojevi.

Na primer, ako dobiješ `"3ab10c8"`, vrati `5`. Dodaj 3 i 10 da bi dobio/dobila 13 jer postoji broj slova čak ni između njih. Zatim oduzmi 8 od 13 jer postoji broj slova koji nije čak ni između rezultata i 8.

# --hints--

`doMath("3ab10c8")` inapaswa kurudisha `5`.

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")` inapaswa kurudisha `2`.

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")` inapaswa kurudisha `12`.

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")` inapaswa kurudisha `15`.

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` inapaswa kurudisha `67`.

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
