---
id: 69bc6cb30c1d112a2e110a04
title: "Vežba 247: Poslednje slovo
challengeType: 28
dashedName: challenge-247
---

# --description--
Kada je dat niz slova, vrati slovo koje je najposlednije u alfabetu.

- Ako su dva ili više slova duplikata na kraju alfabeta, vrati prvo.
- Ignoriši sve nealfabetičke znakove.

# --hints--

`getLastLetter("world")` inapaswa kurudisha `"w"`.

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` inapaswa kurudisha `"W"`.

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` inapaswa kurudisha `"z"`.

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` inapaswa kurudisha `"L"`.

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` inapaswa kurudisha `"T"`.

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
