---
id: 69c5f3d787b1725d5f00c8bd
title: "Zoezi la 259: Mlipuko wa FizzBuzz"
challengeType: 28
dashedName: challenge-259
---

# --description--
Ako date kompletan broj, vratite broj koraka potrebnih da promenite reč `"fizzbuzz"` u niz slova koji ima najmanje `"z"`, kao što je navedeno korišćenjem sledećih pravila:

- Počni sa nizom karaktera `"fizzbuzz"`.
- Za svaki korak, koristi standardna FizzBuzz pravila koristeći poziciju slova u nizu (prvo slovo `"f"` je na poziciji 1).
  - Ako je pozicija slova deljiva sa 3, zameni to slovo sa `"fizz"`
  - Ako je deljivo sa 5, zameni to slovo sa `"buzz"`
  - Ako je deljivo sa 3 i 5, zameni to slovo sa `"fizzbuzz"`

Nakon koraka 1, `"fizzbuzz"` se menja u `"fifizzzbuzzfizzzz"`, što sadrži devet znakova `"z"`.

# --hints--

`explodeFizzbuzz(9)` inapaswa kurudisha `1`.

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)` inapaswa kurudisha `2`.

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)` inapaswa kurudisha `3`.

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)` inapaswa kurudisha `4`.

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)` inapaswa kurudisha `5`.

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)` inapaswa kurudisha `6`.

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)` inapaswa kurudisha `11`.

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)` inapaswa kurudisha `14`.

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```
