---
id: 69c5f3d787b1725d5f00c8b9
title: "Zoezi la 255: Kikosi cha Usafishaji wa Siku ya Dunia"
challengeType: 28
dashedName: challenge-255
---

# --description--
Danas je Svetski dan. Ako dobiješ skup podataka stvari koje si očistio/očistile, vrati ukupan zbroj svojih ocena čišćenja prema pravilima navedenim ispod.

Predmeti koji su izdati biće jedan od sljedećih:

| Polje | Podrazumevana vrednost |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

Predmet retke vrste je predstavljen sa `["rare", value]`. Na primer, `["rare", 80]`. Retki predmeti ne dobijaju bonus za seriju.

- Sekvencijalni bonus: Ako se neki predmet pojavi uzastopno, dobija rastuće bonus poene.
  - Prvi uzastopni: osnovna vrednost
  - Drugi: osnovna vrednost + 1
  - Treći: osnovna vrednost + 2
  - itd.
- Multiplikator petog predmeta: Svaki peti prikupljeni predmet dobija multiplikator.
  - Peti predmet: *2
  - Deseti predmet: *3
  - itd.

- Tumia kizidishi baada ya kuhesabu bonasi zozote.

# --hints--

`getCleanupScore(["bottle", "straw", "shoe", "battery"])` inapaswa kurudisha `44`.

```js
assert.equal(getCleanupScore(["bottle", "straw", "shoe", "battery"]), 44);
```

`getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"])` inapaswa kurudisha `58`.

```js
assert.equal(getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"]), 58);
```

`getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` inapaswa kurudisha `79`.

```js
assert.equal(getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79);
```

`getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` inapaswa kurudisha `358`.

```js
assert.equal(getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358);
```

`getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` inapaswa kurudisha `383`.

```js
assert.equal(getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383);
```

# --seed--

## --seed-contents--

```js
function getCleanupScore(items) {

  return items;
}
```

# --solutions--

```js
function getCleanupScore(items) {
  const values = {
    bottle: 10,
    can: 6,
    bag: 8,
    tire: 35,
    straw: 4,
    cardboard: 3,
    newspaper: 3,
    shoe: 12,
    electronics: 25,
    battery: 18,
    mattress: 38
  };

  let total = 0;
  let prevItem = null;
  let streak = 0;

  for (let i = 0; i < items.length; i++) {
    let item = items[i];
    let base;
    let isRare = Array.isArray(item) && item[0] === "rare";

    if (isRare) {
      base = item[1];
      prevItem = null;
      streak = 0;
    } else {
      base = values[item];
      streak = item === prevItem ? streak + 1 : 0;
      prevItem = item;
    }

    let points = base;
    if (!isRare) points += streak;

    if ((i + 1) % 5 === 0) {
      const multiplier = (i + 1) / 5 + 1;
      points *= multiplier;
    }

    total += points;
  }

  return total;
}
```
