---
id: 69c5f3d787b1725d5f00c8ba
title: "Vežba 256: Najbliži smer vremena"
challengeType: 28
dashedName: challenge-256
---

# --description--
Uzevši u obzir dva trenutka/punkta, odluči da li možešstići od prvog do drugog brže krećući se napred ili nazad.

- Vremenski podaci su prikazani u formatu od 24 sata (`"HH:MM"`)
- Sat se kreće (23:59 prelazi u 00:00 kada napredujete, a 00:00 prelazi u 23:59 kada se vraćate unazad)

Rudisha:

- `"forward"` ako je kretanje unapred kraće.
- `"backward"` ako je kretanje unazad kraće.
- `"equal"` ako oba smera zahtevaju isto vreme.

# --hints--

`getDirection("10:00", "12:00")` inapaswa kurudisha `"forward"`.

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")` inapaswa kurudisha `"backward"`.

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")` inapaswa kurudisha `"equal"`.

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")` inapaswa kurudisha `"forward"`.

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")` inapaswa kurudisha `"backward"`.

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")` inapaswa kurudisha `"equal"`.

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```
