---
id: 69c5f3d787b1725d5f00c8b8
title: "Vežba 254: Reči bez dva broja"
challengeType: 28
dashedName: challenge-254
---

# --description--

Ako je dat niz znakova reči, vrati reči koje imaju broj slova koji nije deljiv sa dva.

- Reči u datom nizu biće razdvojene samo jednim razmakom.
- Rudisha maneno yakiwa yameachwa na nafasi moja tu.

# --hints--

`getOddWords("This is a super good test")` inapaswa kurudisha `"a super"`.

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")` inapaswa kurudisha `"one two three"`.

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")` inapaswa kurudisha `"split rainbow sprinkles top"`.

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")` inapaswa kurudisha `"The quick brown fox the river"`.

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```
