---
id: 69c6ff713a52713463aa7929
title: "Izazov 263: Binarni rečnik"
challengeType: 28
dashedName: challenge-263
---

# --description--
Ukoliko vam se da slovo, odredite da li njegovo binarno predstavljanje od 8 bita može biti pronađeno u sledećoj mreži, horizontalno ili vertikalno u bilo kom pravcu:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

Na primer, `"A"` ima binarno predstavljanje `01000001`, vidljivo u prvom redu od leve ka desnoj.

# --hints--

`isInCrossword("I")` inapaswa kurudisha `true`.

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` inapaswa kurudisha `true`.

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` inapaswa kurudisha `true`.

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` inapaswa kurudisha `true`.

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` inapaswa kurudisha `false`.

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` inapaswa kurudisha `false`.

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` inapaswa kurudisha `false`.

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` inapaswa kurudisha `false`.

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
