---
id: 69c6ff713a52713463aa7928
title: "Zoezi la 262: Mchambuzi wa hoja ya utafutaji ya URL"
challengeType: 28
dashedName: challenge-262
---

# --description--
Ako URL ima niz karaktera za pretraga argumenta, parsiraj taj niz karaktera u objekat (ili rečnik) parova ključ-vrednost.

- Sekvenca znakova upita za pretragu počinje nakon `"?"`,
- Svaki parametar je razdvojen sa `"&"`,
- Svaki par ključ/vrednost je razdvojen sa `"="`

Na primer, ako se ukloni `"https://example.com/search?name=Alice&age=30"`, vratite:

```json
{
  "name": "Alice",
  "age": "30"
}
```

Sve vrednosti moraju biti vraćene kao niz znakova.

# --hints--

`parseUrlQuery("https://example.com/search?name=Alice&age=30")` treba da vrati `{"name": "Alice", "age": "30"}`.

```js
assert.deepEqual(parseUrlQuery("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"});
```

`parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python")` inapaswa kurudisha `{"skill": "programming", "language": "python"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"});
```

`parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2")` inapaswa kurudisha `{"category": "books", "sort": "asc", "page": "2"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"});
```

`parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now")` inapaswa kurudisha `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"});
```

# --seed--

## --seed-contents--

```js
function parseUrlQuery(url) {

  return url;
}
```

# --solutions--

```js
function parseUrlQuery(url) {
  const queryString = url.split('?')[1];
  return Object.fromEntries(
    queryString.split('&').map(param => param.split('='))
  );
}
```
