---
id: 69c6ff713a52713463aa7927
title: "Vežba 261: Brojčane reči"
challengeType: 28
dashedName: challenge-261
---

# --description--
Ako dobiješ kompletan broj od 0 do 99, vrati mu izgovor u engleskom jeziku.

- `0` vraća `"zero"`.
- Brojevi 1-19 imaju svoja jedinstvena imena (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- Množbenici desetice od 20 do 90 imaju svoja imena (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- Brojevi od 21 do 99 koji nisu množbenici desetice pišu se kao dve reči spojene kvačićem. Primeri su `"forty-two"` i `"fifty-three"`.

# --hints--

`getNumberWords(0)` inapaswa kurudisha `"zero"`.

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` inapaswa kurudisha `"ten"`.

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` inapaswa kurudisha `"nineteen"`.

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` inapaswa kurudisha `"thirty"`.

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` inapaswa kurudisha `"fifty-three"`.

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` inapaswa kurudisha `"seven"`.

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` inapaswa kurudisha `"twelve"`.

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` inapaswa kurudisha `"sixty"`.

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` inapaswa kurudisha `"sixty-seven"`.

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` inapaswa kurudisha `"ninety-eight"`.

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
