---
id: 69cfca90e8a0a6d4d6871c50
title: "Zoezi la 266: Habari za asubuhi"
challengeType: 28
dashedName: challenge-266
---

# --description--

Ako je data vremenska oznaka u formatu `"HH:MM"` (24-časovni format), vrati:

- `"Good morning"` u periodima `05:00` do `11:59`
- `"Good afternoon"` u periodima `12:00` do `17:59`
- `"Good evening"` u periodima `18:00` do `21:59`
- `"Good night"` u periodima `22:00` do `04:59`

# --hints--

`getGreeting("06:30")` inapaswa kurudisha `"Good morning"`.

```js
assert.equal(getGreeting("06:30"), "Good morning");
```

`getGreeting("12:00")` inapaswa kurudisha `"Good afternoon"`.

```js
assert.equal(getGreeting("12:00"), "Good afternoon");
```

`getGreeting("21:59")` inapaswa kurudisha `"Good evening"`.

```js
assert.equal(getGreeting("21:59"), "Good evening");
```

`getGreeting("00:01")` inapaswa kurudisha `"Good night"`.

```js
assert.equal(getGreeting("00:01"), "Good night");
```

`getGreeting("11:30")` inapaswa kurudisha `"Good morning"`.

```js
assert.equal(getGreeting("11:30"), "Good morning");
```

# --seed--

## --seed-contents--

```js
function getGreeting(time) {

  return time;
}
```

# --solutions--

```js
function getGreeting(time) {
  const [hours, minutes] = time.split(':').map(Number);
  const total = hours * 60 + minutes;

  if (total >= 300 && total < 720) return 'Good morning';
  if (total >= 720 && total < 1080) return 'Good afternoon';
  if (total >= 1080 && total < 1320) return 'Good evening';
  return 'Good night';
}
```
