---
id: 69cfca90e8a0a6d4d6871c54
title: "Vežba 270: Najduži ponavljajući niz znakova"
challengeType: 28
dashedName: challenge-270
---

# --description--

Ako je dat niz znakova, vrati najduži podniz koji se ponavlja više od jednom.

- Podnizovi mogu da se preklapaju.

# --hints--

`getLongestSubstring("abracadabra")` inapaswa kurudisha `"abra"`.

```js
assert.equal(getLongestSubstring("abracadabra"), "abra");
```

`getLongestSubstring("hello world hello")` inapaswa kurudisha `"hello"`.

```js
assert.equal(getLongestSubstring("hello world hello"), "hello");
```

`getLongestSubstring("mississippi")` inapaswa kurudisha `"issi"`.

```js
assert.equal(getLongestSubstring("mississippi"), "issi");
```

`getLongestSubstring("ha ha ha ha ha ha ha")` inapaswa kurudisha `"ha ha ha ha ha ha"`.

```js
assert.equal(getLongestSubstring("ha ha ha ha ha ha ha"), "ha ha ha ha ha ha");
```

`getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over")` inapaswa kurudisha `"the quick brown fox jumped over"`.

```js
assert.equal(getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over"), "the quick brown fox jumped over");
```

# --seed--

## --seed-contents--

```js
function getLongestSubstring(str) {

  return str;
}
```

# --solutions--

```js
function getLongestSubstring(str) {
  let longest = '';

  for (let len = str.length - 1; len >= 1; len--) {
    for (let i = 0; i <= str.length - len; i++) {
      const sub = str.slice(i, i + len);
      if (str.indexOf(sub) !== str.lastIndexOf(sub)) {
        return sub;
      }
    }
  }

  return longest;
}
```
