---
id: 69cfca90e8a0a6d4d6871c51
title: "Zoezi la 267: Mbadalishaji wa Parsec"
challengeType: 28
dashedName: challenge-267
---

# --description--
U dalekoj galaksiji, parsecima se koriste za merenje vremena i udaljenosti. Ako dobijete kompletan broj parsecima, vratite mu ekvivalentnu vrednost u vreme ili udaljenost.

- Ako je ceo broj neparan, predstavlja vreme. Ako je paran, predstavlja udaljenost.

Koristi sledeće stope konverzije:

| Parsecs | Wakati/Umbali |
| - | - |
| 1 | masaa 2 |
| 2 | miaka 6 ya mwanga |

Vratite izmenjenu vrednost kao ceo broj.

# --hints--

`convertParsecs(1)` inapaswa kurudisha `2`.

```js
assert.equal(convertParsecs(1), 2);
```

`convertParsecs(2)` inapaswa kurudisha `6`.

```js
assert.equal(convertParsecs(2), 6);
```

`convertParsecs(31)` inapaswa kurudisha `62`.

```js
assert.equal(convertParsecs(31), 62);
```

`convertParsecs(88)` inapaswa kurudisha `264`.

```js
assert.equal(convertParsecs(88), 264);
```

`convertParsecs(17)` inapaswa kurudisha `34`.

```js
assert.equal(convertParsecs(17), 34);
```

`convertParsecs(14)` inapaswa kurudisha `42`.

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```
