---
id: 69e2383af7832c8032603b92
title: "Vežba 276: Element koji krši
challengeType: 28
dashedName: challenge-276
---

# --description--
Dostavljen je red podataka kompletnih brojeva sortirano u uzlaznom redosledu, osim jedne komponente koja je na netačnoj poziciji; vratite model te komponente.

- Ako postoji više od jednog elementa koji se može smatriti neispravnom lokacijom, vratite model prvog elementa.

# --hints--

`findOffender([1, 6, 2, 3, 4, 5])` inapaswa kurudisha `1`.

```js
assert.equal(findOffender([1, 6, 2, 3, 4, 5]), 1);
```

`findOffender([1, 2, 3, 5, 4, 5])` inapaswa kurudisha `3`.

```js
assert.equal(findOffender([1, 2, 3, 5, 4, 5]), 3);
```

`findOffender([2, 1])` inapaswa kurudisha `0`.

```js
assert.equal(findOffender([2, 1]), 0);
```

`findOffender([2, 4, 1, 6, 8])` inapaswa kurudisha `2`.

```js
assert.equal(findOffender([2, 4, 1, 6, 8]), 2);
```

`findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])` inapaswa kurudisha `6`.

```js
assert.equal(findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6);
```

# --seed--

## --seed-contents--

```js
function findOffender(arr) {

  return arr;
}
```

# --solutions--

```js
function findOffender(arr) {
  for (let i = 0; i < arr.length; i++) {
    const without = [...arr.slice(0, i), ...arr.slice(i + 1)];
    if (without.every((n, j) => j === 0 || without[j - 1] <= n)) return i;
  }
}
```
