---
id: 69f35a5bb823ed620fcb7cbc
title: "Vežba 283: Preplitanje niski"
challengeType: 28
dashedName: challenge-283
---

# --description--
Ako dobijete dve sekvence slova, vratite novi niz koji spaja njihova slova direktno po redosledu. Ukoliko je jedna sekvenca duža, dodajte preostala slova na kraj.

Počni sa prvim slovom prve sekvence.

# --hints--

`zipStrings("abc", "123")` inapaswa kurudisha `"a1b2c3"`.

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")` inapaswa kurudisha `"abcdefghijklmnopqrstuvwxyz"`.

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")` inapaswa kurudisha `"dnaiyght"`.

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")` treba da vrati `"pjyatvhaosncript"`.

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")` inapaswa kurudisha `"freeCodeCamp"`.

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```
