---
id: 69f8c998d78ad3171a0713bd
title: "Zoezi la 290: Sherehe ya pizza"
challengeType: 28
dashedName: challenge-290
---

# --description--
Dostavite red podataka o radnim satima za danas za sve i vratite broj pizzy koje treba naručiti za pizzu proslavu.

- Podelite ukupan broj sati koji je svako radio sa 3 da biste dobili broj komada pice koje bi trebalo da ima.
- Ne možete jesti mali komad, pa zaokružite broj komada po osobi do sledećeg celog broja.
- Svako dobija najmanje dva komada.
- Svaka pice ima 8 komada. Zaokružite ukupan broj pica do sledeće cele pice.

# --hints--

`getPizzasToOrder([8, 8, 8])` inapaswa kurudisha `2`.

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` inapaswa kurudisha `3`.

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` inapaswa kurudisha `2`.

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` inapaswa kurudisha `3`.

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` inapaswa kurudisha `1`.

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` inapaswa kurudisha `5`.

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```
