---
id: 69f8c998d78ad3171a0713be
title: "Zoezi la 291: Hesabu ya FizzBuzz"
challengeType: 28
dashedName: challenge-291
---

# --description--
Za zadati početni i krajnji broj, izračunaj koliko puta se pojavljuju 'fizz' i 'buzz' u tom opsegu (uključujući te brojeve).

- Brojevi deljivi sa 3 broje se kao fizz.
- Brojevi deljivi sa 5 broje se kao buzz.
- Brojevi deljivi sa 3 i 5 broje se kao i fizz i buzz (oba).

Vrati objekat ili rečnik sa brojevima u formatu: `{ fizz, buzz }`.

# --hints--

`fizzBuzzCount(1, 11)` inapaswa kurudisha `{fizz: 3, buzz: 2}`.

```js
assert.deepEqual(fizzBuzzCount(1, 11), {fizz: 3, buzz: 2});
```

`fizzBuzzCount(14, 41)` inapaswa kurudisha `{fizz: 9, buzz: 6}`.

```js
assert.deepEqual(fizzBuzzCount(14, 41), {fizz: 9, buzz: 6});
```

`fizzBuzzCount(24, 100)` inapaswa kurudisha `{fizz: 26, buzz: 16}`.

```js
assert.deepEqual(fizzBuzzCount(24, 100), {fizz: 26, buzz: 16});
```

`fizzBuzzCount(-635, -14)` inapaswa kurudisha `{fizz: 207, buzz: 125}`.

```js
assert.deepEqual(fizzBuzzCount(-635, -14), {fizz: 207, buzz: 125});
```

`fizzBuzzCount(-5432, 6789)` inapaswa kurudisha `{fizz: 4074, buzz: 2444}`.

```js
assert.deepEqual(fizzBuzzCount(-5432, 6789), {fizz: 4074, buzz: 2444});
```

# --seed--

## --seed-contents--

```js
function fizzBuzzCount(start, end) {

  return start;
}
```

# --solutions--

```js
function fizzBuzzCount(start, end) {
  let fizz = 0, buzz = 0;

  for (let i = start; i <= end; i++) {
    if (i % 3 === 0) fizz++;
    if (i % 5 === 0) buzz++;
  }

  return { fizz, buzz };
}
```
