---
id: 69f8c998d78ad3171a0713b9
title: "Izazov 286: Otvorena pitanja
challengeType: 28
dashedName: challenge-286
---

# --description--
Uzimajući u obzir listu brojeva problema i drugu listu brojeva zahteva za povlačenje promena (Pull Request/PR), vrati listu problema koji su još uvek otvoreni nakon što su svi PR-ovi spojisani.

- Zahtev za PR zatvara problem ako su njihovi brojevi permutacija jedni drugog. Na primer, problem 123 biće zatvoren sa PR-om 231 ili 312.
- Zahtev za PR ne zatvara problem sa potpuno istim brojem. Na primer, PR 123 neće zatvoriti problem 123. Stoga, problem sa svim identičnim brojevima se ne može zatvoriti.
- Bilo koji broj može imati uklonjene vodeće nule. Na primer, PR 201 bi zatvorio problem 12 (012, permutacija od 201). Slično tome, problem 201 biće zatvoren sa PR-om 12.

Vratite preostale otvorene probleme u navedenom redosledu.

# --hints--

`getOpenIssues([123, 234], [231])` inapaswa kurudisha `[234]`.

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])` inapaswa kurudisha `[345, 16]`.

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])` inapaswa kurudisha `[456, 332, 15]`.

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` inapaswa kurudisha `[115, 296, 24]`.

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` inapaswa kurudisha `[95, 395, 754, 296, 709, 237, 1802]`.

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
