---
id: 69f8c998d78ad3171a0713bf
title: "Zoezi la 292: Uwiano mpana wa sura"
challengeType: 28
dashedName: challenge-292
---

# --description--
Ako je dato sekvenca od dve različite dimenzije slike, vrati aspekt (ili omjer stranica) slike koja ima najveći odnos širine na visinu.

- Generisani niz karaktera biće u formatu `"WxH"`, na primer, `"1920x1080"`.
- Omjer stranica je omjer širine na visinu, sveden na najmanji cijeli broj. Na primer, `"1920x1080"` se smanjuje na `"16:9"`.
- Vratite niz karaktera u formatu `"W:H"`, na primer, `"16:9"`.

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` inapaswa kurudisha `"16:9"`.

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` inapaswa kurudisha `"4:3"`.

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` inapaswa kurudisha `"2:1"`.

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` inapaswa kurudisha `"9:16"`.

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` inapaswa kurudisha `"43:18"`.

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` inapaswa kurudisha `"2089:8539"`.

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
