---
id: 69f90c1329a94b37e2a2086d
title: "Zoezi la 294: Mchanganyiko wa mabano ya kawaida"
challengeType: 28
dashedName: challenge-294
---

# --description--

Ako je dat ceo broj `n`, vrati broj validnih kombinacija parova običnih zagrada.

- Validna kombinacija je niz znakova gde svaka otvorena obična zagrada ima odgovarajuću zatvorenu zagradu i nijedna zatvorena zagrada se ne pojavljuje pre svoje odgovarajuće otvorene zagrade.

Na primer, kada uklonite `2`, postoji 2 važeće kombinacije:

```json
(())
()()
```

# --hints--

`getCombinations(2)` inapaswa kurudisha `2`.

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` inapaswa kurudisha `5`.

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` inapaswa kurudisha `42`.

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` inapaswa kurudisha `1430`.

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` inapaswa kurudisha `742900`.

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
