---
id: 6a0dcd03ee4e68698080ef6c
title: "Zoezi la 308: Kagua kadi ya mkopo"
challengeType: 28
dashedName: challenge-308
---

# --description--
Ako vam je dat niz brojeva kreditne kartice, odredite da li je validan broj koristeći sledeću metodu:

- Počni od drugog broja od kraja, udvostručaj svaki drugi broj dok se krećeš ulevo.
- Ako udvostručavanje broja rezultira brojem većim od 9, oduzmi 9.
- Saberi sve brojeve (udvostručene i one koji nisu udvostručeni).
- Ako je zbir deljiv sa 10, broj je validan.

# --hints--

`isValidCard("4532015112830366")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```
