---
id: 6a0dcc730cb92a616f86f0c5
title: "Poglavlje 301: Konačni teret
challengeType: 28
dashedName: challenge-301
---

# --description--
Uzimajući u obzir količinu kašika deterdženta za veš koje imate i skup podataka o kašikama koje ste koristili u prethodnim danima, vratite broj punih dana deterdženta koji vam ostaje.

Izračunajte prosečnu dnevnu potrošnju na osnovu istorije korišćenja i uzmite tu potrošnju svakog dana, počevši od sada.

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` inapaswa kurudisha `5`.

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` inapaswa kurudisha `7`.

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` inapaswa kurudisha `11`.

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` inapaswa kurudisha `12`.

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` inapaswa kurudisha `1`.

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
