---
id: 6a0dcd03ee4e68698080ef68
title: "Zoezi la 304: Mpangilio wa Ratiba ya Safari"
challengeType: 28
dashedName: challenge-304
---

# --description--
Ako vam je dat red podataka koji sadrži najmanje dva opciona jednodnevna centra za obilazak, vratite broj validnih podešavanja raspoređivanja.

Raspored uvek uključuje `"breakfast"`, `"lunch"`, i `"dinner"`, što neće biti prihvaćeno kao argument. Opcione stanice mogu se postaviti bilo gde u rasporedu, uzimajući u obzir sledeća pravila:

- `"breakfast"` je uvek prvi, i mora postojati barem jedna stanica pre `"lunch"`.
- `"lunch"` mora se pojaviti pre `"dinner"`, i mora postojati barem jedna stanica između njih.
- Pored toga, samo jedna opcionalna stanica može biti vidljiva nakon `"dinner"`.

Rudisha idadi ya mipangilio halali.

# --hints--

`getItineraryCount(["library", "park"])` inapaswa kurudisha `2`.

```js
assert.equal(getItineraryCount(["library", "park"]), 2);
```

`getItineraryCount(["library", "park", "arcade"])` inapaswa kurudisha `18`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade"]), 18);
```

`getItineraryCount(["library", "park", "arcade", "store"])` inapaswa kurudisha `120`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store"]), 120);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe"])` inapaswa kurudisha `840`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe"]), 840);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"])` inapaswa kurudisha `55440`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440);
```

# --seed--

## --seed-contents--

```js
function getItineraryCount(stops) {

  return stops;
}
```

# --solutions--

```js
function getItineraryCount(stops) {
  const items = ["lunch", "dinner", ...stops];
  let count = 0;

  function permute(arr) {
    if (arr.length === 0) return [[]];
    return arr.flatMap((item, i) =>
      permute([...arr.slice(0, i), ...arr.slice(i + 1)]).map(p => [item, ...p])
    );
  }

  for (const perm of permute(items)) {
    const lunchIdx = perm.indexOf("lunch");
    const dinnerIdx = perm.indexOf("dinner");

    if (lunchIdx >= dinnerIdx) continue;
    if (lunchIdx < 1) continue;
    if (dinnerIdx - lunchIdx < 2) continue;
    if (perm.length - dinnerIdx - 1 > 1) continue;

    count++;
  }

  return count;
}
```
