---
id: 6a0dcd03ee4e68698080ef66
title: "Zoezi la 302: Kuathirika na Muda wa Safari ya Ndege"
challengeType: 28
dashedName: challenge-302
---

# --description--
Uzevši u obzir grad polaska, grad dolaska, trajanje leta u satima i pravac putovanja, vratite broj sati uticajnih od jet lag-a koje putnik doživljava.

Gradovi će biti uzeti iz sledeće liste koja uključuje vremenske zone UTC:

| Jiji | Mbali |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Ili kuhesabu saa za kuathirika na muda wa safari ya ndege:

1. Pronađite razliku u satima vremenske zone između dva grada.
2. Odredite faktor pravca. Ako putujete `"east"`, on je 1.5, inače je 1.0.
3. Izračunajte prilagođene sate sa formulom: razlika u vremenskoj zoni + (trajanje putovanja * 0.1) * faktor pravca

Zaokruži rezultujuće vrednosti na jednu decimalnu zaostavku.

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` inapaswa kurudisha `6.5`.

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` inapaswa kurudisha `5.8`.

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` inapaswa kurudisha `1.6`.

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` inapaswa kurudisha `4.7`.

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` inapaswa kurudisha `17.5`.

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` inapaswa kurudisha `5.9`.

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` inapaswa kurudisha `9.5`.

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
