---
id: 6a0dcd03ee4e68698080ef6b
title: "Zoezi la 307: Kanuni za Mipangilio ya Miji"
challengeType: 28
dashedName: challenge-307
---

# --description--
Ukoliko vam je data dvodimenzionalna mreža (niz podataka nizova podataka) koja predstavlja raspored gradskih objekata, vratite koordinate svih zgrada koje krše građevinske propise.

Svaka ćelija u mrežici ima jednu od oznaka iz tabele ispod. Struktura krši pravilo ako je bilo koji njen susjed (do četiri), vertikalno ili horizontalno, tip koji ne može biti njen susjed.

| Oznaka | Tip | Ne može biti sused sa |
| - | - | - |
| `"i"` | viwanda | `"R"`, `"I"` |
| `"A"` | kilimo | `"C"` |
| `"R"` | makazi | `"i"`, `"C"` |
| `"I"` | taasisi | `"i"` |
| `"C"` | biashara | `"R"`, `"A"` |
| `""` (prazna niska) | nije ugnježdeno | nema ograničenja |

Vratite sve koordinate ćelija koje krše kao niz parova `[row, col]`, u bilo kom redosledu. Ako nema nikakvih kršenja, vratite prazan niz podataka.

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])` inapaswa kurudisha `[[0, 0], [0, 1]]`.

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])` inapaswa kurudisha `[[0, 1], [1, 1]]`.

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` inapaswa kurudisha `[]`.

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` inapaswa kurudisha `[[0, 1], [0, 2], [0, 3]]`.

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` inapaswa kurudisha `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
