---
id: 6a151d32d772271dd2448c2c
title: "Vežba 310: Britanski engleski u američki engleski"
challengeType: 28
dashedName: challenge-310
---

# --description--
Kada dobijete rečenicu, promenite sve pravopise britanskog engleskog kako bi odgovarali američkom engleskom koristeći sledeću tabelu za pretraživanje i vratite ažuriranu rečenicu:

| Kiingereza cha Uingereza | Kiingereza cha Marekani |
| - | - |
| `"colour"` | `"color"` |
| `"flavour"` | `"flavor"` |
| `"honour"` | `"honor"` |
| `"neighbour"` | `"neighbor"` |
| `"labour"` | `"labor"` |
| `"humour"` | `"humor"` |
| `"centre"` | `"center"` |
| `"fibre"` | `"fiber"` |
| `"defence"` | `"defense"` |
| `"offence"` | `"offense"` |
| `"organise"` | `"organize"` |
| `"recognise"` | `"recognize"` |
| `"analyse"` | `"analyze"` |

- Promene bi trebalo da budu nezavisne od velikih ili malih slova. Na primer, `"Colour"` treba da se promeni u `"Color"`.
- Ulazni podaci mogu sadržati reči sastavljene na osnovu stvarnog pravopisa korena u tabeli, koje takođe treba promeniti. Na primer, `"colouring"` treba da se promeni u `"coloring"`, a `"disorganised"` treba da se promeni u `"disorganized"`.

# --hints--

`britishToAmerican("I love the colour blue.")` inapaswa kurudisha `"I love the color blue."`

```js
assert.equal(britishToAmerican("I love the colour blue."), "I love the color blue.");
```

`britishToAmerican("The fibre optic cable is new.")` inapaswa kurudisha `"The fiber optic cable is new."`

```js
assert.equal(britishToAmerican("The fibre optic cable is new."), "The fiber optic cable is new.");
```

`britishToAmerican("It's an honour to meet someone with such humour.")` inapaswa kurudisha `"It's an honor to meet someone with such humor."`

```js
assert.equal(britishToAmerican("It's an honour to meet someone with such humour."), "It's an honor to meet someone with such humor.");
```

`britishToAmerican("The unrecognised artist analysed his colour palette at the centre.")` inapaswa kurudisha `"The unrecognized artist analyzed his color palette at the center."`

```js
assert.equal(britishToAmerican("The unrecognised artist analysed his colour palette at the centre."), "The unrecognized artist analyzed his color palette at the center.");
```

`britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful.")` inapaswa kurudisha `The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.`

```js
assert.equal(britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful."), "The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.");
```

# --seed--

## --seed-contents--

```js
function britishToAmerican(sentence) {

  return sentence;
}
```

# --solutions--

```js
function britishToAmerican(sentence) {
  const lookup = {
    colour: "color",
    flavour: "flavor",
    honour: "honor",
    neighbour: "neighbor",
    labour: "labor",
    humour: "humor",
    centre: "center",
    fibre: "fiber",
    defence: "defense",
    offence: "offense",
    organise: "organize",
    recognise: "recognize",
    analyse: "analyze"
  };

  let result = sentence;

  for (const [british, american] of Object.entries(lookup)) {
    const regex = new RegExp(british, "gi");
    result = result.replace(regex, match => {
      return match[0] === match[0].toUpperCase()
        ? american[0].toUpperCase() + american.slice(1)
        : american;
    });
  }

  return result;
}
```
