---
id: 6a151d32d772271dd2448c2d
title: "Zoezi la 311: Mchawi wa maneno"
challengeType: 28
dashedName: challenge-311
---

# --description--

Dat je niz kodova reči koje koristiš, izračunaj ukupan broj poena.

Slovo u nizu predstavlja reč:

| Kod | Reč | Kategorija | Osnovni broj poena |
| - | - | - | - |
| `"f"` | Moto | Uharibifu | 3 |
| `"l"` | Umeme | Uharibifu | 3 |
| `"i"` | Barafu | Udhibiti | 2 |
| `"w"` | Upepo | Udhibiti | 2 |
| `"h"` | Tiba | Urejeshaji | 1 |
| `"s"` | Mlinzi | Urejeshaji | 1 |

Indikator kombinacije se određuje prema tome koliko uzastopnih reči je korišćeno iz različitih kategorija:

- Prva reč uvek dobija osnovnu bodovnu vrednost.
- Svaka sledeća reč iz različite kategorije od prethodne povećava množilac za 1.
- Korišćenje reči iz iste kategorije kao prethodna vraća množilac na 1.
- Bod za svaku reč je njena osnovna vrednost pomnožena trenutnim množilacem.

Vrati skup znakova iz niza reči.

# --hints--

`cast("fihwl")` inapaswa kurudisha `33`.

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` inapaswa kurudisha `45`.

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` inapaswa kurudisha `37`.

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` inapaswa kurudisha `50`.

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` inapaswa kurudisha `101`.

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```
