---
id: a26cbbe9ad8655a977e1ceb5
title: Napravi program za pretraživanje najduže reči
challengeType: 26
dashedName: build-a-longest-word-finder-app
---

# --description--

U ovoj laboratoriji, kreiraćete funkciju koja vraća dužinu najdužeg reči u datoj rečenici.

Na primer, u rečenici `"The quick brown fox jumped over the lazy dog"`, najduža reč je `"jumped"`, koja ima dužinu od 6.

**Cilj:** Popunite priče korisnika ispod i prođite sve testove da biste završili laboratoriju.

Priče o korisniku:

1. Moraš kreirati funkciju nazvanu `findLongestWordLength` koja prima niz znakova kao argument.
2. Funkcija treba da vrati dužinu najdužeg reči u nizu znakova.

# --hints--

Trebalo bi da napraviš funkciju koja se zove `findLongestWordLength`.

```js
assert.isFunction(findLongestWordLength);
```

`findLongestWordLength` treba da ima jedan parametar.

```js
assert.lengthOf(findLongestWordLength, 1);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` Trebalo bi da vrati broj.

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` mora da vrati `6`.

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

`findLongestWordLength("May the force be with you")` treba da vrati `5`.

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

`findLongestWordLength("Google do a barrel roll")` treba da vrati `6`.

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

`findLongestWordLength("Googling do a barrel roll")` trebalo bi da vrati `8`.

```js
assert.strictEqual(findLongestWordLength('Googling do a barrel roll'), 8);
```

`findLongestWordLength("What is the average airspeed velocity of an unladen swallow")` trebalo bi vratiti `8`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

`findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")` trebalo bi da vrati `19`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
