---
id: 6989f1ba358acba651d5ee6b
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada ćete inicijalizovati strukturu podataka za skladištenje vaših rezultata. Kreirajte varijablu nazvanu `result` i počnite sa praznim nizom podataka. Ovaj niz će sačuvati sve važeće kombinacije regulacionih izraza koje generišete.

Ispravite vaše informacije za return ili ponovo unesite `result` umesto praznog reda podataka.

# --hints--

Treba da deklariš parametar nazvan `result` u svojoj funkciji `genParentheses`.

```js
assert.match(genParentheses.toString(), /result/);
```

Trebalo bi da inicijalizujete `result` kao prazan niz podataka.

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

Trebalo bi da vratiš `result` na kraju svoje funkcije.

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
