---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Sada morate obraditi dodavanje zatvarajućih zagrada. Glavno pravilo je da možete dodati zatvarajuću zagradu samo ako osigurava balans, to jest, mora postojati više otvarajućih nego zatvarajućih zagrada.

Još unutar sekcije `else`, dodaj drugu tvrdnju od `if` ili proveri da li `closesUsed` je manji od `opensUsed`.

Ako je ovaj uslov tačan, dodajte još jedan novi red na `queue`: `[current + ')', opensUsed, closesUsed + 1]`. Ovo predstavlja stanje nakon dodavanja standardnih zatvarajućih zagrada.

# --hints--

Trebalo bi da imaš drugi deo `if` sa stanjem `closesUsed < opensUsed` unutar tvog dela `else`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

Trebalo bi da pritisnete `[current + ')', opensUsed, closesUsed + 1]` na panelu kada je stanje tačno.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
