---
id: 6989f207306332c06b0add39
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Po metodi širenja po širini, koristićete red za praćenje različitih stanja kada kreirate kombinacije običnih zagrada. Svako stanje će biti predstavljeno kao niz podataka sa tri elementa:

- Sekvenca karakterâ u izgradnji
- Broj otvorenih zagrada korišteno do sada
- Broj zatvorenih zagrada korišteno do sada

Kreiraj parametar nazvan `queue` i počni sa redom podataka koji ima samo jedan red: `['', 0, 0]`. Ovo predstavlja početno stanje sa sekvencijom praznih karaktera i korišćenjem nula zagradi.

# --hints--

Trebalo bi da objaviš parametar nazvan `queue` u svojoj funkciji `genParentheses`.

```js
assert.match(genParentheses.toString(), /queue/);
```

Trebalo bi da inicijalizujete `queue` za red podataka sa redom `['', 0, 0]`.

```js
assert.match(genParentheses.toString(), /queue\s*=\s*\[\s*\['',\s*0,\s*0\]\s*\]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  
  return result;
}
```
