---
id: 6989f2e35e6941e13e445156
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Sada ćeš implementirati glavnu petlju BFS-a. Napravi petlju za `while` koja nastavlja projekat `queue` koji nije pun, odnosno, `queue.length` je veće od `0`.

Unutar tog ciklusa, ukucajte `queue` u konzolu.

# --hints--

Trebalo bi da kreirate krug sa uslovom `while` koji je `queue.length > 0`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

Treba da napišete `queue` u konzoli unutar vaše petlje `while`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
--fcc-editable-region--
  
--fcc-editable-region--
  return result;
}
```
