---
id: 6989f48c3f0292c84f8d8938
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Sada morate odrediti koliko ste izgradili kompletnih skupa standardnih zagrada. Kompletan skup ima dužinu jednaku dvostrukom broju parova (jer svaki par doprinosi otvarajućoj i zatvarajućoj zagradi).

U okviru vašeg ciklusa od `while`, dodajte izjavu `if` ili proverite da li `current.length === 2 * pairs`. Kada je ovaj uslov tačan, dodajte `current` u red podataka `result`.

# --hints--

Trebalo bi da imaš deklaraciju `if` sa stanjem `current.length === 2 * pairs` unutar svog ciklusa `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

Treba da ubacite `current` u red podataka za `result` unutar vaše nove deklaracije `if`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
