---
id: 6989f3632f8b8bdf168aa321
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Unutar vašeg kruga od `while`, koristite `queue.shift()` da biste uklonili i dobili prvi element iz redova. Ovo implementira ponašanje "prvi uđe, prvi izađe" (FIFO) poznato u BFS.

Podeli ovaj red podataka po tri kriterijuma: `current`, `opensUsed`, i `closesUsed`.

# --hints--

Treba da koristiš `queue.shift()` da dobiješ prvi element iz reda.

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

Morate da podelite te rezultate u tri kriterijuma: `current`, `opensUsed`, i `closesUsed`.

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}
```
