---
id: 6989f555728e63ad8b349d84
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Ako trenutna sekvenca karaktera još uvek nije kompletirana, morate istražiti moguće sledeće stanja. Dodajte deo `else` za trenutne podatke iz `if`.

Unutar bloka `else`, ćete se baviti dodavanjem otvarajućih zagrada. Dodajte deklaraciju `if` radi provere da li je `opensUsed` unutar (ili ispod) `pairs`. Ovo osigurava da niste koristili više otvarajućih zagrada nego što je dozvoljeno.

Ako je ovo stanje `true`, gurnite novi red podataka na `queue`: `[current + '(', opensUsed + 1, closesUsed]`. Ovo predstavlja stanje nakon dodavanja standardnih otvarajućih zagrada.

# --hints--

Trebalo bi da imaš deo `else` za svoj informativni `if` unutar ciklusa `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

Trebalo bi da ima informacija za `if` unutar bloka od `else` sa statusom `opensUsed < pairs`.

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

Treba da guraš `[current + '(', opensUsed + 1, closesUsed]` u fasciklu kada je stanje `true`.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
