---
id: 6989ef885db3a3488a13b426
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Pre nego što izvršite glavni algoritam, morate validirati ulaz. Parametar `pairs` mora biti ceo broj, jer predstavlja broj standardnih parova zagrada za generisanje.

Dodajte informaciju za `if` na početak vaše funkcije kako biste proverili da li je `pairs` pun broj. Koristite metodu `Number.isInteger()` za to.

Ako je situacija istinita, ponovi sekvencu slova `The number of pairs should be an integer`.

# --hints--

Trebalo bi da imaš informaciju sa if statusa `!Number.isInteger(pairs)` unutar funkcije `genParentheses`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

Trebalo bi da vratite niz karaktera kada je unos ceo broj.

```js
let result = genParentheses("i");
assert.isString(result);
```

Treba da vratiš niz znakova `The number of pairs should be an integer` kada je ulaz niz znakova.

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  
  return [];
}
```

