---
id: 6989f10a72ad424d4ad6d514
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Stoga, morate potvrditi da je broj parova najmanje jedan, jer ne možete kreirati kombinaciju standardnih zagrada sa nula ili negativan broj parova.

Dodaj drugi deklarativni izraz za `if` ili proveri da li je `pairs` manje od `1`. Ako je ovo stanje tačno, vrati niz znakova `The number of pairs should be at least 1`.

# --hints--

Trebalo bi da imaš drugi deo za `if` sa stanjem `pairs < 1`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

Treba da vratiš niz karaktera `The number of pairs should be at least 1` kada je `pairs` manji od `1`.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  
  return [];
}
```
