---
id: 6989eb65324d97775c4965fe
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, ćete implementirati funkciju koja generiše sve validne kombinacije binarnih zagrada koristeći metod širine prvenog pretraživanja (BFS). Na primer, validna kombinacija binarnih zagrada za par dve je `(())` i `()()`.

Počnite kreiranjem funkcije nazvane `genParentheses` sa jednim parametrom `pairs`. Za sada, vratite prazan niz podataka iz te funkcije.

# --hints--

Treba da definišete funkciju nazvanu `genParentheses`.

```js
assert.isFunction(genParentheses);
```

Vaša funkcija za `genParentheses` treba da ima jedan parametar nazvan `pairs`.

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

Tvoja funkcija treba da vrati prazan red podataka.

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
