---
id: 6735a7370e0ae93a4577c689
title: Napravi raspored novina
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

U ovoj laboratoriji, kreiraćete raspored novina koristeći `CSS` Grid, uključujući koncepte kao što su redovi mreže i kolone mreže.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Tvoja stranica treba da sadrži element `main` sa klasom `newspaper-layout`.
1. Element `.newspaper-layout` treba da sadrži `header` sa klasom `title`, koji sadrži element `h1` sa nazivom novina.
1. Element `.newspaper-layout` treba da sadrži element `article` sa klasom `feature-article` za glavni članak.
1. `.feature-article` inapaswa kujumuisha kipengee cha `h2` kwa kichwa cha makala ikifuatiwa na kipengee cha `p` kwa maudhui ya makala.
1. Dodaj tri elementa `article` za manje članke, sa klasama `small-article1`, `small-article2` i `small-article3`.
1. Svaki manji članak treba da sadrži element `h3` za naslov članka, a zatim element `p` za sadržaj članka.
1. Element `.newspaper-layout` treba da sadrži element `article` sa klasom `secondary-article` za dodatni članak.
1. `.secondary-article` inapaswa kujumuisha kipengee cha `h2` kwa kichwa cha makala ikifuatiwa na kipengee cha `p` kwa maudhui ya makala.
1. Element `.newspaper-layout` treba da sadrži element `figure` sa klasom `cover-image` za prikaz slike koja predstavlja sadržaj novina.
1. Elementi sa klasama `title`, `feature-article`, `secondary-article`, `cover-image`, `small-article1`, `small-article2` i `small-article3` treba da imaju svojstvo `grid-area` postavljeno na isto ime kao njihova klasa.
1. Element `.newspaper-layout` treba da koristi CSS Grid i svojstvo `grid-template-areas` za definisanje rasporeda:
- `.title` treba da se prostire preko sva tri stupca u prvom redu.
- `.feature-article` treba da se prostire preko prva dva stupca u drugom redu, dok se `.cover-image` nalazi u trećem stupcu.
- `.secondary-article` treba da se prostire preko prva dva stupca u trećem redu, dok se `.cover-image` nalazi u trećem stupcu.
- Tri manja članka treba da budu prikazana u četvrtom redu.
1. Element `.newspaper-layout` treba da ima svojstvo `grid-template-columns` koje deli prostor na tri jednaka stupca.
1. Postavi svojstvo `grid-template-rows` elementa `.newspaper-layout` na `auto` za prvi red, a preostali prostor ravnomerno raspodeli na ostale redove.
1. Dodaj razmak između stavki u mreži.
1. Pobrini se da slika unutar elementa `.cover-image` stane u dodeljeni prostor tako što ćeš joj postaviti svojstvo `max-width` na `100%`.

**Napomena:** Uverite se da ste povezali svoj stilski list na HTML i koristili CSS.

# --hints--

Trebalo bi da imate element `main`.

```js
assert.exists(document.querySelector('main'));
```

Vaš element `main` treba da ima svojstvo `class` od `newspaper-layout`.

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

Trebalo bi da imate element sa `header` koji ima svojstvo `class` od `title` unutar vašeg elementa `.newspaper-layout`.

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

Vaš komponenta za `header` mora da uključuje komponentu od `h1`.

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

Vaša komponenta od `.title` mora imati tekst.

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

Treba da postoji element od `article` za svaki članak sa klasama predmeta od `feature-article`, `small-article1`, `small-article2`, `small-article3`, i `secondary-article`, respektivno.

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

`.feature-article` mora da sadrži element `h2` praćen elementom `p`.

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`, `.small-article2`, na `.small-article3` mora da uključuje element od `h3` praćen elementom od `p`.

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

`.secondary-article` mora da uključuje element od `h2`, praćen elementom od `p`.

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

Trebalo bi da imaš element sa klasom objekta `cover-image` unutar `.newspaper-layout`, koji je `figure`.

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

`.cover-image` mora da uključuje element/komponentu od `img`.

```js
assert.exists(document.querySelector(".cover-image img"));
```

Sifa za element `display` treba postaviti na `grid`, a to je `.newspaper-layout`.

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

Element `.newspaper-layout` treba da ima svojstvo `grid-template-areas` u kojem se oblast `title` prostire preko sva tri stupca prvog reda.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

Druga red mrežne strukture treba da sadrži `feature-article` koja se deli na dve kolone i `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

Treći red mrežnog dizajna treba da sadrži `secondary-article` koji se deli na dve kolone sa `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

Četvrti red strukture grida treba da ima `small-article1`, `small-article2`, na `small-article3`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

Vaš selektor za `.newspaper-layout` treba da koristi svojstvo `grid-template-columns` kako bi podelio prostor na tri jednaka dela. Zapamtite da možete koristiti funkciju `repeat` ili ručno definisati svaki deo.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

Vaš selektor za `.newspaper-layout` mora imati atribut `grid-template-rows` postavljen na `auto 1fr 1fr 1fr`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

Treba da dodate razmak između ćelija mreže.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title` trebalo bi da ima atribut `grid-area` od/za `title`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` mora imati svojstvo `grid-area` od `feature-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` treba imati svojstvo `grid-area` od `cover-image`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article` treba da ima svojstvo `grid-area` od `secondary-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1` mora imati svojstvo `grid-area` za `small-article1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2` treba da ima svojstvo `grid-area` od `small-article2`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3` mora imati svojstvo `grid-area` za `small-article3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

`img` unutar `.cover-image` mora imati svojstvo `max-width` od `100%`.

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
