---
id: 66db529d37ad966480ebb633
title: Napravi program za odabir ručka
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

U ovom laboratoriju, razradićeš program koji pomaže u upravljanju opcijama ručka. Radićeš sa skupom podataka o ručku, dodajući i uklanjajući stavke iz skupa podataka, i nasumično birajući jednu od opcija za ručak.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Morate deklarisati varijablu `lunches` i dodeliti joj vrednost praznog niza podataka koji će se koristiti za skladištenje predmeta za ručak.

2. Treba da kreirate funkciju `addLunchToEnd` koja prima niz podataka kao prvi argument i string kao drugi argument. Funkcija treba da:
    - Dodaje string na kraj niza podataka.
    - Ispisuje string `[Lunch Item] added to the end of the lunch menu.` na konzoli, gde je `[Lunch Item]` string prosleđen funkciji.
    - Vraća ažurirani niz podataka.

3. Treba da kreirate funkciju ``addLunchToStart`` koja prima red podataka kao prvi argument i niz znakova kao drugi argument. Funkcija treba da:
    - Dodaje niz znakova na početak reda podataka.
    - Ispisuje niz znakova ``[Lunch Item] added to the start of the lunch menu.`` u konzolu, gde je ``[Lunch Item]`` niz znakova poslat funkciji.
    - Vraća ažurirani red podataka.

4. Treba da kreirate funkciju ``removeLastLunch`` koja prima niz podataka kao argument. Funkcija mora:
   - Ukloniti poslednji element iz niza podataka.
   - Ako je uklanjanje uspešno, ispisati niz znakova ``[Lunch Item] removed from the end of the lunch menu.`` u konzolu, gde je ``[Lunch Item]`` element uklonjen iz niza podataka.
   - Ako je niz podataka prazan, ispisati niz znakova ``"No lunches to remove."`` u konzolu.
   - Vratiti ažurirani niz podataka.

5. Treba da kreirate funkciju ``removeFirstLunch`` koja prima niz podataka kao argument. Funkcija mora:
   - Ukloniti prvi element iz niza podataka.
   - Ako je uklanjanje uspešno, ispisati string ``[Lunch Item] removed from the start of the lunch menu.`` u konzolu, gde je ``[Lunch Item]`` element uklonjen iz niza podataka.
   - Ako je niz podataka prazan, ispisati string ``"No lunches to remove."`` u konzolu.
   - Vratiti ažurirani niz podataka.

6. Trebalo bi da kreirate funkciju `getRandomLunch` koja prima niz podataka kao argument. Funkcija treba da:
   - Izabere nasumičan element iz niza podataka.
   - Ako je uspešno, ispisaće niz znakova `Randomly selected lunch: [Lunch Item]` u konzolu, gde je `[Lunch Item]` nasumičan element iz niza podataka.
   - Ako je niz podataka prazan, ispisaće niz znakova `"No lunches available."` u konzolu.

7. Treba da kreirate funkciju `showLunchMenu` koja prima niz podataka kao svoj argument i:
   - Ako postoje elementi u nizu podataka, napišite sekvencu znakova `Menu items: [Lunch Item], [Lunch Item]...` u konzoli, gde je svaki `[Lunch item]` jedan od elemenata u nizu podataka, po redosledu.
   - Ako je niz podataka prazan, napišite sekvencu znakova `"The menu is empty."` u konzoli.

# --hints--

Trebalo bi da objavite razliku `lunches` i postavite prazan red podataka za skladištenje sastojaka za ručak.

```js
assert.isArray(lunches);
```

Trebalo bi da definišete funkciju `addLunchToEnd`.

```js
assert.isFunction(addLunchToEnd);
```

Funkcija `addLunchToEnd` mora imati dva parametra.

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` treba da napiše sekvencu karaktera `"Tacos added to the end of the lunch menu."` u konzolu.

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` treba da vrati `["Pizza", "Tacos", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

Treba da definišete funkciju `addLunchToStart`.

```js
assert.isFunction(addLunchToStart);
```

Funkcija `addLunchToStart` mora imati dva parametra.

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` mora da unese sekvencu karaktera `"Sushi added to the start of the lunch menu."` u konzolu.

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` treba da vrati `["Pizza", "Burger", "Sushi"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

Trebalo bi da definišete funkciju `removeLastLunch`.

```js
assert.isFunction(removeLastLunch)
```

Funkcija `removeLastLunch` mora imati jedan parametar.

```js
assert.lengthOf(removeLastLunch, 1);
```

Kada je ulazni red podataka prazan, funkcija `removeLastLunch` treba da izpiše sekvencu karaktera `"No lunches to remove."` na konzoli.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])` mora da napiše sekvencu karaktera `"Toast removed from the end of the lunch menu."` na konzoli.

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` bi trebalo da vrati `["Sushi", "Pizza"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

Trebalo bi da definišete funkciju `removeFirstLunch`.

```js
assert.isFunction(removeFirstLunch);
```

Funkcija `removeFirstLunch` mora imati jedan parametar.

```js
assert.lengthOf(removeFirstLunch, 1);
```

Kada je red ulaznih podataka prazan, funkcija `removeFirstLunch` treba da napiše niz znakova `"No lunches to remove."` na konzolu.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])` Mora da unese sekvencu karaktera `"Salad removed from the start of the lunch menu."` na konzoli.

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` treba da vrati `["Pizza", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`, `addLunchToStart`, `removeLastLunch`, na `removeFirstLunch` i treba da vrate isti red koji je poslat kao unos nakon poboljšanja.

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

Trebalo bi da definišete funkciju `getRandomLunch`.

```js
assert.isFunction(getRandomLunch);
```

Funkcija `getRandomLunch` treba da ima jedan parametar.

```js
assert.lengthOf(getRandomLunch, 1);
```

Kada je red ulaznih podataka prazan, funkcija `getRandomLunch` treba da ispise niz znakova `"No lunches available."` na konzoli.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

Kada red podataka nije prazan, funkcija ``getRandomLunch`` mora da napiše sekvencu karaktera u formatu ``Randomly selected lunch: [Lunch Item]`` na konzoli.

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

Funkcija `getRandomLunch` ne bi trebalo da modifikuje red podataka koji je poslat kao argument.

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

Trebalo bi da definišete funkciju `showLunchMenu`.

```js
assert.isFunction(showLunchMenu);
```

Funkcija `showLunchMenu` treba da ima jedan parametar.

```js
assert.lengthOf(showLunchMenu, 1);
```

Kada je red unosnih podataka prazan, funkcija `showLunchMenu` bi trebalo da ispise niz znakova `"The menu is empty."` na konzolu.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])` mora se napisati `"Menu items: Greens, Corns, Beans"` na konzoli.

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` treba da se napiše `"Menu items: Pizza, Burger, Fries, Salad"` u konzolu.

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
