---
id: 5a24c314108439a4d4036143
title: Premesti logiku stanja u Redux
challengeType: 6
forumTopicId: 301429
dashedName: extract-state-logic-to-redux
---

# --description--

Sada si završio deo React, moraš da premestiš logiku koja je implementirana unutar `state` ka Redux-u. Ovo je prvi korak za povezivanje jednostavne aplikacije/programa React sa Redux-om. Jedina funkcija tvog programa je dodavanje nove poruke od korisnika u neuređenu listu. Primer je jednostavan da pokaže kako React i Redux rade zajedno.

# --instructions--

Prvo, definiš tip akcije `ADD` i postavi ga kao const `ADD`. Zatim, definiši kreator akcije `addMessage()` koji kreira akciju dodavanja poruke. Moraćeš da proslediš `message` ovom kreatoru akcije i uključiš poruku u vraćeni `action`.

Zatim kreiraj reducer nazvan `messageReducer()` koji obrađuje stanje poruke. Početno stanje bi trebalo da bude prazan niz podataka. Ovaj reducer bi trebalo da doda poruku u niz poruka sačuvanih u stanju, ili da vrati trenutno stanje. Na kraju, kreiraj svoj Redux store i prosledi mu reducer.

# --hints--

Konstanta `ADD` treba da postoji i da ima vrednost stringa `ADD`.

```js
assert(ADD === 'ADD');
```

Generator akcije `addMessage` mora da vrati nešto sa `type` jednako kao `ADD` i `message` jednako kao poruka koja se prenosi.

```js
assert(
  (function () {
    const addAction = addMessage('__TEST__MESSAGE__');
    return addAction.type === ADD && addAction.message === '__TEST__MESSAGE__';
  })()
);
```

`messageReducer` treba da bude funkcija.

```js
assert(typeof messageReducer === 'function');
```

Skladište treba da postoji i da ima početno stanje postavljeno na prazan niz.

```js
assert(
  (function () {
    const initialState = store.getState();
    return typeof store === 'object' && initialState.length === 0;
  })()
);
```

Slanje `addMessage` akcije prema skladištu treba bezbedno da doda novu poruku u niz poruka sačuvan u stanju.

```js
assert(
  (function () {
    const initialState = store.getState();
    const isFrozen = DeepFreeze(initialState);
    store.dispatch(addMessage('__A__TEST__MESSAGE'));
    const addState = store.getState();
    return isFrozen && addState[0] === '__A__TEST__MESSAGE';
  })()
);
```

`messageReducer` treba da vrati trenutno stanje ako se pozove sa bilo kojom drugom akcijom.

```js
assert(
  (function () {
    const addState = store.getState();
    store.dispatch({ type: 'FAKE_ACTION' });
    const testState = store.getState();
    return addState === testState;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Define ADD, addMessage(), messageReducer(), and store here:
```

# --solutions--

```jsx
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);
```
