---
id: 5a24c314108439a4d403614a
title: Idi napred odavde
challengeType: 6
forumTopicId: 301434
dashedName: moving-forward-from-here
---

# --description--

Čestitam! Završili ste lekcije o React i Redux-u. Postoji još jedna stvar koju treba pomenuti pre nego što nastavite. Obično nećete kreirati programe za React u editoru koda kao ovaj. Ovaj izazov vam daje uvid kako izgleda sintaksa ako radite sa sistemom fajlova na vašem sopstvenom računaru. Kod bi trebalo da izgleda ovako, osim korišćenja `import` (ovo donosi sve zavisnosti koje su vam date u izazovu).

Na kraju, pisanje koda za React sa Reduxom obično zahteva određenu konfiguraciju. To može brzo postati komplikovano. Ako ste zainteresovani za testiranje na sopstvenom računaru, <a href="https://www.freecodecamp.org/news/install-react-with-create-react-app/" target="_blank" rel="noopener noreferrer nofollow">Create React App</a> dolazi gotov i instaliran.

Inače, možete omogućiti Babel kao Preprocesor za JavaScript u CodePen-u, dodajte React i ReactDOM kao eksterne resurse za JavaScript, i radite tamo takođe.

# --instructions--

Napišite ``'Now I know React and Redux!'`` u konzoli.

# --hints--

Ujumbe `Now I know React and Redux!` unapaswa kuandikwa kwenye konsoli.

```js
assert(
    /console\s*\.\s*log\s*\(\s*('|"|`)Now I know React and Redux!\1\s*\)/.test(
      code)
  );
```

# --seed--

## --seed-contents--

```jsx
/*
import React from 'react'
import ReactDOM from 'react-dom'
import { Provider, connect } from 'react-redux'
import { createStore, combineReducers, applyMiddleware } from 'redux'
import thunk from 'redux-thunk'

import rootReducer from './redux/reducers'
import App from './components/App'

const store = createStore(
  rootReducer,
  applyMiddleware(thunk)
);

ReactDOM.render(
  <Provider store={store}>
    <App/>
  </Provider>,
  document.getElementById('root')
);
*/

// Only change code below this line
```

# --solutions--

```jsx
console.log('Now I know React and Redux!');
```
