---
id: 5a24c314108439a4d4036144
title: Koristi Provider da povežeš Redux i React
challengeType: 6
forumTopicId: 301435
dashedName: use-provider-to-connect-redux-to-react
---

# --description--

U prethodnom zadatku, kreirao ste Redux store za upravljanje slojem podataka poruka i kreiranje akcije za dodavanje nove poruke. Sledeći korak je pružanje pristupa React za Redux store i akcije koje treba poslati za ažuriranje. React Redux nudi svoj paket `react-redux` kako bi pomogao u implementaciji ovih funkcija.

React Redux pruža mali API sa dve važne komponente: `Provider` i `connect`. Drugi izazov obrađuje `connect`. `Provider` je omotač iz React Redux-a koji okružuje vašu React aplikaciju. Ovaj omotač omogućava pristup Redux `store` skladištu i `dispatch` akcijama kroz celo stablo vaših komponenti. `Provider` prima dva svojstva: Redux skladište i podkomponente vaše aplikacije. Postavljanje `Provider` komponente oko `App` komponente može izgledati ovako:

```jsx
<Provider store={store}>
  <App/>
</Provider>
```

# --instructions--

Editor koda sada prikazuje sav vaš Redux kod i React iz nekoliko prošlih izazova. Uključuje Redux store, akcije i sekciju `DisplayMessages`. Jedina nova funkcija je sekcija `AppWrapper` na samom dnu. Koristite ovu komponentu višegraga da prikažete `Provider` iz `ReactRedux`, i prosledite Redux store kao prop. Zatim prikažite sekciju `DisplayMessages` kao dete (child). Kada završite, trebalo bi da vidite vašu komponentu React prikazanu na stranici.

**Napomena:** React Redux je ovde dostupan kao globalna promenljiva, tako da možete pristupiti Provider-u pomoću tačkaste notacije. Kod u editoru koristi ovu mogućnost i čuva ga kao konstantu `Provider` kako bi ga koristio u `render` metodi `AppWrapper`.

# --hints--

`AppWrapper` Trebalo bi da prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('AppWrapper').length === 1;
  })()
);
```

`Provider` omotač treba da primi prosleđeni `store` argument, odnosno Redux skladište.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
      return __helpers
        .removeWhiteSpace(code)
        .includes('<Providerstore={store}>');
    })()
  );
```

`DisplayMessages` treba da se prikazuje kao dete komponente `AppWrapper`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('AppWrapper').find('DisplayMessages').length === 1
    );
  })()
);
```

Sekcija od `DisplayMessages` treba da prikaže `h2`, `input`, `button`, na `ul`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h2').length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('ul').length === 1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};



const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
    this.handleChange = this.handleChange.bind(this);
    this.submitMessage = this.submitMessage.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {  
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Render the Provider below this line

  // Change code above this line
};
```

# --solutions--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
 this.handleChange = this.handleChange.bind(this);
 this.submitMessage = this.submitMessage.bind(this);
 }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };  
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Change code below this line
  render() {
    return (
      <Provider store = {store}>
        <DisplayMessages/>
      </Provider>
    );
  }
  // Change code above this line
};
```
