---
id: 66ba762af611169359d9d369
title: Napiši članak za novine
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Treba da postavite `font-size` korenskog HTML elementa na `24px`.
1. Treba da imate element sa klasom `newspaper` koji sadrži sve ostale elemente.
1. Vaš element `.newspaper` treba da ima svojstvo `font-size` sa vrednošću `16px` i font `Open Sans` sa rezervnim fontom `sans-serif`.
1. Unutar elementa `.newspaper` treba da imate najmanje još sedam elemenata: jedan za naziv novina sa klasom `name`, jedan za datum članka sa klasom `date`, jedan za naslov sa klasom `headline`, jedan za podnaslov sa klasom `sub-headline`, jedan za autora sa klasom `author` i dva pasusa, svaki sa klasom `text`. Svi ovi elementi treba da budu popunjeni sadržajem vašeg članka.
1. Element `.name` treba da ima svojstvo `font-size` koje je dvostruko veće od `font-size` korenskog elementa i treba da koristi font `Times New Roman` sa rezervnim fontom `serif`.
1. Elementi `.name` i `.author` treba da koriste CSS kako bi sav tekst bio ispisan velikim slovima.
1. Element `.headline` treba da ima svojstvo `font-size` koje je dvostruko veće od `font-size` njegovog roditeljskog elementa i treba da bude podebljan.
1. Element `.sub-headline` treba da ima svojstvo `font-weight` sa vrednošću `100`, svojstvo `font-size` koje je `1.5` puta veće od `font-size` njegovog roditeljskog elementa i treba da bude ispisan malim velikim slovima.
1. Element `.author` treba da koristi CSS kako bi bio podebljan.
Elementi `.text` treba da imaju svojstvo `text-indent` sa vrednošću `20px`.
Elementi `.text` treba da imaju svojstvo `line-height` koje je dvostruko veće od `font-size` njihovog roditeljskog elementa.
Prvo slovo elemenata `.text` treba da bude podebljano i dvostruko veće od `font-size` njihovog roditeljskog elementa. Za ovo koristite selektor `::first-letter`.

# --hints--

Trebalo bi da imate selektor `html` koji postavlja njegov atribut `font-size` na `24px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

Trebalo bi da ima element sa klasom `newspaper`.

```js
assert.exists(document.querySelector('.newspaper'));
```

Trebalo bi da ima selektor `.newspaper` koji postavlja atribute njegovih elemenata `font-size` na `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

Trebalo bi da imaš selektor `.newspaper` koji postavlja atribute njegovih elemenata `font-family` na `'Open Sans', sans-serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

Trebalo bi da imaš element sa klasom `name` unutar tvog elementa `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

Vaše polje `.name` treba da ima ime vaše novine.

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

Trebalo bi da imaš element sa klasom `date` unutar tvog elementa `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

Polje vaše `.date` mora imati datum.

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

Treba da ima element sa klasom `headline` unutar vašeg elementa `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

Vaš element za `.headline` mora imati naslov vašeg članka.

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

Trebalo bi da imaš element sa klasom `sub-headline` unutar svog elementa `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

Kao element vašeg `.sub-headline`, treba da ima mali naslov za vaš članak.

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

Trebalo bi da imaš element sa klasom `author` unutar tvog elementa `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

Vaše polje `.author` treba da ima ime autora vašeg članka.

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

Trebalo bi da imaš najmanje dva paragrafa elementa, svaki sa klasom `text`, unutar tvog elementa `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

Vaši elementi za `.text` moraju imati tekst vašeg članka.

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

Trebalo bi da imaš selektor `.name` koji postavlja atribute njegovih elemenata `font-size` na `2rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

Trebalo bi da imaš selektor `.name` koji postavlja atribute njegovih elemenata za `font-family` na `'Times New Roman', serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

Element za `.name` treba da ima svojstvo `text-transform` postavljeno na `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Element `.author` mora imati atribut `text-transform` postavljen na `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Trebalo bi da imaš birač za `.headline` koji postavlja atribute njegovih elemenata `font-size` kao `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

Trebalo bi da imaš selektor `.headline` koji postavlja atribute njegovih elemenata `font-weight` na `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

Trebalo bi da ima selektor `.sub-headline` koji postavlja atribute njegovih elemenata za `font-weight` na `100`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

Trebalo bi da imaš birač za `.sub-headline` koji postavlja atribute njegovih elemenata `font-size` na `1.5em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

Trebalo bi da imaš selektor `.sub-headline` koji postavlja atribute njegovih elemenata `font-style` na `italic`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

Trebalo bi da imaš birač za `.author` koji postavlja atribute njegovih elemenata `font-weight` na `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

Trebalo bi da imaš selektor `.text` koji postavlja atribute za njegove elemente `text-indent` na `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

Trebalo bi da imate birnik `.text` koji postavlja atribute njegovih elemenata `line-height` na `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

Trebalo bi da imate selektor `.text::first-letter` koji postavlja atribute za njegove elemente `font-weight` kao `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

Trebalo bi da imaš selektor `.text::first-letter` koji postavlja atribute za njegove elemente `font-size` na `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
