---
id: 69008059d43d6751385883c1
title: Pronađi greške u operacijama inkrementa i dekrementa u neispravnom programu
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

Tvoj prijatelj te je zamolio da mu pomogne da pronađe greške u svom kodu. Uči kako da koristi inkrementalni i dekrementalni operator, ali dobija neočekivane rezultate u konzoli.

Zadatak je da pomognete u identifikaciji grešaka i ispravite kod kako bi rezultati prikazani na konzoli odgovarali očekivanim rezultatima.

Možete pregledati prethodne lekcije o operacijama povećanja i smanjenja, zajedno sa napomenama o prefiksima i sufiksima.

Priče o korisniku:

1. Trebalo bi da imate parametar nazvan `updatedCounter` koji tačno pruža vrednost `11` korišćenjem operatora sabiranja. Ažurirajte matematički izraz koristeći ispravnu notaciju prefiksa ili sufiksa umesto da direktno postavite vrednost.
2. Trebalo bi da imate parametar nazvan `finalScore` koji tačno pruža vrednost `8` korišćenjem operatora sabiranja. Ažurirajte matematički izraz koristeći ispravnu notaciju prefiksa ili sufiksa umesto da direktno postavite vrednost.
3. Trebalo bi da imate parametar nazvan `updatedCoins` koji tačno pruža vrednost `2` korišćenjem operatora oduzimanja. Ažurirajte matematički izraz koristeći ispravnu notaciju prefiksa ili sufiksa umesto da direktno postavite vrednost.
4. Trebalo bi da imate parametar nazvan `newHealth` koji tačno pruža vrednost `7` korišćenjem operatora oduzimanja. Ažurirajte matematički izraz koristeći ispravnu notaciju prefiksa ili sufiksa umesto da direktno postavite vrednost.

# --hints--

Trebalo bi da imaš kriterijum nazvan `updatedCounter` koji koristi prefiksnu notifikaciju za dodavanje (`++counter`).

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter` mora biti jednako `11`.

```js
assert.strictEqual(updatedCounter, 11);
```

Trebalo bi da imaš parametar koji se zove `finalScore` i koristi notifikaciju za završni prefiks dodavanja (`score++`).

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore` mora biti jednako sa `8`.

```js
assert.strictEqual(finalScore, 8);
```

Trebalo bi da imate kriterijum nazvan `updatedCoins` koji koristi notu prefiksa za smanjenje (`--coins`).

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins` mora biti jednako `2`.

```js
assert.strictEqual(updatedCoins, 2);
```

Trebalo bi da imaš parametar nazvan `newHealth` koji koristi notifikaciju za sufiksno smanjenje (`health--`).

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth` treba biti jednako `7`.

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
