---
id: 66cc14ca14e65e7a9ebecb80
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Sada, nakon što vaš kalkulator može da sabere dva broja, došlo je vreme za predstavljanje funkcije koja oduzima dva broja.

Objavi funkciju nazvanu `calculateDifference` koja prima dva parametra, `num1` i `num2`.

Unutar funkcije, vrati razliku između dve parametre.

# --hints--

Trebalo bi da imaš funkciju koja se zove `calculateDifference`.

```js
assert.isFunction(calculateDifference);
```

Vaša funkcija `calculateDifference` treba da ima dva parametra pod nazivima `num1` i `num2`.

```js
const regex = __helpers.functionRegex('calculateDifference', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Vaša funkcija `calculateDifference` bi trebalo da vrati razliku između dva parametra.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
