---
id: 66cc0a9e06e00b75d6782be9
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

U prethodnim lekcijama, naučili ste kako da pozovete funkciju. Pozivanje funkcije znači izvršiti kod koji se nalazi unutar te funkcije.

Evo podsetnika o tome kako pozvati funkciju:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Povećajte izjavu za `console.log`, unutar njega, biće funkcionalni deo od `addTwoAndSeven`.

Sada bi trebalo da vidite zbir `2` i `7`, prikazan na konzoli.

# --hints--

Trebalo bi da imaš `console.log` u svom kodu.

```js
assert.match(code, /console\.log(.*)/);
```

Trebalo bi da pozoveš funkciju `addTwoAndSeven` unutar `console.log`.

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
