---
id: 66cc172af7159a7c67804544
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

Sada je vreme da dodate funkciju množenja na svoj kalkulator.

Objavi funkciju nazvanu `calculateProduct` koja prima dva parametra, `num1` i `num2`.

U toj funkciji, vrati rezultat sa dve parametre.

Zatim, dodajte `console.log` koji poziva funkciju `calculateProduct` sa argumentima `13` i `5`.

# --hints--

Trebalo bi da imaš funkciju nazvanu `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

Tvoja funkcija za `calculateProduct` mora imati dva parametra koja se zovu `num1` i `num2`.

```js
const regex = __helpers.functionRegex('calculateProduct', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Tvoja funkcija ``calculateProduct`` bi trebalo da vrati rezultat sa dve parametre.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

Trebalo bi da imaš `console.log` koji poziva funkciju `calculateProduct` sa argumentima `13` i `5`.

```js
assert.match(code, /console\.log\s*\(\s*calculateProduct\s*\(\s*13\s*,\s*5\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
