---
id: 66cc21d23238dc8240a8a182
title: Korak 16
challengeType: 1
dashedName: step-16
---

# --description--

Biće dobro ako tvoj kalkulator računa kvadrat broja. Kvadrat broja je broj pomnožen sam sobom.

Da bi izračunali kvadrat broja u JavaScript, možete koristiti metodu <dfn>Math.pow()</dfn>. Metoda `Math.pow()` prima dva argumenta: osnovni broj i stepen.

Takođe možete koristiti parametar snage (`**`) da izračunate kvadrat broja.

Evo primera:

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

Objavi funkciju nazvanu `calculateSquare` koja uzima argument `num` i vraća kvadrat od `num`.

# --hints--

Trebalo bi da ima funkciju koja se zove `calculateSquare`.

```js
assert.isFunction(calculateSquare);
```

Vaša funkcija za `calculateSquare` treba da uzme parametar `num`.

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Funkcija tvoja za `calculateSquare` treba da vrati kvadrat od `num`.

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
