---
id: 66cc0de2920bec775f610424
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Iako funkcije rade kako je očekivano, ima mnogo ponavljanja u vašem kodu.

Pošto gradiš kalkulator, ne želiš kreirati funkciju za svaku moguću kombinaciju brojeva koje želiš da dodaš zajedno.

Biće bolje kreirati jednu funkciju koja se može ponovo koristiti i koja može sabrati bilo koji par brojeva.

Ukloni sav kod koji si do sada napisao.

U sledećem koraku, ćete pregledati kako da radite sa parametrima i argumentima u funkcijama.

# --hints--

Nema funkciju pod nazivom `addTwoAndSeven`.

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

Treba da ima funkcionalni poziv za `addTwoAndSeven`.

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

Nema funkciju koja se zove `addThreeAndFour`.

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

Nedostaje mu funkcionalni poziv za `addThreeAndFour`.

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
