---
id: 66cc06012e46aa6bc9b8c001
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

Na ovoj radionici, naučićete kako da radite sa funkcijama prilikom kreiranja programa kalkulatora.

U prethodnim lekcijama, naučili ste kako da deklarišete funkcije ovako:

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

Počni kreiranjem funkcije koja se zove `addTwoAndSeven`.

Možete izabrati da koristite sintaksu obične funkcije ili sintaksu strelične funkcije.

# --hints--

Trebalo bi da objaviš funkciju nazvanu `addTwoAndSeven`.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
