---
id: 66cc1a3a39aef47d6473cb2f
title: Korak 13
challengeType: 1
dashedName: step-13
---

# --description--

Sljedeći korak je povećanje performansi deljenja na vašem kalkulatoru.

Objavi funkciju nazvanu `calculateQuotient` koja prima dva parametra, `num1` i `num2`.

Unutar te funkcije, vrati vrednost sa dve parametre.

Zatim, dodaj `console.log` koja poziva funkciju `calculateQuotient` sa argumentima `7` i `11`.

# --hints--

Trebalo bi da imaš funkciju nazvanu `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

Tvoja funkcija za `calculateQuotient` mora imati dva parametra koja se zovu `num1` i `num2`.

```js
const regex = __helpers.functionRegex('calculateQuotient', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Tvoja funkcija ``calculateQuotient`` treba da vrati vrednost sa dve parametre.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

Trebalo bi da imaš `console.log` koji poziva funkciju `calculateQuotient` sa argumentima `7` i `11`.

```js
assert.match(code, /console\.log\s*\(\s*calculateQuotient\s*\(\s*7\s*,\s*11\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
