---
id: 66cc095dfe1682753d2ab030
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Koncept koji sam naučio/la u prethodnim časovima bio je kako vratiti vrednost iz funkcije.

Evo podsetnika o tome kako vratiti vrednost iz funkcije:

```js
function myFunction() {
  return "Hello World";
}
```

Unutar vaše funkcije za `addTwoAndSeven`, vratite vrednost od `2` na `7`.

# --hints--

Tvoja funkcija za `addTwoAndSeven` treba da vrati zbir od `2` i `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Tvoja funkcija sa `addTwoAndSeven` treba da vrati zbir koristeći operator za `+`.

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--

--fcc-editable-region--
}
```
