---
id: 66cc0ba5881acb7692cfc4de
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada, je vreme da dodamo još jednu funkciju.

Objavi funkciju nazvanu `addThreeAndFour` koja rešava zbir od `3` i `4`.

Zatim pozovite funkciju `addThreeAndFour` unutar `console.log` da biste videli rezultate.

# --hints--

Trebalo bi da deklarišete funkciju nazvanu `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

Tvoja funkcija za `addThreeAndFour` treba da vrati zbir od `3` i `4`.

```js
assert.equal(addThreeAndFour(), 7);
```

Tvoja funkcija ``addThreeAndFour`` treba da vrati zbir koristeći specijalistu za ``+``.

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

Trebalo bi da pozoveš funkciju `addThreeAndFour` unutar `console.log`.

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
