---
id: 66cc12fa504b0479dac479a0
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Prednost funkcija je u tome što se mogu pozvati sa različitim argumentima, što vam omogućava da ponovo koristite isti kod sa različitim vrednostima.

Ovde su primeri poziva funkcija i različitih argumenata za sekvencu znakova:

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

Dodaj `console.log`, funkciju nazvanu `calculateSum` na listu `2` na `5`.

# --hints--

Trebalo bi da imaš `console.log` koji poziva funkciju `calculateSum` sa argumentima `2` i `5`.

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
