---
id: 66cc0f1ae40802781b2ea972
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

U prethodnim lekcijama, naučili ste kako da radite sa parametrima funkcije.

Funkcijski parametar je parametar deklarisan u deklaraciji funkcije i funkcioniše kao identifikator.

Ovde je primer funkcije sa parametrom:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Objavi funkciju nazvanu `calculateSum` koja prima dva parametra, `num1` i `num2`.

# --hints--

Trebalo bi da objavite funkciju nazvanu `calculateSum`.

```js
assert.isFunction(calculateSum);
```

Tvoja funkcija za `calculateSum` treba da primi dva parametra, `num1` i `num2`.

```js
const regex = __helpers.functionRegex('calculateSum', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
