---
id: 66cc2739f687e484d50bb6f1
title: Korak 18
challengeType: 1
dashedName: step-18
---

# --description--

Poslednji deo funkcije vašeg kalkulatora će biti funkcija kvadratnog korena.

Kvadrat je broj pomnožen sobom. Dakle, kvadratni koren je broj koji, kada se pomnoži sobom, daje originalan broj.

Primer, kvadratni koren iz `9` je `3` jer `3 * 3 = 9`.

Da biste dobili kvadratni koren broja u JavaScript, možete koristiti metodu <dfn>Math.sqrt()</dfn>.

Evo primera:

```js
// result: 3
Math.sqrt(9);
```

Objavi funkciju nazvanu `calculateSquareRoot` koja prima argument `num` i vraća kvadratni koren od `num`.

# --hints--

Trebalo bi da imaš funkciju nazvanu `calculateSquareRoot`.

```js
assert.isFunction(calculateSquareRoot);
```

Vaša funkcija `calculateSquareRoot` treba da primi parametar za `num`.

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Tvoja funkcija za `calculateSquareRoot` treba da vrati kvadratni koren iz `num`.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
