---
id: 66cc10958d72e57907f898b2
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Unutar vaše funkcije za `calculateSum`, moraćete da vratite zbir od dva parametra.

# --hints--

Tvoja funkcija za `calculateSum` treba da vrati zbir dva parametra.

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
