---
id: 66cdf76685e4cb5a8726e27b
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Da biste više testirali svoju funkciju za `calculateSum`, moraćete da je pozovete sa različitim argumentima.

Poziva drugi poziv od `console.log` koji pokreće funkciju `calculateSum` sa argumentima `10` i `10`.

Pod tim pozivom od `console.log`, dodaj još jedan poziv od `console.log` koji poziva funkciju/metodu `calculateSum` sa argumentima `5` i `5`.

# --hints--

Treba da imate poziv za `console.log` koji poziva funkciju `calculateSum` sa argumentima `10` i `10`.

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

Trebalo bi da ima poziv za `console.log` koji poziva funkciju `calculateSum` sa argumentima `5` i `5`.

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));
--fcc-editable-region--

--fcc-editable-region--
```
