---
id: 68e76b1ce8216249d04df670
title: Izgradi čas o planetama
challengeType: 27
dashedName: build-a-planet-class
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba da kreirate klasu objekata pod nazivom `Planet`.
1. Klasa `Planet` treba da ima metodu `__init__` koja:
   - Ima četiri parametra: `self`, `name`, `planet_type`, i `star`.
   - Vraća `TypeError` sa porukom `name, planet type, and star must be strings` ako je bilo koji poslat argument tipa različit od niza karaktera.
   - Vraća `ValueError` sa porukom `name, planet_type, and star must be non-empty strings` ako je bilo koji poslat argument prazan niz karaktera.
   - Postavite poslate vrednosti na atribute primerka elementa `name`, `planet_type`, i `star`.
1. Klasa `Planet` treba da ima metodu `orbit` koja vraća niz karaktera u formatu `{name} is orbiting around {star}...`.
1. Klasa `Planet` treba da ima metodu `__str__` koja vraća niz karaktera u formatu `Planet: {name} | Type: {planet_type} | Star: {star}`.
1. Treba da kreirate tri instance klase `Planet` pod nazivima `planet_1`, `planet_2`, i `planet_3`.
1. Morate ispisati svaki objekat planete da biste videli rezultat metode `__str__`.
1. Treba pozvati metodu `orbit` na svakom objektu planete i ispisati rezultate.

# --hints--

Treba da kreiraš klasu objekata koja se zove `Planet`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_class('Planet')
`
    )
})
```

Klasa `Planet` treba da ima metodu `__init__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__init__')

`
    )
})
```

Putanja za `__init__` treba da ima četiri parametra: `self`, `name`, `planet_type`, na `star`, u ovom redosledu.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').find_function('__init__').has_args('self, name, planet_type, star')

`
    )
})
```

Putanja za `__init__` treba da vrati `TypeError` sa porukom `name, planet type, and star must be strings` ako prvi argument nije tip niza karaktera.

```js
({
    test: () => runPython(
`
try:
  planet = Planet(8, "planet", "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Putanja `__init__` bi trebalo da vrati `TypeError` sa porukom `name, planet type, and star must be strings` ako drugi argument koji je poslat nije tip stringa.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", [], "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Puta `__init__` bi trebalo da vrati `TypeError` sa porukom `name, planet type, and star must be strings` ako treći poslat argument nije tip sekvence karaktera.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", 0)
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Putanja za `__init__` treba da izbaci `ValueError` sa porukom `name, planet_type, and star must be non-empty strings` ako je prvi poslati post niz praznih karaktera.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("", "planet", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Putanja `__init__` treba da vrati `ValueError` sa porukom `name, planet_type, and star must be non-empty strings` ako je drugi argument poslat prazan niz znakova.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Putanja za `__init__` treba da izbaci `ValueError` sa porukom `name, planet_type, and star must be non-empty strings` ako je treći poslat komentar niz praznih znakova.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", "")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Treba da staviš vrednost `name` u `self.name` putem `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.name').is_equivalent('self.name = name')
`
    )
})
```

Trebalo bi da postaviš vrednost `planet_type` na `self.planet_type` putem `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.planet_type').is_equivalent('self.planet_type = planet_type')
`
    )
})
```

Trebalo bi da postaviš vrednost `star` na `self.star` putem `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.star').is_equivalent('self.star = star')
`
    )
})
```

Klasa za `Planet` bi trebalo da ima metodu od `orbit`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('orbit')
`
    )
})
```

Putanja za `orbit` treba da restaurira sekvencu karaktera u formatu `{name} is orbiting around {star}...`.

```js
({
    test: () => runPython(
`
pl = Planet('planetino', 'type', 'star')
assert pl.orbit() == 'planetino is orbiting around star...'
`
    )
})
```

Klasa za `Planet` treba da ima metodu za `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__str__')
`
    )
})
```

Putanja za `__str__` treba da ponovo postavi sekvencu karaktera u format od `Planet: {name} | Type: {planet_type} | Star: {star}`.

```js
({
    test: () => runPython(
`
assert str(Planet("abcd", "bcdea", "jhkl")) == "Planet: abcd | Type: bcdea | Star: jhkl"
`
    )
})
```

Treba da kreiraš tri primere klase `Planet` pod nazivima `planet_1`, `planet_2`, na `planet_3`.

```js
({
    test: () => runPython(
`
assert isinstance(planet_1, Planet)
assert isinstance(planet_2, Planet)
assert isinstance(planet_3, Planet)
`
    )
})
```

Treba da ispišete sve o planeti kako biste videli rezultate puta za `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1)')
assert _Node(_code).has_call('print(planet_2)')
assert _Node(_code).has_call('print(planet_3)')
`
    )
})
```

Trebaće ti da pozoveš putanju `orbit` za svaki planetarni objekat i ispišeš rezultate.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1.orbit())')
assert _Node(_code).has_call('print(planet_2.orbit())')
assert _Node(_code).has_call('print(planet_3.orbit())')
`
    )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class Planet:
    def __init__(self, name, planet_type, star):
        args = (name, planet_type, star)
        if not all(isinstance(arg, str) for arg in args):
            raise TypeError("name, planet type, and star must be strings")
        if not all(args):
            raise ValueError('name, planet_type, and star must be non-empty strings')
        self.name = name
        self.planet_type = planet_type
        self.star = star   # The star the planet orbits

    def orbit(self):
        return f'{self.name} is orbiting around {self.star}...'

    def __str__(self):
        return f'Planet: {self.name} | Type: {self.planet_type} | Star: {self.star}'


# Test the Planet class
planet_1 = Planet('Earth', 'terrestrial', 'Sun')
planet_2 = Planet('Jupiter', 'gas giant', 'Sun')
planet_3 = Planet('Kepler-452b', 'super-Earth', 'Kepler-452')

print(planet_1.orbit()) # Earth is orbiting around Sun...
print(planet_1) # Planet: Earth | Type: terrestrial | Star: Sun

print(planet_2.orbit()) # Jupiter is orbiting around Sun...
print(planet_2) # Planet: Jupiter | Type: gas giant | Star: Sun

print(planet_3.orbit()) # Kepler-452b is orbiting around Kepler-452...
print(planet_3) # Planet: Kepler-452b | Type: super-Earth | Star: Kepler-452
```
