---
id: 587d7fb3367417b2b2512bfc
title: Dodaj detalje u svoj package.json
challengeType: 2
forumTopicId: 301522
dashedName: add-a-description-to-your-package-json
---

# --description--

Sledeći deo fajla package.json je polje `description`; gde se unose kratki, ali informativni detalji o vašem projektu.

Ako jednog dana planirate da objavite paket na npm-u, ovo je niz karaktera koji treba da prodate svoju ideju korisniku kada odluči da instalira ili ne instalira vaš paket. Međutim, to nije jedina upotreba opisa; to je dobar način za sažimanje onoga što projekat radi. Izuzetno je važno u svakom projektu Node.js da pomogne drugim programerima, budućim održivačima ili čak vama samoj u budućnosti da brzo razumeju projekat.

Bez obzira šta planirate za svoj projekat, opisi su veoma preporučljivi. Evo primera:

```json
"description": "A project that does something awesome",
```

# --instructions--

Dodajte `description` u fajl package.json vašeg projekta.

**Napomena:** Zapamtite da koristite dvostruke navodne znakove za imena polja (") i zarez (,) za razdvajanje polja.

# --hints--

package.json mora imati validan ključ za "description"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.description, '"description" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
