---
id: 587d7fb5367417b2b2512c04
title: Uklonite paket iz vaše zavisnosti
challengeType: 2
forumTopicId: 301530
dashedName: remove-a-package-from-your-dependencies
---

# --description--

Sada si pokušao nekoliko načina za upravljanje zavisnostima tvog projekta korišćenjem sekcije `dependencies` u fajlu `package.json`. Takođe si uključio eksterne pakete postavljanjem u fajl i čak reći `npm`-u koje verzije želiš, koristeći specijalne znakove kao što su tilda ili krez.

Ali šta ako želite da uklonite eksterni paket koji više ne želite? Možda ste već otkrili, samo uklonite par ključ-vrednost koji odgovara tom paketu iz vaših zavisnosti.

Ovim istim načinom se koristi i za uklanjanje drugog dela u vašem package.json.

# --instructions--

Uklonite paket `@freecodecamp/example` iz vaših zavisnosti.

**Napomena:** Proverite da imate tačan broj zareza nakon uklanjanja tog paketa.

# --hints--

`"dependencies"` ne bi trebalo da sadrži `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.notProperty(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" still includes "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
