---
id: 587d7fb4367417b2b2512bff
title: Dodaj verziju u svoj package.json
challengeType: 2
forumTopicId: 301525
dashedName: add-a-version-to-your-package-json
---

# --description--

`version` je jedan od potrebnih delova za vašu package.json datoteku. Ovaj deo opisuje trenutnu verziju vašeg projekta. Evo primera:

```json
"version": "1.2.0",
```

# --instructions--

Dodaj `version` u fajl package.json tvog projekta.

# --hints--

package.json mora da ima važeći ključ "version"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.version, '"version" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
