---
id: 587d7fb4367417b2b2512bfd
title: Dodajte ključne reči u svoj package.json
challengeType: 2
forumTopicId: 301526
dashedName: add-keywords-to-your-package-json
---

# --description--

Sehemu ya `keywords` je deo gde možete opisati svoj projekat koristeći relevantne ključne reči. Evo primera:

```json
"keywords": [ "descriptive", "related", "words" ],
```

Kao što vidite, ovaj deo je strukturisan kao red podataka sekvence slova koji su unutar dvostrukih navodnika.

# --instructions--

Dodajte red sa odgovarajućom sekvencijom karaktera u deo `keywords` u fajl `package.json` vašeg projekta.

Moj ključni izraz bi trebalo da bude "freecodecamp".

# --hints--

package.json mora imati važeći ključ za "keywords"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.keywords, '"keywords" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Deo sa "keywords" treba da bude red podataka

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.isArray(data.keywords, '"keywords" is not an array');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

"keywords" treba da uključuje "freecodecamp"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.include(
        data.keywords,
        'freecodecamp',
        '"keywords" does not include "freecodecamp"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
