---
id: 587d7fb4367417b2b2512c00
title: Vaš projekat ima eksterne pakete iz npm-a
challengeType: 2
forumTopicId: 301527
dashedName: expand-your-project-with-external-packages-from-npm
---

# --description--

Jedan od najvećih razloga za korišćenje menadžera paketa je njihovo moćno upravljanje zavisnostima. Umesto da ručno proveravate da li ste dobili sve zavisnosti svaki put kada pokrenete projekat na novom računaru, npm automatski vam postavlja sve. Ali kako npm zna tačno šta vaš projekat zahteva? Pogledajte deo vaše package.json datoteke: `dependencies`

U ovom odeljku, paketi koji su potrebni vašem projektu skladište se korišćenjem sledeće strukture:

```json
"dependencies": {
  "package-name": "version",
  "express": "4.14.0"
}

```

# --instructions--

Povećaj verziju `1.1.0` paketa `@freecodecamp/example` u sekciji `dependencies` vašeg fajla `package.json`.

**Napomena:** `@freecodecamp/example` je lažni paket koji se koristi kao alat za učenje.

# --hints--

`"dependencies"` trebalo bi da sadrži `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Verzija za `"@freecodecamp/example"` treba da bude `"1.1.0"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^[\^\~]?1\.1\.0/,
        'Wrong version of "@freecodecamp/example" installed. It should be 1.1.0'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
