---
id: 587d7fb5367417b2b2512c03
title: Koristite znak krez (^) za korišćenje najnovije verzije zavisnosti
challengeType: 2
forumTopicId: 301531
dashedName: use-the-caret-character-to-use-the-latest-minor-version-of-a-dependency
---

# --description--

Kao što tilde, koju smo naučili u prethodnom izazovu, dozvoljava npm da instalira najnoviji PATCH za zavisnost, znak caret-a (`^`) takođe omogućava npm da instalira buduća ažuriranja. Razlika je u tome što će znak caret-a dozvoliti MINOR i PATCH ažuriranja.

Trenutna verzija za `@freecodecamp/example` treba da bude `~1.2.13`, što omogućava npm-u instaliranje najnovije verzije od `1.2.x`. Ako umesto toga koristite znak kapeše (^) kao prefiks verzije, npm će biti dozvoljen da se ažurira na bilo koju verziju od `1.x.x`.

```json
"package": "^1.3.8"
```

Ovo će omogućiti ažuriranja za bilo koju verziju paketa `1.x.x`.

# --instructions--

Koristite znak caret (`^`) kao prefiks verzije `@freecodecamp/example` u vašoj zavisnosti i dozvolite npm-u da ga ažurira na bilo koju noviju MINOR verziju.

**Napomena:** Brojevi verzije ne smeju se menjati.

# --hints--

`"dependencies"` treba da uključuje `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Verzija od `"@freecodecamp/example"` bi trebalo da odgovara `"^1.x.x"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\^1\./,
        'Wrong version of "@freecodecamp/example". It should be ^1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
