---
id: 587d7fb4367417b2b2512bfe
title: Dodajte licencu u vaš `package.json`
challengeType: 2
forumTopicId: 301523
dashedName: add-a-license-to-your-package-json
---

# --description--

Seo/Deo `license` je deo koji informiše korisnike o tome šta im je dozvoljeno da urade sa vašim projektom.

Popularne licence za projekte sa otvorenim izvorom uključuju MIT i BSD. Nije potrebna informacija o licenci, a zakoni o autorskom pravu u mnogim zemljama automatski vam daju vlasništvo nad onim što kreirate. Međutim, uvek je dobra praksa da jasno definišete šta korisnici mogu, a šta ne smeju da rade. Evo primera sekcije licence:

```json
"license": "MIT",
```

# --instructions--

Popunite deo `license` u fajl `package.json` vašeg projekta, kako smatrate prikladnim.

# --hints--

package.json mora imati važeći ključ za "license"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.license, '"license" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
