---
id: 587d7fb5367417b2b2512c02
title: Koristite prefiks tilde ili uvek koristite najnoviju verziju zavisnosti/popravaka
challengeType: 2
forumTopicId: 301532
dashedName: use-the-tilde-character-to-always-use-the-latest-patch-version-of-a-dependency
---

# --description--

U poslednjem izazovu, rekao si npm-u da uključi samo određenu verziju paketa. To je važan način zaključavanja zavisnosti ako moraš osigurati da su različiti delovi tvog projekta kompatibilni jedni sa drugima. Ali u mnogim slučajevima, ne želiš da propustiš ispravke grešaka jer često uključuju važne sigurnosne zakrpe i (nadamo se) ne pokreću stvari dok to radiš.

Da biste dozvolili da se npm zavisnost ažurira na najnoviju PATCH verziju, možete postaviti prefiks verzije zavisnosti sa znakom tilde (`~`). Evo primera kako dozvoliti ažuriranja za bilo koju verziju `1.3.x`.

```json
"package": "~1.3.8"
```

# --instructions--

U fajlu package.json, vaša trenutna pravila o tome kako npm može da ažurira `@freecodecamp/example` je korišćenje specifične verzije (`1.2.13`). Ali sada, želite da dozvolite najnoviju verziju `1.2.x`.

Koristi prefiks tilde (`~`) da instaliraš verziju `@freecodecamp/example` u svoje zavisnosti, i dozvoli npm-u da je ažurira na bilo koju novu _patch_ verziju.

**Napomena:** Brojevi vlasničke verzije ne bi trebalo menjati.

# --hints--

`"dependencies"` mora da uključuje `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Verzija od `"@freecodecamp/example"` mora biti u skladu sa `"~1.2.13"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\~1\.2\.13/,
        'Wrong version of "@freecodecamp/example". It should be ~1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
