---
id: 6140c7e645d8e905819f1dd4
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Počnite sa osnovnim šablonom. Dodajte svoju izjavu za `DOCTYPE`, vaš element `html` čiji je jezik postavljen na engleski, vaše elemente `head` i `body`.

Dodaj svoj element `meta` na tačno `charset`, tvoje polje `title`, na elementu `link` za fajl `./styles.css`.

Stavi `title` postati `Ferris Wheel`.

# --hints--

Vaš kod treba da ima prefiks `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE/gi);
```

Treba da staviš razmak nakon deklaracije `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Trebalo bi da definišete tip dokumenta kao `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Trebalo bi da zatvorite deklaraciju `DOCTYPE` sa etiketom za zatvaranje `>` nakon tipa.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Vaš element za `html` mora imati oznaku za otvaranje sa atributom elementa `lang` sa vrednošću `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Element vašeg `html` treba da ima zatvarajući tag.

```js
assert.match(code, /<\/html\s*>/);
```

Tvoje saopštenje za `DOCTYPE` trebalo bi biti na početku tvog HTML.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Trebalo bi da imaš etiketu za otvaranje `head`.

```js
assert.match(code, /<head\s*>/i);
```

Trebalo bi da ima oznaku za zatvaranje `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

Trebalo bi da imaš etiketu za otvaranje `body`.

```js
assert.match(code, /<body\s*>/i);
```

Trebalo bi da ima oznaku za zatvaranje `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

Tvoj element ``body`` treba da bude prvi element nakon elementa ``head``.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Vaši elementi za `head` i `body` treba da budu braća/sestre.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Tvoj element za `head` treba da bude prvi element unutar tvog elementa za `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Tvoj element `head` mora biti unutar elementa `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Tvoj element `body` treba da bude unutar elementa `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Tvoj kod mora imati element `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Vaš element za `meta` mora imati svojstvo elementa `charset` sa vrednošću `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Vaš element od `meta` mora da bude unutar elementa `head`.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Vaš kod mora imati karakteristiku `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tvoj projekat treba da ima zaglavlje `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Zapamtite, velika i mala slova zajedno sa pravopisom su važne za naslov.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Vaš kod mora imati element `link`.

```js
assert.match(code, /<link/);
```

Trebalo bi da imate jedan element `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Tvoj element `link` treba da bude unutar tvog elementa `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Vaš element `link` mora imati svojstvo elementa `rel` sa vrednošću `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Vaš element `link` mora imati atribut elementa `href` sa vrednošću `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
