---
id: 6140d10d50636e14695013b2
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Kreiraj selektor ili ciljaj svoj drugi element od `.line`. Postavi atribut za `transform` da bude `rotate(60deg)`.

Zapamtite da vam svojstvo `transform` omogućava da kontrolišete oblik elementa. U ovom slučaju, korišćenjem vrednosti `rotate(60deg)` rotiraće se element oko svog dela `transform-origin` za 60 stepeni u smeru kazaljke na satu.

# --hints--

Moraš da kreiraš selektor za `.line:nth-of-type(2)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)'));
```

Vaš selektor za `.line:nth-of-type(2)` mora imati svojstvo `transform` postavljeno na `rotate(60deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)')?.transform === 'rotate(60deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
