---
id: 6140cbeec34e970dfe75e710
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Napravite čarobnjak za vašu komponentu `.wheel`. Počnite postavljanjem `border` na `2px solid black`, `border-radius` na `50%`, i `margin-left` na `50px`.

# --hints--

Trebalo bi da imate birač za `.wheel`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

Tvoj selektor za `.wheel` mora imati svojstvo `border` postavljeno na `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

Tvoj selektor za `.wheel` mora imati svojstvo `border-radius` postavljeno na `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

Tvoj selektor za `.wheel` mora imati svojstvo `margin-left` postavljeno na `50px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
