---
id: 6140c9d35015ae0ba0c250e8
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Dodaj `div` unutar tvog elementa od `body` i polja od `class` od `wheel`.

U novom `div`, dodaj šest komponenti od `span` sa `class` postavljeno na `line`, i šest komponenti od `div` sa `class` postavljeno na `cabin`.

# --hints--

Trebalo bi da kreirate novi `div` unutar vašeg elementa `body`.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert(divs?.length === 1);
```

`div` tvoj unutar elementa od `body` treba imati `class` od `wheel`.

```js
assert(document.querySelector('body div')?.classList?.contains('wheel'));
```

Vaš element za `.wheel` treba da ima šest elemenata `span` unutar sebe.

```js
assert(document.querySelectorAll('.wheel span')?.length === 6);
```

Šest vaših elemenata `span` unutar elementa `.wheel` treba da ima `class` od `line`.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert(spans?.every(span => span?.classList?.contains('line')));
assert(document.querySelectorAll('.line')?.length === 6);
```

Vaša komponenta za `.wheel` mora da sadrži šest elemenata od `div` unutar sebe.

```js
assert(document.querySelectorAll('.wheel div')?.length === 6);
```

Šest vaših elemenata `div` unutar elementa `.wheel` treba da ima `class` od `cabin`.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert(divs?.every(div => div?.classList?.contains('cabin')));
assert(document.querySelectorAll('.cabin')?.length === 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
