---
id: 6140cd32d018ed0f600eefce
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Postavi svojstvo ``position`` za selektor ``.wheel`` na ``absolute``. Postavi ``height`` i ``width`` na oba na ``55vw``.

# --hints--

Vaš selektor za `.wheel` mora imati svojstvo od `position`, postavljeno na `absolute`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.position === 'absolute');
```

Tvoj selektor za `.wheel` mora imati svojstvo `height` postavljeno na `55vw`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.height === '55vw');
```

Vaš selektor za `.wheel` mora imati svojstvo `width` postavljeno na `55vw`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.width === '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
