---
id: 6140d0069049f5139d78da40
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Atribut `transform-origin` koristi se za postavljanje sekcije gde se promene CSS primenjuju oko. Na primer, kada izvršite `rotate` (što ćete učiniti kasnije u ovom projektu), `transform-origin` određuje sekciju gde se element okreće oko.

Za selektor `.line`, postavi svojstvo `transform-origin` sa vrednošću `0% 0%`. Ovo će pomeriti originalni deo za `0%` od leve strane i `0%` od gornje, čineći ga gornjim levim uglom elementa.

# --hints--

Vaš selektor za `.line` mora imati svojstvo `transform-origin` postavljeno na `0% 0%`.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert(transformOrigin === '0% 0%' || transformOrigin === '0% 0% 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
