---
id: 617ace7d831f9c16a569b38a
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Izaberite vaš selektor za `.wheel` svojstvo od `max-height` do svojstva `max-width`, sve postavljeno na `500px`.

# --hints--

Vaš selektor za `.wheel` mora imati svojstvo `max-height` postavljeno na `500px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxHeight === '500px');
```

Selektor za `.wheel` mora imati svojstvo od `max-width` postavljeno na `500px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxWidth === '500px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
}
--fcc-editable-region--
```
