---
id: 587d824e367417b2b2512c57
title: Proverite da li je nešto primer komponente
challengeType: 2
forumTopicId: 301605
dashedName: test-if-an-object-is-an-instance-of-a-constructor
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`#instanceOf` potvrđuje da je stavak primer komponente.

# --instructions--

Unutar `tests/1_unit-tests.js` pod testom sa oznakom `#18` u paketu `Objects`, promeni svaki `assert` u `assert.instanceOf` ili `assert.notInstanceOf` da bi test prošao (trebalo bi da se proceni kao `true`). Ne menjajte argumente koji se prosleđuju za potvrdu.

# --hints--

Sve testiranje mora proći.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Treba da izaberete ispravan put za prvu verifikaciju – `instanceOf` naspram `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInstanceOf',
  'myCar is not an instance of Plane'
);
```

Treba da izaberete ispravan način za drugu verifikaciju – `instanceOf` naspram `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'instanceOf',
  'airlinePlane is an instance of Plane'
);
```

Trebate izabrati ispravnu metodu za trostruku autentifikaciju - `instanceOf` naspram `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'instanceOf',
  'everything is an Object in JavaScript...'
);
```

Treba da izaberete ispravan način za četvrti identifikator – `instanceOf` naspram `notInstanceOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'notInstanceOf',
  'myCar.wheels is not an instance of String'
);
```
