---
id: 587d824c367417b2b2512c4e
title: Proverite da li je neka vrednost manja ili bar jednaka drugoj
challengeType: 2
forumTopicId: 301606
dashedName: test-if-one-value-is-below-or-at-least-as-large-as-another
---

# --description--

Kao podsetnik, ovaj projekat je baziran na početnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

Unutar `tests/1_unit-tests.js`, pod testom sa oznakom `#9` u paketu `Comparisons`, promenite svaki `assert` u `assert.isBelow` ili `assert.isAtLeast` kako bi test prošao (trebalo bi da se evaluira kao `true`). Ne menjajte argumente poslati za potvrdu.

# --hints--

Sve testiranje mora proći.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Trebalo bi da izaberete ispravan način za prvu verifikaciju - `isBelow` naspram `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtLeast',
  '5 is at least (>=) 5'
);
```

Treba da izaberete ispravan način za sekundnu verifikaciju - `isBelow` naspram `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isAtLeast',
  '2 * Math.random() is at least 0'
);
```

Trebalo bi vam da izaberete tačnu metodu za treću autentifikaciju - `isBelow` naspram `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isBelow', '1 is smaller than 2');
```

Trebalo bi da izaberete ispravan put za četvrti identifikacioni kod - `isBelow` naspram `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isBelow',
  '2/3 (0.6666) is smaller than 1'
);
```
