---
id: 587d824d367417b2b2512c50
title: Pokušaj ako je vrednost red podataka
challengeType: 2
forumTopicId: 301600
dashedName: test-if-a-value-is-an-array
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

Unutar `tests/1_unit-tests.js`, pod testom sa oznakom `#11` u redu `Arrays`, promenite svaku `assert` u `assert.isArray` ili `assert.isNotArray` kako bi test prošao (treba da se evaluira kao `true`). Nemojte menjati argumente koji se prosleđuju za potvrdu.

# --hints--

Sve provere moraju da prođu.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Treba da izaberete ispravan način za prvu verifikaciju – `isArray` naspram `isNotArray`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isArray',
  'String.prototype.split() returns an Array'
);
```

Trebalo bi da izaberete ispravan način za drugu verifikaciju - `isArray` naspram `isNotArray`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=10');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isNotArray',
  'Array.prototype.indexOf() returns a number'
);
```
