---
id: 587d824b367417b2b2512c47
title: Proverite da li je parametar ili funkcija definisan
challengeType: 2
forumTopicId: 301602
dashedName: test-if-a-variable-or-function-is-defined
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

U okviru `tests/1_unit-tests.js` pod testom sa oznakom `#2` u suiti `Basic Assertions`, promenite svaki `assert` u `assert.isDefined()` ili `assert.isUndefined()` kako bi test prošao (trebalo bi da se evaluira kao `true`). Ne menjajte argumente koji se šalju za potvrdu.

# --hints--

Sve testiranje mora proći.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Trebalo bi da izaberete ispravan način za prvu verifikaciju – `isDefined` naspram `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'isDefined',
    'Null is not undefined'
  );
```

Treba da izaberete ispravan način za dvostruku verifikaciju – `isDefined` naspram `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isUndefined',
    'Undefined is undefined'
  );
```

Trebalo bi da izaberete ispravan metod za treću verifikaciju - `isDefined` naspram `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isDefined',
    'A string is not undefined'
  );
```
