---
id: 587d8250367417b2b2512c5d
title: Pokreni funkcionione testove koristeći bezglavi pretraživač
challengeType: 2
forumTopicId: 301595
dashedName: run-functional-tests-using-a-headless-browser
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Na stranici postoji formular za unos podataka. On šalje podatke na krajnji deo `PUT /travellers` kao AJAX zahtev.

Kada se zahtev uspešno završi, kod klijenta dodaje `<div>` sa informacijama u odgovoru na DOM.

Evo primera kako da se koristi Zombie.js za interakciju sa formularima:

```js
test('Submit the surname "Polo" in the HTML form', function (done) {
  browser.fill('surname', 'Polo').then(() => {
    browser.pressButton('submit', () => {
      browser.assert.success();
      browser.assert.text('span#name', 'Marco');
      browser.assert.text('span#surname', 'Polo');
      browser.assert.elements('span#dates', 1);
      done();
    });
  });
});
```

Prvo, `fill` metoda objekta `browser` popunjava polje `surname` u formi vrednošću `'Polo'`. `fill` vraća obećanje, pa se nakon njega koristi `then`.

Unutar poziva povratne funkcije za `then`, metoda `pressButton` objekta `browser` koristi se za pozivanje slušaoca događaja (event listener) forme `submit`. Metoda `pressButton` je asinhrona.

Zatim, kada se primi odgovor sa AJAX zahteva, izvode se nekoliko validacija kako bi se potvrdilo:

1. Status odgovora je `200`
2. Tekst unutar elementa `<span id='name'></span>` odgovara na `'Marco'`
3. Tekst unutar elementa `<span id='surname'></span>` odgovara na `'Polo'`
4. Postoji element `1` `<span id='dates'></span>`

Konačno, poziv za povratak `done` je aktiviran, što je potrebno zbog ne-istovremenog testiranja.

# --instructions--

Unutar `tests/2_functional-tests.js`, u testu `'Submit the surname "Colombo" in the HTML form'` (`// #5`), automatizujte sledeće:

1. Popunite formu sa imenom porodice `Colombo`
2. Kliknite na dugme za slanje

Unutar poziva povratka na `pressButton`:

1. Potvrdite da je status OK `200`
2. Potvrdite da je tekst unutar elementa `span#name` to je `'Cristoforo'`
3. Potvrdite da je tekst unutar elementa `span#surname` to je `'Colombo'`
4. Potvrdite da element (v) `span#dates` postoji i da je njegov broj `1`

Ne zaboravi da ukloniš poziv za `assert.fail()`.

# --hints--

Sve pokušaje treba da prođe.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Potvrdite da je zahtev sa bezglavičnog pretraživača uspešan.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

Treba da potvrdiš da je tekst unutar elementa `span#name` sadržaj `'Cristoforo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Cristoforo\1/);
```

Treba da potvrdite da je tekst unutar elementa `span#surname` na mestu `'Colombo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Colombo\1/);
```

Trebalo bi da potvrdite da element `span#dates` postoji i da je njegov broj 1.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```
