---
id: 587d824e367417b2b2512c58
title: Pokreni testiranje funkcija na API endpointima koristeći Chai-HTTP
challengeType: 2
forumTopicId: 301593
dashedName: run-functional-tests-on-api-endpoints-using-chai-http
---

# --description--

Kao podsjetnik, ovaj projekat je napravljen na osnovu početnog projekta kopiranog iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Mocha vam omogućava da testirate asinhrone operacije, kao što su pozivi na životni ciklus za API, koristeći ekstenziju nazvanu `chai-http`.

Sledeći je primer testa koji koristi `chai-http` za grupu nazvanu `'GET /hello?name=[name] => "hello [name]"'`.

```js
suite('GET /hello?name=[name] => "hello [name]"', function () {
  test('?name=John', function (done) {
    chai
      .request(server)
      .keepOpen()
      .get('/hello?name=John')
      .end(function (err, res) {
        assert.equal(res.status, 200, 'Response status should be 200');
        assert.equal(res.text, 'hello John', 'Response should be "hello John"');
        done();
      });
  });
});
```

Ovaj test koristi zahtev za `GET` na serveru sa imenom kao pretraga argumenta URL (`?name=John`). U funkciji povratnog poziva za putanju `end`, objekat odgovora (`res`) prima se i ima svojstvo `status`.

Prvi red `assert.equal` proverava da li je stanje isto kao `200`. Drugi red `assert.equal` proverava da li niz karaktera odgovora (`res.text`) je isti kao `"hello John"`.

Pia, proverite parametar `done` u funkciji povratnog poziva testa. Pozivanje bez argumenata na kraju testa je važno da signalizira da je asinhrotna operacija završena.

Na kraju, proverite putanju `keepOpen` odmah nakon putanje `request`. Obično pokrećete svoje testove sa komandne linije, ili kao deo procesa direktnog povezivanja, i mogli biste dozvoliti `chai-http` da direktno pokrene i zaustavi vaš server.

Međutim, testovi koji se pokreću kada pošaljete link vašeg projekta zahtevaju da je vaš server aktivan, tako da morate koristiti metodu `keepOpen` kako biste sprečili da `chai-http` zaustavi vaš server.

# --instructions--

U okviru `tests/2_functional-tests.js`, izmenite test `'Test GET /hello with no name'` (`// #1`) tako da proverava ili `status` i `text` odgovora ili testira upit. Ne menjajte argumente koji se prosleđuju assert funkcijama.

Nije potrebno postojanje pretrage za URL. Bez upita/pretrage imena URL, sistem vraća odgovor od `hello Guest`.

# --hints--

Sve probne stvari moraju proći

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Trebalo bi da testiraš za `res.status` == 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Trebalo bi da testiraš za `res.text` == `'hello Guest'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /('|")hello Guest\1/);
```
