---
id: 587d824c367417b2b2512c4d
title: Uporedite karakteristike dva elementa
challengeType: 2
forumTopicId: 301588
dashedName: compare-the-properties-of-two-elements
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

U okviru `tests/1_unit-tests.js`, pod testom sa oznakom `#8` u paketu `Comparisons`, promenite svaki `assert` u `assert.isAbove` ili `assert.isAtMost` kako bi se test prošao (trebalo bi da se proceni kao `true`). Ne menjajte argumente koji se prosleđuju za potvrdu.

# --hints--

Svi testovi moraju proći.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Treba da izaberete ispravnu metodu za prvu verifikaciju - `isAbove` naspram `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtMost',
  '5 is at most (<=) 5'
);
```

Treba da izaberete ispravan način za dvostruku verifikaciju - `isAbove` naspram `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'isAbove', '1 is greater than 0');
```

Trebalo bi da izaberete ispravan put za treću verifikaciju - `isAbove` naspram `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'isAbove',
  'Math.PI = 3.14159265 is greater than 3'
);
```

Trebalo bi da izaberete ispravan način za četvrti identifikator - `isAbove` naspram `isAtMost`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=7');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isAtMost',
  '1 - Math.random() is > 0 and <= 1. It is atMost 1 !'
);
```
