---
id: 587d824c367417b2b2512c4f
title: Pokušajte ako je vrednost unutar specifičnog raspona
challengeType: 2
forumTopicId: 301598
dashedName: test-if-a-value-falls-within-a-specific-range
---

# --description--

Kao podsetnik, ovaj projekat je izgrađen na osnovu osnovnog projekta preuzetog iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

```javascript
.approximately(actual, expected, delta, [message])
```

Potvrđuje se da je `actual` jednako `expected`, unutar opsega +/- `delta`.

# --instructions--

Unutar `tests/1_unit-tests.js` ispod testa sa oznakom `#10` u redu `Comparisons`, promeni svaki `assert` u `assert.approximately` da bi test prošao (treba da se proceni kao `true`).

Izaberite raznolikost niske razine (parametar 3) kako bi eksperiment uvek uspeo. Trebalo bi da bude ispod 1.

# --hints--

Sve testiranje mora proći.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Treba da izaberete odgovarajuću raznolikost za prvu verifikaciju - `approximately(actual, expected, range)`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'approximately');
assert.equal(
  data.assertions[0].args[2],
  0.5,
  "weirdNumbers(0.5) is in the range (0.5, 1.5]. It's within 1 +/- 0.5"
);
```

Treba da izaberete odgovarajući tip za drugu verifikaciju - `approximately(actual, expected, range)`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'approximately');
assert.equal(
  data.assertions[1].args[2],
  0.8,
  "weirdNumbers(0.2) is in the range (0.2, 1.2]. It's within 1 +/- 0.8"
);
```
