---
id: 587d824d367417b2b2512c53
title: Proveri da li sekvenca karaktera ima podsekvencu
challengeType: 2
forumTopicId: 301597
dashedName: test-if-a-string-contains-a-substring
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`include()` na `notInclude()` takođe funkcioniše sa nizovima karaktera! `include()` Potvrđuje da stvarna sekvenca ima očekivanu podsekvencu.

# --instructions--

Unutar ``tests/1_unit-tests.js`` ispod testa sa oznakom ``#14`` u redu ``Strings``, promenite svaki ``assert`` u ``assert.include`` ili ``assert.notInclude`` kako bi test prošao (treba biti ocenjen kao ``true``). Ne menjajte argumente koji se šalju za potvrdu.

# --hints--

Sve provere moraju da prođu.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Trebalo bi da izaberete ispravan put za prvu verifikaciju - `include` naspram `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'include',
  "'Arrow' contains 'row'..."
);
```

Trebalo bi da izaberete ispravan način za sekundnu verifikaciju - `include` naspram `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=13');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notInclude',
  "... a 'dart' doesn't contain a 'queue'"
);
```
