---
id: 587d824b367417b2b2512c4b
title: Koristite tri znakovi jednakosti da potvrdite strogo jednakost
challengeType: 2
forumTopicId: 301610
dashedName: use-the-triple-equals-to-assert-strict-equality
---

# --description--

Kao podsetnik, ovaj projekat je izgrađen na početni projekat preuzet iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`strictEqual()` upoređuje stavke koristeći `===`.

# --instructions--

Unutar `tests/1_unit-tests.js` pod testom sa oznakom `#6` u suiti `Equality`, promenite svako `assert` u `assert.strictEqual` ili `assert.notStrictEqual` kako bi test prošao (trebalo bi da se evaluira kao `true`). Ne menjajte argumente koji se prosleđuju za potvrdu.

# --hints--

Svi testovi moraju da prođu.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Treba da izaberete ispravan put za prvu verifikaciju - `strictEqual` naspram `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'notStrictEqual',
    'with strictEqual the type must match'
  );
```

Trebalo bi da izaberete ispravan način za dvostruku verifikaciju - `strictEqual` naspram `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'strictEqual', '3*2 = 6...');
```

Treba da izaberete ispravan put za treću verifikaciju – `strictEqual` naspram `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'strictEqual',
    "6 * '2' is 12. Types match !"
  );
```

Trebalo bi da izaberete ispravan put za četvrti identifikator - `strictEqual` naspram `notStrictEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=5');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[3].method,
    'notStrictEqual',
    'Even if they have the same elements, the Arrays are notStrictEqual'
  );
```
