---
id: 587d824f367417b2b2512c59
title: Pokreni funkcionalne testove na API endpointima koristeći Chai-HTTP II
challengeType: 2
forumTopicId: 301592
dashedName: run-functional-tests-on-api-endpoints-using-chai-http-ii
---

# --description--

Kao podsjetnik, ovaj projekat je izgrađen na osnovu baznog projekta preuzet od <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

Unutar `tests/2_functional-tests.js`, promenite test `'Test GET /hello with your name'` (`// #2`) da potvrdite `status` i `text` odgovora, kako bi test propao.

Pošaljite svoje ime kao parametar pretrage URL-a dodavanjem `?name=<your_name>` na putanju podataka. Endpoint vraća odgovor za `'hello <your_name>'`.

# --hints--

Sve probne stvari moraju proći

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Trebalo bi da pokušate sa `res.status` == 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Treba ti da pokušaš sa `res.text` == `'hello <your_name>'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=1');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /hello [\w\d_-]/);
```
