---
id: 587d824d367417b2b2512c52
title: Proveri da li je vrednost niz slova
challengeType: 2
forumTopicId: 301599
dashedName: test-if-a-value-is-a-string
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`isString` i `isNotString` potvrđuje da stvarna vrednost je niz slova.

# --instructions--

Unutar ``tests/1_unit-tests.js`` pod testom sa oznakom ``#13`` u redu od ``Strings``, promenite svaki ``assert`` u ``assert.isString`` ili ``assert.isNotString`` da bi test prošao (treba da se evaluira kao ``true``). Ne menjajte argumente koji su prosleđeni za potvrdu.

# --hints--

Sve provere moraju da prođu.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Treba da izaberete ispravan način za prvu verifikaciju – `isString` naspram `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isNotString',
  'A float number is not a string'
);
```

Treba da izaberete ispravan način za drugu verifikaciju - `isString` naspram `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isString',
  'environment vars are strings (or undefined)'
);
```

Treba da izaberete ispravan put za treću verifikaciju – `isString` naspram `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isString', 'A JSON is a string');
```
