---
id: 587d824a367417b2b2512c46
title: Naučite kako funkcionišu validacije za JavaScript
challengeType: 2
forumTopicId: 301589
dashedName: learn-how-javascript-assertions-work
---

# --description--

Rad na ovim izazovima će zahtevati pisanje vašeg koda korišćenjem jednog od sledećih metoda:

- Kopiraj <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">ovaj repozitorijum od GitHub</a> i završi ove izazove interno.
- Koristi graditelj sajta koji ti se sviđa za završetak projekta. Uverite se da uključujete sve fajlove iz našeg repozitorijuma od GitHub.

# --instructions--

U okviru `tests/1_unit-tests.js`, pod testom sa oznakom `#1` u paketu (ili: *sui*) `Basic Assertions`, promenite svaki `assert` u `assert.isNull` ili `assert.isNotNull` kako bi test prošao (treba da bude procenjeno kao `true`). Ne menjajte argumente koji se prosleđuju za potvrdu.

# --hints--

Sve provere moraju da prođu.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Trebalo bi da izaberete ispravan način za prvu verifikaciju - `isNull` naspram `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNull', 'Null is null');
```

Treba da izaberete ispravnu metodu za sekundnu verifikaciju - `isNull` naspram `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isNotNull', '1 is not null');
```
