---
id: 587d824b367417b2b2512c4a
title: Koristite dva znaka jednakosti da potvrdite jednakost
challengeType: 2
forumTopicId: 301609
dashedName: use-the-double-equals-to-assert-equality
---

# --description--

Kao podsjetnik, ovaj projekat je izgrađen na osnovu originalnog projekta preuzetog iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`equal()` upoređuje stavke koristeći `==`.

# --instructions--

Unutar `tests/1_unit-tests.js` ispod testa sa oznakom `#5` u kolekciji `Equality`, promenite svako `assert` u `assert.equal` ili `assert.notEqual` kako bi test prošao (treba da se proceni kao `true`). Ne menjajte argumente koji se prosleđuju za validaciju.

# --hints--

Sve testiranje mora da prođe.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Treba da izaberete ispravan način za prvu autentifikaciju - `equal` protiv `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'equal',
    'Numbers are coerced into strings with == '
  );
```

Treba da izaberete ispravnu metodu za drugu verifikaciju - `equal` naspram `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notEqual',
    ' == compares object references'
  );
```

Trebalo bi da izaberete ispravan način za treću verifikaciju - `equal` naspram `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'equal',
    "6 * '2' is 12 ! It should be equal to '12'"
  );
```

Treba da izaberete ispravnu metodu za četvorostruku autentifikaciju - `equal` naspram `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[3].method, 'notEqual', "6 + '2' is '62'...");
```
