---
id: 587d824e367417b2b2512c56
title: Proverite da li je vrednost određeni tip strukture podataka
challengeType: 2
forumTopicId: 301601
dashedName: test-if-a-value-is-of-a-specific-data-structure-type
---

# --description--

Kao podsetnik, ovaj projekat je baziran na početnom projektu uzet iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`#typeOf` potvrđuje da je tip vrednosti niz karaktera koji je dostavljen, kako određeno od strane `Object.prototype.toString`.

# --instructions--

Unutar `tests/1_unit-tests.js` pod testom sa oznakom `#17` u redu `Objects`, promenite svaki `assert` u `assert.typeOf` ili `assert.notTypeOf` kako bi se test uspešno izvršio (treba biti ocenjeno kao `true`). Ne menjajte argumente koji se prosleđuju za potvrdu.

# --hints--

Sve testiranje mora proći.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Trebalo bi da izaberete ispravan način za prvu verifikaciju – `typeOf` naspram `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'typeOf',
  'myCar is typeOf Object'
);
```

Morate odabrati ispravan način za sekundarnu verifikaciju - `typeOf` naspram `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'typeOf',
  'Car.model is a String'
);
```

Trebalo bi da izaberete ispravan put za treću verifikaciju – `typeOf` naspram `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'notTypeOf',
  'Plane.wings is a Number (not a String)'
);
```

Treba da izaberete ispravan način za četvrti identifikator – `typeOf` naspram `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'typeOf',
  'Plane.engines is an Array'
);
```

Treba da izaberete ispravan metod za petu verifikaciju – `typeOf` naspram `notTypeOf`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=16');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[4].method,
  'typeOf',
  'Car.wheels is a Number'
);
```
