---
id: 587d824c367417b2b2512c4c
title: Proveri duboku jednakost sa .deepEqual i .notDeepEqual
challengeType: 2
forumTopicId: 301587
dashedName: assert-deep-equality-with--deepequal-and--notdeepequal
---

# --description--

Kao podsetnik, ovaj projekat je baziran na početnom projektu uzet iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`deepEqual()` Potvrđuje da dve stavke imaju duboku jednakost.

# --instructions--

Unutar `tests/1_unit-tests.js` pod testom sa oznakom `#7` u paketu/suiti `Equality`, promenite svako `assert` u `assert.deepEqual` ili `assert.notDeepEqual` kako bi test prošao (treba da se proceni kao `true`). Ne menjajte argumente koji se prosleđuju za potvrdu.

# --hints--

Sve testiranje mora proći.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
})();
```

Trebalo bi da izaberete ispravan način za prvu verifikaciju - `deepEqual` naspram `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'deepEqual',
    'The order of the keys does not matter'
  );
})();
```

Trebalo bi da izaberete ispravan način za drugu verifikaciju - `deepEqual` naspram `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notDeepEqual',
    'The position of elements within an array does matter'
  );
})();
```
