---
id: 587d824b367417b2b2512c49
title: Pokušaj utvrđivanja istine
challengeType: 2
forumTopicId: 301596
dashedName: test-for-truthiness
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`isTrue()` pokuša vrednost tačno/netačno `true` i `isNotTrue()` prolazi kada im se da bilo šta osim vrednosti tačno ili netačno od `true`.

```js
assert.isTrue(true, 'This will pass with the boolean value true');
assert.isTrue('true', 'This will NOT pass with the string value "true"');
assert.isTrue(1, 'This will NOT pass with the number value 1');
```

`isFalse()` na `isNotFalse()` takođe su dostupni, i rade kao prave (ili stvarne), osim što traže stvarnu ili ne-stvarnu vrednost od `false`.

# --instructions--

U okviru `tests/1_unit-tests.js`, unutar testa sa oznakom `#4` u suiti `Basic Assertions`, promenite svaki `assert` u `assert.isTrue` ili `assert.isNotTrue` kako bi test prošao (trebalo bi da se proceni kao `true`). Ne menjajte argumente poslati za potvrdu.

# --hints--

Sva iskušenja moraju da prođu.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Treba da izaberete ispravan način za prvu verifikaciju – `isTrue` naspram `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isTrue', 'True is true');
```

Treba da izaberete ispravan način za drugu verifikaciju – `isTrue` naspram `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isTrue',
    'Double negation of a truthy value is true'
  );
```

Treba da izaberete ispravan način za treću verifikaciju – `isTrue` naspram `isNotTrue`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=3');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isNotTrue',
    'A truthy object is not true - neither is a false one'
  );
```
