---
id: 587d824f367417b2b2512c5b
title: Izvrši funkcionalno testiranje na odgovoru od API-ja koristeći Chai-HTTP IV - metod PUT
challengeType: 2
forumTopicId: 301591
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iv---put-method
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Ovaj zadatak je isti kao prethodni.

Sada znaš kako da testiraš zahtev za `PUT`; tvoj je red da ga uradiš od početka.

# --instructions--

Unutar `tests/2_functional-tests.js`, promeni pokušaj za `'Send {surname: "da Verrazzano"}'` (`// #4`) da koristi metode od `put` i `send` za testiranje endpointa `'/travellers'`.

Pošalji sledeći JSON zajedno sa svojim PUT zahtevom:

```json
{
  "surname": "da Verrazzano"
}
```

Proverite sledeće u pozivu povratka za `request.end`:

1. `status` trebalo bi biti `200`
2. `type` trebalo bi biti `application/json`
3. `body.name` trebalo bi biti `Giovanni`
4. `body.surname` trebalo bi biti `da Verrazzano`

Prati podešavanja verifikacije iznad – oslanjamo se na to. Takođe, uverite se da uklonite `assert.fail()` čim završite.

# --hints--

Svi testovi moraju proći

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Trebalo bi da isprobaš `res.status` da li je 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Trebalo bi da isprobaš `res.type` sa `'application/json'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

Trebalo bi da isprobaš `res.body.name` da bude `'Giovanni'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Giovanni\1/);
```

Trebalo bi da isprobate `res.body.surname` sa `'da Verrazzano'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")da Verrazzano\1/);
```
