---
id: 587d824d367417b2b2512c54
title: Koristite običan jezik za testiranje sekvence slova
challengeType: 2
forumTopicId: 301608
dashedName: use-regular-expressions-to-test-a-string
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`match()` Potvrđuje da stvarna vrednost odgovara standardnom izrazu drugog argumenta.

# --instructions--

Unutar `tests/1_unit-tests.js` ispod testa sa oznakom `#15` u redu `Strings`, promenite svaki `assert` u `assert.match` ili `assert.notMatch` kako bi test prošao (trebalo bi da se proceni kao `true`). Ne menjajte argumente koji se prosleđuju za potvrdu.

# --hints--

Svi testovi moraju da prođu.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Trebalo bi da izaberete ispravan način za prvu verifikaciju – `match` naspram `notMatch`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'match',
  "'# name:John Doe, age:35' matches the regex"
);
```

Treba da izaberete ispravnu metodu za sekundarnu verifikaciju - `match` naspram `notMatch`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notMatch',
  "'# name:Paul Smith III, age:twenty-four' does not match the regex (the age must be numeric)"
);
```
