---
id: 587d824d367417b2b2512c51
title: Proverite da li red podataka sadrži element
challengeType: 2
forumTopicId: 301603
dashedName: test-if-an-array-contains-an-item
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

Unutar `tests/1_unit-tests.js` ispod testa sa oznakom `#12` u redu `Arrays`, promenite svaki `assert` u `assert.include` ili `assert.notInclude` kako bi test prošao (trebalo bi da evaluira na `true`). Ne menjajte argumente poslati za potvrdu.

# --hints--

Sve provere moraju da prođu.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Trebalo bi da izaberete ispravan put za prvu verifikaciju - `include` naspram `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInclude',
  "It's summer in july..."
);
```

Trebalo bi da izaberete ispravan način za dvostruku verifikaciju - `include` naspram `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'include',
  'JavaScript is a back-end language !!'
);
```
