---
id: 587d824e367417b2b2512c55
title: Proveri kao da nešto ima osobinu
challengeType: 2
forumTopicId: 301604
dashedName: test-if-an-object-has-a-property
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`property` dokazuje da stvarna stvar ima određene karakteristike.

# --instructions--

Unutar `tests/1_unit-tests.js`, pod testom sa oznakom `#16` u paketu `Objects`, promenite svaki `assert` u `assert.property` ili `assert.notProperty` kako bi test prošao (trebalo bi da se proceni kao `true`). Ne menjajte argumente koji su prosleđeni za potvrdu.

# --hints--

Svi testovi moraju da prođu.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Treba da izaberete ispravan način za prvu verifikaciju - `property` naspram `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notProperty',
  'A car has not wings'
);
```

Treba da izaberete ispravan način za drugu verifikaciju - `property` naspram `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'property',
  'planes have engines'
);
```

Treba da izaberete ispravan način za treću verifikaciju - `property` naspram `notProperty`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=15');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'property', 'Cars have wheels');
```
