---
id: 587d824f367417b2b2512c5a
title: Pokreni funkcionalne testove na API odgovoru koristeći Chai-HTTP III - PUT metodu
challengeType: 2
forumTopicId: 301590
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iii---put-method
---

# --description--

Kao podsetnik, ovaj projekat je baziran na početnom projektu uzet iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Kada pokušate da pošaljete zahtev `PUT`, često šaljete i podatke sa njim. Podaci koje unosite zajedno sa vašim zahtevom `PUT` nazivaju se telo zahteva.

Da pošaljete zahtev za `PUT` i nešto sa JSON ka odredištu `'/travellers'`, možete koristiti dodatak `chai-http` i metode `put` i `send`:

```js
chai
  .request(server)
  .keepOpen()
  .put('/travellers')
  .send({
    "surname": [last name of a traveller of the past]
  })
  ...
```

I putanja podataka odgovara sa:

```json
{
  "name": [first name],
  "surname": [last name],
  "dates": [birth - death years]
}
```

Pogledajte server kod za različite odgovore u segmentu `'/travellers'`.

# --instructions--

Unutar `tests/2_functional-tests.js`, promeni pokušaj za `'Send {surname: "Colombo"}'` (`// #3`) i koristi metode od `put` i `send` da testira sufiks/završetak `'/travellers'`.

Pošalji sledeći objekat JSON zajedno sa vašim PUT zahtevom:

```json
{
  "surname": "Colombo"
}
```

Pogledajte sljedeće unutar poziva za povratak `request.end`:

1. `status` treba da bude `200`
2. `type` treba da bude `application/json`
3. `body.name` treba da bude `Cristoforo`
4. `body.surname` treba da bude `Colombo`

Prati podešavanje potvrde gore – računamo na to. Takođe, obavezno ukloni `assert.fail()` kada završiš.

# --hints--

Sve pokušaje treba da prođe.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Trebalo bi da pokušaš da je `res.status` 200.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Trebalo bi da probaš `res.type` da bude `'application/json'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

Trebalo bi da pokušate `res.body.name` da bude `'Cristoforo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Cristoforo\1/);
```

Trebalo bi da pokušaš da bude `res.body.surname` i `'Colombo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")Colombo\1/);
```
