---
id: 587d824b367417b2b2512c48
title: Koristi Assert.isOK i Assert.isNotOK
challengeType: 2
forumTopicId: 301607
dashedName: use-assert-isok-and-assert-isnotok
---

# --description--

Kao podsetnik, ovaj projekat je baziran na početnom projektu uzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`isOk()` meri stvarna vrednost, i `isNotOk()` meri lažna vrednost.

Da biste saznali više o istinitim i lažnim vrednostima, probajte naš izazov <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-algorithm-scripting/falsy-bouncer" target="_blank" rel="noopener noreferrer nofollow">Falsy Bouncer</a>.

# --instructions--

Unutar `tests/1_unit-tests.js`, pod testom sa labelom `#3` u suiti `Basic Assertions`, promenite svaki `assert` u `assert.isOk()` ili `assert.isNotOk()` kako bi test prošao (trebalo bi da se evaluira kao `true`). Ne menjajte argumente koji se prosleđuju za potvrdu.

# --hints--

Sve provere moraju da prođu.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Morate izabrati ispravnu metodu za prvu autentifikaciju - `isOk` naspram `isNotOk`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNotOk', 'Null is falsy');
```

Treba da izaberete ispravan način za sekundnu verifikaciju – `isOk` naspram `isNotOk`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isOk', 'A string is truthy');
```

Treba da izaberete ispravan put za treću verifikaciju – `isOk` naspram `isNotOk`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[2].method, 'isOk', 'true is truthy');
```
