---
id: 5f8884f4c46685731aabfc41
title: Pokrenite funkcionalna testiranja koristeći Headless Browser II
challengeType: 2
forumTopicId: 301594
dashedName: run-functional-tests-using-a-headless-browser-ii
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

Unutar `tests/2_functional-tests.js`, tokom testa `'Submit the surname "Vespucci" in the HTML form'` (`// #6`), automatski izvrši sledeće:

1. Popunite formular sa imenom plemena `Vespucci`
2. Kliknite na dugme za slanje

Unutar poziva za vraćanje `pressButton`:

1. Potvrdite da je status OK `200`
2. Potvrdite da je tekst unutar elementa `span#name` je `'Amerigo'`
3. Potvrdite da je tekst unutar elementa `span#surname` je `'Vespucci'`
4. Potvrdite da element(element) `span#dates` postoji i da je njegov broj `1`

Ne zaboravi da ukloniš poziv `assert.fail()`.

# --hints--

Svi testovi moraju proći.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Trebalo bi da potvrdite da je zahtev iz headless pretraživača uspeo.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

Trebalo bi da potvrdite da je tekst unutar elementa `span#name` na mestu `'Amerigo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Amerigo\1/);
```

Trebalo bi da potvrdite da je tekst unutar elementa `span#surname` na mestu `'Vespucci'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Vespucci\1/);
```

Trebalo bi da potvrdite da element `span#dates` postoji i da je njegov broj 1.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```
