---
id: 669fcb5106eeb13ab7afcd09
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, naučićete se da radite sa semantic HTML kreiranjem posvećene stranice za blog za gospodina Whiskersa mačka.

Prvi deo koji ćeš napraviti je <dfn>zaglavlje</dfn> stranice.

Element ``header`` se koristi za predstavljanje uvodnog sadržaja, kao što su navigacija po stranici i druge uvodne informacije.

Ovde je primer koji koristi element `header`:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

Unutar elementa `body`, dodaj element `header`.

# --hints--

Mora da ima etiketu pakovanja za `header`.

```js
assert.match(code, /<header>/i);
```

Trebalo bi da imaš oznaku za zatvaranje `header`.

```js
assert.match(code, /<\/header>/i);
```

Tvoj element za `header` mora biti unutar elementa `body`.

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
