---
id: 669fdfc9a5e3da42d2376609
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Sada si dodao sekciju o nama, pokušaj da klikneš na link `About` da vidiš kako stranica skokom ide do tog dela.

Sledeći deo na blog stranici biće lista članaka koji govore o gospodinu Whiskers-u.

Dodaj drugi element od `section` koji ima `id` i koji je postavljen na `"posts"`.

Unutar elementa `section`, dodaj element `h2` koji sadrži tekst `Posts`.

# --hints--

Tvoj kod treba da ima oznaku otvaranja `section`.

```js
assert.match(code, /<section.*>/i);
```

Vaš kod bi trebalo da ima zatvarajući tag od `section`.

```js
assert.match(code, /<\/section>/i);
```

Element vašeg `section` treba da ima `id` postavljen na `"posts"`.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

Element vašeg `section` treba biti unutar elementa vašeg `main`.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Tvoj element za `h2` treba da bude unutar elementa `section`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

Vaš element za `h2` treba da sadrži tekst od `Posts`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
