---
id: 669fd6fd12918e3de87854d4
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Na vašem blogu, treba da postoji način za korisnike da navigiraju različitim sekcijama na stranici.

Element ``nav`` se koristi za pružanje navigacionih linkova ka drugom mestu u dokumentu ili ka drugom delu na veb stranici. Često ćete videti element ``nav`` koji se koristi za meni ili sadržajnu tabelu.

Ovo je primer korišćenja funkcije `nav`:

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

Ispod vašeg elementa za `figure`, dodajte element `nav` sa elementom `ul` koji je interno podešen.

Unutar elementa `ul`, dodaj tri elementa od `li`.

# --hints--

Vaša komponenta `nav` treba da ima otvarajući label/tag.

```js
assert.match(code, /<nav\>/);
```

Vaša komponenta za `nav` mora imati zatvarajući tag.

```js
assert.match(code, /<\/nav\>/);
```

Vaš element `nav` mora biti unutar vašeg elementa `header`.

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

Vaš komponenta za `ul` treba da ima etiketu za otvaranje.

```js
assert.match(code, /<ul\>/);
```

Tvoj element `ul` treba da ima zatvarajući tag.

```js
assert.match(code, /<\/ul\>/);
```

Element vašeg `ul` mora biti unutar elementa vašeg `nav`.

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

Vaša komponenta za `ul` mora da ima tri elementa od `li`.

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
