---
id: 669fd003cb89ee3c2402e041
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

U ovom uvodičnom sadržaju, treba da prikažete sliku gospodina Whiskersa zajedno sa opisom slike.

Ispod elementa `h1`, počnite dodavanjem elementa `figure`.

Unutar elementa `figure`, dodajte element `img`.

Svojstvo elementa `src` od `img` treba da ima vrednost `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`, a tekst/string `alt` treba da ima vrednost `"a cat in the garden"`.

Ispod vašeg elementa sa `img`, dodajte `figcaption` sa tekstom `Mr. Whiskers in the Garden`.

# --hints--

Tvoj element za `figure` treba da ima otvorenu oznaku.

```js
assert.match(code, /<figure\>/);
```

Vaš element `figure` treba da ima zatvarački tag.

```js
assert.match(code, /<\/figure\>/);
```

Tvoj element `figure` treba da se nalazi unutar elementa `header`.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

Trebalo bi da imaš komponentu `img`.

```js
assert.exists(document.querySelector("img"));
```

Vaš element `img` treba da bude unutar vašeg elementa `figure`.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

Vaš element za `img` mora imati svojstvo elementa `src` postavljeno na `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Vaš element za `img` mora imati svojstvo elementa `alt`, postavljeno na `"a cat in the garden"`.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

Vaš element `figcaption` trebao bi imati oznaku za otvaranje.

```js
assert.match(code, /<figcaption\>/);
```

Tvoj komponent za `figcaption` treba da ima zatvarajući tag/label.

```js
assert.match(code, /<\/figcaption\>/);
```

Element vašeg `figcaption` treba da bude unutar elementa vašeg `figure`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

Tvoj element `figcaption` treba da sadrži tekst `Mr. Whiskers in the Garden`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

Element `figcaption` treba da bude ispod elementa `img`. Trenutno su postavljeni pogrešnim redosledom.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
