---
id: 669fde2081f65141ad703fe4
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Prvi deo na stranici će biti sekcija "O nama". Taj deo će predstaviti gospodina Whiskersa i dati korisnicima ideju o tome o čemu se radi na ovom blogu.

Unutar tvog elementa `main`, dodaj element `section` koji ima svojstvo elementa `id` postavljeno na `"about"`.

Unutar elementa `section`, dodaj `h2` sa tekstom `About`.

# --hints--

Tvoj element `section` treba da ima otvarajući tag.

```js
assert.match(code, /<section.*>/);
```

Komponenta vašeg `section` mora imati zatvarajući tag.

```js
assert.match(code, /<\/section\>/);
```

Vaš element za `section` mora imati svojstvo elementa `id` postavljeno na `"about"`.

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

Vaš element `section` treba da bude unutar vašeg elementa `main`.

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Tvoja komponenta od `h2` mora imati etiketu za otvaranje.

```js
assert.match(code, /<h2\>/);
```

Vaša funkcija za `h2` treba da ima zatvarajući atribut.

```js
assert.match(code, /<\/h2\>/);
```

Tvoj element `h2` treba da sadrži tekst `About`. Proveri pravopis.

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

Vaša komponenta za `h2` mora da bude unutar vaše komponente `section`.

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
