---
id: 669fdc11c9b0263fe0814a7a
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Unutar svakog elementa od `li`, bit će vam potreban anker element.

Za prvi jangkoranički element, tekst bi trebalo da bude `About` i vrednost atributa za `href` bi trebalo da bude `"#about"`. Hash znak (#) ispred `about` predstavlja ime `id`, koje će biti dodato kasnije u projekat.

Za drugi jangkorni element, tekst treba da bude `Posts` i vrednost atributa od `href` treba da bude `"#posts"`.

Za treći element sidra, tekst treba da bude ``Contact`` i vrednost atributa za ``href`` treba da bude ``"#contact"``.

# --hints--

Tvoj prvi jangkorni element treba da sadrži tekst `About` unutar tvog prvog elementa `li`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

Tvoja prva komponenta sidra mora imati svojstvo `href` postavljeno na `"#about"`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

Vaš drugi anker element bi trebalo da sadrži tekst od `Posts` unutar vašeg drugog elementa od `li`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

Drugi atribut noktića treba da ima svojstvo od `href` postavljeno na `"#posts"`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

Treći element vašeg sidra mora da sadrži tekst `Contact` unutar trećeg elementa `li`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

Tvoj treći atribut anker-a mora imati svojstvo `href` postavljeno na `"#contact"`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
