---
id: 66a334a1a7cca6354999f9bf
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Za prvi blog post, koristićete element <dfn>article</dfn>.

Element za `article` predstavlja nezavisni sadržaj na veb stranici.

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

Ispod elementa `h2`, dodaj element `article`.

Unutar elementa `article`, dodaj element `h3` sa tekstom `Mr. Whiskers' First Day Home`.

Razlog za korišćenje elementa `h3` ovde je održavanje ispravne hijerarhije naslova. Pošto je naslov odeljka sa objavama element `h2`, sledeći nivo u hijerarhiji treba da bude element `h3`.

# --hints--

Tvoj kod mora imati otvarajuću oznaku od `article`.

```js
assert.match(code, /<article>/i);
```

Tvoj kod bi trebalo da ima zatvarajući tag od `article`.

```js
assert.match(code, /<\/article>/i);
```

Vaš element `article` mora biti unutar vašeg elementa `section` i `id`, ukoliko je postavljeno na `posts`.

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

Vaš kod bi trebalo da ima oznaku otvaranja `h3`.

```js
assert.match(code, /<h3>/i);
```

Tvoj kod treba da sadrži zatvarajući tag `h3`.

```js
assert.match(code, /<\/h3>/i);
```

Vaš element za `h3` treba da sadrži tekst `Mr. Whiskers' First Day Home`.

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Day Home");
```

Vaš element `h3` treba da bude unutar Vašeg elementa `article`.

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
