---
id: 66a49f685961e997e337cab1
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

`header` Biće zaduženo za prikaz glavnog naslova, slike i navigacije stranice za blog.

Unutar elementa `header`, dodaj `h1` sa tekstom `Welcome to Mr. Whiskers' Blog Page!`.

# --hints--

Trebalo bi da imate etiketu pakovanja za `h1`.

```js
assert.match(code, /<h1>/i);
```

Trebalo bi da imaš etiketu za zatvaranje `h1`.

```js
assert.match(code, /<\/h1>/i);
```

Vaše polje `h1` mora imati tekst `Welcome to Mr. Whiskers' Blog Page!`. Proverite da li postoji bilo kakva pravopisna greška.

```js
const h1Text = document.querySelector("h1")?.innerText.trim();
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

Element vašeg `h1` treba biti postavljen unutar elementa vašeg `header`.

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
--fcc-editable-region--
      
--fcc-editable-region--
    </header>
  </body>
</html>
```
