---
id: 6391d1a4f7ac71efd0621380
title: Izgradite projekat stranice sa receptima
challengeType: 25
dashedName: top-build-a-recipe-project
---

# --description--

Sajt će sadržavati glavni demonstrativni deo koji će imati linkove ka nekim receptima. Sajt neće imati veoma dobar izgled nakon završetka.

**Priče korisnika:**

1. Vaša stranica recepta treba da ima oznaku `DOCTYPE`.
1. Vaša stranica recepta bi trebalo da sadrži element sa `html` koji ima elemente od `head` i `body` kao decu.
1. Trebalo bi da imate element `title` unutar elementa sa `head` koji ima tekst `The Odin Recipes`.
1. Trebalo bi da vidite element `h1` sa tekstom `Creamy Chocolate Fudge`.
1. Treba da postoji slika povezana sa atributom `alt`.
1. Ispod slike treba da postoji element `h2` sa tekstom `Description`.
1. Trebalo bi da vidite nekoliko pasusa ispod `Description` koji opisuju recept.
1. Treba da postoji element `h2` sa tekstom `Ingredients`.
1. Ispod naslova sekcije `Ingredients` treba da postoji neuređeni spisak sastojaka potrebnih za recept.
1. Ispod spiska sastojaka dodajte još jedan naslov sekcije pod nazivom `Steps`.
1. Trebalo bi da vidite uređeni spisak sa nekoliko koraka potrebnih za završetak recepta.
1. Ispod koraka treba da postoji element `h2` sa tekstom `More Recipes`.
1. Trebalo bi da vidite dva sastojka iz drugog recepta unutar neuređenog spiska koji sadrži dva stavke liste sa ankor elementima.
1. Ovi anker elementi treba da imaju atribut `href` sa vrednošću `#`.

# --hints--

Trebalo bi da ima oznaku `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+?html\s*?>/gi));
```

Trebalo bi da imaš element sa karakteristikama `html`, `head` i `body`.

```js
const html = document.querySelectorAll('html')[0];
const head = document.querySelectorAll('html > head')[0];
const body = document.querySelectorAll('html > body')[0];

assert(html && head && body);
```

Trebalo bi da imaš element `title` unutar elementa `head` sa tekstom `The Odin Recipes`.

```js
assert(document.querySelectorAll('HEAD > TITLE')[0]?.innerText == 'The Odin Recipes');
```

Trebalo bi da imaš element `h1` unutar svog elementa `body` sa tekstom `Creamy Chocolate Fudge`.

```js
assert(document.querySelectorAll('BODY > H1')[0]?.innerText == 'Creamy Chocolate Fudge');
```

Trebalo bi da imaš sliku sa karakteristikom `alt`.

```js
const img = document.querySelectorAll('IMG')[0];

assert(img && img.alt !='' && img.src != '')
```

Trebalo bi da imaš element sa tekstom `h2` koji sadrži tekst `Description`.

```js
const h2 = document.querySelectorAll('H2')[0];

assert(h2?.innerText == 'Description');
```

Trebalo bi da ima barem dva elementa od `p` koja opisuju recept.

```js
const paragraphs = document.querySelectorAll('P');

assert(paragraphs?.length > 1);
```

Trebalo bi da imaš element sa tekstom `h2` koji sadrži `Ingredients`.

```js
const h2 = document.querySelectorAll('H2')[1];

assert(h2?.innerText == 'Ingredients');
```

Trebalo bi da ima element `ul` ugrađen unutar elemenata `li` koji prikazuju linkove.

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "Ingredients") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);
        } else {
            assert(false);
        }
    }
});
```

Trebalo bi da imaš element sa tekstom `h2` koji sadrži tekst `Steps`.

```js
const h2 = document.querySelectorAll('H2')[2];

assert(h2?.innerText == 'Steps');
```

Trebalo bi da imaš element sa `ol` koji sadrži elemente `li` koji prikazuju korake recepta.

```js
const orderedList = document.querySelectorAll('OL')[0];
const listItems = document.querySelectorAll('OL > LI');

assert(orderedList && listItems && listItems.length > 1);
```

Trebalo bi da imaš element sa tekstom `h2` koji sadrži `More Recipes`.

```js
const h2 = document.querySelectorAll('H2')[3];

assert(h2?.innerText == 'More Recipes');
```

Trebalo bi imati funkciju `ul`, integrisanu sa funkcijama `li` koje prikazuju karakteristike `a` koje se povezuju sa drugim receptima.

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "More Recipes") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);

            listItems.forEach(listItem => {
                const anchor = listItem.querySelector("A");

                assert(anchor);
            });
        } else {
            assert(false);
        }
    }
});
```

Tvoji anker elementi koji se povezuju sa receptima moraju imati atribut `href` sa vrednošću `#`.

```js
const anchorTags = document.querySelectorAll("a");

const allAnchorsHaveHrefHash = [...anchorTags].every(function(anchorTag) {
  return anchorTag.hasAttribute("href") && anchorTag.getAttribute("href") === "#";
});

assert(allAnchorsHaveHrefHash && anchorTags.length > 0); 
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>The Odin Recipes</title>
  </head>
  <body>
    <h1>Creamy Chocolate Fudge</h1>
    <img src="https://cdn.freecodecamp.org/curriculum/odin-project/build-a-recipe-page/build-a-recipe-page-01.jpg" alt="A delicious chocolate fudge dessert">
    <h2>Description</h2>
    <p>This recipe is for a rich and creamy chocolate fudge that is sure to satisfy your sweet tooth. It's perfect for a special occasion or as a tasty treat for any time of the year.</p>
    <p>This recipe is easy to follow and only requires a few simple ingredients. With just a few steps, you'll be able to create a delicious dessert that everyone will love.</p>
    <h2>Ingredients</h2>
    <ul>
      <li>1 cup sugar</li>
      <li>1/2 cup unsalted butter</li>
      <li>1/4 cup milk</li>
      <li>1/4 cup cocoa powder</li>
      <li>1/4 cup chocolate chips</li>
      <li>1/4 tsp salt</li>
      <li>1 tsp vanilla extract</li>
    </ul>
    <h2>Steps</h2>
    <ol>
      <li>In a medium saucepan, melt the butter over medium heat.</li>
      <li>Add the sugar, milk, cocoa powder, and salt to the saucepan and stir until well combined.</li>
      <li>Bring the mixture to a boil, stirring constantly, and then reduce the heat to low and simmer for 5 minutes.</li>
      <li>Remove the saucepan from the heat and stir in the chocolate chips and vanilla extract until the chocolate is melted and the mixture is smooth.</li>
      <li>Pour the fudge into a greased 8-inch square pan and let it cool completely before cutting into squares.</li>
    </ol>
    <h2>More Recipes</h2>
    <ul>
      <li><a href="#">Peanut Butter Cookies</a></li>
      <li><a href="#">Lemon Bars</a></li>
      <li><a href="#">Chocolate Chip Pancakes</a></li>
    </ul>
  </body>
</html>
```

```css

```
