---
id: 619665c9abd72906f3ad30f9
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Izgradit ćete srećnog Flappy Pingvina i dalje vežbati CSS transformacije i animacije tokom procesa.

Počnite povezivanjem vašeg fajla za stilove na stranicu.

# --hints--

Vaš kod bi trebalo da ima `link` element.

```js
assert.exists(document.querySelector('link'));
```

Vaš `link` element treba da bude unutar vašeg `head` elementa.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Tvoj element `link` treba da ima atribut `rel` sa vrednošću `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Element vašeg `link` treba da ima atribut `href` sa vrednošću `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
