---
id: 68e707257ea91c2e12d489fb
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, vežbaće kako dodati prilagođene stilove na radio dugmad kreiranjem obrasca za sastanak roditelja i nastavnika. HTML šablon je dostavljen za vas.

Počnite dodavanjem elementa `main` sa klasom nazvanom `container`.

# --hints--

You should have a `main` element.

```js
assert.exists(document.querySelector("main"));
```

Your `main` element should have a class called `container`.

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
