---
id: 68ea8666b8b473455f5b33ef
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Prvi deo obrasca će se fokusirati na informacije o učeniku, kao što su njegovo ime i razred.

Unutar elementa `form`, dodajte element `fieldset`. Unutar tog elementa `fieldset`, dodajte element `legend` sa tekstom `Student Information`. 

Ispod elementa `legend`, dodajte element `label` sa tekstom `Full Name: ` i atributom `for` sa vrednošću `student-name`.

# --hints--

You should have a `fieldset` element.

```js
assert.exists(document.querySelector("fieldset"));
```

Your `fieldset` element should be inside the `form` element.

```js
assert.exists(document.querySelector("form fieldset"));
```

You should have a `legend` element.

```js
assert.exists(document.querySelector("legend"));
```

Your `legend` element should be inside the `fieldset` element.

```js
assert.exists(document.querySelector("fieldset legend"));
```

Your `legend` element should have the text `Student Information`.

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

You should have a `label` element.

```js
assert.exists(document.querySelector("label"));
```

Your `label` element should have the text `Full Name: `. Don't forget the space after the colon.

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

Your `label` element should have a `for` attribute.

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

Your `label` element's `for` attribute should have the value `student-name`.

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

Your `label` element should be below the `legend` element.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
