---
id: 68ea8a5db73b155708992687
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Sledeći korak je dodavanje povezanog elementa `input` za informacije o učeniku.

Počnite dodavanjem elementa `input` sa atributom `type` postavljenim na `"text"`. Zatim dodajte atribut `name` i atribut `id`, oba postavljena na `student-name`.

Nakon toga, dodajte atribut `placeholder` postavljenog na `E.g., Jane Doe`. I konačno, dodajte atribut `required`.

# --hints--

You should have an `input` element.

```js
assert.exists(document.querySelector("fieldset input"));
```

Your `input` element should have a `type` attribute set to `"text"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Tvoj element `input` treba da ima atribut `name` postavljen na `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

Your `input` element should have an `id` attribute set to `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

Your `input` element should have a `placeholder` attribute set to `E.g., Jane Doe`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

Your `input` element should have a `required` attribute.

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
