---
id: bad87fee1348bd9aec908856
title: Etikete za dugmad Bootstrap
challengeType: 0
forumTopicId: 18222
dashedName: label-bootstrap-buttons
---

# --description--

Kao što smo postavili etikete u naše polja, želimo da postavimo etikete i na naša dugmad.

Dodajte tekst svakom vašem elementu `button` koji odgovara njegovoj ID opciji.

# --hints--

Vaš element `button` sa id-om `target1` treba da ima tekst `#target1`.

```js
assert.match(document.querySelector('#target1')?.textContent,/#target1/gi)
```

Vaš element `button` sa id-om `target2` treba da ima tekst `#target2`.

```js
assert.match(document.querySelector('#target2')?.textContent,/#target2/gi)
```

Vaš element `button` sa id-om `target3` treba da ima tekst `#target3`.

```js
assert.match(document.querySelector('#target3')?.textContent,/#target3/gi)
```

Vaš element `button` sa id-om `target4` treba da ima tekst `#target4`.

```js
assert.match(document.querySelector('#target4')?.textContent,/#target4/gi)
```

Vaš element `button` sa id-om `target5` treba da ima tekst `#target5`.

```js
assert.match(document.querySelector('#target5')?.textContent,/#target5/gi)
```

Vaš element `button` sa id-om `target6` treba da ima tekst `#target6`.

```js
assert.match(document.querySelector('#target6')?.textContent,/#target6/gi)
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1"></button>
        <button class="btn btn-default target" id="target2"></button>
        <button class="btn btn-default target" id="target3"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4"></button>
        <button class="btn btn-default target" id="target5"></button>
        <button class="btn btn-default target" id="target6"></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```