---
id: bad87fee1348bd9aec908746
title: Stavić našu stranicu unutar Bootstrap container-fluid div
challengeType: 0
forumTopicId: 18198
dashedName: house-our-page-within-a-bootstrap-container-fluid-div
---

# --description--

Sada osigurajte da se sav sadržaj vaše stranice prilagođava uređaju za mobilni telefon.

Umotajte vaš element `h3` unutar elementa `div` sa klasom `container-fluid`.

# --hints--

Vaš `div` element treba da ima klasu `container-fluid`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container-fluid'));
```

Svaki vaš `div` element mora imati zatvarajući tag.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);

assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

Vaš `h3` element treba da bude smešten unutar `div` elementa.

```js
const divElement = document.querySelector('div');
const divChildren = divElement?.querySelectorAll(`:scope ${'h3'}`)
assert.lengthOf(divChildren,1);
```

# --seed--

## --seed-contents--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```

# --solutions--

```html
<div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
</div>
```