---
id: bad87fee1348bd9aec908846
title: Kreiranje naslova Bootstrap-a
challengeType: 0
forumTopicId: 16812
dashedName: create-a-bootstrap-headline
---

# --description--

Hajde da nešto napravimo od nule kako bismo vežbali naše znanje HTML-a, CSS-a i Bootstrap-a.

Napradićemo jQuery Playground, koji ćemo uskoro koristiti u našim jQuery izazovima.

Da počnemo, kreirajte `h3` element sa tekstom `jQuery Playground`.

Stilizujte svoj `h3` element klasom Bootstrap-a `text-primary`, i centrirajte ga koristeći klasu Bootstrap-a `text-center`.

# --hints--

Trebalo bi da dodate `h3` element na svoju stranicu.

```js
assert.lengthOf(document.querySelectorAll('h3'),1);
```

Vaš `h3` element treba da ima zatvarajući tag.

```js
assert.match(code,/<\/h3>/g);
assert.match(code,/<h3/g);
assert.equal( code.match(/<\/h3>/g).length , code.match(/<h3/g).length);
```

Vaš `h3` element treba da bude stilizovan korišćenjem klase `text-primary`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-primary'));
```

Vaš `h3` element treba da bude centriran korišćenjem klase `text-center`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-center'));
```

Vaš `h3` element treba da ima tekst `jQuery Playground`.

```js
assert.match(document.querySelector('h3')?.textContent, /jquery(\s)+playground/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```