---
id: 6909ed80e1c195f49564edac
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

Na ovoj radionici ćete naučiti kako da radite sa uslovnim izjavama i operatorima poređenja.

Počnite deklarisanjem varijable pod nazivom `truthyOrFalsy` i dodeljujte joj booleanski vrednost `true`.

# --hints--

You should have a variable named `truthyOrFalsy`.

```js
assert.exists(truthyOrFalsy);
```

You should assign `true` to your `truthyOrFalsy` variable.

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
