---
id: 690a51806fb743bc896c3af7
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Kao što ste naučili u prethodnoj lekciji, booleani su tip podataka sa samo vrednostima `true` i `false`.

Truthy vrednost je vrednost koja se evaluira na `true` u booleanskom kontekstu. Primeri truthy vrednosti su:

- Ne-prazni nizovi znakova
- Brojevi različiti od nule
- Booleanski `true`

I druge vrednosti koje ćete naići kasnije.

Možete koristiti funkciju `Boolean()` za proveru truthiness vrednosti. Na primer, `Boolean("Hello World!")` će vratiti `true` jer je `"Hello World!"` truthy.

Sada, pozovite `Boolean(truthyOrFalsy)` i logujte rezultat u konzolu.

# --hints--

You should log `Boolean(truthyOrFalsy)` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
