---
id: 690a562fd611e8155599ce98
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Promenite dodelu za `truthyOrFalsy` postavljanjem njegove vrednosti na `"freeCodeCamp"`. I dalje ćete videti `true` logovan u konzolu jer je neprzizan niz istinitog (truthy) vrednosti.

# --hints--

Your `truthyOrFalsy` variable should have the value `"freeCodeCamp"`.

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
