---
id: 690a0f7a5c22c15222a6aa5c
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Sada kreiraj još jednu `if` deklaraciju koja koristi `isTimmyAGamer` kao uslov. Unutar tela vaše `if` deklaracije, loguj `"Timmy loves to play World of Warcraft."` na konzolu.

Napomena da nećeš videti ništa logovano na konzolu jer je uslov vaše nove `if` deklaracije netačan.

# --hints--

You should have two `if` statements in your code.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

Your second `if` statement should use `isTimmyAGamer` as its condition.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

You should log `"Timmy loves to play World of Warcraft."` to the console within your second `if` statement.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
