---
id: 690a06833d3e83b5226ade78
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Izjavni uslov se koristi za izvršavanje koda u zavisnosti od uslova. Sastoji se od ključne reči ``if``, praćene uslovom unutar zagrada i telom unutar vitičastih zagrada. Evo primera izjave ``if``:

```js
if (condition) {
  console.log("condition is truthy");
}
```

Kada je ``condition`` istinitog vrednosti (ne-prazan niz znakova, ne-nulti broj, ``true``, itd.), izvršava se kod unutar tela izjave ``if``.

Ispod vaše varijable, kreirajte izjavu ``if`` koja koristi ``hasDeveloperJob`` kao svoj uslov. Unutar tela vaše izjave ``if``, logujte ``"Timmy is employed as a developer."`` u konzolu.

# --hints--

You should have an `if` statement.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

Your `if` statement should use `hasDeveloperJob` as its condition.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

You should log `"Timmy is employed as a developer."` to the console within your `if` statement.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;
--fcc-editable-region--

--fcc-editable-region--
```
