---
id: 690a57095c4afb27bbe751b3
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Falsy vrednosti su vrednosti koje se evaluiraju na `false` u boolovskom kontekstu. Primeri falsy vrednosti su:

- `""` (prazan string)
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

Svi ćete naići na sve ovo tokom ovog kursa.

Za sada, promenite dodelu za `truthyOrFalsy` postavljanjem njegove vrednosti na prazan string. Nakon toga, videćete da je `false` logovan u konzolu.

# --hints--

You should assign an empty string to `truthyOrFalsy`.

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
