---
id: 690a3ad44874a928a2fe3077
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Operatori poređenja vam omogućavaju da uporedite vrednosti i napišete složenije uslovne izjave.

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

Na primer, kod iznad loguje `"Happy Birthday!"` na konzolu samo kada su vrednosti promenljivih `today` i `birthday` jednake.

Kreirajte treću `if` izjavu. Za njen uslov, koristite operator `>=` da proverite da li je `timmyAge` veće ili jednako `16`.

Unutar tela vaše nove `if` izjave, logujte `"Timmy is old enough to drive."` na konzolu.

# --hints--

You should have a third `if` statement in your code.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

Your third `if` statement should use `timmyAge >= 16` as its condition.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{.+\}/s);
```

You should log `"Timmy is old enough to drive."` to the console inside your third `if` statement.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

--fcc-editable-region--

--fcc-editable-region--
```
