---
id: 65821c1c40ccac44adf38e6b
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Kreiraj prazan spisak pod nazivom `expenses`. Koristićeš ga za čuvanje svake od tvojih upotreba.

# --hints--

Trebalo bi da objavite novi parametar nazvan `expenses`.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

Treba da postavite vrednost praznog niza u parametar `expenses`.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
