---
id: 65823e87c9741750a22085a7
title: Korak 28
challengeType: 20
dashedName: step-28
---

# --description--

Sledeće, definiši funkciju nazvanu `filter_expenses_by_category` koja prima dva parametra: `expenses` i `category`. Koristi `pass` da popuni telo funkcije.

# --hints--

Trebalo bi da definišete funkciju nazvanu `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(filter_expenses_by_category)
  `))
})
```

Vaša funkcija od `filter_expenses_by_category` treba da primi `expenses` i `category` kao parametre.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(filter_expenses_by_category))
    sig == '(expenses, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
