---
id: 65823ff0d4b991510fade1a8
title: Korak 29
challengeType: 20
dashedName: step-29
---

# --description--

Unutar funkcije `filter_expenses_by_category`, zamenite `pass` sa funkcijom `lambda`. Koristite `expense` kao parametar i procenite poređenje između vrednosti ključa `'category'` iz rečnika `expense` i `category` koristeći operator jednakosti.

# --hints--

Trebali biste kreirati funkciju `lambda` koja koristi parametar nazvan `expense` i koji procenjuje matematički izraz `expense['category'] == category` unutar funkcije `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

Ne sme da ima `pass` unutar funkcije `filter_expenses_by_category`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
