---
id: 65822bd82d708c4895080c35
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

U Python, važno je znati da isti tip navodnika koji se koristi za definisanje sekvence karaktera ne može se koristiti unutar nje. Na primer, sekvenca karaktera `'I'm a string!'` nije validna. Da biste koristili navodnik unutar te sekvence karaktera, morate:

- Izbegavanje tog navodnika postavljanjem kvačice ispred njega: `'I\'m a string!'`
- Ili koristite dvostruke navodnike da definišete niz znakova: `"I'm a string!"` (preporučeno).

Možete dobiti vrednost iz rečnika koristeći njegove ključeve. Morate koristiti notaciju uglastih zagradi i postaviti ključeve unutar kvadratnih zagrada:

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

Trenutno vrši interpolaciju rečnika `expense` u vašem f-stringu. Izmenite izraz u f-stringu da biste dobili vrednost ključa `'amount'` i ključeve `'category'` u rečniku `expense`.

# --hints--

Trebalo bi da prosledite `f'Amount: {expense["amount"]}, Category: {expense["category"]}'` u poziv `print()`. Zapamtite da koristite dvostruke navodne znakove unutar f-stringa sa jednim i suprotnim navodnim znakom.

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
