---
id: 6582324b5d37fe4b622bc3c0
title: Korak 19
challengeType: 20
dashedName: step-19
---

# --description--

Lambda funkcija je kratka, bezimena funkcija u Python, pogodna za jednostavne zadatke koji se izvode samo jednom. Deklariše se glavnom ključnom rečju `lambda` i koristi sledeću sintaksu:

```py
lambda x: expr
```

U navedenom primeru, `x` predstavlja parametar koji se koristi u matematičkom izrazu `expr`, i funkcioniše kao bilo koji parametar u običnoj funkciji. `expr` je matematički izraz koji se procenjuje i vraća kada se pozove lambda funkcija.

Kreirajte parametar nazvan `test` i postavite mu vrednost na lambda funkciju koja prima parametar `x` i vraća `x * 2`.

# --hints--

Trebalo bi da deklarišete parametar koji se zove `test`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

Treba da postaviš vrednost `lambda x: x * 2` za svoj parametar `test`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
