---
id: 65821fcc010c3245718f2a06
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Parametar `expenses` vaše funkcije `add_expense` će biti lista upotrebe. Želite da budete u mogućnosti dodati stavke na kraj vaše liste. Stoga ćete koristiti metodu liste `.append()`.

Dodaj poziv putanje `.append()` u listu `expenses`. Ne prosleđujte bilo koji argument za `.append()` trenutno.

# --hints--

Trebalo bi da imaš `expenses.append()` u svojoj funkciji.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

Ne bi trebalo da imaš `pass` u svojoj funkciji.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
