---
id: 658225d560369446ece5518b
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Dodaj još jedan par ključ-vrednost u rečnik koji dodaje na listu `expense`. Koristi niz znakova `'category'` kao ključ, i parametar `category` kao vrednost.

# --hints--

Trebalo bi dodati drugi par ključ-vrednost u rečnik koji prosleđujete kao argument putem `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict) and len(node[0].tree.keys) == 2
`)) })
```

Trebalo bi da prosledite `{'amount': amount, 'category': category}` kao argument putem `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount, 'category': category})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
