---
id: 65823634a308cb4d6bcc5fa6
title: Korak 21
challengeType: 20
dashedName: step-21
---

# --description--

Lambda funkcije se mogu lako kombinovati sa `map()`, koja izvršava definisanu funkciju za svaki element u kolekciji objekata, kao što je lista:

```py
map(lambda x: x * 2, [1, 2, 3])
```

Funkcija koja se izvršava prosljeđuje se kao prvi argument, a ponovljiva kao drugi argument.

Rezultat gornjeg primera će biti `[2, 4, 6]`, gde je svaki element u listi prosleđenoj na `map()` pomnožen sa dva pomoću lambda funkcije.

Promenite svoj poziv `print()` ili odštampajte rezultate poziva `map()` uz `test` kao prvi argument, i `[2, 3, 5, 8]` kao drugi argument. Još nećete moći da vidite čitljive rezultate.

# --hints--

Trebalo bi da imaš `print(map(test, [2, 3, 5, 8]))` u svom kodu.

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
