---
id: 65822934c8a0454803e3ef12
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Sledeće, će se prikazati detalji za svaku upotrebu.

Unutar petlje od `for`, promeni `pass` na poziv `print()` i koristi sledeći f-string: `f'Amount: {expense}, Category: {expense}'`.

# --hints--

Trebalo bi da ispišete `f'Amount: {expense}, Category: {expense}'` u vašoj `for` petlji/ciklusu.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

Ne bi trebalo da imaš `pass` unutar tela tvog kruga.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
