---
id: 6582267ef39f5047411c1743
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Definišite funkciju nazvanu `print_expenses` koja prima jedan parametar `expenses`. Ova funkcija će kasnije biti korišćena za prikaz svakog troška u vašoj listi.

Popuni telo vaše nove funkcije i izraz za `pass`.

# --hints--

Trebalo bi da definišete novu funkciju koja se zove `print_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

Tvoja funkcija `print_expenses` treba da primi jedan parametar `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
