---
id: 658246c28575d653d1f89b59
title: Korak 33
challengeType: 20
dashedName: step-33
---

# --description--

Petlja za `while` je drugi tip petlje koja izvršava deo koda samo ako je određeni uslov `True`. Petlja se završava kada je uslov `False`:

```py
while condition:
    <code>
```

Ispod liste `expenses`, kreiraj petlju za `while`. Koristi `True` kao uslov i ispiši niz znakova `'\nExpense Tracker'` unutar tela petlje da prikaže zaglavlje programa.

# --hints--

Treba da napraviš ciklus za `while` koristeći `True` kao stanje unutar svoje funkcije `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Treba da ispišete `'\nExpense Tracker'` u vašem `while` ciklusu.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
