---
id: 65822308ded1ba4632f66aa3
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Rečnik je drugi tip podataka ugrađen u Python. Rečnik je kolekcija podataka u strukturi parova ključ-vrednost. Rečnik je definisan vitičastim zagradama (`{}`) i sadrži parove ključ-vrednost razdvojene zarezima. Svaki ključ je praćen dvotačkom (`:`) i vrednošću:

```py
{'amount': 50.0, 'category': 'Food'}
```

U prethodnom primeru, `'amount'` i `'category'` su funkcije, a `50.0` i `'Food'` su njihove odgovarajuće vrednosti.

Ukreiraj rečnik sa ključem `'amount'` i vrednošću parametra `amount`, zatim prosledite svoj novi rečnik pozivom `.append()`.

# --hints--

Treba da proslediš rječnik kao argument na način `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

Trebalo bi da prosledite `{'amount': amount}` kao argument putem `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
