---
id: 65823cfc74aa564ffc460489
title: Korak 26
challengeType: 20
dashedName: step-26
---

# --description--

Sada, izvrši poziv za `map()` koristeći svoju funkciju `lambda` kao prvi argument i listu `expenses` kao drugi argument.

# --hints--

Trebalo bi da pozovete funkciju za `map()` unutar funkcije `total_expenses`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("total_expenses").find_calls("map")) == 1`)) })
```

Treba da prosledite svoju funkciju `lambda` kao prvi argument za poziv `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[0].is_equivalent("lambda expense: expense['amount']")`)) })
```

Trebalo bi da prosledite listu `expenses` kao drugi argument pozivu `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    lambda expense: expense['amount']
--fcc-editable-region--

expenses = []
```
