---
id: 658212ba98182f3e855e85f9
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

U ovom projektu, kreiraćeš jednostavan, ali funkcionalan tracker upotrebe, koristeći Python.

Počnite sa deklarisanjem funkcije nazvane `add_expense` koja prima tri parametra: `expenses`, `amount` i `category`. Koristite ključnu reč `pass` za popunjavanje tela funkcije.

# --hints--

Trebalo bi da koristite glavnu reč `def` za deklarisanje funkcije pod nazivom `add_expense`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

Tvoja funkcija za `add_expense` mora da primi tri parametra: `expenses`, `amount` i `category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
