---
id: 658238f7604f154ea9a23e1e
title: Korak 23
challengeType: 20
dashedName: step-23
---

# --description--

Funkcija ``sum()`` vraća zbir stavki sadržanih u iterabilnom objektu koji mu je prosleđen kao argument. Koristićete ``sum()`` zajedno sa funkcijama ``map()`` i ``lambda`` da biste dobili ukupnu cenu.

Trenutno, napravite mali pokušaj i promenite vaš trenutni poziv ``print()`` zamenom poziva ``list()`` sa pozivom funkcije ``sum()``, prosleđujući trenutni poziv ``map()`` kao argument.

# --hints--

Trebalo bi da odštampate `sum(map(test, [2, 3, 5, 8]))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(list(map(test, [2,3,5,8])))
--fcc-editable-region--

expenses = []
```
