---
id: 6582575b8089f85b8b92d7c8
title: Korak 42
challengeType: 20
dashedName: step-42
---

# --description--

Nakon što dobijete detalje o troškovima, morate pozvati funkciju `add_expense` da dodate novi trošak u listu `expenses`.

Nakon dobijanja `amount` na `category` koristeći `input()`, poziva funkciju `add_expense`, prosleđujući tri argumenta: `expenses`, `amount` i `category`.

- `expenses` je prazna lista kreirana u glavnoj funkciji prvobitno u ovom projektu.
- `amount` je iznos troška.
- `category` je kategorija troška.

# --hints--

Trebaće ti da pozoveš funkciju `add_expense()` unutar tvog deklaracije `if`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_calls("add_expense")) == 1`)) })
```

Trebalo bi da prođeš tri tačke za svoj poziv ka `add_expense()`: `expenses`, `amount`, na `category`. Redosled je važan.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("add_expense(expenses, amount, category)")`)) })
```

Vaš poziv za `add_expense()` treba da dođe nakon podele od `amount` i `category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("amount = float(input('Enter amount: '))\\ncategory = input('Enter category: ')\\nadd_expense(expenses, amount, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            
--fcc-editable-region--
```
