---
id: 65824561f3478e5371a33ae5
title: Korak 32
challengeType: 20
dashedName: step-32
---

# --description--

Sledeći korak je definisanje glavne funkcije, koja će biti deo tačke ulaska programa za praćenje korišćenja unosa.

Definišite funkciju sa imenom `main` bez parametara. Popunite telo te funkcije sa listom `expenses` koju ste kreirali na početku ovog projekta. Koristićete ovu listu za čuvanje zapisa o korišćenju.

# --hints--

Treba da definišete funkciju pod nazivom `main()` bez parametara.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

Treba da premestite listu od `expenses` unutar tela funkcije/metode `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
