---
id: 6582352cbdbcde4d0c4f7b0b
title: Korak 20
challengeType: 20
dashedName: step-20
---

# --description--

Da pozoveš lambda funkciju, možeš da koristiš standardnu sintaksu funkcije i obične zagrade nakon imena parametra.

Pozovi svoju funkciju za `test` lambda i prosledi `3` kao argument. Zatim, ispiši rezultat.

# --hints--

Trebalo bi da imaš `print(test(3))` u svom kodu.

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
