---
id: 65822773f4c60a479ec89f66
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

U funkciji `print_expenses`, kreiraj petlju `for` koja iterira kroz svaki element u listi `expenses`. Koristi `expense` kao varijablu petlje i ubaci `pass` unutar tela petlje.

# --hints--

Morate kreirati petlju za `for` koja iterira/ponavlja `expenses` unutar vaše funkcije `print_expenses`. Zapamtite da koristite `pass` unutar tela petlje.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

Trebalo bi da koristiš `expense` kao parametar ciklusa.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

Trebalo bi da imaš `pass` samo unutar tela svog kruga.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
