---
id: 65824dfdb6815d563b2d3256
title: Korak 37
challengeType: 20
dashedName: step-37
---

# --description--

Funkcija `input()` koja uzima unos korisnika i vraća taj unos u obliku stringa.

Unutar vaše petlje `while`, pozovite funkciju `input()` prosljeđujući niz znakova `'Enter your choice: '` kao argument, i postavite vrednost rezultata u parametar nazvan `choice`.

# --hints--

Trebalo bi da deklarišete parametar nazvan `choice` unutar vašeg bloka `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].has_variable("choice")
`)) })
```

Treba da postavite vrednost rezultata funkcije ``input()``, sa nizom karaktera ``'Enter your choice: '`` kao njegov argument, u vaš parametar ``choice``.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_variable("choice").is_equivalent("choice = input('Enter your choice: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
--fcc-editable-region--
```
