---
id: 65824111a09164518320088d
title: Korak 30
challengeType: 20
dashedName: step-30
---

# --description--

Funkcija `filter()` omogućava vam da izaberete stavke iz ponovljivog izvora, kao što je lista, na osnovu rezultata funkcije:

```py
filter(my_function, my_list)
```

`filter()` uzima funkciju kao svoj prvi argument i ponovljivu vrednost kao svoj drugi argument. Vraća indeksator, koji je posebna stvar koja vam omogućava da ponavljate elemente kolekcije, poput liste.

Rezultat gore navedenog primera je indikator koji sadrži stavke od `my_list`, koje `my_function` vraća `True`.

Unutar funkcije `filter_expenses_by_category`, pozovi `filter()` prosleđujući funkciju `lambda` koju ste napisali u prethodnom koraku kao prvi argument i listu od `expenses` kao drugi argument.

# --hints--

Trebalo bi da izvršiš poziv `filter()` unutar funkcije `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

Trebalo bi da prosledite `lambda expense: expense['category'] == category` kao prvi argument pozivu `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

Trebalo bi da prosledite `expenses` kao drugi argument pozivu `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
