---
id: 65823040bdb32949d4c1a96f
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Biće vam potrebna funkcija za izračunavanje ukupnih troškova.

Definiši funkciju nazvanu `total_expenses` koja prima jedan parametar `expenses`. Popuni telo funkcije sa izrazom `pass` za sada.

# --hints--

Trebalo bi da definišete funkciju nazvanu `total_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

Tvoja funkcija za `total_expenses` mora da primi jedan parametar od `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
