---
id: 66314f2add0a8b57da5d6d91
title: Korak 22
challengeType: 20
dashedName: step-22
---

# --description--

Trebalo bi da vidite nešto poput `<map object at 0xd273a8>` odštampano u konzoli, što je reprezentacija niza karaktera objekta mape vraćenog sa `map()`.

Da biste dobili čitljive rezultate, morate nešto promeniti sa mape u listu. Učinite to prosleđivanjem poziva `map()` kao argumenta funkciji `list()`.

# --hints--

Treba da proslediš poziv `map()` kao argument funkciji `list()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
