---
id: 6669539c1379793f9cb8917c
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Lista za Python je dinamička, što znači da se vrednost stavki u listi može menjati. Možete promeniti vrednost elementa koristeći notaciju uglastih zagrada.

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

Ovo čini `example_list` da ima vrednost od `[4, 'oh', 6, 7]`.

Promeni prvo polje od `my_list` u `0`, zatim objavi listu da provjeriš vrijednost.

# --hints--

Nemojte menjati redosled liste.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Prvi element `my_list` treba da bude 0.

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

Trebalo bi da promeniš prvi deo koristeći `my_list[0] = 0`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

Trebalo bi da odštampate listu.

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

Trebalo bi da odštampate listu nakon što promenite prvi element.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
