---
id: 66694fc4bba24f33ca01fa5b
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Python ima nekoliko metoda za listu. Kao metode za dodavanje ili uklanjanje stavki sa liste.

Možete dodati stav na kraju liste koristeći metodu `append()`. Na primer:

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list` sad je `[4, 5, 6, 7]`.

Pokušajte da koristite putanju `append()` za dodavanje `3` na `my_list`. Zatim objavite listu.

# --hints--

Poslednji element od `my_list` treba biti `3`.

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list` Trebalo bi da ima dužinu od 3.

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

Ne smete menjati deklaraciju liste.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Trebalo bi da odštampate listu.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

Trebalo bi da ispišete listu nakon korišćenja `append()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
