---
id: 66694eaaeb7d253049fa7883
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Da biste dodali stavke u listu, možete ih napisati unutar uglastih zagrada, razdvajajući vrednosti zarezima.

Dodaj broj `1` na `2` unutar liste.

# --hints--

Prvi element `my_list` treba da bude `1`.

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

Drugi element od `my_list` bi trebalo da bude `2`.

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

`my_list` trebalo bi da ima dva elementa unutar sebe.

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

Treba da uneseš vrednost `[1, 2]` direktno u `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
