---
id: 666951d2e4b27e3a3c72ca5f
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Takođe možete dobiti samo jedan element da biste dobili njegovu vrednost.

Liste imaju nulti obrazac, kao što je niz slova koji vidite. To znači da prvi element postoji u obrascu `0`, drugi element postoji u obrascu `1` i tako dalje.

Za dohvaćanje elemenata koristi se pisanje u uglatim zagradama. Na primer, `example_list[1]` je dohvaćanje elementa iz primjera `1`, drugi element, od `example_list`.

Objavi prvi element od `my_list`.

# --hints--

Trebalo bi da odštampate prvi deo liste koristeći `print(my_list[0])`.

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
