---
id: 66694bfa435cef241f6f9a68
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Pre nego što počnete sa projektom, naučićete o listama.

Za razliku od brojeva i sekvenci slova, <dfn>lista</dfn> je vrsta podataka koja funkcioniše kao kontejner za druge vrednosti:

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

Lista ima karakteristiku kvadratnih zagrada koje okružuju sve vrednosti, i zareza između vrednosti, kao što je: `["A", "happy", "list"]`. Ako lista nema nijednu vrednost, onda je prazna lista: `[]`.

Lista može imati različite tipove podataka: `[1, "Up", ["Down", "Twice"]]`. Ovo uključuje sve moguće tipove podataka. Takođe može sadržati drugu listu unutar sebe!

Kreiraj parametar nazvan `my_list` i dodeli mu vrednost prazne liste.

# --hints--

Trebalo bi da imaš parametar nazvan `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

Polje `my_list` mora imati vrednost praznog niza/liste.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
