---
id: 666ab82c874add3c95145213
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Metoda `pop()` može se koristiti za uklanjanje elementa iz liste. Podrazumevano, uklanja poslednji element iz liste. Možete proslediti objekat kao argument na ovaj način, i to će ukloniti element koji se nalazi u tom objektu.

```py
fruits_list = ["cherry", "lemon", "tomato", "apple", "orange"]

fruits_list.pop(2)

print(fruits_list) # ["cherry", "lemon", "apple", "orange"]
```

U ovom slučaju, `fruits_list.pop(2)` uklanja element koji se nalazi u šablonu `2` liste.

Koristi `pop()` da ukloniš poslednji element iz `my_list`, a zatim objavi `my_list`.

# --hints--

Treba da koristiš `pop()`.

```js
({
    test: () => runPython(`
    assert len(_Node(_code).find_calls('pop')) > 0
    `)
})
```

`my_list` treba da bude `[0, 1, 2]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2]
    `)
})
```

Treba da odštampate `my_list` nakon korišćenja `pop()`.

```js
({
    test: () => runPython(`
    code = _Node(_code)
    assert(
        any(
            code.is_ordered(pop, 'print(my_list)') for pop in ['my_list.pop()', 'my_list.pop(3)']
        )
    )
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)


--fcc-editable-region--
```
