---
id: 666aae6a5d30a71f1fd7749f
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Metoda za `insert` može dodati element na bilo kojoj poziciji u listi. Prvi argument je pozicija gde treba dodati element, a drugi argument je element koji se dodaje. Na primer, ovako se dodaje novi element na treću poziciju od `example_list`:

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

Ako zavisiš od `insert()`, dodaj `1` u `my_list` na odgovarajuće mesto radi kumulativne kalkulacije, zatim ispiši listu.

# --hints--

Trebalo bi da dodate `1` na `my_list` za primer `1` koristeći `insert()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

`my_list` treba da ima 4 komponente.

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

`my_list` mora biti jednako `[0, 1, 2, 3]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

Treba da ispišete listu nakon dodavanja nove stavke.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
