---
id: 66be24cb4144f955b6bcc550
title: Izgradi stranicu sa kartama za igru
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

Ne smete ukloniti svojstvo `background-color` i njegovu vrednost iz CSS pravila `.container`.
1. Unapaswa kuwa na kipengele cha `main` chenye kitambulisho cha `playing-cards`.
1. Unutar elementa `#playing-cards` treba da imate najmanje tri `div` elementa, od kojih svaki ima klasu `card`.
1. Unutar svakog `.card` elementa treba da postoje tri `div` elementa: prvi sa klasom `left`, drugi sa klasom `middle`, a treći sa klasom `right`.
1. Element `#playing-cards` treba da koristi flexbox kako bi centrirao svoju decu horizontalno, omogućio prelamanje i postavio razmak od `20px` između njih.
1. Svaki `.card` element treba da koristi flexbox za raspoređivanje svoje dece pomoću `space-between`.
Vaše CSS pravilo `.container` treba da ima svojstvo `padding` postavljeno na `10px 0`.
1. Svaki `.middle` element treba da koristi flexbox svojstva da se poravna na sredinu svog roditeljskog kontejnera.
1. Svaki `.right` element treba da koristi flexbox svojstva da se poravna na kraj svog roditeljskog kontejnera.
1. Svaki `.middle` element treba da koristi flexbox za prikaz svoje dece u koloni.

Evo nekih znakova koje možete kopirati i zalepiti za korišćenje na vaše kartice ako želite: `♠`, `♣`, `♥`, na `♦`.

# --hints--

Trebalo bi da imaš element sa `main` koji ima `id` od `playing-cards`.

```js
assert.exists(document.querySelector('main#playing-cards'));
```

Trebalo bi da imaš najmanje tri elementa `div` sa klasom objekta `card` unutar tvog elementa `#playing-cards`.

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

Svaki vaš element za `.card` mora imati širinu i visinu.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

Svaki tvoj element sa `.card` mora imati tri elementa `div` kao decu.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

Trebalo bi da imaš `div` sa klasom objekta `left` unutar svakog elementa `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

Treba da imaš `div` sa klasom objekta `middle` unutar svakog elementa `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

Trebalo bi da imaš `div` sa klasom objekta `right` unutar svakog elementa `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

Trebalo bi da imaš selektor `#playing-cards` koji postavlja `display` njegovih elemenata na `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

Trebalo bi da imaš birač za `#playing-cards` koji postavlja `justify-content` njegovih elemenata na `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

Trebalo bi da imaš selektor `#playing-cards` koji postavlja `flex-wrap` njegovih elemenata na `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

Trebalo bi da imaš selektor za `#playing-cards` koji postavlja `gap` njegovih elemenata na `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

Trebalo bi da imaš selektor `.card` koji postavlja `display` njegovih elemenata na `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

Trebalo bi da imaš birač za `.card` koji postavlja `justify-content` njegovih elemenata na `space-between`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

Trebalo bi da imaš selektor za `.left` koji postavlja `align-self` njegovih elemenata na `flex-start`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

Trebalo bi da imaš birač za `.middle` koji postavlja `align-self` njegovih elemenata na `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

Treba da imaš selektor za `.right` koji postavlja `align-self` njegovih elemenata na `flex-end`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

Trebalo bi da imaš birač za `.middle` koji postavlja `display` njegovih elemenata na `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

Trebalo bi da imaš `flex-direction` postavljen kao `column` u svojoj klasi `.middle`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
