---
id: 66b62d0ad68488dd76228d6c
title: Pronađi greške u promeni pozadinske boje nasumično
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---

# --description--

Camperbot uči JavaScript i pokušao je da izgradi Nasumičnu promenu pozadinske boje. Međutim, napravio je nekoliko grešaka na putu.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove kako bi laboratorij radio ispravno.

**Priče korisnika:**

1. Treba da ispravite sintaksičke greške u nizu podataka za `darkColorsArr`.
1. Treba da ispravite `ReferenceError` koji se baca kada `math` vraća unutar funkcije `getRandomIndex`.
1. Treba da zaokružite `randomIndex` na najbliži ceo broj.
1. Treba da ispravite vrednost parametra `body` kako bi koristio tačno ime puta i popravili vezano `TypeError`.
1. Treba da ispravite `bgHexCodeSpanElement` kako bi tačno selektovao element koji ima `id` od `bg-hex-code`.
1. Treba da promenite parametar `color` unutar `changeBackgroundColor` tako da bude postavljen na slučajnu boju iz niza podataka za `darkColorsArr`.
1. Treba da ažurirate parametar `btn` kako bi ciljao na element dugmeta. Otvorite `index.html` da biste videli tačno ime `id` za ovo dugme.
1. Treba da ispravite liniju koja dodaje slušalicu događaja `click` za `btn` prosleđivanjem referencu funkcije `changeBackgroundColor` na taj način.
1. Vaša pozadinska boja bi trebalo da se nasumično promeni u jednu od boja koje se nalaze u `darkColorsArr` svaki put kada se dugme `#btn` klikne.

# --hints--

Potrebno je ispraviti sintaksičke greške u redu podataka za `darkColorsArr`. Imajte na umu da elementi reda podataka treba da budu razdvojeni zarezima.

```js
assert.isArray(darkColorsArr);
```

Treba da ispraviš grešku sa velikim slovom u liniji `math.random()`.

```js
assert.match(getRandomIndex.toString(), /\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*/);
```

Treba da smanjiš `darkColorsArr.length * Math.random()` na najbliži ceo broj.

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

Trebalo bi da ažuriraš vrednost svog parametra `body` ili `queryselector` koristeći camel case.

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

Treba da popraviš `"bg-hex-code"` ili tačno referenciraš `id` od `bg-hex-code`. Zapamti da selektori za `id` u CSS počinju sa hash znakom (`#`).

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

Trebalo bi da pozoveš funkciju `getRandomIndex` unutar linije `darkColorsArr[getRandomIndex]`.

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

Morate da ispravite ime `id` linije `"#click-btn"` ili ga uskladite sa tačnim imenom `id` u fajlu `index.html`.

```js
assert.match(btn?.id, /btn/);
```

Ne smeš da pozoveš funkciju `changeBackgroundColor`. Umesto toga, trebao bi da proslediš referencu na funkciju putem `addEventListener`.

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];
  
  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())
  
  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)
  
  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];
  
    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)
  
  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```
