---
id: 587d7dbd367417b2b2512bb5
title: Strukturiranje CSS-a i Sass-a
challengeType: 0
forumTopicId: 301457
dashedName: nest-css-with-sass
---

# --description--

Sass omogućava strukturiranje unutar pravila CSS, što je ključna metoda za organizovanje stilskog lista.

Obično, svaki element se usmerava na različit red kako bi se formatirao stil, kao što sledi:

```scss
article {
  height: 200px;
}

article p {
  color: white;
}

article ul {
  color: blue;
}
```

Za veliki projekat, fajl CSS će imati mnogo linija i mnogo pravila. Ovde interno formatiranje može pomoći u organizovanju vašeg koda postavljanjem pravila stila za decu unutar odgovarajućih roditeljskih elemenata:

```scss
article {
  height: 200px;

  p {
    color: white;
  }

  ul {
    color: blue;
  }
}

```

# --instructions--

Koristite gore prikazanu metodu raspoređivanja da ponovo rasporedite pravila za CSS za sve decu elementa `.blog-post`. U cilju testiranja, `h1` treba da dođe pre elementa `p`.

# --hints--

Vaš kod bi trebalo da ponovo konfiguriše pravila za CSS ili `h1` na `p`, postavljena unutar roditeljskog elementa `.blog-post`.

```js
assert(
  code.match(
    /\.blog-post\s*?{\s*?h1\s*?{\s*?text-align:\s*?center;\s*?color:\s*?blue;\s*?}\s*?p\s*?{\s*?font-size:\s*?20px;\s*?}\s*?}/gi
  )
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  .blog-post {

  }
  h1 {
    text-align: center;
    color: blue;
  }
  p {
    font-size: 20px;
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  .blog-post {
    h1 {
      text-align: center;
      color: blue;
    }
    p {
      font-size: 20px;
    }
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```
