---
id: 587d7dbe367417b2b2512bb9
title: Koristi @for za kreiranje petlje Sass-a
challengeType: 0
forumTopicId: 301462
dashedName: use-for-to-create-a-sass-loop
---

# --description--

Komanda za `@for` dodaje stilove u ciklus, veoma slično ciklusu od `for` u JavaScript.

`@for` koristi se na dva načina: "do kraja" ili "od početka do kraja". Glavna razlika je u tome što "od početka **do** kraja" *ne uključuje* poslednji broj kao deo brojanja, dok "do kraja **pored**" *uključuje* poslednji broj kao deo brojanja.

Evo primera od početka do kraja:

```scss
@for $i from 1 through 12 {
  .col-#{$i} { width: 100%/12 * $i; }
}
```

Deo `#{$i}` je sintaksa za spajanje varijable (`i`) sa tekstom kako bi se kreirao niz znakova. Kada se Sass fajl konvertuje u CSS, izgleda ovako:

```scss
.col-1 {
  width: 8.33333%;
}

.col-2 {
  width: 16.66667%;
}

...

.col-12 {
  width: 100%;
}
```

Ovo je moćan način za kreiranje mrežnog rasporeda. Sada imate dvanaest opcija širine redova dostupnih kao klase CSS.

# --instructions--

Napiši komandu `@for` koja uzima parametar `$j` koji ide od 1 **do** 6.

Inapaswa kuunda madarasa 5 yanayoitwa `.text-1` hadi `.text-5`, ambapo kila moja ina `font-size` iliyowekwa kuwa 15px mara nambari ya kielezo.

# --hints--

Msimbo wako unapaswa kutumia amri ya `@for`.

```js
assert(code.match(/@for /g));
```

Vaša klasa za `.text-1` treba da ima `font-size` od 15px.

```js
assert($('.text-1').css('font-size') == '15px');
```

Tvoja klasa za `.text-2` treba da ima `font-size` od 30px.

```js
assert($('.text-2').css('font-size') == '30px');
```

Vaša klasa za `.text-3` treba da ima `font-size` od 45px.

```js
assert($('.text-3').css('font-size') == '45px');
```

Vaša klasa za `.text-4` treba da ima `font-size` od 60px.

```js
assert($('.text-4').css('font-size') == '60px');
```

Darasa lako la `.text-5` treba da ima `font-size` od 75px.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>

@for $i from 1 through 5 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

---

```html
<style type='text/scss'>

@for $i from 1 to 6 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
