---
id: 587d7fa5367417b2b2512bbd
title: Razširi skup CSS stilova za drugi element
challengeType: 0
forumTopicId: 301456
dashedName: extend-one-set-of-css-styles-to-another-element
---

# --description--

Sass ima karakteristiku nazvanu `extend` koja olakšava preuzimanje pravila za CSS iz jednog elementa i izgradnju na njemu u drugom.

Na primer, donji deo pravila za CSS definiše klasu `.panel`. Ima `background-color`, `height` i `border`.

```scss
.panel{
  background-color: red;
  height: 70px;
  border: 2px solid green;
}
```

Sada želite drugi panel nazvan `.big-panel`. Ima osnovne kriterijume slične `.panel`, ali zahteva i `width` i `font-size`. Moguće je kopirati i zalepiti početna pravila za CSS iz `.panel`, ali kod postaje repetitivan kada dodate više tipova panela. Komanda `extend` je jednostavan način za ponovnu upotrebu pravila dizajniranih za jednu funkciju, a zatim dodavanje na drugu:

```scss
.big-panel{
  @extend .panel;
  width: 150px;
  font-size: 2em;
}
```

`.big-panel`će imati parametre slične `.panel` uz nove stilove.

# --instructions--

Kreiraj klasu za `.info-important` koja proširuje `.info` i takođe ima set od `background-color` koji je postavljen na magenta.

# --hints--

Vaš kurs za `info-important` mora imati `background-color`, postavljen na `magenta`.

```js
assert(
  code.match(
    /\.info-important\s*?{[\s\S]*background-color\s*?:\s*?magenta\s*?;[\s\S]*}/gi
  )
);
```

Vaša klasa za `info-important` treba da koristi `@extend` kako bi nasledila stil od klase `info`.

```js
assert(
  code.match(/\.info-important\s*?{[\s\S]*@extend\s*?.info\s*?;[\s\S]*/gi)
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }




</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }
  .info-important{
    @extend .info;
    background-color: magenta;
  }



</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```
