---
id: 587d7dbe367417b2b2512bb8
title: Koristite `@if` i `@else` za dodavanje logike vašim stilovima
challengeType: 0
forumTopicId: 301463
dashedName: use-if-and-else-to-add-logic-to-your-styles
---

# --description--

Komanda za `@if` u Sass-u je važna za testiranje specifičnih slučajeva - radi kao komanda za `if` u JavaScript.

```scss
@mixin make-bold($bool) {
  @if $bool == true {
    font-weight: bold;
  }
}
```

Kao što je u JavaScript, komande za `@else if` i `@else` testiraju više uslova:

```scss
@mixin text-effect($val) {
  @if $val == danger {
    color: red;
  }
  @else if $val == alert {
    color: yellow;
  }
  @else if $val == success {
    color: green;
  }
  @else {
    color: black;
  }
}
```

# --instructions--

Kreiraj mixin nazvan `border-stroke` koji prima parametar `$val`. Mixin mora da proveri sledeće uslove koristeći naredbe za `@if`, `@else if` i `@else`:

```scss
light - 1px solid black
medium - 3px solid black
heavy - 6px solid black
```

Ako je vrednost parametra `$val` različita od `light`, `medium`, ili `heavy`, onda svojstvo `border` treba da se postavi na `none`.

# --hints--

Tvoj kod mora da deklarše mixin nazvan `border-stroke` koji ima parametar nazvan `$val`.

```js
assert(code.match(/@mixin\s+?border-stroke\s*?\(\s*?\$val\s*?\)\s*?{/gi));
```

Vaš Mixin bi trebalo da ima komandu za `@if` da proveri da li `$val` je `light`, i postavi `border` kao `1px solid black`.

```js
assert(
  code.match(
    /@if\s+?\$val\s*?===?\s*?light\s*?{\s*?border\s*?:\s*?1px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Tvoj Mixin treba da ima komandu za `@else if` proveravanje da li `$val` je `medium`, i postavljanje `border` na `3px solid black`.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?medium\s*?{\s*?border\s*?:\s*?3px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Vaš mixin treba da ima komandu za `@else if` da proveri da li `$val` je `heavy`, i postavlja `border` na `6px solid black`.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?heavy\s*?{\s*?border\s*?:\s*?6px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Mixin yako inapaswa kuwa na amri ya `@else` kuweka `border` kuwa `none`.

```js
assert(code.match(/@else\s*?{\s*?border\s*?:\s*?none\s*?;\s*?}/gi));
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```

# --solutions--

```html
<style type='text/scss'>
  @mixin border-stroke($val) {
    @if $val == light {
      border: 1px solid black;
    }
    @else if $val == medium {
      border: 3px solid black;
    }
    @else if $val == heavy {
      border: 6px solid black;
    }
    @else {
      border: none;
    }
  }


  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```
