---
id: 587d7dbd367417b2b2512bb6
title: Kreiranje CSS-a koji se može ponovno koristiti korišćenjem mixina
challengeType: 0
forumTopicId: 301455
dashedName: create-reusable-css-with-mixins
---

# --description--

U Sass-u, <dfn>mixin</dfn> je grupa deklaracija CSS koje se mogu ponovo koristiti u celom stilskom listu. Definicija počinje sa sintaksom `@mixin` za at, praćena običnim imenom. Koristi se mixin koristeći sintaksu `@include` za at.

```scss
@mixin reset-list {
  margin: 0;
  padding: 0;
  list-style: none;
}

nav ul {
  @include reset-list;
}
```

Kreira:

```css
nav ul {
  margin: 0;
  padding: 0;
  list-style: none;
}
```

Mixin-i takođe mogu da uzmu argumente, što omogućava prilagođavanje njihovog ponašanja. Ti argumenti su potrebni prilikom korišćenja mixina.

```scss
@mixin prose($font-size, $spacing) {
  font-size: $font-size;
  margin: 0;
  margin-block-end: $spacing;
}

p {
  @include prose(1.25rem, 1rem);
}

h2 {
  @include prose(2.4rem, 1.5rem);
}
```

Možete tvrditi da su opcione tako što im date podrazumevane vrednosti za parametre.

```scss
@mixin text-color($color: black) {
  color: $color;
}

p {
  @include text-color(); /* color: black */
}

nav a {
  @include text-color(orange);
}
```

# --instructions--

Napišite mixin pod nazivom `shape` sa sledećim 3 parametra: `$w`, `$h`, i `$bg-color`.

Koristi mixin od `shape` da elementu `#square` širinu i visinu od `50px`, sa bojom pozadine `red`. Za element `#rect-a` dodaj širinu od `100px`, visinu od `50px`, i pozadinsku boju `blue`. Na kraju, za element `#rect-b` dodaj širinu od `50px`, visinu od `100px`, i pozadinsku boju `orange`.

# --hints--

Unapaswa kutangaza mixin inayoitwa `shape` yenye vigezo 3: `$w`, `$h`, na `$bg-color`.

```js
assert.match(code, /@mixin\s+shape\s*\(\s*\$w,\s*\$h,\s*\$bg-color\s*\)\s*{/gi);
```

Mixin yako inapaswa kujumuisha sifa ya `width` inayotumia kigezo `$w`.

```js
assert.match(__helpers.removeWhiteSpace(code), /width:\$w;/gi);
```

Mixin yako inapaswa kujumuisha sifa ya `height` inayotumia kigezo `$h`.

```js
assert.match(__helpers.removeWhiteSpace(code), /height:\$h;/gi);
```

Vaš Mixin treba da uključuje svojstvo `background-color` koje koristi parametar `$bg-color`.

```js
assert.match(__helpers.removeWhiteSpace(code), /background-color:\$bg\-color;/gi);
```

Treba da promeniš stilove unutar selektora `#square` za poziv miksina `shape` koristeći glavni selektor `@include`. Postavi širinu i visinu `50px`, i boju pozadine `red`.

```js
assert.match(code, /#square\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*50px\s*,\s*red\s*\)\s*;\s*}/gi);
```

Treba da promeniš stilove unutar selektora `#rect-a` za poziv miksina `shape` koristeći glavnu ključnu reč `@include`. Postavi širinu `100px`, visinu `50px`, i boju pozadine `blue`.

```js
assert.match(code, /#rect-a\s*{\s*@include\s+shape\s*\(\s*100px\s*,\s*50px\s*,\s*blue\s*\)\s*;\s*}/gi);
```

Trebalo bi da promenite stilove unutar selektora `#rect-b` za poziv miksina `shape` koristeći glavnu klasu `@include`. Postavite širinu `50px`, visinu `100px` i boju pozadine `orange`.

```js
assert.match(code, /#rect-b\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*100px\s*,\s*orange\s*\)\s*;\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
#square {
  width: 50px;
  height: 50px;
  background-color: red;
}

#rect-a {
  width: 100px;
  height: 50px;
  background-color: blue;
}

#rect-b {
  width: 50px;
  height: 100px;
  background-color: orange;
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```

# --solutions--

```html
<style type='text/scss'>
@mixin shape($w, $h, $bg-color) {
  width: $w;
  height: $h;
  background-color: $bg-color;
}

#square {
  @include shape(50px, 50px, red);
}

#rect-a {
  @include shape(100px, 50px, blue);
}

#rect-b {
  @include shape(50px, 100px, orange);
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```
