---
id: 587d7dbf367417b2b2512bbc
title: Razdvoj svoje stilove na manje delove koristeći partiale
challengeType: 0
forumTopicId: 301459
dashedName: split-your-styles-into-smaller-chunks-with-partials
---

# --description--

<dfn>Partials</dfn> u Sass-u su zasebne datoteke koje sadrže delove koda za CSS. One se unose i koriste u drugim Sass fajlovima. Ovo je dobar način za kombinovanje sličnog koda u modulima kako bi ostao organizovan.

Nazivi partiala počinju sa znakom podznakse (`_`) koji govori Sass-u da je to mali deo CSS i da ga ne konvertuje u fajl CSS. Takođe, Sass fajlovi završavaju sa ekstenzijom fajla `.scss`. Da biste uključili kod koji se nalazi u partialu u drugi Sass fajl, koristite komandu `@import`.

Na primer, ako su svi tvoji mixini sačuvani u partial fajlu pod nazivom "\_mixins.scss", a potrebni su u fajlu "main.scss", ovako ga koristiš u glavnom fajlu:

```scss
@import 'mixins'
```

Zapamtite da u deklaraciji podvučnika i ekstenzije fajla nisu potrebni za `import` - Sass razume da je to partial. Kada se partial uključi u fajl, sve varijable, mixini i drugi kod su dostupni za upotrebu.

# --instructions--

Napišite deklaraciju `@import` za uključivanje partial-a pod nazivom `_variables.scss` u fajl main.scss.

# --hints--

Tvoj kod bi trebalo da koristi komandu `@import`, i nemoj uključiti podvučnik u ime fajla.

```js
assert(code.match(/@import\s+?('|")variables\1/gi));
```

# --seed--

## --seed-contents--

```html
<!-- The main.scss file -->
```

# --solutions--

```html
@import 'variables'
```
