---
id: 587d7dbf367417b2b2512bbb
title: Koristi stil dok uslov nije ispunjen sa @while
challengeType: 0
forumTopicId: 301454
dashedName: apply-a-style-until-a-condition-is-met-with-while
---

# --description--

Komanda za `@while` je opcija sa funkcionalnošću sličnom ciklusu JavaScript `while`. Kreira pravila za CSS dok ne bude zadovoljen određeni uslov.

Izazov sa `@for` dao je primer kreiranja jednostavnog mrežnog sistema/grid sistema. Ovo takođe može da radi sa `@while`.

```scss
$x: 1;
@while $x < 13 {
  .col-#{$x} { width: 100%/12 * $x;}
  $x: $x + 1;
}
```

Prvo, definiši parametar `$x` i postavi mu vrednost na 1. Zatim, koristi komandu `@while` za kreiranje mrežnog sistema *kada* `$x` je manji od 13. Nakon postavljanja pravila CSS za `width`, `$x` se povećava za 1 kako bi se izbegao beskonačni ciklus.

# --instructions--

Tumia `@while` kuunda mfululizo wa madarasa yenye `font-size` tofauti.

Imate 5 različitih klasa od `text-1` do `text-5`. Zatim postavite `font-size` na `15px` broj trenutnog modela. Obavezno izbegavajte beskonačnu petlju!

# --hints--

Msimbo wako unapaswa kutumia amri ya `@while`.

```js
assert(code.match(/@while /g));
```

Vaš kod treba da koristi parametar modela koji počinje sa vrednošću od 1.

```js
assert(code.match(/\$.*:\s*?1;/gi));
```

Msimbo wako unapaswa kuongeza thamani ya kigezo cha kuhesabu.

```js
assert(code.match(/\$(.*)\s*?:\s*\$\1\s*\+\s*1\s*;/gi));
```

Darasa lako la `.text-1` linapaswa kuwa na `font-size` ya `15px`.

```js
assert($('.text-1').css('font-size') == '15px');
```

Darasa lako la `.text-2` linapaswa kuwa na `font-size` ya `30px`.

```js
assert($('.text-2').css('font-size') == '30px');
```

Vaš razred za `.text-3` mora imati `font-size` od `45px`.

```js
assert($('.text-3').css('font-size') == '45px');
```

Vaša klasa za `.text-4` bi trebalo da ima `font-size` od `60px`.

```js
assert($('.text-4').css('font-size') == '60px');
```

Darasa lako za `.text-5` treba da ima `font-size` od `75px`.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>
  $x: 1;
  @while $x < 6 {
    .text-#{$x}{
      font-size: 15px * $x;
    }
    $x: $x + 1;
  }
</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
