---
id: 587d7dbd367417b2b2512bb4
title: Sačuvajte podatke koristeći Sass parametre
challengeType: 0
forumTopicId: 301460
dashedName: store-data-with-sass-variables
---

# --description--

Jedna karakteristika Sass-a koja se razlikuje od CSS je što koristi varijable. Deklariše se i postavlja za skladištenje podataka, kao što su JavaScript.

U JavaScript, varijable se deklarišu korišćenjem ključnih reči `let` i `const`. U Sass-u, varijable počinju sa `$` praćeno imenom varijable.

Evo nekoliko primera:

```scss
$main-fonts: Arial, sans-serif;
$headings-color: green;
```

Koristeći kriterijume:

```scss
h1 {
  font-family: $main-fonts;
  color: $headings-color;
}
```

Jedan primjer gdje su parametri važni je kada mnogo elemenata mora imati istu boju. Ako se ta boja promijeni, jedini dio koda koji treba izmijeniti je vrijednost parametra.

# --instructions--

Kreiraj parametar `$text-color` i postavi vrednost `red`. Zatim promeni vrednost atributa `color` na `.blog-post` i `h2` u parametar `$text-color`.

# --hints--

Tvoj kod mora imati Sass svojstvo deklarisano za `$text-color` i vrednost od `red`.

```js
assert(code.match(/\$text-color\s*:\s*?red\s*;/g));
```

Msimbo wako unapaswa kutumia kigezo `$text-color` kubadilisha sifa ya `color` kwa vipengele vya `.blog-post` na `h2`.

```js
assert(code.match(/color\s*:\s*\$text-color\s*;?/g));
```

Element vam `.blog-post` bi trebalo da ima `color` u crvenoj boji.

```js
assert($('.blog-post').css('color') == 'rgb(255, 0, 0)');
```

Vipengele vyako vya `h2` vinapaswa kuwa na `color` ya rangi nyekundu.

```js
assert($('h2').css('color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>


  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: red;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  $text-color: red;

  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: $text-color;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```
