---
id: 587d7dbf367417b2b2512bba
title: Koristi `@each` kroz stavke u listi
challengeType: 0
forumTopicId: 301461
dashedName: use-each-to-map-over-items-in-a-list
---

# --description--

Poslednji zadatak pokazao je kako direktiva ``@for`` koristi početnu i krajnju vrednost za iteriranje određeni broj puta. Sass takođe pruža direktivu ``@each`` koja iterira kroz svaki element u listi ili mapi. U svakoj iteraciji, promenljiva dobija trenutnu vrednost iz liste ili mape.

```scss
@each $color in blue, red, green {
  .#{$color}-text {color: $color;}
}
```

Ramani ima malo drugačiju sintaksu. Evo primera:

```scss
$colors: (color1: blue, color2: red, color3: green);

@each $key, $color in $colors {
  .#{$color}-text {color: $color;}
}
```

Zapamtite da je parametar `$key` potreban za referenciranje ključeva u mapi. Inače, prikupljeni CSS bi imali `color1`, `color2`... unutra. Oba primera koda iznad konvertuju se u sledeći CSS:

```scss
.blue-text {
  color: blue;
}

.red-text {
  color: red;
}

.green-text {
  color: green;
}
```

# --instructions--

Napišite komandu za `@each` koja prolazi kroz listu: `blue, black, red` i postavite svaki parametar na klasu `.color-bg`, gde se deo `color` menja za svaki element u odgovarajućoj boji. Svaka klasa takođe mora da postavi `background-color` za odgovarajuću boju.

# --hints--

Tvoj kod bi trebalo da koristi instrukciju `@each`.

```js
assert(code.match(/@each /g));
```

Darasa lako la `.blue-bg` linapaswa kuwa na `background-color` ya bluu.

```js
assert($('.blue-bg').css('background-color') == 'rgb(0, 0, 255)');
```

Darasa lako la `.black-bg` linapaswa kuwa na `background-color` ya nyeusi.

```js
assert($('.black-bg').css('background-color') == 'rgb(0, 0, 0)');
```

Tvoja klasa za `.red-bg` mora imati `background-color` crveni.

```js
assert($('.red-bg').css('background-color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

# --solutions--

```html
<style type='text/scss'>

  @each $color in blue, black, red {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

---

```html
<style type='text/scss'>

  $colors: (color1: blue, color2: black, color3: red);

  @each $key, $color in $colors {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```
