---
id: 68f890b73bbfc9efa7379e7b
title: PronaД‘ite prisiljene tipove greЕЎaka u aplikaciji sa greЕЎkama
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

Tek ste se pridruЕѕili lokalnoj web razvojnoj radnici/platformi, i vaЕЎ prvi zadatak je da oДЌistite kod sa greЕЎkama koji su ostavili prethodni programeri.

PokuЕЎavali su da izvrЕЎe jednostavne aritmetiДЌke operacije, ali neЕЎto je poЕЎlo naopako, rezultati nemaju smisla. VaЕЎ zadatak je pregledati, pronaД‡i greЕЎke i ispraviti kod tako da radi kako treba i da bude lako ДЌitljiv.

**Cilj:** Popunite korisniДЌke priДЌe navedene ispod i proД‘ite sve testove da biste zavrЕЎili laboratoriju.

**PriДЌe korisnika:**

1. Trebalo bi da imaЕЎ parametar nazvan `firstResult` koji precizno dodaje brojeve `5` i `10` kako bi proizveo vrednost `15`.
2. Trebalo bi da imaЕЎ parametar nazvan `secondResult` koji precizno oduzima `5` od `8` kako bi proizveo vrednost `3`.
3. Trebalo bi da imaЕЎ parametar nazvan `thirdResult` koji vraД‡a vrednost `6` dodavanjem dva broja. Izmeni trenutnu istinu ili laЕѕnu vrednost koja se koristi u izrazu za proraДЌun i broj.
4. Trebalo bi da imaЕЎ parametar nazvan `fourthResult` koji vraД‡a vrednost `8` dodavanjem dva broja. Izmeni trenutnu istinu ili laЕѕnu vrednost koja se koristi u izrazu za proraДЌun i broj.
5. Trebalo bi da imaЕЎ parametar nazvan `fifthResult` koji precizno mnoЕѕi dva broja kako bi dobio proizvod `20`. Izmeni trenutni niz znakova koja se koristi u izrazu za proraДЌun i broj.
6. Trebalo bi da imaЕЎ parametar nazvan `sixthResult` koji precizno dodaje dva broja kako bi proizveo vrednost `22`. Izmeni `null` koja se trenutno koristi u izrazu za proraДЌun i broj.

# --hints--

Trebalo bi da imaЕЎ parametar nazvan `firstResult`.

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult` mora imati numeriДЌku vrednost `15`.

```js
assert.strictEqual(firstResult, 15);
```

Trebalo bi da imaЕЎ parametar koji se zove `secondResult`.

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult` treba da ima numeriДЌku vrednost `3`.

```js
assert.strictEqual(secondResult, 3);
```

Trebalo bi da imate parametar koji se zove `thirdResult`.

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult` treba da ima numeriДЌku vrednost `6`.

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult` ne bi trebalo da sadrЕѕi vrednost taДЌno ili netaДЌno (`true` ili `false`).

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Trebalo bi da imaЕЎ parametar nazvan `fourthResult`.

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult` mora imati numeriДЌku vrednost `8`.

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult` ne bi smelo da sadrЕѕi vrednost istinito ili laЕѕno (`true` ili `false`).

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Trebalo bi da imaЕЎ parametar nazvan `fifthResult`.

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult` mora imati numeriДЌku vrednost `20`.

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult` Ne bi trebalo da ukljuДЌuje mnoЕѕenje sekvence karaktera (primer "10" * 2).

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

Trebalo bi da imaЕЎ parametar nazvan `sixthResult`.

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult` mora imati numeriДЌku vrednost `22`.

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult` ne bi trebalo da ukljuДЌuje `null` u svoj matematiДЌki izraz.

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
