---
id: 5b7d72c338cd7e35b63f3e14
title: Poboljsaj Kompatibilnost Uz Browser Fallback Vrednosti
challengeType: 0
forumTopicId: 301087
dashedName: improve-compatibility-with-browser-fallbacks
---

# --description--

Kada radite sa CSS-om, mozete naici na probleme kompatibilnosti izmedju pregledača. Zato je vazno obezbediti fallback vrednosti kako biste izbegli moguce probleme.

Kada pregledač cita CSS stranice, ignorise svojstva koja ne prepoznaje ili ne podrzava. Na primer, ako koristite CSS promenljivu za boju pozadine, Internet Explorer ce je ignorisati jer ne podrzava CSS promenljive. U tom slucaju pregledač koristi neku drugu dostupnu vrednost tog svojstva. Ako je nema, vraca se na podrazumevanu vrednost, sto cesto nije pozeljno.

To znaci da fallback mozete dodati vrlo jednostavno: navedite podrzaniju vrednost odmah pre postojece deklaracije. Tako stariji pregledači imaju sta da koriste, dok noviji pregledači primenjuju kasniju deklaraciju u nizu.

# --instructions--

Izgleda da se promenljiva koristi za postavljanje boje pozadine klase `.red-box`. Poboljsajte kompatibilnost dodavanjem jos jedne `background` deklaracije odmah pre postojece, sa vrednoscu `red`.

# --hints--

Vase `.red-box` pravilo treba da sadrzi fallback `background` postavljen na `red`, odmah pre postojece `background` deklaracije.

```js
const spacelessCode = __helpers.removeWhiteSpace(__helpers.removeCssComments(code));
assert.match(
  spacelessCode,
  /\.red-box{background:(red|#ff0000|#f00|rgb\(255,0,0\)|rgb\(100%,0%,0%\)|hsl\(0,100%,50%\));background:var\(--red-color\);height:200px;width:200px;}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  :root {
    --red-color: red;
  }
  .red-box {

    background: var(--red-color);
    height: 200px;
    width:200px;
  }
</style>
<div class="red-box"></div>
```

# --solutions--

```html
<style>
  :root {
    --red-color: red;
  }
  .red-box {
    background: red;
    background: var(--red-color);
    height: 200px;
    width:200px;
  }
</style>
<div class="red-box"></div>
```
