---
id: bad87fee1348bd9aedf08805
title: Koristi CSS Selektore za Stilizovanje Elemenata
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKMBT2'
forumTopicId: 18349
dashedName: use-css-selectors-to-style-elements
---

# --description--

Koriscenjem CSS-a, postoji mnogo CSS svojstava koje mozete upotrebiti da promenite izgled elemenata na stranici.

Kada ste napisali `<h2 style="color: red;">CatPhotoApp</h2>`, stilizovali ste `h2` element inline CSS-om, odnosno Cascading Style Sheets.

To je jedan nacin da stilizujete element, ali postoji bolji nacin za rad sa CSS-om.

Na vrhu koda, napravite `style` sekciju ovako:

```html
<style>
</style>
```

Unutar te sekcije mozete napisati <dfn>CSS selektor</dfn> za sve `h2` elemente. Na primer, ako zelite da svi `h2` elementi budu crveni, dodajte pravilo ovako:

```html
<style>
  h2 {
    color: red;
  }
</style>
```

Vazno je da imate otvarajucu i zatvarajucu viticastu zagradu (`{` i `}`) oko stilskih pravila svakog elementa. Takodje treba da stil deklaracija bude izmedju otvarajuceg i zatvarajuceg `style` taga. Na kraju, dodajte tacka-zarez na kraju svakog stilskog pravila.

# --instructions--

Obrisite `style` atribut sa `h2` elementa i umesto toga napravite CSS `style` sekciju. Dodajte potreban CSS da svi `h2` elementi budu plavi.

# --hints--

`style` atribut treba da bude uklonjen sa `h2` elementa.

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

Treba da napravite `style` element.

```js
assert.exists(document.querySelector('style:not(.fcc-hide-header)'));
assert.isAtLeast(document.querySelectorAll('style:not(.fcc-hide-header)').length, 1);
```

Vas `h2` element treba da bude plav.

```js
const h2Element = document.querySelector('h2');
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

`h2` deklaracija u stylesheet-u treba da bude validna i da ima tacka-zarez i zatvarajucu zagradu.

```js
assert.match(__helpers.removeCssComments(code), /h2\s*\{\s*color\s*:.*;\s*\}/g);
```

Svi `style` elementi treba da budu validni i da imaju zatvarajuce tagove.

```js
assert.match(__helpers.removeHtmlComments(code), /<\/style>/g);

const closingElementLength = __helpers.removeHtmlComments(code).match(/<\/style>/g).length;
const openingElementsLength = __helpers.removeHtmlComments(code).match(/<style((\s)*((type|media|scoped|title|disabled)="[^"]*")?(\s)*)*>/g).length;
assert.strictEqual(closingElementLength, openingElementsLength);
```

# --seed--

## --seed-contents--

```html
<h2 style="color: red;">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  h2 {
    color: blue;
  }
</style>
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
