---
id: bad87fee1348bd9aedf08736
title: Stilizuj HTML Body Element
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB77PHW'
forumTopicId: 18313
dashedName: style-the-html-body-element
---

# --description--

Sada krenimo iz pocetka i razgovarajmo o CSS nasledjivanju.

Svaka HTML stranica ima `body` element.

# --instructions--

Mozemo potvrditi da `body` element postoji ovde tako sto cemo mu postaviti crnu `background-color` vrednost.

To mozemo uraditi dodavanjem sledeceg u nas `style` element:

```css
body {
  background-color: black;
}
```

# --hints--

Vas `body` element treba da ima crnu `background-color` vrednost.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color'];

assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Vase CSS pravilo treba da ima ispravnu strukturu sa otvarajucom i zatvarajucom viticastom zagradom.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

Vase CSS pravilo treba da se zavrsava tacka-zarezom.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

# --seed--

## --seed-contents--

```html
<style>

</style>
```

# --solutions--

```html
<style>
body {
  background-color: black;
}
</style>
```
