---
id: bad87fee1348bd9aedf08808
title: Definisi Kako Fontovi Treba da Degradiraju
challengeType: 0
videoUrl: 'https://scrimba.com/c/cpVKBfQ'
forumTopicId: 18304
dashedName: specify-how-fonts-should-degrade
---

# --description--

Postoji nekoliko osnovnih fontova dostupnih u svim pregledacima. Ove genericke font porodice ukljucuju `monospace`, `serif` i `sans-serif`.

Kada jedan font nije dostupan, mozete reci pregledacu da "degradira" na drugi font.

Na primer, ako zelite da element koristi `Helvetica`, ali da degradira na `sans-serif` kada `Helvetica` nije dostupna, napisacete ovako:

```css
p {
  font-family: Helvetica, sans-serif;
}
```

Nazivi generickih font porodica nisu osetljivi na velika/mala slova. Takodje, ne zahtevaju navodnike jer su CSS kljucne reci.

# --instructions--

Za pocetak, postavite `monospace` font za `h2` element, tako da sada ima dva fonta - `Lobster` i `monospace`.

U prethodnom izazovu uvezli ste `Lobster` font koristeci `link` tag. Sada zakomentarisite taj `Lobster` import (koristeci HTML komentare koje ste ranije naucili) tako da vise nije dostupan. Posmatrajte kako `h2` element degradira na `monospace` font.

**Napomena:** Ako imate `Lobster` font instaliran na svom racunaru, necete primetiti degradaciju jer pregledac moze pronaci taj font lokalno.

# --hints--

Vas `h2` element treba da koristi `Lobster` font.

```js
const h2Element = document.querySelector('h2');
const fontFamily = window.getComputedStyle(h2Element)['font-family']; 
assert.match(fontFamily, /^"?lobster/i);
```

Vas `h2` element treba da degradira na `monospace` font kada `Lobster` nije dostupan.

```js
assert.match(__helpers.removeCssComments(code), /\s*h2\s*\{\s*font-family\s*\:\s*(\'|"|)Lobster\1\s*,\s*monospace\s*;?\s*\}/gi);
```

Treba da zakomentarisete Google poziv za `Lobster` font stavljanjem `<!--` ispred njega.

```js
assert.match(code, /<!--[^fc]/gi);
```

Treba da zatvorite komentar dodavanjem `-->`.

```js
assert.match(code, /[^fc]-->/gi);
```

# --seed--

## --seed-contents--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<!--<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">-->
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster, monospace;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
