---
id: bad87fee1348bd9aedf04756
title: Pregazi Stilove u Naknadnom CSS-u
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDQug'
forumTopicId: 18253
dashedName: override-styles-in-subsequent-css
---

# --description--

Nasa `pink-text` klasa je pregazila CSS deklaraciju naseg `body` elementa.

Upravo smo pokazali da nase klase mogu pregaziti CSS `body` elementa. Sledece logicno pitanje je: kako da pregazimo nasu `pink-text` klasu?

# --instructions--

Napravite dodatnu CSS klasu pod nazivom `blue-text` koja elementu daje plavu boju. Uverite se da je deklarisana ispod `pink-text` klase.

Primeni klasu `blue-text` na svoj `h1` element zajedno sa klasom `pink-text`, pa da vidimo koja ce pobediti.

Koriscenje vise klasa na HTML elementu radi se tako sto ih razdvojite razmakom, ovako:

```html
class="class1 class2"
```

**Napomena:** Nije vazno kojim redosledom su klase navedene unutar HTML elementa.

Medjutim, redosled deklaracije `class` pravila u delu `<style>` je bitan. Druga deklaracija uvek ima prednost nad prvom. Zato sto je `.blue-text` deklarisana druga, ona pregazi svojstva `.pink-text`.

# --hints--

Vas `h1` element treba da ima klasu `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Vas `h1` element treba da ima klasu `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

I `blue-text` i `pink-text` treba da budu na istom `h1` elementu.

```js
assert.isTrue(document.querySelector('.pink-text').classList.contains('blue-text'));
```

Vas `h1` element treba da bude plave boje.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }

  .blue-text {
    color: blue;
  }  
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```
