---
id: bad87fee1348bd9aedf08807
title: Uvezi Google Font
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9MRsJ'
forumTopicId: 18200
dashedName: import-a-google-font
---

# --description--

Pored navodjenja standardnih fontova dostupnih na mnogim operativnim sistemima, mozemo navesti i nestandardne web fontove za nas sajt. Na internetu postoji mnogo izvora web fontova. U ovom primeru fokusiramo se na Google Fonts biblioteku.

Google Fonts je besplatna biblioteka web fontova koju mozete koristiti u CSS-u pozivanjem URL-a tog fonta.

Dakle, hajde da uvezemo i koristimo Google font (napomena: ako je Google blokiran u vasoj zemlji, ovaj izazov treba preskociti).

Da biste uvezli Google font, mozete kopirati URL fonta iz Google Fonts biblioteke i nalepiti ga u svoj HTML. U ovom izazovu uvozimo font `Lobster`. Kopirajte sledeci isecak koda i nalepite ga na vrh editora (pre otvaranja `style` elementa):

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
```

Sada mozete koristiti `Lobster` font u CSS-u koristeci `Lobster` kao FAMILY_NAME, kao u sledecem primeru:

```css
font-family: FAMILY_NAME, GENERIC_NAME;
```

GENERIC_NAME je opciono i predstavlja fallback font ako prethodno navedeni font nije dostupan. Ovo ce biti obradjeno u sledecem izazovu.

Nazivi font-family vrednosti su osetljivi na velika/mala slova i treba ih staviti pod navodnike ako sadrze razmake. Na primer, navodnici su potrebni za `"Open Sans"`, ali nisu potrebni za `Lobster`.

# --instructions--

Uvezite font `Lobster` na svoju stranicu. Zatim, koristeci selektor elementa, postavite `Lobster` kao `font-family` za `h2` element.

# --hints--

Treba da uvezete `Lobster` font.

```js
assert.exists(document.querySelector('link[href*="googleapis" i]'));
```

Vas `h2` element treba da koristi `Lobster` font.

```js
const h2 = document.querySelector('h2'); 
const fontFamily = window.getComputedStyle(h2)['font-family']; 
assert.match(fontFamily, /lobster/i);
```

Treba da koristite samo `h2` selektor elementa da promenite font.

```js
assert.match(__helpers.removeHtmlComments(code), /\s*[^\.]h2\s*\{\s*font-family\s*:\s*('|"|)Lobster\1\s*(,\s*('|"|)[a-z -]+\3\s*)?(;\s*\}|\})/gi);
```

Vas `p` element i dalje treba da koristi `monospace` font.

```js
const paragraphElement = document.querySelector('p');
const fontFamily = window.getComputedStyle(paragraphElement)['font-family']; 
assert.match(fontFamily, /monospace/i);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }

  h2 {
    font-family: Lobster;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
