---
id: bad87fee1348bd8aedf06756
title: Pregazi Deklaracije Klasa Stilizovanjem ID Atributa
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpDhB'
forumTopicId: 18251
dashedName: override-class-declarations-by-styling-id-attributes
---

# --description--

Upravo smo pokazali da pregledaci citaju CSS od vrha ka dnu redosledom deklaracija. To znaci da ce, kada postoji konflikt, pregledac primeniti poslednju CSS deklaraciju. Primetite da cak i kada bismo stavili `blue-text` pre `pink-text` u klasama na `h1` elementu, i dalje bi se gledao redosled deklaracija, a ne redosled navodjenja klasa.

Ali to nije sve. Postoje i drugi nacini da pregazite CSS. Secate li se `id` atributa?

Hajde da pregazimo vase `pink-text` i `blue-text` klase i da `h1` element obojimo narandzasto tako sto cemo mu dodeliti `id` i stilizovati taj `id`.

# --instructions--

Dodelite svom `h1` elementu `id` atribut sa vrednoscu `orange-text`. Podsetnik, stil za `id` izgleda ovako:

```html
<h1 id="orange-text">
```

Zadrzite klase `blue-text` i `pink-text` na `h1` elementu.

Napravite CSS deklaraciju za `orange-text` `id` u svom `style` elementu. Evo primera kako to izgleda:

```css
#brown-text {
  color: brown;
}
```

**Napomena:** Nije vazno da li ovaj CSS deklarisete iznad ili ispod klase `pink-text`, jer `id` atribut uvek ima veci prioritet.

# --hints--

Vas `h1` element treba da ima klasu `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Vas `h1` element treba da ima klasu `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Vas `h1` element treba da ima `id` `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'),'orange-text');
```

Treba da postoji samo jedan `h1` element.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Vas `orange-text` `id` treba da ima CSS deklaraciju.

```js
assert.match(__helpers.removeCssComments(code), /#orange-text\s*{/gi);
```

`h1` ne treba da ima `style` atribut.

```js
assert.notMatch(__helpers.removeHtmlComments(code), /<h1.*style.*>/gi);
```

Vas `h1` element treba da bude narandzaste boje.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 165, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
  #orange-text {
    color: orange;
  }  
</style>
<h1 id="orange-text"  class="pink-text blue-text">Hello World!</h1>
```
