---
id: bad87fee1348bd9aedf08756
title: Daj Jednom Stilu Veci Prioritet od Drugog
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ8wnHv'
forumTopicId: 18258
dashedName: prioritize-one-style-over-another
---

# --description--

Ponekad ce vasi HTML elementi dobiti vise stilova koji su u konfliktu.

Na primer, vas `h1` element ne moze istovremeno biti i zelen i roze.

Pogledajmo sta se desava kada napravimo klasu koja boji tekst u roze i zatim je primenimo na element. Da li ce nasa klasa *pregaziti* `body` CSS svojstvo `color: green;`?

# --instructions--

Napravite CSS klasu pod nazivom `pink-text` koja daje elementu roze boju.

Dodelite svom `h1` elementu klasu `pink-text`.

# --hints--

Vas `h1` element treba da ima klasu `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Vas `<style>` treba da ima CSS klasu `pink-text` koja menja `color`.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*\{\s*color\s*:\s*.+\s*;?\s*\}/g);
```

Vas `h1` element treba da bude roze boje.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
</style>
<h1>Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```
