---
id: bad87fee1348bd9aedf08746
title: Nasledi Stilove iz Body Elementa
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bmdtR'
forumTopicId: 18204
dashedName: inherit-styles-from-the-body-element
---

# --description--

Sada smo potvrdili da svaka HTML stranica ima `body` element i da `body` moze biti stilizovan pomocu CSS-a.

Zapamtite, `body` element mozete stilizovati kao bilo koji drugi HTML element, a svi ostali elementi ce naslediti njegove stilove.

# --instructions--

Prvo napravite `h1` element sa tekstom `Hello World`.

Zatim postavite svim elementima na stranici boju `green` dodavanjem `color: green;` u stil deklaraciju `body` elementa.

Na kraju, dodelite svom `body` elementu `monospace` font dodavanjem `font-family: monospace;` u stil deklaraciju `body` elementa.

# --hints--

Treba da napravite `h1` element.

```js
assert.isNotEmpty(document.querySelectorAll('h1'));
```

Vas `h1` element treba da sadrzi tekst `Hello World`.

```js
assert.match(
  document.querySelector('h1').textContent,
  /hello world/i
);
```

Vas `h1` element treba da ima zatvarajuci tag.

```js
const commentlessCode = __helpers.removeHtmlComments(code);
assert.match(commentlessCode, /<\/h1>/g);
assert.match(commentlessCode, /<h1/g);
assert.lengthOf(commentlessCode.match(/<\/h1>/g), commentlessCode.match(/<h1/g).length);
```

Vas `body` element treba da ima `color` svojstvo sa vrednoscu `green`.

```js
const bodyElement = document.querySelector('body');
const color = window.getComputedStyle(bodyElement)['color']; 
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Vas `body` element treba da ima `font-family` svojstvo sa vrednoscu `monospace`.

```js
const bodyElement = document.querySelector('body');
const fontFamily = window.getComputedStyle(bodyElement)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

Vas `h1` element treba da nasledi `monospace` font od `body` elementa.

```js
const h1Element = document.querySelector('h1');
const fontFamily = window.getComputedStyle(h1Element)['font-family'];
assert.include(fontFamily.toLowerCase(), "monospace");
```

Vas `h1` element treba da nasledi `green` boju od `body` elementa.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }

</style>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }

</style>
<h1>Hello World!</h1>
```
