---
id: bad82fee1322bd9aedf08721
title: Razumi Apsolutne Naspram Relativnih Jedinica
challengeType: 0
videoUrl: 'https://scrimba.com/c/cN66JSL'
forumTopicId: 301089
dashedName: understand-absolute-versus-relative-units
---

# --description--

U prethodnim izazovima margina i padding elemenata postavljani su u pikselima (`px`). Pikseli su tip jedinice duzine koji govori pregledaču kako da meri i rasporedi element. Pored `px`, CSS nudi vise razlicitih jedinica duzine.

Dva osnovna tipa jedinica duzine su apsolutne i relativne jedinice. Apsolutne jedinice su vezane za fizicke mere. Na primer, `in` i `mm` predstavljaju inče i milimetre. Apsolutne jedinice na ekranu priblizno odgovaraju stvarnoj meri, uz razlike zbog rezolucije ekrana.

Relativne jedinice, kao `em` i `rem`, zavise od neke druge vrednosti duzine. Na primer, `em` zavisi od velicine fonta elementa. Ako ga koristite za `font-size`, racuna se u odnosu na velicinu fonta roditelja.

**Napomena:** Postoji vise relativnih jedinica koje zavise od velicine viewport-a. One se obradjuju u delu Responsive Web Design Principles.

# --instructions--

Dodajte svojstvo `padding` elementu sa klasom `red-box` i postavite ga na `1.5em`.

# --hints--

Vasa `red-box` klasa treba da ima svojstvo `padding`.

```js
const redBox =document.querySelector('.red-box'); 
const style = window.getComputedStyle(redBox); 

assert.notEqual(style['padding-top'], '0px');
assert.notEqual(style['padding-right'], '0px');
assert.notEqual(style['padding-bottom'], '0px');
assert.notEqual(style['padding-left'], '0px');
```

Vasa `red-box` klasa treba da postavi `padding` na 1.5em.

```js
assert.match(__helpers.removeCssComments(code), /\.red-box\s*?{[\s\S]*padding\s*:\s*?1\.5em/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;

  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;
    padding: 1.5em;
  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```
