---
id: bad87fee1348bd9aedf08721
title: Koristi Hex Kod za Mesanje Boja
challengeType: 0
videoUrl: 'https://scrimba.com/c/cK89PhP'
forumTopicId: 18359
dashedName: use-hex-code-to-mix-colors
---

# --description--

Podsetimo se: hex kodovi koriste 6 heksadecimalnih cifara za predstavljanje boja, po dve za crvenu (R), zelenu (G) i plavu (B).

Od ove tri ciste boje (crvena, zelena i plava), mozemo menjati kolicinu svake i napraviti preko 16 miliona boja.

Na primer, narandzasta boja je cista crvena pomesana sa malo zelene i bez plave. U hex kodu to je `#FFA500`.

Cifra `0` je najniza vrednost u hex kodu i predstavlja potpuno odsustvo boje.

Cifra `F` je najvisa vrednost u hex kodu i predstavlja najvecu mogucu jacinu boje.

# --instructions--

Zamenite nazive boja u nasem `style` delu njihovim odgovarajucim hex kodovima.

<table><tbody><tr><th>Boja</th><th>Heksadecimalni kod</th></tr><tr><td>Dodger Blue</td><td><code>#1E90FF</code></td></tr><tr><td>Zelena</td><td><code>#00FF00</code></td></tr><tr><td>Narandžasta</td><td><code>#FFA500</code></td></tr><tr><td>Crvena</td><td><code>#FF0000</code></td></tr></tbody></table>

# --hints--

Vas `h1` sa tekstom `I am red!` treba da ima crvenu `color` vrednost.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Hex kod za crvenu treba da se koristi umesto reci `red`.

```js
assert.match(code, /\.red-text\s*?{\s*?color\s*:\s*?(#FF0000|#F00)\s*?;?\s*?}/gi);
```

Vas `h1` sa tekstom `I am green!` treba da ima zelenu `color` vrednost.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

Hex kod za zelenu treba da se koristi umesto reci `green`.

```js
assert.match(code, /\.green-text\s*?{\s*?color\s*:\s*?(#00FF00|#0F0)\s*?;?\s*?}/gi);
```

Vas `h1` sa tekstom `I am dodger blue!` treba da ima dodger blue `color` vrednost.

```js
const blueText = document.querySelector('.dodger-blue-text');
const color = window.getComputedStyle(blueText)['color']; 
assert.strictEqual(color, 'rgb(30, 144, 255)');
```

Hex kod za dodger blue treba da se koristi umesto reci `dodgerblue`.

```js
assert.match(code, /\.dodger-blue-text\s*?{\s*?color\s*:\s*?#1E90FF\s*?;?\s*?}/gi);
```

Vas `h1` sa tekstom `I am orange!` treba da ima narandzastu `color` vrednost.

```js
const orangeText = document.querySelector('.orange-text');
const color = window.getComputedStyle(orangeText)['color']; 

assert.strictEqual(color, 'rgb(255, 165, 0)');
```

Hex kod za narandzastu treba da se koristi umesto reci `orange`.

```js
assert.match(code, /\.orange-text\s*?{\s*?color\s*:\s*?#FFA500\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: black;
  }
  .green-text {
    color: black;
  }
  .dodger-blue-text {
    color: black;
  }
  .orange-text {
    color: black;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #FF0000;
  }
  .green-text {
    color: #00FF00;
  }
  .dodger-blue-text {
    color: #1E90FF;
  }
  .orange-text {
    color: #FFA500;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```
