---
id: bad87fee1348bd9aede08718
title: Koristi RGB Vrednosti za Bojenje Elemenata
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkp2fr'
forumTopicId: 18369
dashedName: use-rgb-values-to-color-elements
---

# --description--

Jos jedan nacin da predstavite boje u CSS-u je koriscenje `RGB` vrednosti.

`RGB` vrednost za crnu boju izgleda ovako:

```css
rgb(0, 0, 0)
```

`RGB` vrednost za belu boju izgleda ovako:

```css
rgb(255, 255, 255)
```

Umesto koriscenja sest hexadecimalnih cifara kao kod hex koda, u `RGB` formatu navodite jacinu svake boje brojevima od 0 do 255.

Ako izracunate, dve cifre za jednu boju daju 16 puta 16, sto je ukupno 256 vrednosti. Zato `RGB`, koji broji od nule, ima isti broj mogucih vrednosti kao i hex kod.

Evo primera kako menjate boju pozadine `body` elementa u narandzastu koristeci njen RGB kod.

```css
body {
  background-color: rgb(255, 165, 0);
}
```

# --instructions--

Hajde da zamenimo hex kod u boji pozadine naseg `body` elementa RGB vrednoscu za crnu: `rgb(0, 0, 0)`.

# --hints--

Vas `body` element treba da ima crnu pozadinu.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Treba da koristite `rgb` da postavite crnu boju pozadine `body` elementa.

```js
assert.match(__helpers.removeCssComments(code), /rgb\s*\(\s*0\s*,\s*0\s*,\s*0\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #F00;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: rgb(0, 0, 0);
  }
</style>
```
