---
id: bad87fee1348bd9aedf08719
title: Koristi Skraceni Hex Kod
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpKAm'
forumTopicId: 18338
dashedName: use-abbreviated-hex-code
---

# --description--

Mnogi ljudi se osecaju preoptereceno mogucnoscu izbora vise od 16 miliona boja. Takodje je tesko pamtiti hex kodove. Srecom, mozete ih skratiti.

Na primer, hex kod za crvenu boju `#FF0000` moze se skratiti na `#F00`. Ovaj skraceni oblik koristi jednu cifru za crvenu, jednu za zelenu i jednu za plavu.

Ovo smanjuje ukupan broj mogucih boja na oko 4.000. Ali pregledaci ce `#FF0000` i `#F00` tumaciti kao potpuno istu boju.

# --instructions--

Hajde, pokusajte da koristite skracene hex kodove da obojite odgovarajuce elemente.

<table><tbody><tr><th>Boja</th><th>Kratki heksadecimalni kod</th></tr><tr><td>Cijan</td><td><code>#0FF</code></td></tr><tr><td>Zelena</td><td><code>#0F0</code></td></tr><tr><td>Crvena</td><td><code>#F00</code></td></tr><tr><td>Fuksija</td><td><code>#F0F</code></td></tr></tbody></table>

# --hints--

Vas `h1` element sa tekstom `I am red!` treba da ima crvenu `color` vrednost.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Skraceni hex kod za crvenu treba da se koristi umesto `#FF0000`.

```js
assert.match(code,/\.red-text\s*?{\s*?color\s*:\s*?#F00\s*?;?\s*?}/gi);
```

Vas `h1` element sa tekstom `I am green!` treba da ima zelenu `color` vrednost.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

Skraceni hex kod za zelenu treba da se koristi umesto `#00FF00`.

```js
assert.match(__helpers.removeCssComments(code), /\.green-text\s*?{\s*?color\s*:\s*?#0F0\s*?;?\s*?}/gi);
```

Vas `h1` element sa tekstom `I am cyan!` treba da ima cyan `color` vrednost.

```js
const cyanText = document.querySelector('.cyan-text');
const color = window.getComputedStyle(cyanText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 255)');
```

Skraceni hex kod za cyan treba da se koristi umesto `#00FFFF`.

```js
assert.match(__helpers.removeCssComments(code), /\.cyan-text\s*?{\s*?color\s*:\s*?#0FF\s*?;?\s*?}/gi);
```

Vas `h1` element sa tekstom `I am fuchsia!` treba da ima fuchsia `color` vrednost.

```js
const fuchsiaText = document.querySelector('.fuchsia-text');
const color = window.getComputedStyle(fuchsiaText)['color']; 
assert.strictEqual(color,'rgb(255, 0, 255)');
```

Skraceni hex kod za fuchsia treba da se koristi umesto `#FF00FF`.

```js
assert.match(__helpers.removeCssComments(code), /\.fuchsia-text\s*?{\s*?color\s*:\s*?#F0F\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: #000000;
  }
  .fuchsia-text {
    color: #000000;
  }
  .cyan-text {
    color: #000000;
  }
  .green-text {
    color: #000000;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #F00;
  }
  .fuchsia-text {
    color: #F0F;
  }
  .cyan-text {
    color: #0FF;
  }
  .green-text {
    color: #0F0;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```
