---
id: bad87fee1348bd9aedf06756
title: Pregazi Deklaracije Klase Inline Stilovima
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDRha'
forumTopicId: 18252
dashedName: override-class-declarations-with-inline-styles
---

# --description--

Dakle, potvrdili smo da id deklaracije pregaze deklaracije klase, bez obzira gde su deklarisane u CSS-u vaseg `style` elementa.

Postoji jos nacina za pregazivanje CSS-a. Secate li se inline stilova?

# --instructions--

Koristite inline stil da pokusate da nas `h1` element bude beo. Podsetnik, inline stil izgleda ovako:

```html
<h1 style="color: green;">
```

Ostavite klase `blue-text` i `pink-text` na svom `h1` elementu.

# --hints--

Vas `h1` element treba da ima klasu `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Vas `h1` element treba da ima klasu `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Vas `h1` element treba da ima id `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

Vas `h1` element treba da ima inline stil.

```js
assert.exists(document.querySelector('h1[style]'));
```

Vas `h1` element treba da bude bele boje.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
