---
id: bad87fee1348bd9aedf08726
title: Koristi Hex Kod za Specificne Boje
challengeType: 0
videoUrl: 'https://scrimba.com/c/c8W9mHM'
forumTopicId: 18350
dashedName: use-hex-code-for-specific-colors
---

# --description--

Da li ste znali da postoji vise nacina za predstavljanje boja u CSS-u? Jedan od njih se zove heksadecimalni kod, ili skraceno hex kod.

Uobicajeno koristimo <dfn>decimalne cifre</dfn>, odnosno brojeve u osnovi 10, koji koriste simbole od 0 do 9 za svaku cifru. <dfn>Heksadecimalni</dfn> (ili <dfn>hex</dfn>) sistem je brojcani sistem osnove 16. To znaci da koristi 16 razlicitih simbola. Kao i kod decimalnog sistema, simboli 0-9 predstavljaju vrednosti od nula do devet, a zatim A, B, C, D, E, F predstavljaju vrednosti od deset do petnaest. Ukupno, 0 do F mogu predstavljati cifre u heksadecimalnom sistemu, sto daje 16 mogucih vrednosti. Vise o <a href="https://www.freecodecamp.org/news/hexadecimal-number-system/" target="_blank" rel="noopener noreferrer nofollow">heksadecimalnim brojevima procitajte ovde</a>.

U CSS-u mozemo koristiti 6 heksadecimalnih cifara za predstavljanje boja, po dve za crvenu (R), zelenu (G) i plavu (B). Na primer, `#000000` je crna i ujedno najniza moguca vrednost. Vise o <a href="https://www.freecodecamp.org/news/rgb-color-html-and-css-guide/#whatisthergbcolormodel" target="_blank" rel="noopener noreferrer nofollow">RGB sistemu boja procitajte ovde</a>.

```css
body {
  color: #000000;
}
```

# --instructions--

Zamenite rec `black` u boji pozadine naseg `body` dela i predstavite je njenim hex kodom, `#000000`.

# --hints--

Vas `body` deo treba da ima crnu `background-color` vrednost.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Hex kod za crnu treba da se koristi umesto reci `black`.

```js
assert.match(code, /body\s*{(([\s\S]*;\s*?)|\s*?)background.*\s*:\s*?#000(000)?((\s*})|(;[\s\S]*?}))/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: #000000;
  }
</style>
```
