---
id: bad87fee1348bd9aedf07756
title: Pregazi Sve Ostale Stilove Koristeci Important
challengeType: 0
videoUrl: 'https://scrimba.com/c/cm24rcp'
forumTopicId: 18249
dashedName: override-all-other-styles-by-using-important
---

# --description--

Cestitamo. Pokazali smo da inline stilovi pregaze sve CSS deklaracije iz vaseg `style` elementa.

Ali sacekajte. Postoji jos jedan poslednji nacin za pregazivanje CSS-a. To je najjaci nacin od svih. Pre nego sto to uradimo, hajde da kazemo zasto biste uopste zeleli da pregazite CSS.

U mnogim slucajevima koristicete CSS biblioteke. One mogu slucajno pregaziti vas CSS. Kada apsolutno morate osigurati da element ima tacno odredjeni CSS, mozete koristiti `!important`.

Vratimo se na deklaraciju nase `pink-text` klase. Setite se da je nasu `pink-text` klasu pregazila kasnija deklaracija klase, id deklaracija i inline stil.

# --instructions--

Dodajte kljucnu rec `!important` na deklaraciju boje u svojoj `pink-text` klasi da 100% osigurate da ce vas `h1` biti roze boje.

Primer kako to da uradite je:

```css
color: red !important;
```

# --hints--

Vas `h1` element treba da ima klasu `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Vas `h1` element treba da ima klasu `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Vas `h1` element treba da ima id `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

Vas `h1` element treba da ima inline stil `color: white`.

```js
const commentessCode = __helpers.removeHtmlComments(code);
assert.match(commentessCode, /<h1.*style/gi);
assert.match(commentessCode, /<h1.*style.*color\s*?:/gi);
```

Vasa deklaracija klase `pink-text` treba da ima kljucnu rec `!important` da pregazi sve ostale deklaracije.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*?\{[\s\S]*?color:.*pink.*!important\s*;?[^\.]*\}/g);
```

Vas `h1` element treba da bude roze boje.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color'];
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink !important;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
