---
id: bad88fee1348bd9aedf08825
title: Podesi Unutrasnji Razmak Elementa
challengeType: 0
videoUrl: 'https://scrimba.com/c/cED8ZC2'
forumTopicId: 301083
dashedName: adjust-the-padding-of-an-element
---

# --description--

Hajde sada da na kratko ostavimo Cat Photo App i naucimo vise o stilizovanju HTML-a.

Mozda ste vec primetili, ali svi HTML elementi su u osnovi mali pravougaonici.

Tri vazna svojstva kontrolisu prostor oko svakog HTML elementa: `padding`, `border` i `margin`.

`padding` elementa kontrolise kolicinu prostora izmedju sadrzaja elementa i njegovog `border` okvira.

Ovde vidimo da su plavi i crveni okvir smesteni unutar zutog okvira. Primetite da crveni okvir ima veci `padding` od plavog.

Kada povecate `padding` plavog okvira, povecava se razmak izmedju teksta i ivice koja ga okruzuje.

# --instructions--

Promenite `padding` svog plavog okvira tako da odgovara `padding` vrednosti crvenog okvira.

# --hints--

Vasa `blue-box` klasa treba da postavi `padding` na `20px`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 10px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
