---
id: 68e2c945cc1d8e778152be31
title: Izgradnja profila igrača
challengeType: 27
dashedName: build-a-player-interface
---

# --description--

**Cilj:** Popunite priče korisnika ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Treba da definišete abstraktnu klasu nazvanu `Player` koja nasleđuje od klase `abc.ABC`.
2. Klasa `Player` treba da ima metodu `__init__` koja postavlja:

- Svojstvo elementa za `moves` da bude prazna lista.
- Svojstvo elementa za `position` da bude `(0, 0)`.
- Svojstvo elementa za `path` da bude lista sa početnom pozicijom.

1. Klasa za `Player` mora imati metodu nazvanu `make_move`, koja:

- Koristi `random.choice` da dobije nasumičnu smernicu iz atributa elementa `moves` (definisanog u klasi stvarnih objekata).
   - Dodaje vrednost iz definisane smernice na trenutnu poziciju i modifikuje atribut elementa `position`.
   - Dodaje novi tuple `position` na atribut elementa `path`.
   - Vraća novi `position`.

1. Klasa objekata `Player` mora imati metod apstrakcije pod nazivom `level_up` koja će biti implementirana u konkretnim klasama.
1. Treba definisati klasu objekata `Pawn` koja nasleđuje od klase objekata `Player`.
1. Klasa objekata `Pawn` mora koristiti `super()` za pozivanje metode roditelja `__init__`, a zatim postaviti atribut `moves` na listu tuple elemenata koji predstavljaju koordinate za `x, y`.
1. Svaki koordinatni tuple treba da predstavlja kretanje jedinice `1` u sledećim pravcima: gore, dole, levo, desno.
1. Klasa objekata `Pawn` mora implementirati apstraktnu metodu `level_up` dodavanjem drugih kretanja atributu `moves`. Dodata kretanja treba da predstavljaju četiri kretanja susednog pravca (na primer, jedinica `1` dole zajedno sa jedinicom `1` levo).

**Napomena:** Standardni biblioteka moduli treba uvesti bez korišćenja aliasa. Testovi povezani sa klasom `Player` će propasti dok se klasa `Pawn` ne može inicijalizovati.

# --hints--

Trebalo bi da imaš klasu objekata koja se zove `Player`.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Player")`) })
```

Klasa objekata za `Player` mora da nasledi od klase objekata `ABC` iz modula `abc`.

```js
({ test: () => runPython(`
from abc import ABC
import ast
cls = _Node(_code).find_class("Player")
assert cls.inherits_from("ABC") or cls.inherits_from("abc.ABC")
`) })
```

Klasa za stavke `Player` mora imati metodu od `__init__`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("__init__")`) })
```

Putanja za `Player` od `__init__` bi trebalo da ima jedan parametar `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("__init__").has_args("self")`) })
```

Putanja za `Player` od `__init__` treba da postavi `self.moves` kao praznu listu.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.moves = []") or init.has_stmt("self.moves = list()")
`) })
```

Putanja za `Player` za `__init__` treba da postavi `self.position` na `(0, 0)`.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.position = (0, 0)")
`) })
```

Putanja za `Player` za `__init__` treba da postavi `self.path` kao listu sa prvom pozicijom.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.path = [self.position]")
`) })
```

Klasa objekata za `Player` bi trebalo da ima metodu `make_move`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("make_move")`) })
```

Putanja za `Player` od `make_move` treba da ima jedan parametar `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("make_move").has_args("self")`) })
```

Putanja `Player` za `make_move` treba da koristi `random.choice` kako bi dobila nasumičnu orijentaciju iz svojstva elementa `moves`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").block_has_call("choice", "make_move")`) })
```

Putanja `Player` za `make_move` treba da ispravi svojstvo elementa `position` dodavanjem koordinata nasumično izabranog pravca.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.position == (1, 2)
  p.make_move()
  assert p.position == (2, 4)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

Put `Player` od `make_move` treba da doda novi tuple `position` na svojstvo elementa `path`.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.path == [(0, 0), (1, 2)]
  p.make_move()
  assert p.path == [(0, 0), (1, 2), (2, 4)]

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

Putanja za `Player` od `make_move` bi trebalo da vrati svojstvo elementa `position` koje je ažurirano.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  assert p.make_move() == (1, 2)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

Klasa objekata `Player` treba da ima metodu `level_up`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("level_up")`) })
```

Putanja za `Player` koja je za `level_up` treba da ima jedan parametar `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("level_up").has_args("self")`) })
```

Način za `Player` od `level_up` mora biti način apstrakcije.

```js
({ test: () => runPython(`
target = _Node(_code).find_class("Player").find_function("level_up")
assert target.has_decorators("abstractmethod") or target.has_decorators("abc.abstractmethod")
`) })
```

Klasa objekata ``Player`` bi trebalo da bude klasa apstrakcije.

```js
({ test: () => runPython(`
try:
  Player()
except TypeError as e:
  assert str(e) == "Can't instantiate abstract class Player with abstract method level_up"
else:
  assert False, "Player class should not be instantiable"
`) })
```

Trebalo bi da imaš klasu objekata koja se zove `Pawn`.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Pawn")`) })
```

Klasa objekata `Pawn` treba da nasledi od klase objekata `Player`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").inherits_from("Player")`) })
```

Klasa za `Pawn` mora imati metodu `__init__`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("__init__")`) })
```

Putanja za `Pawn` od `__init__` mora imati jedan parametar `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_args("self")`) })
```

Putanja za `Pawn` od `__init__` mora da pozove putanju `__init__` roditelja koristeći akciju `super`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_stmt("super().__init__()")`) })
```

Putanja od `Pawn` za `__init__` treba da postavi atribut elementa `moves` na listu tupola koji predstavljaju koordinate `x, y`, gde svaki koordinatni tupel predstavlja kretanje jedinice `1` u sledećim pravcima: gore, dole, levo, desno.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]
assert isinstance(p.moves, list)
assert len(p.moves) == 4
assert all(move in p.moves for move in moves)
`) })
```

Klasa objekata za `Pawn` bi trebalo da ima metodu za `level_up`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("level_up")`) })
```

Putanja za `Pawn` od `level_up` treba da doda četiri bočna hareketa jedinice `1` na svojstvo elementa `moves`.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1), (1, 1), (1, -1), (-1, 1), (-1, -1)]
p.level_up()
assert isinstance(p.moves, list)
assert len(p.moves) == 8
assert all(move in p.moves for move in moves)
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
from abc import ABC, abstractmethod
from random import choice


class Player(ABC):
    def __init__(self):
        self.moves = []
        self.position = (0, 0)
        self.path = [self.position]

    def make_move(self):
        x, y = self.position
        move = choice(self.moves)
        new_x = move[0] + x
        new_y = move[1] + y
        self.position = (new_x, new_y)
        self.path.append(self.position)
        return self.position

    @abstractmethod
    def level_up(self):
        pass


class Pawn(Player):
    def __init__(self):
        super().__init__()
        self.moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]

    def level_up(self):
        self.moves.extend([(1, 1), (1, -1), (-1, 1), (-1, -1)])
```
