---
id: 680900675ae3d54ee19590c9
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Unutar vaĹˇeg komponenta za `select`, morate kreirati opcije za korisnika da izabere svoj izvor napajanja. `select` treba da ima sedam opcija navedenih ispod:

| Izvor napajanja                  |
|-------------------------------|
| `Select one`                    |
| `Bitten by a strange creature`  |
| `Radioactive exposure`          |
| `Science experiment`           |
| `Alien heritage `               |
| `Ancient artifact discovery`    |
| `Other`                         |

Prvi izbor (`Select one`) mora imati prazan niz znakova kao svojstvo elementa za `value` (tj., `value=''`). Ovo funkcioniĹˇe kao zamena u UI.

# --hints--

Trebalo bi da kreirate prvi element sa tekstom `Select one` unutar vaĹˇeg elementa `select` u `option`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[0]?.textContent, "Select one");
```

Trebalo bi mu dati prvi element sa atributom elemenata `option` i atributom elemenata `value` sa praznom nizom karaktera.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[0]?.getAttribute('value'), '');
```

Trebalo bi da kreirate drugi element `option` sa tekstom `Bitten by a strange creature` unutar vaĹˇeg elementa `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[1]?.textContent, "Bitten by a strange creature");
```

Trebalo bi da kreiraĹˇ treÄ‡i element za `option` sa tekstom `Radioactive exposure` unutar tvog elementa `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[2]?.textContent, "Radioactive exposure");
```

Treba da kreirate ÄŤetvrti element za `option` sa tekstom `Science experiment` unutar vaĹˇeg elementa od `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[3]?.textContent, "Science experiment");
```

Trebalo bi da kreiraĹˇ peto polje za `option` sa tekstom `Alien heritage` unutar tvog polja `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[4]?.textContent, "Alien heritage");
```

Treba da kreiraĹˇ Ĺˇestu komponentu za `option` sa tekstom `Ancient artifact discovery` unutar svoje komponente `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[5]?.textContent, "Ancient artifact discovery");
```

Trebali biste kreirati sedmi element od `option` koji sadrĹľi tekst `Other` unutar vaĹˇeg elementa `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[6]?.textContent, "Other");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
          <label>
            Hero Name
            <input
              type='text'
              value={heroName}
              onChange={e => setHeroName(e.target.value)}
            />
          </label>
          <label>
            Real Name
            <input
              type='password'
              value={realName}
              onChange={e => setRealName(e.target.value)}
            />
          </label>
        </div>
          <label className='section column'>
            How did you get your powers?
            --fcc-editable-region--
            <select>
              
            </select>
            --fcc-editable-region--
          </label>
      </form>
    </div>
  )
};
```
