---
id: 587d825c367417b2b2512c8f
title: Implementiraj merge sort
challengeType: 1
forumTopicId: 301614
dashedName: implement-merge-sort
---

# --description--

Još jedan uobičajeni algoritam prosečnog slučaja za sortiranje je merge sort. Kao i quick sort, merge sort takođe koristi pristup podeli-pa-vladaj, odnosno rekurzivno sortiranje niza podataka. Koristi činjenicu da je lako sortirati dva niza podataka ako je svaki od njih već unapred sortiran. Ali počinjemo samo sa jednim nizom kao ulazom, pa kako dobijamo dva sortirana niza odatle? Pa, možemo podeliti originalni ulaz više puta dok ne stignemo do osnovnog slučaja - malog niza sa samo jednim elementom. Niz sa jednim elementom je po prirodi sortiran, tako da možemo početi da spajamo. Ovo spajanje će ukloniti rekurzivni poziv koji je podelio originalni niz, i na kraju će proizvesti konačan sortirani niz sa svim elementima. Koraci merge sorta, dakle, su:

**1)** Podeli niz ulaznih podataka više puta na pola dok ne ostane samo mali niz sa jednim elementom.

**2)** Spoji svaki mali sortirani niz zajedno da bi se proizveo konačni sortirani niz.

Merge sort je dobar način sortiranja, sa vremenskom složenošću *O(nlog(n))*. Ovaj algoritam je popularan zato što je efikasan i lak za implementaciju.

S druge strane, ovo će biti poslednji algoritam za sortiranje koji ćemo naučiti ovde. Međutim, kasnije u delu o strukturama podataka stabala ćemo opisati heap sort, drugi dobar način sortiranja koji zahteva binarno stablo u svom izvršavanju.

**Uputstva:** Napiši funkciju `mergeSort` koja prima niz celih brojeva kao ulaz i vraća niz tih brojeva sortiran od najmanjeg do najvećeg. Dobar način da se ovo implementira je pisanje jedne funkcije, na primer `merge`, koja se bavi spajanjem dva sortirana niza, i druge funkcije, na primer `mergeSort`, koja se bavi rekurzijom koja proizvodi male nizove sa jednim elementom koji se prosleđuju spajanju. Srećno!

# --hints--

`mergeSort` treba da bude funkcija.

```js
assert.isFunction(mergeSort);
```

`mergeSort` treba da vrati sortirani niz (od najmanjeg do najvećeg).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    mergeSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`mergeSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` treba da vrati niz koji nije promenjen osim redosleda.

```js
assert.sameMembers(
  mergeSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`mergeSort` ne sme da koristi metod koji je ugrađen u `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    mergeSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function mergeSort(array) {
  if (array.length === 1) {
    return array;
  } else {
    const splitIndex = Math.floor(array.length / 2);
    return merge(
      mergeSort(array.slice(0, splitIndex)),
      mergeSort(array.slice(splitIndex))
    );
  }

  // Merge two sorted arrays
  function merge(array1, array2) {
    let merged = [];
    while (array1.length && array2.length) {
      if (array1[0] < array2[0]) {
        merged.push(array1.shift());
      } else if (array1[0] > array2[0]) {
        merged.push(array2.shift());
      } else {
        merged.push(array1.shift(), array2.shift());
      }
    }

    // After looping ends, one array is empty, and other array contains only
    // values greater than all values in `merged`
    return [...merged, ...array1, ...array2];
  }
}

mergeSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]);
```
