---
id: 587d8259367417b2b2512c85
title: Implementirajte selection sort
challengeType: 1
forumTopicId: 301616
dashedName: implement-selection-sort
---

# --description--

Ovde ćemo implementirati selection sort. Selection sort radi biranjem najmanje vrednosti u listi i zamenom njenog mesta sa prvom vrednošću u listi. Zatim počinje od drugog mesta, bira najmanju vrednost u preostaloj listi i menja njeno mesto sa drugim elementom. Ovaj postupak se ponavlja kroz listu i menja elemente sve dok se ne dođe do kraja liste. Sada je lista sortirana. Selection sort ima vremensku složenost kvadratnog reda u svim slučajevima.

**Uputstva**: Napišite funkciju `selectionSort` koja prima niz celih brojeva kao ulaz i vraća niz ovih celih brojeva sortiranih od najmanjeg do najvećeg.

# --hints--

`selectionSort` treba da bude funkcija.

```js
assert.isFunction(selectionSort);
```

`selectionSort` treba da vrati sortiran niz (od najmanjeg do najvećeg).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}

assert.isTrue(
  isSorted(
    selectionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`selectionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` treba da vrati niz koji se ne menja osim redosleda.

```js
assert.sameMembers(
  selectionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` ne treba da koristi ugrađenu `.sort()` metodu.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function selectionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
