---
id: 587d8259367417b2b2512c86
title: Implementirajte insertion sort
challengeType: 1
forumTopicId: 301613
dashedName: implement-insertion-sort
---

# --description--

Sledeći algoritam sortiranja koji ćemo analizirati je insertion sort. Ovaj algoritam radi kreiranjem sortiranog niza na početku liste. Počinje sa sortiranim nizom od prvog elementa. Zatim razmatra sledeći element i pomera ga unazad unutar sortiranog niza sve dok ne zauzme svoje sortirano mesto. Nastavlja kroz listu i pomera nove elemente unazad unutar sortiranog dela sve dok ne dođe do kraja. Ovaj algoritam ima vremensku složenost kvadratnog reda u prosečnom i najgorem slučaju.

**Uputstva:** Napišite funkciju `insertionSort` koja prima niz celih brojeva kao ulaz i vraća niz ovih celih brojeva sortiranih od najmanjeg do najvećeg.

# --hints--

`insertionSort` treba da bude funkcija.

```js
assert.isFunction(insertionSort);
```

`insertionSort` treba da vrati sortiran niz (od najmanjeg do najvećeg).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    insertionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`insertionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` treba da vrati niz koji se ne menja osim redosleda.

```js
assert.sameMembers(
  insertionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort([5, 4, 33, 2, 8])` treba da vrati `[2, 4, 5, 8, 33]`.

```js
assert.deepEqual(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` ne treba da koristi ugrađenu `.sort()` metodu.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function insertionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
