---
id: 587d825a367417b2b2512c89
title: Implementirajte quick sort
challengeType: 1
forumTopicId: 301615
dashedName: implement-quick-sort
---

# --description--

Ovde ćemo nastaviti sa algoritmom sortiranja srednjeg nivoa: quick sort. Quick sort je efikasan, rekurzivni pristup "podeli i vladaj" (divide-and-conquer) za sortiranje niza podataka. U ovom pristupu, vrednost pivot-a se bira u početnom nizu. Niz se deli u dva manja niza sa vrednostima manjim i većim od vrednosti pivot-a. Zatim spajamo rezultate rekurzivnog poziva quick sort algoritma na oba manja niza. Ovo se nastavlja sve dok ne dođemo do osnovnog slučaja praznog niza ili niza sa jednim elementom, koji vraćamo. Smanjenje rekurzivnih poziva nam vraća sortiran niz.

Quick sort je veoma efikasan algoritam sortiranja, koji pruža performanse *O(nlog(n))* u proseku. Takođe je jednostavan za implementaciju. Ove osobine ga čine popularnim i korisnim algoritmom za sortiranje.

**Uputstva:** Napišite funkciju `quickSort` koja prima niz celih brojeva kao ulaz i vraća niz ovih celih brojeva sortiranih od najmanjeg do najvećeg. Iako je izbor pivot vrednosti važan, bilo koji pivot će biti dovoljan za naše svrhe ovde. Za jednostavnost, prvi ili poslednji element se može koristiti.

# --hints--

`quickSort` treba da bude funkcija.

```js
assert(typeof quickSort == 'function');
```

`quickSort` treba da vrati sortiran niz (od najmanjeg do najvećeg).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    quickSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`quickSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` treba da vrati niz koji se ne menja osim redosleda.

```js
assert.sameMembers(
  quickSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quickSort` ne treba da koristi ugrađenu `.sort()` metodu.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quickSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function quickSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function quickSort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quickSort(lesser), ...equal, ...quickSort(greater)];
  }
}
```
