---
id: 61abc7ebf3029b56226de5b6
title: Implementirajte binarno pretraživanje
challengeType: 1
forumTopicId: 487618
dashedName: implement-binary-search
---

# --description--

Binarno pretraživanje je algoritam sa efikasnošću **O(log(n))** za pretraživanje sortiranog niza podataka radi pronalaženja elementa. Radi koristeći sledeće korake:

1. Pronađite srednju `value` sortiranog niza. Ako `value == target` vratite `true` (Vrednost je pronađena i pretraživanje je završeno).
1. Ako je srednja `value < target`, pretražite desnu polovinu niza u sledećem poređenju.
1. Ako je srednja `value > target`, pretražite levu polovinu niza u sledećem poređenju.
1. Ako nakon pretraživanja celog niza vrednost ne postoji, vratite `false` (Niz je pretražen i vrednost ne postoji u nizu).

Kao što vidite, delite niz podataka više puta, što vam daje efikasnost log(n). U ovom izazovu, želimo da pokažete svoj rad - kako ste došli do ciljane vrednosti... putanju koju ste koristili!

# --instructions--

Napišite funkciju `binarySearch` koja implementira algoritam binarnog pretraživanja nad nizom podataka, vraćajući putanju koju ste koristili (svako poređenje srednje vrednosti) da biste pronašli cilj u nizu.

Funkcija prima sortiran niz celih brojeva i ciljanu vrednost kao ulaz. Vraća niz sadržaj (po redu) srednjih vrednosti koje ste pronašli pri svakoj podeli početnog niza sve dok ne pronađete ciljanu vrednost. Ciljana vrednost treba da bude poslednji element vraćenog niza. Ako vrednost ne postoji, vratite niz karaktera `Value Not Found`.

Na primer, `binarySearch([1,2,3,4,5,6,7], 5)` vratiće `[4,6,5]`.

U ovom izazovu, prilikom deljenja, MORATE koristiti `Math.floor()` pri deljenju: `Math.floor(x/2)`. Ovo će dati doslednu putanju koja se može testirati.

**Napomena:** Sledeći niz će se koristiti u testovima:

```js
const testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
```

# --hints--

`binarySearch` treba da bude funkcija.

```js
assert.isFunction(binarySearch);
```

`binarySearch(testArray, 0)` treba da vrati `[13, 5, 2, 0]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 0), [13, 5, 2, 0]);
```

`binarySearch(testArray, 1)` treba da vrati `[13, 5, 2, 0, 1]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 1), [13, 5, 2, 0, 1]);
```


`binarySearch(testArray, 2)` treba da vrati `[13, 5, 2]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 2), [13, 5, 2]);
```

`binarySearch(testArray, 6)` treba da vrati niz karaktera `Value Not Found`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.strictEqual(binarySearch(_testArray, 6), 'Value Not Found');
```

`binarySearch(testArray, 11)` treba da vrati `[13, 5, 10, 11]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 11), [13, 5, 10, 11])
```

`binarySearch(testArray, 13)` treba da vrati `[13]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 13), [13]);
```

`binarySearch(testArray, 70)` treba da vrati `[13, 19, 22, 49, 70]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 70), [13, 19, 22, 49, 70]);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let arrayPath = [];
  return arrayPath;
}
```



# --solutions--

```js
let binarySearch = (searchList, value) => {
  let arrayPath = [];

  // set initial L - M - R
  let left = 0;
  let right = searchList.length - 1;
  let middle = Math.floor(right / 2);

  // if first comparison finds value
  if (searchList[middle] == value) {
    arrayPath.push(searchList[middle]);
    return arrayPath;
  }

  while (searchList[middle] !== value) {
    // add to output array
    arrayPath.push(searchList[middle]);

    // not found
    if (right < left) {
      return 'Value Not Found';
    }
    // value is in left or right portion of array
    // update L - M - R
    if (searchList[middle] > value) {
      right = middle - 1;
      middle = left + Math.floor((right - left) / 2);
    } else {
      left = middle + 1;
      middle = left + Math.floor((right - left) / 2);
    }

    // if found update output array and exit
    if (searchList[middle] == value) {
      arrayPath.push(searchList[middle]);

      break;
    }
  }
  return arrayPath;
};
```
