---
id: a3f503de51cfab748ff001aa
title: Parovi elemenata
challengeType: 1
forumTopicId: 301617
dashedName: pairwise
---

# --description--

Ako vam je dat niz `arr`, pronađite parove elemenata čiji je zbir jednak drugom argumentu `arg` i vratite zbir njihovih indeksa.

Možete koristiti više parova sa istim brojevima ali sa različitim indeksima. Svaki element treba da koristi najmanji dostupni indeks. Kada je jednom element izabran, ne može se ponovo koristiti za spajanje sa drugim elementom. Na primer, `pairwise([1, 1, 2], 3)` formira par `[2, 1]` koristeći 1 sa indeksom 0 umesto 1 sa indeksom 1, zato što 0+2 < 1+2.

Na primer `pairwise([7, 9, 11, 13, 15], 20)` vraća `6`. Parovi čiji je zbir 20 su `[7, 13]` i `[9, 11]`. Možemo zapisati niz sa indeksima i njihovim vrednostima.

<div style='margin-left: 2em;'>

| Kielezo | 0 | 1 | 2  | 3  | 4  |
| ------- | - | - | -- | -- | -- |
| Thamani | 7 | 9 | 11 | 13 | 15 |

</div>

Ispod ćemo uzeti odgovarajuće indekse i sabrati ih.

<div style='margin-left: 2em;'>

7 + 13 = 20 → Indeksi 0 + 3 = 3  
9 + 11 = 20 → Indeksi 1 + 2 = 3  
3 + 3 = 6 → Vratiti `6`

</div>

# --hints--

`pairwise([1, 4, 2, 3, 0, 5], 7)` treba da vrati 11.

```js
assert.deepEqual(pairwise([1, 4, 2, 3, 0, 5], 7), 11);
```

`pairwise([1, 3, 2, 4], 4)` treba da vrati 1.

```js
assert.deepEqual(pairwise([1, 3, 2, 4], 4), 1);
```

`pairwise([1, 1, 1], 2)` treba da vrati 1.

```js
assert.deepEqual(pairwise([1, 1, 1], 2), 1);
```

`pairwise([0, 0, 0, 0, 1, 1], 1)` treba da vrati 10.

```js
assert.deepEqual(pairwise([0, 0, 0, 0, 1, 1], 1), 10);
```

`pairwise([], 100)` treba da vrati 0.

```js
assert.deepEqual(pairwise([], 100), 0);
```

# --seed--

## --seed-contents--

```js
function pairwise(arr, arg) {
  return arg;
}

pairwise([1,4,2,3,0,5], 7);
```

# --solutions--

```js
function pairwise(arr, arg) {
  var sum = 0;
  arr.forEach(function(e, i, a) {
    if (e != null) {
      var diff = arg-e;
      a[i] = null;
      var dix = a.indexOf(diff);
      if (dix !== -1) {
        sum += dix;
        sum += i;
        a[dix] = null;
      }
    }
  });
  return sum;
}

pairwise([1,4,2,3,0,5], 7);