---
id: 587d7b83367417b2b2512b37
title: Razumeti razliku između konzole freeCodeCamp-a i konzole pretraživača
challengeType: 1
forumTopicId: 301193
dashedName: understanding-the-differences-between-the-freecodecamp-and-browser-console
---

# --description--
Možda ste primetili da neke izazove freeCodeCamp-a imaju sopstvene konzole. Ova konzola radi malo drugačije od konzole pretraživača.

Postoji mnogo načina da se koristi `console` za prikazivanje poruke. `log`, `warn` i `clear` su neki od tih načina. FreeCodeCamp konzola će prikazati samo poruku `log`, dok će konzola pretraživača prikazati sve poruke. Kada napravite izmene u svom kodu, automatski će se pokrenuti i pokazati memoriju/istoriju. FreeCodeCamp konzola se briše svaki put kada izvršite svoj kod.

# --instructions--
Prvo, otvori konzolu pretraživača da vidiš logove. Da bi to uradio, možeš desnim klikom na gornju navigacionu traku freeCodeCampa i kliknuti `inspect` u većini pregledača. Zatim pronađi karticu `console` u otvaranom prozoru.

Nakon toga, koristi `console.log` da napiše istoriju promena za `output`. Proveri obe konzole kako bi video log. Na kraju, koristi `console.clear` nakon što tvoj log obriše konzolu pretraživača. Proveri razliku između te dve konzole.

# --hints--
Trebalo bi da koristiš `console.log()` za štampanje promena od `output`.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/console\.log\(output\)/));
```

Trebalo bi da koristite `console.clear()` za brisanje konzole pretraživača.

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/console.clear\(\)/)
);
```

Treba da obrišeš sadržaj konzole nakon svog logovanja.

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/console\.log\(output\)[\s\S]*console.clear\(\)/)
);
```

# --seed--

## --seed-contents--

```js
let output = "Get this to show once in the freeCodeCamp console and not at all in the browser console";

```

# --solutions--

```js
let output = "Get this to show once in the freeCodeCamp console and not at all in the browser console";

console.log(output);
console.clear();
```
