---
id: 587d7b85367417b2b2512b39
title: Detektujte nedostatak standardnih zagradi za otvaranje i zatvaranje nakon poziva funkcije
challengeType: 1
forumTopicId: 301185
dashedName: catch-missing-open-and-closing-parenthesis-after-a-function-call
---

# --description--
Kada funkcija ili metoda ne prima argumente, možda zaboravite da uključite standardne otvorene i zatvorene vitičaste zagrade (prazno) prilikom pozivanja. Često se rezultat poziva funkcije skladišti u varijablu za dalju korišćenje u vašem kodu. Ova greška se može otkriti tako što ćete vrednosti varijabli (ili njihove tipove) napisati u konzolu i videti da je jedna postavljena kao referenca funkcije, umesto vrednosti koju bi funkcija trebalo da vrati.

Kriterijumi u sledećem primeru su različiti:

```js
function myFunction() {
  return "You rock!";
}
let varOne = myFunction;
let varTwo = myFunction();
```

Hapa `varOne` je funkcija `myFunction`, a na `varTwo` je sekvenca slova `You rock!`.

# --instructions--
Ispravite kod ili parametar `result` da mu se postavi vrednost koja je vraćena pozivom funkcije `getNine`.

# --hints--
Vaš kod treba da ispravi parametar `result` ili da je postavljen na broj koji funkcija `getNine` vraća.

```js
assert(result == 9);
```

Tvoj kod treba da pozove funkciju `getNine`.

```js
assert(__helpers.removeJSComments(code).match(/getNine\(\)/g).length == 2);
```

# --seed--

## --seed-contents--

```js
function getNine() {
  let x = 6;
  let y = 3;
  return x + y;
}

let result = getNine;
console.log(result);
```

# --solutions--

```js
function getNine() {
 let x = 6;
 let y = 3;
 return x + y;
}

let result = getNine();
console.log(result);
```
