---
id: 587d7b86367417b2b2512b3c
title: Budi oprezan kada ponovo inicijalizuješ promenljive unutar petlje
challengeType: 1
forumTopicId: 301194
dashedName: use-caution-when-reinitializing-variables-inside-a-loop
---

# --description--
Ponekad je važno sačuvati podatke, povećati pokazivač ili resetovati parametre unutar ciklusa. Problem se može pojaviti tamo gde parametri treba da se ponovo inicijalizuju, ali to ne rade, ili obrnuto. Ovo je posebno opasno ako slučajno resetujete parametar koji se koristi za konačno stanje, što uzrokuje beskonačan ciklus.

Pisanje vrednosti parametara svakog ciklusa korišćenjem `console.log()` može otkriti neispravno ponašanje povezano sa resetovanjem, ili neuspešnim resetovanjem parametra.

# --instructions--
Funkcija ovde bi trebalo da kreira dvodimenzionalni niz podataka sa redovima `m` i kolonama `n`, popunjenima nulama. Nažalost, ne daje očekivani rezultat jer parametar `row` nije ponovo inicijalizovan (resetovan na prazan niz) u spoljnjoj petlji. Ispravite kod tako da vrati ispravan niz podataka od 3x2 nula, sličan `[[0, 0], [0, 0], [0, 0]]`.

# --hints--
Vaš kod bi trebalo da postavi parametar `matrix` kao matricu podataka sa 3 reda i po 2 nule u svakom.

```js
assert(JSON.stringify(matrix) == '[[0,0],[0,0],[0,0]]');
```

Parameter `matrix` treba da ima 3 linije.

```js
assert(matrix.length == 3);
```

Parametar `matrix` treba da ima 2 reda u svakoj liniji.

```js
assert(
  matrix[0].length == 2 && matrix[1].length === 2 && matrix[2].length === 2
);
```

`zeroArray(4,3)` Treba da vrati red podataka koji ima 4 elementa i po 3 nulte vrednosti, za svaki.

```js
assert(JSON.stringify(zeroArray(4,3)) == '[[0,0,0],[0,0,0],[0,0,0],[0,0,0]]');
```

# --seed--

## --seed-contents--

```js
function zeroArray(m, n) {
  // Creates a 2-D array with m rows and n columns of zeroes
  let newArray = [];
  let row = [];
  for (let i = 0; i < m; i++) {
    // Adds the m-th row into newArray

    for (let j = 0; j < n; j++) {
      // Pushes n zeroes into the current row to create the columns
      row.push(0);
    }
    // Pushes the current row, which now has n zeroes in it, to the array
    newArray.push(row);
  }
  return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```

# --solutions--

```js
function zeroArray(m, n) {
 // Creates a 2-D array with m rows and n columns of zeroes
 let newArray = [];
 for (let i = 0; i < m; i++) {
   let row = [];
   // Adds the m-th row into newArray

   for (let j = 0; j < n; j++) {
     // Pushes n zeroes into the current row to create the columns
     row.push(0);
   }
   // Pushes the current row, which now has n zeroes in it, to the array
   newArray.push(row);
 }
 return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```
