---
id: 587d7b84367417b2b2512b35
title: Gundua Majina ya Vigezo na Vitendakazi Vilivyoandikwa Vibaya
challengeType: 1
forumTopicId: 301186
dashedName: catch-misspelled-variable-and-function-names
---

# --description--
Metode `console.log()` i `typeof` su dva glavna načina za proveru međuvrednosti i tipova programskih rezultata. Sada je vreme da se fokusiramo na uobičajene vrste grešaka. Jedan problem na nivou sintakse kojem se mogu odnositi brzi programeri je greška u pravopisu.

Promene pozicije slova, nedostajanje ili pogrešno korišćenje velikih slova u imenu parametra ili funkcije će uzrokovati da pretraživač traži nešto što ne postoji – i to će prijaviti kao grešku referenciranja. Imena parametara i funkcija u JavaScript paze na velika i mala slova.

# --instructions--
Ispravite dve greške u pravopisu u kodu kako bi izračunavanje `netWorkingCapital` funkcionisalo.

# --hints--
Proverite pravopis dva parametra koja se koriste u izračunavanju netWorkingCapital, izlaz konzole bi trebalo da prikaže "Net working capital is: 2".

```js
assert(netWorkingCapital === 2);
```

Nema nijednog kriterijuma pogrešno napisanog u kodu.

```js
assert(!__helpers.removeJSComments(code).match(/recievables/g));
```

Parametar ``receivables`` mora biti deklarisan i korišćen odgovarajuće u kodu.

```js
assert(__helpers.removeJSComments(code).match(/receivables/g).length == 2);
```

Nema nijednog parametra koji je pogrešno napisan u kodu.

```js
assert(!__helpers.removeJSComments(code).match(/payable;/g));
```

Parametar `payables` treba biti deklarisan i korišćen odgovarajuće u kodu.

```js
assert(__helpers.removeJSComments(code).match(/payables/g).length == 2);
```

# --seed--

## --seed-contents--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = recievables - payable;
console.log(`Net working capital is: ${netWorkingCapital}`);
```

# --solutions--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = receivables - payables;
console.log(`Net working capital is: ${netWorkingCapital}`);
```
