---
id: 587d7b84367417b2b2512b37
title: Gundua Matumizi Mchanganyiko ya Alama za Nukuu Moja na Maradufu
challengeType: 1
forumTopicId: 301188
dashedName: catch-mixed-usage-of-single-and-double-quotes
---

# --description--
JavaScript omogućava korišćenje jednostrukih navodnika (`'`) i dvostrukih navodnika (`"`) za prikaz sekvence slova. Odlučivanje koji koristiti obično zavisi od lične preferencije, osim u nekim slučajevima.

Imati dve opcije je dobro kada sekvenca slova sadrži skraćenice ili drugi deo teksta unutar navodnika. Molimo vas da budete pažljivi i ne zatvarajte sekvencu slova previše rano, jer to uzrokuje sintaktičku grešku.

Evo nekoliko primera kombinovanja navodnika:

```js
const grouchoContraction = "I've had a perfectly wonderful evening, but this wasn't it.";
const quoteInString = "Groucho Marx once said 'Quote me as saying I was mis-quoted.'";
const uhOhGroucho = 'I've had a perfectly wonderful evening, but this wasn't it.';
```

Prva dva primera su tačna, ali treći nije tačan.

Naravno, nije u redu koristiti samo jedan stil navodnika. Možete izbeći navodnike unutar sekvence slova koristeći povratnik (`\`):

```js
const allSameQuotes = 'I\'ve had a perfectly wonderful evening, but this wasn\'t it.';
```

# --instructions--
Ispravite niz karaktera ili koristite različite navodne znakove za vrednost `href`, ili izbegnite postojeće znake. Obuhvatite ceo niz dvostrukim navodnicima.

# --hints--
Vaš kod treba da izmeni navodne oznake koje okružuju vrednosti `href` `#Home` menjanjem ili izbegavanjem.

```js
assert(__helpers.removeJSComments(code).match(/<a href=\s*?('|\\")#Home\1\s*?>/g));
```

Tvoj kod treba da sačuva dvostruke navodne tačke koje okružuju ceo niz.

```js
assert(__helpers.removeJSComments(code).match(/"<p>.*?<\/p>";/g));
```

# --seed--

## --seed-contents--

```js
let innerHtml = "<p>Click here to <a href="#Home">return home</a></p>";
console.log(innerHtml);
```

# --solutions--

```js
let innerHtml = "<p>Click here to <a href=\"#Home\">return home</a></p>";
console.log(innerHtml);
```
