---
id: 587d7b84367417b2b2512b34
title: Koristi `typeof` za proveru tipa parametra
challengeType: 1
forumTopicId: 18374
dashedName: use-typeof-to-check-the-type-of-a-variable
---

# --description--
Možete koristiti `typeof` za proveru strukture podataka ili tipa parametra. Ovo je važno prilikom ispravljanja grešaka kada koristite različite tipove podataka. Ako mislite da sabirate dva broja, ali je jedan niz slova, rezultat može biti nepouzdan. Greške tipa mogu se pojaviti u proračunima ili pozivima funkcija. Budite pažljivi, posebno kada koristite i radite sa eksternim podacima u formatu JavaScript Object Notation (JSON).

Evo nekoliko primera korišćenja operatora `typeof`:

```js
console.log(typeof "");
console.log(typeof 0);
console.log(typeof []);
console.log(typeof {});
```

Redosledom, konzola će prikazati sekvencu karakterâ `string`, `number`, `object`, na `object`.

JavaScript prepoznaje sedam tipova primitivnih podataka (nepromenljivih): `Boolean`, `Null`, `Undefined`, `Number`, `String`, `Symbol` (novo u ES6), i `BigInt` (novo u ES2020), i jedan tip za mutabilne objekte: `Object`. Napomena da u JavaScript, nizovi podataka su tehnički tip objekta.

# --instructions--
Dodajte dve deklaracije za `console.log()` provere `typeof` za dva parametra `seven` u kodu `three`.

# --hints--
Vaš kod bi trebalo da koristi `typeof` u dve deklaracije za `console.log()` kako bi proverio tip parametara.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\s*\(typeof[\( ].*\)?\)/g).length == 2);
```

Vaš kod bi trebalo da koristi `typeof` za proveru tipa parametra `seven`.

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]seven\)?/g));
```

Vaš kod bi trebalo da koristi `typeof` za proveru tipa parametra `three`.

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]three\)?/g));
```

# --seed--

## --seed-contents--

```js
let seven = 7;
let three = "3";
console.log(seven + three);
// Only change code below this line
```

# --solutions--

```js
let seven = 7;let three = "3";console.log(typeof seven);
console.log(typeof three);
```
