---
id: 587d7b84367417b2b2512b36
title: 'Pronađite standardne tabe, tabe, okrugle tabe i neotvorene quote-ove'
challengeType: 1
forumTopicId: 301190
dashedName: catch-unclosed-parentheses-brackets-braces-and-quotes
---

# --description--
Još jedna sintaksička greška na koju treba obratiti pažnju je što sve obične zagrade, vitičaste zagrade, okrugle zagrade i navodnike moraju imati zatvarajući par. Zaboravljanje dela se dešava češće kada menjate postojeći kod i ubacujete elemente iste vrste para. Takođe, budite pažljivi kada postavljate unutar drugih kodnih zagrada, kao što je dodavanje funkcije povratne poziva (callback) kao argumenta na način.

Jedan način da izbegnete ovu grešku je što čim napišete otvarajuću zagradu, odmah uključite zatvarački par, pa stavite hint za pisanje između njih i nastavite sa pisanjem koda. Srećom, većina modernih kod editora automatski generiše drugi deo para.

# --instructions--
Ispravite dve par grešaka u kodu.

# --hints--
Vaš kod treba da popravi nedostajući deo reda podataka.

```js
assert(__helpers.removeJSComments(code).match(/myArray\s*?=\s*?\[\s*?1\s*?,\s*?2\s*?,\s*?3\s*?\];/g));
```

Tvoj kod bi trebalo da popravi nedostajući deo putanje za `.reduce()`. Izlaz konzole bi trebalo da pokaže da `Sum of array values is: 6`.

```js
assert(arraySum === 6);
```

# --seed--

## --seed-contents--

```js
let myArray = [1, 2, 3;
let arraySum = myArray.reduce((previous, current =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```

# --solutions--

```js
let myArray = [1, 2, 3];
let arraySum = myArray.reduce((previous, current) =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```
