---
id: 587d7b83367417b2b2512b33
title: Koristite JavaScript konzolu za proveru vrednosti parametra
challengeType: 1
forumTopicId: 18372
dashedName: use-the-javascript-console-to-check-the-value-of-a-variable
---

# --description--
I Chrome i Firefox imaju dobre konzole za JavaScript, poznate takođe kao alati za developere, za popravljanje vaših grešaka od JavaScript.

Možete pronaći alate za developera u meniju vašeg Chrom-a ili konzolu mreže u meniju FireFox-a. Ako koristite drugi pretraživač ili mobilni telefon, snažno vam preporučujemo da pređete na desktop verziju FireFox-a ili Chrom-a.

Metoda `console.log()`, koja "ispisuje" rezultate onoga što se nalazi unutar njenih običnih zagrada na konzolu, biće najvažniji alat za debugovanje (otklanjanje grešaka). Postavljanje je u ključnim delovima vašeg koda može vam pokazati međuvrednosti promenljivih. Dobra praksa je da imate ideju kakvi bi trebalo biti rezultati pre nego što proverite šta jesu. Imatanje sekcija za proveru svojih proračuna kroz ceo kod pomoći će smanjiti mesto gde se greška nalazi.

Ovo je primer za štampanje sekvence karaktera `Hello world!` na konzoli:

```js
console.log('Hello world!');
```

# --instructions--
Koristite putanju `console.log()` za štampanje vrednosti parametra `a` gde je prikazan u kodu.

# --hints--
Vaš kod bi trebalo da koristi `console.log()` za proveru vrednosti parametra `a`.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\(a\)/g));
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 1;
a++;
// Only change code below this line


let sumAB = a + b;
console.log(sumAB);
```

# --solutions--

```js
var a = 5; console.log(a);
```
