---
id: 587d7b85367417b2b2512b3a
title: Argumenti prosleđeni u pogrešnom redosledu prilikom pozivanja funkcije
challengeType: 1
forumTopicId: 301184
dashedName: catch-arguments-passed-in-the-wrong-order-when-calling-a-function
---

# --description--
Nastavljajući sa raspravom o pozivanju funkcija, sledeća greška koju treba uzeti u obzir je situacija kada se argumenti funkcije proslede u pogrešnom redosledu. Ako su argumenti različitih tipova, kao što je niz podataka i ceo broj, to može izazvati grešku tokom izvršavanja (runtime error). Ako su argumenti istog tipa (poput svih celih brojeva, na primer), onda će logika koda biti besmislenost. Uverite se da pružate sve potrebne argumente u ispravnom redosledu kako biste izbegli ove probleme.

# --instructions--
Funkcija ``raiseToPower`` uzima osnovni broj i podigne mu stepen. Nažalost, nije pozvana ispravno - popravi kod tako da vrednost ``power`` bude 8 kako je očekivano.

# --hints--
Vaš kod bi trebalo da modifikuje parametar `power` ili da bude jednak 2 na stepen 3, a ne 3 na stepen 2.

```js
assert(power == 8);
```

Vaš kod treba da koristi ispravan redosled argumenata za poziv funkcije `raiseToPower`.

```js
assert(__helpers.removeJSComments(code).match(/raiseToPower\(\s*?base\s*?,\s*?exp\s*?\);/g));
```

# --seed--

## --seed-contents--

```js
function raiseToPower(b, e) {
  return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(exp, base);
console.log(power);
```

# --solutions--

```js
function raiseToPower(b, e) {
 return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(base, exp);
console.log(power);
```
