---
id: 66b102ee0bdbad34a9f85ba0
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

Tokom ove radionice, naučićete o sekvencama karaktera kroz kreiranje Chatbota za Učitelja.

Početi, dodajte izjavu za `console`, sa porukom o `"Hi there!"`.

# --hints--

Trebalo bi da imate deklaraciju `console`.

```js
assert.match(code, /console\.log(.*)/);
```

Konsola bi trebalo da prikaže poruku od `"Hi there!"`.

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
