---
id: 66b59be6ab830800c4df9146
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Sada, je vreme za korišćenje parametara koje si kreirao u prethodnom koraku.

U prethodnim lekcijama, naučili ste kako da radite sa šablonskim karakterima kao ovim:

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

Počnite kreiranjem parametra za `sentence`.

Ako se oslanjate na sintaksu karakter(a), postavite vrednost niza karaktera `Today, you will learn about [topic variable goes here] in [subject variable goes here].` u parametar `sentence`.

Zameni identifikatore `[topic variable goes here]` i `[subject variable goes here]` za parametre `topic` i `subject`, uz sintaksu `${}`.

Konačno, napišite parametar za `sentence` u konzolu.

# --hints--

Trebalo bi da imaš kriterijum `sentence`.

```js
assert.isNotNull(sentence);
```

Vaše polje `sentence` mora biti niz slova.

```js
assert.isString(sentence);
```

Trebalo bi da koristiš sintaksu znakova za pločice (placeholder characters) da postaviš vrednost niza znakova `Today, you will learn about [topic variable goes here] in [subject variable goes here].` u parametar `sentence`. Konsultuj primer ako ti je potrebna pomoć.

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

Trebalo bi da napišete parametar `sentence` u konzoli.

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
