---
id: 66b59829dba144ff1351220f
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Sada je vreme za kreiranje pozdrava koristeći varijablu `botName`.

Tokom početnih lekcija, naučili ste kako da spojite niz slova koristeći šablonske slova kao što sledi:

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

Započnite kreiranjem varijable pod nazivom `greeting`.

Zatim, koristeći sintaksu doslovnih stringova, postavi vrednost sekvence karaktera koja kaže `My name is`, praćeno varijablom za `botName`, i završava se tačkom (`.`).

Konačno, napiši varijablu `greeting` u konzolu.

# --hints--

Trebalo bi da imaš varijablu nazvanu `greeting`.

```js
assert.isNotNull(greeting);
```

Trebalo bi da postaviš vrednost niza karaktera na tvoju varijablu `greeting`.

```js
assert.isString(greeting);
```

Treba koristiti doslovne znakove da spoji niz znakova `My name is` i varijablu `botName`, praćen tačkom (`.`).

```js
assert.equal(greeting, "My name is teacherBot.");
```

Trebalo bi da imaš izlaz (deklaraciju) `console` koji ispisuje varijablu `greeting` na konzoli.

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
