---
id: 66b5928b86e907fdfbf34e56
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Sada je vreme da damo robotu ime.

Kreiraj parametar nazvan `botName` i postavi mu vrednost kao niz slova `"teacherBot"`.

# --hints--

Trebalo bi da imaš parametar nazvan `botName`.

```js
assert.isNotNull(botName);
```

Vaš parametar `botName` mora biti niz slova.

```js
assert.isString(botName);
```

Vaš parametar za `botName` treba da ima vrednost od `"teacherBot"`.

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
