---
id: 66b59b12a745e10011158f7b
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sljedeći korak je kreiranje još manjeg skupa parametara koji će se koristiti u budućim porukama bota.

Kreiraj parametar nazvan `subject` i postavi vrednost za niz znakova `"JavaScript"`.

Zatim kreiraj parametar nazvan `topic` i postavi mu vrednost niza karaktera `"strings"`.

# --hints--

Trebalo bi da imate parametar nazvan `subject`.

```js
assert.isNotNull(subject);
```

Trebalo bi da unesete vrednost sekvence slova u vaš parametar `subject`.

```js
assert.isString(subject);
```

Treba da postavite niz karaktera `"JavaScript"` u vaš parametar `subject`.

```js
assert.equal(subject, "JavaScript");
```

Trebalo bi da imaš parametar koji se zove `topic`.

```js
assert.isNotNull(topic);
```

Trebalo bi da unesete vrednost sekvence slova u vaš parametar `topic`.

```js
assert.isString(topic);
```

Treba da unesete sekvencu karaktera `"strings"` u vaš parametar `topic`.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
