---
id: 66b6fdb76441c738719039fa
title: Korak 13
challengeType: 1
dashedName: step-13
---

# --description--

U prethodnim časovima, naučili ste kako da pronađete posledovno slovo u sekvenci slova, kao što sledi:

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` Uvek će vam dati broj zadnje verzije/modela za sekvencu slova.

Kreiraj novi parametar pod nazivom ``lastCharacter`` i postavi mu vrednost poslednjeg karaktera iz parametra ``subject``.

Zatim, napiši vrednost parametra za `lastCharacter` u konzolu.

# --hints--

Trebalo bi da imate parametar nazvan `lastCharacter`.

```js
assert.isNotNull(lastCharacter);
```

Treba da postavite vrednost poslednjeg karaktera u parametar `subject`, i u parametar `lastCharacter`. Konsultujte se primera ako vam je potrebna pomoć.

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

Trebalo bi da napišete vrednost parametra `lastCharacter` na `console`.

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
