---
id: 66b6d482bbb9e12f2e5ee1ae
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Za ovaj sledeći deo projekta, bot će naučiti kako da radi sa karakteristikom niza znakova `length`.

Počnite kreiranjem novog parametra nazvanog `strLengthIntro`.

Kada koristite sintaksu *template literal*-a, postavite vrednost niza karaktera `Here is an example of using the length property on the word [subject].` za parametar `strLengthIntro`.

Zameni `[subject]` sa parametrom `subject` kao što si ranije uradio/uradila.

Konačno, upišite parametar `strLengthIntro` u konzolu.

# --hints--

Trebalo bi da imaš parametar koji se zove `strLengthIntro`.

```js
assert.isNotNull(strLengthIntro);
```

Vaš parametar za `strLengthIntro` treba da bude niz slova.

```js
assert.isString(strLengthIntro);
```

Trebalo bi da koristiš sintaksu šablonnog literala za postavljanje vrednosti niza znakova `Here is an example of using the length property on the word [subject goes here].` za parametar `strLengthIntro`.

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

Trebalo bi da napišete parametar `strLengthIntro` u konzoli.

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
