---
id: 66b6f586767a1534f3097353
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Sada je vreme da se dođe do drugog slova iz promena `subject`.

Počnite dodavanjem deklaracije ``console`` koja generiše poruku ``Here is an example of accessing the second letter in the word [subject].``. Zapamtite da zamenite ``[subject]`` sa stvarnom vrednošću promenljive/izmene ``subject`` i koristite ispravnu sintaksu za *template literal*.

Zatim dodajte još jednu deklaraciju od `console` koja je izvedena sa drugog karaktera promene `subject`, koristeći zapisivanje u uglastim zagradama i tačan indeks broj.

# --hints--

Izjava o `console` treba da emituje poruku `Here is an example of accessing the second letter in the word [subject].`

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

Treba da ima druga deklaracija za `console`, dobijena iz drugog slova promene `subject`, koristeći zagradni zapis i indikator `1`.

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
