---
id: 66b6e62423e8a031d6c1f03d
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Sada je vreme za dobijanje dužine sekvence karaktera `topic`.

Možete koristiti šablonske karaktere unutar deklaracija `console`, kao što je ovo:

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

Počnite slanjem poruke `Here is an example of using the length property on the word [topic].` u konzoli.

Zapamti da zameniš `[topic]` sa parametrom `topic`, i koristi ispravnu sintaksu za znakove kvorca kao što si uradio u prethodnim koracima.

Zatim, dodajte drugu deklaraciju za `console.log` koja ispisuje dužinu niza znakova `topic` na konzoli.

# --hints--

Trebalo bi da koristiš sintaksu literarnih karaktera da izbaci poruku `Here is an example of using the length property on the word [topic].` u `console`.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

Treba da imate drugu deklaraciju od `console` da pružite dužinu parametra `topic` na `console`.

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
