---
id: 66b6e39031393e30f2c48d0f
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Da biste dobili dužinu niza karaktera, možete koristiti svojstvo <dfn>length</dfn>. Ovo svojstvo vraća broj karaktera u nizu karaktera.

Evo primera:

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

Počnite korišćenjem svojstva `length` da dobijete dužinu niza znakova od `subject` i napišite tu vrednost u konzolu.

# --hints--


Trebalo bi da koristiš atribut `length` da dobiješ dužinu niza karaktera od `subject`.

```js
assert.match(code, /subject\.length/);
```

Treba da unesete dužinu niza znakova za `subject` u konzolu.

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
