---
id: 66b6efddeca35833cd6f0b03
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

U prethodnim lekcijama, naučili ste kako da dobijete karaktere u sekvenci karaktera, kao što sledi:

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

Zapamtite da brojevi modela počinju sa `0`, tako da će prvi slovo u nizu slova uvek biti na modelu `0`.

Počni dodavanjem drugog izjave za `console`.

U deklaraciji `console`, izvucite prvo slovo promene od `subject` korišćenjem zagradnog zapisa i tačnog indeksa.

# --hints--

Vaš izraz za `console` treba da izvuče prvo slovo promene `subject` koristeći uglaste zagrade i indikator `0`.

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
