---
id: 66b6ffb42b88e33943788abf
title: Korak 14
challengeType: 1
dashedName: step-14
---

# --description--

U poslednjem delu radionice, ćete proučiti kako pronaći obrazac podstringa unutar niza znakova. Zapamtite da je podstring deo celog niza znakova.

Počnite kreiranjem parametra nazvanog `learningIsFunSentence` i postavite mu vrednost niza slova `"Learning is fun."`.

# --hints--

Trebalo bi da imaš parametar koji se zove `learningIsFunSentence`.

```js
assert.isNotNull(learningIsFunSentence);
```

`learningIsFunSentence` Trebalo bi da ima vrednost sekvence slova.

```js
assert.isString(learningIsFunSentence);
```

Treba da postaviš vrednost `"Learning is fun."` u parametar `learningIsFunSentence`.

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
