---
id: 66b6e80d6c3f0b329c360283
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Sledeći deo ove radionice je proveravanje dostupnosti karaktera iz sekvence karaktera.

Počnite slanjem poruke `Here is an example of accessing the first letter in the word [subject].` u konzoli.

Zapamti da zameniš `[subject]` sa parametrom `subject` i koristi ispravnu sintaksu template literal kao što si uradio u prethodnim koracima.

# --hints--

Tvoja rečenica sa `console` treba da iznese poruku `Here is an example of accessing the first letter in the word [subject].`

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+first\s+letter\s+in\s+the\s+word\s+\${subject}\.`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

--fcc-editable-region--

--fcc-editable-region--
```
