---
id: 66b70b8a611cbf3bcc5c6c5f
title: Korak 16
challengeType: 1
dashedName: step-16
---

# --description--

U prethodnim lekcijama, naučili ste kako da koristite metodu `indexOf` kao što sledi:

```js
const sentence = "I love to learn.";

// returns index 2
console.log(sentence.indexOf("love"));

// returns -1
console.log(sentence.indexOf("hate"));
```

Imajte na umu da metoda `indexOf` vraća indeks prve pozicije pojavljivanja podstringa unutar niza karaktera. Ako podstring nije pronađen, vraća `-1`.

Dodaj novu deklaraciju za `console` koja je dobijena korišćenjem metode `indexOf` na izmene `learningIsFunSentence` za pretraživanje obrasca pozicije malog dela niza karaktera `"Learning"`.

# --hints--

`console` treba da koristi metodu `indexOf` tokom izmene `learningIsFunSentence` kako bi pronašao obrazac lokacije malog segmenta sekvence slova `"Learning"`.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(Learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

--fcc-editable-region--

--fcc-editable-region--
```
