---
id: 66bb6a9c2dd58b73cd759034
title: Izgradite program za svesku bodovanja
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Trebalo bi da imate funkciju pod nazivom `getAverage` koja prima niz podataka ispitnih ocena kao parametar i vraća prosečnu ocenu.
1. Trebalo bi da imate funkciju pod nazivom `getGrade` koja prima ocenu učenika kao parametar i vraća niz slova koji predstavlja slovnu ocenu u zavisnosti od oцене. Evo ocena i njihovih odgovarajućih slovnih oznaka:

| Oblast oznake | Most |
| ----------- | ------- |
| `100` | `"A+"` |
| `90 - 99` | `"A"` |
| `80 - 89` | `"B"` |
| `70 - 79` | `"C"` |
| `60 - 69` | `"D"` |
| `0 - 59` | `"F"` |

1. Moraće da postoji funkcija nazvana `hasPassingGrade` koja prima ocenu kao parametar i vraća `true` ili `false` u zavisnosti od toga da li ocena predstavlja prolaznu ocenu.
1. Funkcija `hasPassingGrade` treba da koristi funkciju `getGrade` kako bi dobila slovo ocenu, i da je koristi za određivanje da li je ocena prolazna. Prolazna ocena je bilo šta drugačije od `"F"`.
1. Moraće da postoji funkcija nazvana `studentMsg` koja prima niz ocena i ocenu studenta kao parametre. Funkcija treba da vrati sekvencu slova sa strukturom:

- `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`, ako je student položio kurs.
    - `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`, ako je student dobio nepoložavku za kurs.

Zameni `average-goes-here` sa prosekom svih ocena i `grade-goes-here` sa ocenom učenika.
Koristi `getAverage` da dobiješ prosečnu ocenu i `getGrade` da dobiješ razred/ocenu učenika.

# --hints--

Trebalo bi da imaš funkciju koja se zove `getAverage`.

```js
assert.isFunction(getAverage);
```

Tvoja funkcija za `getAverage` treba da vrati broj.

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` bi trebalo da vrati `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` treba da vrati `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` treba da vrati `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage([10, 20, 30, 40, 55, 65, 75, 83])` treba vratiti `47.25`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

`getAverage([10, 20, 30, 40, 50, 60, 70, 97])` trebalo bi vratiti `47.125`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

Tvoja funkcija `getAverage` treba da vrati prosečnu ocenu iz rezultata ispita.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

Trebalo bi da imaš funkciju koja se zove `getGrade`.

```js
assert.isFunction(getGrade);
```

Tvoja funkcija za `getGrade` treba da vrati `"A+"` ako je vrednost `100`.

```js
assert.strictEqual(getGrade(100), "A+");
```

Tvoja funkcija za `getGrade` mora da vrati `"A"` ako je vrednost između `90` i `99`.

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

Tvoja funkcija za `getGrade` treba da vrati `"B"` ako je vrednost između `80` i `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Vaša funkcija za `getGrade` treba da vrati `"C"` ako je vrednost između `70` i `79`.

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

Funkcija za `getGrade` treba da vrati `"D"` ako je vrednost između `60` i `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Tvoja funkcija za `getGrade` mora vratiti `"F"` ako je vrednost između `0` i `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```

Trebalo bi da imaš funkciju koja se zove `hasPassingGrade`.

```js
assert.isFunction(hasPassingGrade);
```

Tvoja funkcija za `hasPassingGrade` treba da vrati vrednost istina ili laž.

```js
assert.isBoolean(hasPassingGrade(100));
```

Tvoja funkcija za `hasPassingGrade` treba da vrati `true` ako je daraja `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

Funkcija tvoje `hasPassingGrade` treba da vrati `false` ako je most `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

Vaša funkcija od `hasPassingGrade` treba da vrati `true` za sve prolazne nivoe.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

Trebalo bi da imaš funkciju nazvanu `studentMsg`.

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` Trebalo bi da vrati sledeću poruku: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` trebalo bi da vrati sledeću poruku: `"Class average: 50.8. Your grade: A+. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

`studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)` treba da vrati sledeću poruku: `"Class average: 48.25. Your grade: B. You passed the course."`.

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

`studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)` Trebalo bi da vrati sledeću poruku: `"Class average: 45.625. Your grade: C. You passed the course."`.

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

Funkcija ti [`studentMsg`] bi trebalo da vrati odgovarajuću poruku na osnovu ocene učenika i proseka razreda.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
