---
id: 587d8249367417b2b2512c41
title: Konverter metara i imperialnih jedinica
challengeType: 4
forumTopicId: 301570
dashedName: metric-imperial-converter
---

# --description--

Jenga programu ya JavaScript ya full-stack ambayo inaendana kiutendaji na hii: <a href="https://metric-imperial-converter.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://metric-imperial-converter.freecodecamp.rocks/</a>. Rad na ovom projektu će uključivati pisanje vašeg koda korišćenjem jednog od sledećih metoda:

- Nakili <a href="https://github.com/freeCodeCamp/boilerplate-project-metricimpconverter/" target="_blank" rel="noopener noreferrer nofollow">hii repo ya GitHub</a> na ukamilishe mradi wako kwa ndani ya kompyuta yako.
- Koristite bilo koji graditelj veb sajtova koji vam je omiljen da biste dovršili projekat. Obavezno uključite sve fajlove iz našeg GitHub repozitorijuma.

**Kumbuka:** Majaribio ya mradi huu hayafanyi kazi wakati unatumia `glitch.com`.

# --instructions--

- Popunite potrebnu logiku konverzije u `/controllers/convertHandler.js`
- Popunite potrebne putanje podataka u `/routes/api.js`
- Kopirajte fajl `sample.env` u `.env` i podesite parametre odgovarajuće.
- Da pokrenete test direktno, dodajte `NODE_ENV=test` u vaš fajl `.env`
- Da pokrenete test na konzoli, koristite komandu `npm run test`

Unesite sljedeće eksperimente u `tests/1_unit-tests.js`:

- `convertHandler` treba da precizno pročita unos celog broja.
- `convertHandler` treba da precizno pročita unos decimalnog broja.
- `convertHandler` treba da precizno pročita unos frakcije.
- `convertHandler` treba da precizno pročita unos decimalne frakcije.
- `convertHandler` treba da vrati grešku za dva dela (tj. `3/2/3`).
- `convertHandler` treba da koristi po default vrednost broj `1` kada nije unet nijedan brojac.
- `convertHandler` treba da precizno pročita svaku validnu jedinicu unosa.
- `convertHandler` treba da vrati grešku za nevažeću jedinicu unosa.
- `convertHandler` treba da vrati odgovarajuću jedinicu za povratak za svaku validnu jedinicu unosa.
- `convertHandler` treba da precizno vrati jedinicu napisanu rečima za svaku validnu jedinicu unosa.
- `convertHandler` treba da pretvori `gal` u `L` sa preciznošću.
- `convertHandler` treba da pretvori `L` u `gal` sa preciznošću.
- `convertHandler` treba da pretvori `mi` u `km` sa preciznošću.
- `convertHandler` treba da pretvori `km` u `mi` sa preciznošću.
- `convertHandler` treba da pretvori `lbs` u `kg` sa preciznošću.
- `convertHandler` treba da pretvori `kg` u `lbs` sa preciznošću.

Napišite sljedeće testove u `tests/2_functional-tests.js`:

- Izmena validnog unosa kao `10L`: zahtev od `GET` ka `/api/convert`.
- Izmena nevažećeg unosa kao `32g`: zahtev od `GET` ka `/api/convert`.
- Izmena nevažećeg broja kao `3/7.2/4kg`: zahtev od `GET` ka `/api/convert`.
- Izmena nevažećeg broja i jedinice kao `3/7.2/4kilomegagram`: zahtev od `GET` ka `/api/convert`.
- Izmena bez broja kao `kg`: zahtev od `GET` ka `/api/convert`.

# --hints--

Unaweza kutoa mradi wako mwenyewe, si lazima URL ya mfano.

```js
  assert(
    !/.*\/metric-imperial-converter\.freecodecamp\.rocks/.test(
      code
    )
  );
```

Možete `GET` `/api/convert` za jedan parametar sa prihvatljivim brojem i jedinicom menjanjem. (Podsetnik: Podelite unos traženjem obrasca prvog slova koji predstavlja početak jedinice)

```js

```

Možete zamijeniti `'gal'` sa `'L'` i obrnuto. (1 gal do 3,78541 L)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 3.78541);
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10gal');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 37.8541);
    assert.equal(data2.returnUnit, 'L');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.26417);
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10l');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 2.64172);
    assert.equal(data4.returnUnit, 'gal');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Možete konvertovati `'lbs'` u `'kg'` i obrnuto. (od 1 lbs do 0.453592 kg)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1lbs');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 0.45359);
    assert.equal(data1.returnUnit, 'kg');
    const response2 = await fetch(code + '/api/convert?input=10lbs');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 4.53592);
    assert.equal(data2.returnUnit, 'kg');
    const response3 = await fetch(code + '/api/convert?input=1kg');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 2.20462);
    assert.equal(data3.returnUnit, 'lbs');
    const response4 = await fetch(code + '/api/convert?input=10kg');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 22.04624);
    assert.equal(data4.returnUnit, 'lbs');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Možete konvertovati `'mi'` u `'km'` i obrnuto. (1 mi do 1,60934 km)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 1.60934);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=10mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 16.0934);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(code + '/api/convert?input=1km');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.62137);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(code + '/api/convert?input=10km');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 6.21373);
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Sve primljene jedinice moraju biti prihvaćene sa velikim i malim slovima, ali kada se vrate, biće u malim slovima u `initUnit` i `returnUnit`, osim litra, što je predstavljeno velikim `'L'`.

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.initUnit, 'gal');
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10L');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.initUnit, 'L');
    assert.equal(data2.returnUnit, 'gal');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.initUnit, 'L');
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10KM');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.initUnit, 'km');
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Ako je jedinica mere nevažeća, biće vraćena `'invalid unit'`.

```js
  try {
    const response = await fetch(code + '/api/convert?input=1min');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid unit' || data === 'invalid unit');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Ako broj nije validan, biće vraćen `'invalid number'`.

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2gal'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid number' || data === 'invalid number');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Ako su jedinica i svi brojevi važeći, biće vraćeno `'invalid number and unit'`.

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2min'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(
      data.error === 'invalid number and unit' ||
        data === 'invalid number and unit'
    );
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Možete koristiti razlomke, decimalne brojeve ili oba u parametru (tj. 5, 1/2, 2.5/6), ali ako nije ništa navedeno, koristiće se 1 kao podrazumevana vrednost.

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.approximately(data1.initNum, 1, 0.001);
    assert.approximately(data1.returnNum, 1.60934, 0.001);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=1/5mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.approximately(data2.initNum, 1 / 5, 0.1);
    assert.approximately(data2.returnNum, 0.32187, 0.001);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(
      code + '/api/convert?input=1.5/7km'
    );
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.approximately(data3.initNum, 1.5 / 7, 0.001);
    assert.approximately(data3.returnNum, 0.13315, 0.001);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(
      code + '/api/convert?input=3/2.7km'
    );
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.approximately(data4.initNum, 3 / 2.7, 0.001);
    assert.approximately(data4.returnNum, 0.69041, 0.001);
    assert.equal(data4.returnUnit, 'mi');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Vaš rezultat će imati `initNum`, `initUnit`, `returnNum`, `returnUnit`, na `string` koji opisuju jedinice u formatu `'{initNum} {initUnitString} converts to {returnNum} {returnUnitString}'` i rezultate zaokružene na 5 decimala.

```js
  try {
    const response = await fetch(code + '/api/convert?input=2mi');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.json();
    assert.equal(data.initNum, 2);
    assert.equal(data.initUnit, 'mi');
    assert.approximately(data.returnNum, 3.21868, 0.001);
    assert.equal(data.returnUnit, 'km', 'returnUnit did not match');
    assert.equal(data.string, '2 miles converts to 3.21868 kilometers');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Svi 16 testova jedinica su završeni i uspešni.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const unitTests = getTests.filter(test => {
      return !!test.context.match(/Unit Tests/gi);
    });
    assert.isAtLeast(unitTests.length, 16, 'At least 16 tests passed');
    unitTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Sva pet funkcionalnih testova je završena i prolazi.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const functTests = getTests.filter(test => {
      return !!test.context.match(/Functional Tests/gi);
    });
    assert.isAtLeast(functTests.length, 5, 'At least 5 tests passed');
    functTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```
