---
id: 66b36358ed4f261d64840c24
title: Izgraditi tabelu dostupnosti
challengeType: 25
dashedName: build-an-availability-table
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Treba da imaš tabelu sa najmanje tri kolone i pet redova, uključujući zaglavlja.
1. Ćelije u prvom redu treba da budu zaglavlja tabele koja prikazuju dane u nedelji.
1. Ćelije u prvoj koloni treba da budu zaglavlja tabele sa klasom `time` i da kao tekst sadrže vremenske oznake.
1. Sve ćelije sa podacima treba da imaju klasu `available-#`, gde je `#` broj od `0` do `5` koji označava broj dostupnih osoba u tom terminu.
1. U selektoru `:root` definiši šest CSS promenljivih pod nazivom `--color#`, gde je `#` broj od `0` do `5`, i dodeli svakoj vrednost boje. Koristi ove promenljive za postavljanje boje pozadine odgovarajućih elemenata `.available-#`.
1. U selektoru `:root` definiši CSS promenljive `--solid-border` i `--dashed-border`. Koristi ih za stilizovanje donjih ivica ćelija sa podacima, naizmenično primenjujući punu i isprekidanu ivicu po redovima. Tim redovima dodeli klasu `sharp` ili `half` radi stilizovanja.
1. Treba da imaš `div` element sa `id` vrednošću `legend`. On treba da sadrži `span` element sa tekstom `Availability` i `div` element sa `id` vrednošću `legend-gradient`.
1. Elementu `#legend-gradient` treba da dodeliš linearni gradijent koji prelazi kroz sve boje od `--color0` do `--color5` sa oštrim prelazima.

**Napomena:** Uverite se da ste povezali vaš stilski list (stylesheet) u vaš HTML kako biste koristili vaš CSS.

# --hints--

Trebalo bi da imaš element `table`.

```js
assert.isNotNull(document.querySelector('table'));
```

Trebalo bi da ima samo jedan element `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Tvoja tabela bi trebalo da ima barem redove `5`.

```js
assert.isAtLeast(document.querySelectorAll('tr')?.length, 5);
```

Tvoja tabela bi trebalo da ima najmanje red `3`.

```js
const rows = [...document.querySelectorAll('tr')];
assert.isNotEmpty(rows);
rows.forEach((row) => {
    assert.isAtLeast(row.children.length, 3);
})
```

Prvi red u vašoj tabeli mora imati barem elemente `th`.

```js
assert.isAtLeast(document.querySelectorAll('tr:first-child>th').length, 3)
```

Prvi red u vašoj tabeli ne bi trebalo da sadrži elemente `td`.

```js
const firstRow = document.querySelector('tr');
assert.isNotNull(firstRow);
assert.isEmpty(firstRow.querySelectorAll('td'));
```

Trebalo bi da imate najmanje dve linije sa klasom `sharp`.

```js
assert.isAtLeast(document.querySelectorAll('tr[class~="sharp"]').length, 2);
```

Trebalo bi da ima najmanje dve linije sa klasom `half`.

```js
assert.isAtLeast(document.querySelectorAll('tr[class~="half"]').length, 2);
```

Polja `th` u prvom redu treba da imaju `class` od `time` i moraju imati privremene vrednosti.

```js
const firstColumnData = [...document.querySelectorAll('tr th:first-child')];
assert.isNotEmpty(firstColumnData);
for (let cell of firstColumnData) {
    if (cell.innerText.match(/\d/)?.length) {
        assert.isTrue(cell.classList.contains('time'));
    }
}
```

Trebalo bi da ima najmanje šest elemenata od `th` sa klasom `time` koji nose privremene vrednosti.

```js
const timeClass = document.querySelectorAll('.time');
assert.isAtLeast(timeClass.length, 4);
for (let cell of timeClass) {
    assert.match(cell.innerText, /\d/)
}
```

Svi elementi od `td` moraju imati `class` sa `available-#`, gde `#` je broj od `0` do `5`.

```js
const cells = [...document.querySelectorAll('tr[class="sharp"]>td, tr[class="half"]>td')];
assert.isNotEmpty(cells);
for (let cell of cells) {
    const classString = [...cell.classList].join(' ');
    assert.lengthOf(classString.match(/(?<=\s|^)available-[0-5](?=\s|$)/gm), 1);
}
```

Trebalo bi da ima najmanje osam elemenata za `.available-#`, gde `#` je broj između `0` i `5`.

```js
assert.isAtLeast(document.querySelectorAll('td[class*="available-"]').length, 8);
```

Trebalo bi da definišete parametar za `--color0` u vašem root selektoru.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color0'))
```

Treba da definišete parametar `--color1` u vašem root selektoru.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color1'))
```

Treba da definišete parametar `--color2` u vašem root selektoru.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color2'))
```

Treba da definišete parametar `--color3` u vašem korenu (root) selektoru.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color3'));
```

Trebalo bi da definišete parametar `--color4` u vašem root selektoru.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color4'));
```

Treba da definišete parametar `--color5` u vašem root selektoru.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color5'));
```

Trebalo bi da koristiš `--color0` za postavljanje pozadinske boje elemenata `.available-0`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-0');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color0)');
```

Trebalo bi da koristiš `--color1` za postavljanje pozadinske boje elemenata `.available-1`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-1');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color1)');
```

Trebalo bi da koristiš `--color2` za postavljanje pozadinske boje elemenata `.available-2`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-2');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color2)');
```

Trebalo bi da koristiš `--color3` za postavljanje pozadinske boje elemenata `.available-3`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-3');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color3)');
```

Trebalo bi da koristiš `--color4` za postavljanje pozadinske boje elemenata `.available-4`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-4');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color4)');
```

Trebalo bi da koristiš `--color5` za postavljanje pozadinske boje elemenata `.available-5`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-5');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color5)');
```

Treba da definišete parametar `--solid-border` u vašem root selektoru.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--solid-border'));
```

Morate definisati parametar `--dashed-border` u vašem korijenskom selektoru.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--dashed-border'));
```

Trebalo bi da ciljate na elemente `td` koji su deca elemenata `.sharp`.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.sharp td') || new __helpers.CSSHelp(document).getStyle('.sharp>td');
assert.exists(tds);
```

Treba koristiti `--solid-border` da postavite `bottom-border` za elemente `td` koji su potomci elemenata `.sharp`.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.sharp td') || new __helpers.CSSHelp(document).getStyle('.sharp>td');
assert.exists(tds);
assert.strictEqual(tds.borderBottom, 'var(--solid-border)');
```

Treba da ciljate na elemente `td` koji su potomci elemenata `.half`.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.half td') || new __helpers.CSSHelp(document).getStyle('.half>td');
assert.exists(tds);
```

Trebalo bi da koristiš `--dashed-border` za postavljanje `bottom-border` elemenata od `td` elemenata od `.half`.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.half td') || new __helpers.CSSHelp(document).getStyle('.half>td');
assert.exists(tds);
assert.strictEqual(tds.borderBottom, 'var(--dashed-border)');
```

Trebalo bi da imaš element sa `div` koji ima `id` od `legend`.

```js
assert.equal(document.querySelector('#legend')?.tagName, 'DIV');
```

Treba da imaš `span` element sa tekstom `Availability` unutar elementa `#legend`.

```js
assert.equal(document.querySelector('#legend span')?.textContent.trim(), 'Availability');
```

Trebalo bi da imaš komponentu `div` sa `id` od `legend-gradient` unutar svog `#legend`.

```js
assert.equal(document.querySelector('#legend #legend-gradient')?.tagName, 'DIV');
```

Trebalo bi da postaviš pozadinsku sliku `#legend-gradient` koja je mešavina lebdećih boja.

```js
const legendGradient = new __helpers.CSSHelp(document).getStyle('#legend-gradient');
assert.exists(legendGradient); 
assert.include(legendGradient.backgroundImage, "linear-gradient(")
```

Treba da imaš prelazak od jedne boje ka sledećoj boji kao oštrica linija za `#legend-gradient` tvoju. Proveri da su tvoji procenti celobrojni.

```js
const legendGradient = new __helpers.CSSHelp(document).getStyle('#legend-gradient');
assert.exists(legendGradient);
const gradientStr = legendGradient.backgroundImage;

const COLOR_COUNT = 6;
const EXPECTED_SHORTHAND_STOPS = 6;
const EXPECTED_EXPANDED_STOPS = 12;

for (let i = 0; i < COLOR_COUNT; i++) {
  assert.include(gradientStr, `--color${i}`, `gradient should include --color${i}`);
}

const GRADIENT_PATTERNS = {
  SHORTHAND: /var\(\s*--color[0-5]\s*\)\s+\d+%\s+\d+%/g,
  EXPANDED: /var\(\s*--color[0-5]\s*\)\s+\d+%(?!\s+\d+%)/g
};

const matchGradientFormat = (gradientStr, pattern) => gradientStr.match(pattern);

const shorthandMatches = matchGradientFormat(gradientStr, GRADIENT_PATTERNS.SHORTHAND);
const expandedMatches = matchGradientFormat(gradientStr, GRADIENT_PATTERNS.EXPANDED);

const colorPattern = /var\(\s*--color(\d)\s*\)/g;
const allColorMatches = [...gradientStr.matchAll(colorPattern)];
const colorNumbers = allColorMatches.map(match => parseInt(match[1]));

if (shorthandMatches?.length === EXPECTED_SHORTHAND_STOPS) {
  for (let i = 0; i < COLOR_COUNT; i++) {
    const count = colorNumbers.filter(num => num === i).length;
    assert.equal(count, 1, `In shorthand format, --color${i} should appear exactly once, but it appears ${count} times`);
  }
  
  for (let i = 0; i < COLOR_COUNT; i++) {
    const idx = colorNumbers.indexOf(i);
    assert.isAtLeast(idx, 0, `--color${i} should be in the gradient`);
    if (i > 0) {
      const prevIdx = colorNumbers.indexOf(i - 1);
      assert.isBelow(prevIdx, idx, `--color${i - 1} should appear before --color${i}`);
    }
  }
} else if (expandedMatches?.length === EXPECTED_EXPANDED_STOPS) {
  for (let i = 0; i < COLOR_COUNT; i++) {
    const count = colorNumbers.filter(num => num === i).length;
    assert.equal(count, 2, `In expanded format, --color${i} should appear exactly twice, but it appears ${count} times`);
  }
  
  for (let i = 0; i < COLOR_COUNT; i++) {
    const firstIdx = colorNumbers.indexOf(i);
    const lastIdx = colorNumbers.lastIndexOf(i);
    assert.equal(lastIdx - firstIdx, 1, `--color${i} should appear in consecutive positions for hard lines`);
    
    if (i > 0) {
      const prevLastIdx = colorNumbers.lastIndexOf(i - 1);
      assert.isBelow(prevLastIdx, firstIdx, `--color${i - 1} should complete before --color${i} starts`);
    }
  }
  
  const percentPattern = /var\(\s*--color\d\s*\)\s+(\d+)%/g;
  const percentMatches = [...gradientStr.matchAll(percentPattern)];
  const percentages = percentMatches.map(match => parseInt(match[1]));
  
  for (let i = 1; i < percentages.length - 1; i += 2) {
    assert.equal(percentages[i], percentages[i + 1], `For hard lines, percentage at position ${i} (${percentages[i]}%) should equal percentage at position ${i + 1} (${percentages[i + 1]}%) where colors meet`);
  }
} else {
  assert.fail('Gradient must use either shorthand format (6 stops with two percentages each) or expanded format (12 stops with one percentage each)');
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Availability Table</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Availability Table</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <table>
            <tbody>
                <tr class="sharp">
                    <th rowspan="2" class="time">9 AM</th>
                    <th>Monday</th>
                    <th>Tuesday</th>
                    <th>Wednesday</th>
                    <th>Thursday</th>
                    <th>Friday</th>
                </tr>
                <tr class="half">
                    <td class="available-0"></td>
                    <td class="available-1"></td>
                    <td class="available-2"></td>
                    <td class="available-2"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">10 AM</th>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>

                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">11 AM</th>
                    <td class="available-0"></td>
                    <td class="available-2"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-0"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-1"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">12 PM</th>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-3"></td>
                    <td class="available-2"></td>
                    <td class="available-0"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-0"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">1 PM</th>
                    <td class="available-2"></td>
                    <td class="available-3"></td>
                    <td class="available-3"></td>
                    <td class="available-2"></td>
                    <td class="available-0"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">2 PM</th>
                    <td class="available-0"></td>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">3 PM</th>
                    <td class="available-3"></td>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-4"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">4 PM</th>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-4"></td>
                    <td class="available-4"></td>
                </tr>
                <tr class="half">
                    <td class="available-3"></td>
                    <td class="available-0"></td>
                    <td class="available-4"></td>
                    <td class="available-2"></td>
                    <td class="available-5"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time">5 PM</th>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-5"></td>
                    <td class="available-3"></td>
                    <td class="available-4"></td>
                </tr>
                <tr>
                    <td></td>
                    <td></td>
                    <td></td>
                    <td></td>
                    <td></td>
                </tr>
            </tbody>
        </table>
        <div id="legend">
            <span id="legend-title">Availability</span>
            <div id="legend-line"><span>0</span>
                <div id="legend-gradient"></div><span>5+</span>
            </div>
        </div>
    </main>
</body>

</html>
```

```css
:root {
    font-size: 16px;
    font-family: Arial, Helvetica, sans-serif;
    --color0: hsl(0, 0%, 100%);
    --color1: hsl(120, 95%, 90%);
    --color2: hsl(120, 95%, 80%);
    --color3: hsl(120, 95%, 65%);
    --color4: hsl(120, 95%, 40%);
    --color5: hsl(120, 95%, 25%);
    --solid-border: 0.1rem solid black;
    --dashed-border: 0.09rem dashed black;
}

* {
    box-sizing: border-box;
    margin: 0;
    padding: 0;
}

body {
    background-color: hsl(150, 30%, 75%);
    width: 100vw;
    height: 100vh;
}

main {
    width: 100%;
    height: 100%;
    display: flex;
    flex-direction: column;
    align-items: center;
    justify-content: space-evenly;
}

table {
    border-collapse: collapse;
    table-layout: fixed;
    width: 32rem;
}

tr {
    height: 1.1rem;
}

tr :first-child {
    width: 4rem;
}

td {
    border-right: var(--solid-border);
    border-left: var(--solid-border);
}

.sharp td {
    border-bottom: var(--solid-border);
}

.sharp th:not([class*="time"]) {
    border-bottom: var(--solid-border);
}

.half td {
    border-bottom: var(--dashed-border);
}

tbody :last-child td {
    border: 0;
}

.time {
    text-align: right;
    border: 0;
    padding: 0.4rem;
}

.available-0 {
    background-color: var(--color0);
}

.available-1 {
    background-color: var(--color1);
}

.available-2 {
    background-color: var(--color2);
}

.available-3 {
    background-color: var(--color3);
}

.available-4 {
    background-color: var(--color4);
}

.available-5 {
    background-color: var(--color5);
}

#legend {
    display: flex;
    flex-direction: column;
    align-items: center;
    width: 18rem;
    height: 3.5rem;
    text-align: center;
}

#legend span {
    display: inline-block;
    width: 5rem;
    height: 2rem;
}

#legend-gradient {
    width: 100%;
    height: 60%;
    border: var(--solid-border);
    background-image: linear-gradient(90deg,
            var(--color0) 0% 17%,
            var(--color1) 17% 34%,
            var(--color2) 34% 50%,
            var(--color3) 50% 67%,
            var(--color4) 67% 83%,
            var(--color5) 83% 100%);
}

#legend-line {
    width: 60%;
    height: 100%;
    display: flex;
}
```
