---
id: 587d7fa6367417b2b2512bc3
title: Izaberite skup elemenata koristeći D3
challengeType: 6
forumTopicId: 301490
dashedName: select-a-group-of-elements-with-d3
---

# --description--
D3 takođe ima način da `selectAll()` izabere grupu elemenata. Vraća niz čvorova HTML za sve elemente u dokumentu koji odgovaraju unosnom nizu karaktera. Evo primera kako izabrati sve oznake sidroglišča (anchor tag) u dokumentu:

```js
const anchors = d3.selectAll('a');
```

Kao metoda `select()`, `selectAll()` podržava sekvencijalno povezivanje metoda, i možete ga koristiti zajedno sa drugim metodama.

# --instructions--
Izaberite sve oznake za `li` u tekstu i promenite njihov tekst sekvencom slova `list item` spajanjem puta `.text()`.

# --hints--
Postoji 3 elementa `li` na stranici, i tekst u svakom mora da kaže `list item`. Kapitalizacija i razmak moraju biti potpuno identični.

```js
const listItems = document.querySelectorAll('li');
assert.lengthOf(listItems, 3);
for (let i = 0; i < listItems.length; i++) {
  assert.match(listItems[i]?.textContent, /list item/g);
}
```

Vaš kod treba da pristupi objektu `d3`.

```js
assert.match(code, /d3/g);
```

Vaš kod bi trebalo da koristi metod `selectAll`.

```js
assert.match(code, /\.selectAll/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    d3.selectAll('li').text('list item');
  </script>
</body>
```
