---
id: 587d7fac367417b2b2512bdb
title: Postavite domen i opseg vrednosti u merenje
challengeType: 6
forumTopicId: 301491
dashedName: set-a-domain-and-a-range-on-a-scale
---

# --description--
Po defaultu, metrike koriste identitetnu relaciju. To znači da ulazna vrednost odgovara izlaznoj vrednosti. Međutim, metrike mogu biti efikasnije i privlačnije.

Set podataka ima vrednost od 50 do 480. Ovo je informacija o unosu za merenje, takođe poznata kao <dfn>domena</dfn>.

Želite da rasporedite ove tačke duž ose `x` na SVG-u, između jedinica 10 i jedinica 500. Ovo je izveštaj o rezultatima, takođe poznat kao <dfn>red vrednosti</dfn>.

Putanje za `domain()` i `range()` postavljaju ove vrednosti za meru. Oba puta uzimaju red podatke sa najmanje dva elementa kao argumente. Evo primera:

```js
scale.domain([50, 480]);
scale.range([10, 500]);
scale(50);
scale(480);
scale(325);
scale(750);
d3.scaleLinear();
```

Po defaultu, sledeće vrednosti će biti prikazane na konzoli: `10`, `500`, `323.37`, i `807.67`.

Treba znati da mera koristi linearnu relaciju između vrednosti domena i opsega vrednosti kako bi odredila rezultat za određeni broj. Minimalna vrednost u domenu (50) odgovara minimalnoj vrednosti (10) u opsegu vrednosti.

# --instructions--
Kreiraj specifikaciju i postavi njegov domen na `[250, 500]` i red vrednosti na `[10, 150]`.

**Napomena:** Možete povezati putanje `domain()` sa `range()` prilikom promena `scale`.

# --hints--
Vaš kod treba da koristi metodu `domain()`.

```js
assert.match(code, /\.domain/g);
```

`domain()` ya `scale` inapaswa kuwekwa kuwa `[250, 500]`.

```js
assert.deepEqual(scale.domain(), [250, 500]);
```

Tvoj kod bi trebalo da koristi metod `range()`.

```js
assert.match(code, /\.range/g);
```

`range()` ya `scale` inapaswa kuwekwa kuwa `[10, 150]`.

```js
assert.deepEqual(scale.range(), [10, 150]);
```

Tekst u `h2` treba da bude `-102`.

```js
assert.strictEqual(document.querySelector('h2')?.textContent, '-102');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line
    const scale = d3.scaleLinear();



    // Add your code above this line
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const scale = d3.scaleLinear();
    scale.domain([250, 500]);
    scale.range([10, 150]);
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
