---
id: 587d7fad367417b2b2512bdf
title: Dodajte osu na crtež slike
challengeType: 6
forumTopicId: 301472
dashedName: add-axes-to-a-visualization
---

# --description--
Drug način za poboljšanje grafikona tačaka je dodavanje X ose i Y ose.

D3 ima dve metode, `axisLeft()` i `axisBottom()`, za prikaz Y-ose i X-ose, po odvojeno. Ovde je primer kreiranja X-ose na osnovu `xScale` iz prethodnih izazova:

```js
const xAxis = d3.axisBottom(xScale);
```

Sledeći korak je prikazivanje ose na SVG-u. Da biste to uradili, možete koristiti opštu sekciju SVG-a, element `g`. `g` predstavlja grupu. Za razliku od `rect`, `circle` i `text`, osa je samo crvena linija kada se prikaže. Pošto je jednostavna figura, korišćenje `g` funkcioniše. Poslednji korak je korišćenje atributa `transform` za postavljanje ose na ispravno mesto u SVG-u. Inače, linija bi bila prikazana pored granice SVG-a i nije vidljiva. SVG podržava različite tipove `transforms`, ali postavljanje ose zahteva `translate`. Kada se koristi na elementu `g`, pomera celu grupu u bočno i dole za navedanu količinu. Evo primera:

```js
const xAxis = d3.axisBottom(xScale);

svg
  .append('g')
  .attr('transform', 'translate(0, ' + (h - padding) + ')')
  .call(xAxis);
```

Kod pomenut iznad postavlja x-osu ispod SVG. Zatim se prosleđuje kao argument putem `call()`. Y-osa radi na isti način, osim što je argument `translate` u formatu `(x, 0)`. Budući da je `translate` sekvenca karaktera na način `attr()` iznad, možete koristiti konkatenaciju za kombinovanje vrednosti promena za njegove argumente.

# --instructions--
Dijagram rasiporednih podataka sada ima X osu. Kreiraj Y osu u transformaciji sa imenom `yAxis` koristeći metodu `axisLeft()`. Zatim prikaži osu koristeći funkciju `g`. Osiguraj da koristiš svojstvo `transform` za translaciju ose za količinu jedinica unutrašnjeg prostora sa desne strane, i jedinice `0` dole. Zapamti da izvršiš `call()` na osi.

# --hints--
Vaš kod treba da koristi metodu `axisLeft()` i `yScale`, prosleđene kao argumente.

```js
assert.match(code, /\.axisLeft\(yScale\)/g);
```

Komponenta za `g` Y ose mora imati atribut `transform` da transformiše osu za `(60, 0)`.

```js
const yAxis = document.querySelectorAll('g')[10];
assert.match(yAxis?.getAttribute('transform'), /translate\(60\s*,\s*0\)/g);
```

Vaš kod treba da pozove `yAxis`.

```js
assert.match(code, /\.call\(\s*yAxis\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);
    // Add your code below this line
    const yAxis = undefined;
    // Add your code above this line

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);

    const yAxis = d3.axisLeft(yScale);

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    svg
      .append('g')
      .attr('transform', 'translate(' + padding + ',0)')
      .call(yAxis);
  </script>
</body>
```
