---
id: 587d7fac367417b2b2512bde
title: Koristite preddefinisani parametar za postavljanje elemenata
challengeType: 6
forumTopicId: 301494
dashedName: use-a-pre-defined-scale-to-place-elements
---

# --description--
Nakon postavljanja dimenzija, vreme je za ponovno crtanje markera/dijagrama. Dimenzije su poput procesnih staza koje pretvaraju sirove podatke iz `x` i `y` u odgovarajuće vrednosti i prikazuju ih tačno na SVG-u. One osiguravaju da su podaci unutar oblasti crteža na ekranu.

Postavljate koordinate atributne vrednosti za SVG oblik koristeći funkciju merenja. Ovo uključuje atribute `x` i `y` za elemente `rect` ili `text`, ili `cx` i `cy` za `circles`. Evo primera:

```js
shape.attr('x', d => xScale(d[0]));
```

Dimenzije postavljaju atribute koordinata oblika kako bi postavile markere podataka na SVG. Ne morate koristiti dimenzije kada prikazujete stvarnu vrednost podataka, na primer, u putanji `text()` za privremenu napomenuku ili oznaku.

# --instructions--
Koristite `xScale` i `yScale` da oblikujete `circle` na `text` u SVG. Za `circles`, koristite dimenzije da postavite atribute `cx` i `cy`. Takođe postavite prečnik `5`.

Za elemente `text`, koristite dimenzije za postavljanje svojstava `x` i `y`. Etikete bi trebalo da budu blago poravnane u odnosu na desnu stranu znaka. Da biste to uradili, dodajte `10` na vrednost podataka `x` pre nego što mu date `xScale`.

# --hints--
Vaš kod bi trebalo da ima 10 komponenti za `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

Prvi element `circle` bi trebalo da ima vrednost `cx` blizu `91` na vrednosti `cy` blizu `368` nakon korišćenja mera. Takođe bi trebalo da ima vrednost `r` od `5`.

```js
const firstCircle = document.querySelectorAll('circle')[0];
assert.strictEqual(Math.round(firstCircle?.getAttribute('cx')), 91);
assert.strictEqual(Math.round(firstCircle?.getAttribute('cy')), 368);
assert.strictEqual(Math.round(firstCircle?.getAttribute('r')), 5);
```

Element drugi za `circle` treba da ima vrednost `cx` otprilike `159` sa vrednošću `cy` otprilike `181` nakon korišćenja merača. Takođe treba da ima vrednost `r` od `5`.

```js
const secondCircle = document.querySelectorAll('circle')[1];
assert.strictEqual(Math.round(secondCircle?.getAttribute('cx')), 159);
assert.strictEqual(Math.round(secondCircle?.getAttribute('cy')), 181);
assert.strictEqual(Math.round(secondCircle?.getAttribute('r')), 5);
```

Treći element `circle` bi trebalo da ima vrednost `cx` približno `340` na vrednosti `cy` približno `329` nakon korišćenja mera. Takođe bi trebalo da ima vrednost `r` od `5`.

```js
const thirdCircle = document.querySelectorAll('circle')[2];
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cx')), 340);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cy')), 329);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('r')), 5);
```

Četvrti element `circle` treba da ima vrednost od `cx` približno `131` sa vrednošću `cy` približno `60` nakon korišćenja merenja. Takođe treba da ima vrednost od `r` od `5`.

```js
const fourthCircle = document.querySelectorAll('circle')[3];
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cy')), 60);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('r')), 5);
```

Peti element od `circle` treba da ima vrednost `cx` približno `440` sa vrednošću `cy` približno `237` nakon korišćenja merenja. Takođe treba da ima vrednost `r` od `5`.

```js
const fifthCircle = document.querySelectorAll('circle')[4];
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cx')), 440);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cy')), 237);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('r')), 5);
```

Šesti element od `circle` treba da ima vrednost `cx` približno `271` i vrednost od `cy` približno `306` nakon korišćenja merenja. Takođe bi trebalo da ima vrednost `r` od `5`.

```js
const sixthCircle = document.querySelectorAll('circle')[5];
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cx')), 271);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cy')), 306);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('r')), 5);
```

Sedma komponenta `circle` bi trebalo da ima vrednost `cx` blizu `361` na vrednosti `cy` blizu `351` nakon korišćenja mera. Takođe bi trebalo da ima vrednost `r` od `5`.

```js
const seventhCircle = document.querySelectorAll('circle')[6];
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cx')), 361);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cy')), 351);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('r')), 5);
```

Osmi element od `circle` treba da ima vrednost od `cx` približno `261` sa vrednošću `cy` približno `132` nakon izvođenja merenja. Takođe treba da ima vrednost od `r` za `5`.

```js
const eighthCircle = document.querySelectorAll('circle')[7];
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cx')), 261);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cy')), 132);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('r')), 5);
```

Deveti element `circle` bi trebalo da ima vrednost `cx` približno `131` i vrednost `cy` približno `144` nakon primene merenja. Takođe bi trebalo da ima vrednost `r` od `5`.

```js
const ninthCircle = document.querySelectorAll('circle')[8];
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cy')), 144);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('r')), 5);
```

Deseti element `circle` treba da ima vrednost `cx` približno `79` na vrednosti `cy` približno `326` nakon merenja. Takođe treba da ima vrednost `r` od `5`.

```js
const tenthCircle = document.querySelectorAll('circle')[9];
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cx')), 79);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cy')), 326);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('r')), 5);
```

Tvoj kod treba da ima 10 elemenata od `text`.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

Prvi deo bi trebalo da ima vrednost `x` približno `100` na vrednosti `y` približno `368` nakon korišćenja merenja.

```js
const firstLabel = document.querySelectorAll('text')[0];
assert.strictEqual(Math.round(firstLabel?.getAttribute('x')), 100);
assert.strictEqual(Math.round(firstLabel?.getAttribute('y')), 368);
```

Druga oznaka bi trebalo da ima vrednost `x` otprilike `168` od vrednosti `y` otprilike `181` nakon korišćenja mera.

```js
const secondLabel = document.querySelectorAll('text')[1];
assert.strictEqual(Math.round(secondLabel?.getAttribute('x')), 168);
assert.strictEqual(Math.round(secondLabel?.getAttribute('y')), 181);
```

Treća oznaka treba da ima vrednost od `x` približno `350` sa vrednošću `y` približno `329` nakon korišćenja merača.

```js
const thirdLabel = document.querySelectorAll('text')[2];
assert.strictEqual(Math.round(thirdLabel?.getAttribute('x')), 350);
assert.strictEqual(Math.round(thirdLabel?.getAttribute('y')), 329);
```

Četvrti label treba da ima vrednost od `x` približno `141` sa vrednošću od `y` približno `60` nakon korišćenja merenja.

```js
const fourthLabel = document.querySelectorAll('text')[3];
assert.strictEqual(Math.round(fourthLabel?.getAttribute('x')), 141);
assert.strictEqual(Math.round(fourthLabel?.getAttribute('y')), 60);
```

Peti label treba da ima vrednost `x` približno `449` na vrednosti `y` približno `237` nakon korišćenja merenja.

```js
const fifthLabel = document.querySelectorAll('text')[4];
assert.strictEqual(Math.round(fifthLabel?.getAttribute('x')), 449);
assert.strictEqual(Math.round(fifthLabel?.getAttribute('y')), 237);
```

Šesti label treba da ima vrednost od `x` približno `280` na vrednosti od `y` približno `306` nakon korišćenja merenja.

```js
const sixthLabel = document.querySelectorAll('text')[5];
assert.strictEqual(Math.round(sixthLabel?.getAttribute('x')), 280);
assert.strictEqual(Math.round(sixthLabel?.getAttribute('y')), 306);
```

Sedma oznaka treba da ima vrednost od `x` otprilike `370` od vrednosti `y` otprilike `351` nakon korišćenja merenja.

```js
const seventhLabel = document.querySelectorAll('text')[6];
assert.strictEqual(Math.round(seventhLabel?.getAttribute('x')), 370);
assert.strictEqual(Math.round(seventhLabel?.getAttribute('y')), 351);
```

Osma oznaka mora imati vrednost `x` približno `270` na vrednosti od `y` približno `132` nakon upotrebe merenja.

```js
const eighthLabel = document.querySelectorAll('text')[7];
assert.strictEqual(Math.round(eighthLabel?.getAttribute('x')), 270);
assert.strictEqual(Math.round(eighthLabel?.getAttribute('y')), 132);
```

Oznaka devet treba da ima vrednost od `x` približno `140` na vrednosti `y` približno `144` nakon korišćenja merača.

```js
const ninthLabel = document.querySelectorAll('text')[8];
assert.strictEqual(Math.round(ninthLabel?.getAttribute('x')), 140);
assert.strictEqual(Math.round(ninthLabel?.getAttribute('y')), 144);
```

Deseti label treba da ima vrednost `x` otprilike `88` na vrednosti `y` otprilike `326` nakon korišćenja mera.

```js
const tenthLabel = document.querySelectorAll('text')[9];
assert.strictEqual(Math.round(tenthLabel?.getAttribute('x')), 88);
assert.strictEqual(Math.round(tenthLabel?.getAttribute('y')), 326);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1]);
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));
  </script>
</body>
```
