---
id: 587d7fa8367417b2b2512bcc
title: Prikaži oblike koristeći SVG
challengeType: 6
forumTopicId: 301485
dashedName: display-shapes-with-svg
---

# --description--
Poslednja faza kreirala je element za `svg` sa definisanom širinom i visinom, koji se pojavio jer je imao `background-color` korišćen unutar oznake `style`. Kod je napravio prostor za definisanu širinu i visinu.

Sledeći korak je kreiranje oblika za postavljanje u oblast `svg`. Postoje nekoliko oblika podržanih u SVG-u, kao što su pravougaonik i krug. Oni se koriste za prikaz podataka. Na primer, pravougaoni oblik (`<rect>`) u SVG-u može kreirati liniju na linijskom grafikonu.

Kada postavite figuru u području `svg`, možete definisati gde je postavljena koristeći koordinate `x` i `y`. Početna tačka (0, 0) nalazi se u gornjem levom uglu. Pozitivne vrednosti `x` pomeraju figuru udesno, a pozitivne vrednosti `y` pomeraju figuru nadole od početne tačke.

Da biste postavili oblik u sredinu `svg` sa širinom od 500 i visinom od 100 iz poslednjeg izazova, koordinata za `x` biće 250, a koordinata za `y` biće 50.

SVG oznaka za `rect` ima četiri atributa. Postoje koordinati `x` i `y` za lokaciju gde je postavljeno u području `svg`. Takođe ima `height` i `width` za definisanje veličine.

# --instructions--
Postavi oblik `rect` u `svg` koristeći `append()`, i dodeli mu svojstvo `width` sa vrednošću `25` i svojstvo `height` sa vrednošću `100`. Takođe, dodeli `rect` svojstva za `x` i `y`, svako sa vrednošću `0`.

# --hints--
Vaš dokument bi trebalo da ima jedan element od `rect`.

```js
assert.lengthOf(document.querySelectorAll('rect'), 1);
```

Element za `rect` trebalo bi da ima svojstvo `width` sa vrednošću `25`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('width'), '25');
```

Atribut `rect` treba da ima atribut `height` sa vrednošću `100`.

```js
assert.strictEqual(
  document.querySelector('rect')?.getAttribute('height'),
  '100'
);
```

Element za `rect` treba da ima svojstvo `x` sa vrednošću `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('x'), '0');
```

Element `rect` treba da ima svojstvo `y` sa vrednošću `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('y'), '0');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h)
      .append('rect')
      .attr('width', 25)
      .attr('height', 100)
      .attr('x', 0)
      .attr('y', 0);
  </script>
</body>
```
