---
id: 587d7fa7367417b2b2512bc7
title: Promeni stil prema podacima
challengeType: 6
forumTopicId: 301479
dashedName: change-styles-based-on-data
---

# --description--
D3 se tiče predstavljanja i vizualizacije podataka. Možda želite da promenite stil elemenata u zavisnosti od podataka.
Na primer, možda želite da obojite tačku podacima plavom ako je vrednost ispod 20, a crvenom inače. Možete koristiti callback funkciju na mestu `style()` i uključiti logiku uslova. Callback funkcija koristi parametar `d` za predstavljanje tačke podataka:

```js
selection.style('color', d => {});
```

Putanja za `style()` nije ograničena samo postavljanjem `color` - može se koristiti i sa drugim svojstvima od CSS.

# --instructions--
Dodajte putanju ``style()`` u kod koji se nalazi u editoru kako biste uslovno postavili ``color`` za komponente ``h2``. Napišite funkciju povratnog poziva (callback) tako da ako je vrednost podataka manja od 20, vrati crveno, inače vrati zeleno.

**Napomena:** Možete koristiti logiku if-else, ili ternarni operator.

# --hints--

Prvi `h2` treba da ima `color` postavljen na crvenu.

```js
const headingTwoFirst = document.querySelector('h2');
assert.exists(headingTwoFirst); 
const color = window.getComputedStyle(headingTwoFirst)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Drugi `h2` treba da ima `color` postavljen na zelenu.

```js
const headingTwoSecond = document.querySelectorAll('h2')[1];
assert.exists(headingTwoSecond); 
const color = window.getComputedStyle(headingTwoSecond)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Treći `h2` treba da ima `color` postavljen na zelenu.

```js
const headingTwoThird = document.querySelectorAll('h2')[2];
assert.exists(headingTwoThird); 
const color = window.getComputedStyle(headingTwoThird)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

`h2` ya nne inapaswa kuwa na `color` ya nyekundu.

```js
const headingTwoFourth = document.querySelectorAll('h2')[3];
assert.exists(headingTwoFourth); 
const color = window.getComputedStyle(headingTwoFourth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Peti `h2` treba da ima `color` postavljen na zelenu.

```js
const headingTwoFifth = document.querySelectorAll('h2')[4];
assert.exists(headingTwoFifth); 
const color = window.getComputedStyle(headingTwoFifth)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Šesti `h2` treba da ima `color` postavljen na crvenu.

```js
const headingTwoSixth = document.querySelectorAll('h2')[5];
assert.exists(headingTwoSixth); 
const color = window.getComputedStyle(headingTwoSixth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Sedmi `h2` treba da ima `color` postavljen na zelenu.

```js
const headingTwoSeventh = document.querySelectorAll('h2')[6];
assert.exists(headingTwoSeventh); 
const color = window.getComputedStyle(headingTwoSeventh)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

`h2` ya nane inapaswa kuwa na `color` ya nyekundu.

```js
const headingTwoEighth = document.querySelectorAll('h2')[7];
assert.exists(headingTwoEighth); 
const color = window.getComputedStyle(headingTwoEighth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Deveti `h2` treba da ima `color` postavljen na crvenu.

```js
const headingTwoNinth = document.querySelectorAll('h2')[8];
assert.exists(headingTwoNinth); 
const color = window.getComputedStyle(headingTwoNinth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('color', d => (d < 20 ? 'red' : 'green'));
  </script>
</body>
```
