---
id: 587d7fac367417b2b2512bdc
title: Koristite funkcije `d3.max` i `d3.min` da biste pronašli maksimalnu i minimalnu vrednost u skupu podataka.
challengeType: 6
forumTopicId: 301496
dashedName: >-
  use-the-d3-max-and-d3-min-functions-to-find-minimum-and-maximum-values-in-a-dataset
---

# --description--
Načini D3 `domain()` i `range()` postavljaju te informacije na vašem nivou u zavisnosti od podataka. Postoji nekoliko načina da se to uradi lakše.

Često kada definišete domen, želite da koristite minimalne i maksimalne vrednosti unutar skupa podataka. Pokušaj ručnog određivanja ovih nivoa, posebno u velikom skupu podataka, može izazvati greške.

D3 ima dva načina - `min()` i `max()` za vraćanje ovih informacija. Evo primera:

```js
const exampleData = [34, 234, 73, 90, 6, 52];
d3.min(exampleData);
d3.max(exampleData);
```

Skup podataka može sadržati ugnježdeni nizove podataka, kao što su parovi koordinata `[x, y]` koji se nalaze u primeru raspršenog grafikona. U tom slučaju, morate reći D3 kako da izračuna maksimum i minimum. Srećom, metode `min()` i `max()` sve uzimaju funkciju povratne pozive (callback function). U ovom primeru, argument funkcije povratne pozive `d` je namenjen trenutnom unutrašnjem nizu. Funkcija povratne pozive mora vratiti element iz unutrašnjeg niza (vrednost `x` ili `y`) za koji želite da izračunate maksimum ili minimum. Evo primera kako dobiti vrednosti minima i maksimuma za niz niza podataka:

```js
const locationData = [
  [1, 7],
  [6, 3],
  [8, 3]
];
const minX = d3.min(locationData, d => d[0]);
```

`minX` ima vrednost od `1`.

# --instructions--
Red `positionData` ima podredove koordinata x, y i z. Koristite funkciju D3 da dobijete najvišu koordinatnu vrednost za z (treća vrednost) iz tih redova i unesite je u parametar `output`.

# --hints--
Tekst u `h2` bi trebalo da bude `8`.

```js
assert.strictEqual(output, 8);
assert.strictEqual(document.querySelector('h2')?.textContent, '8');
```

Vaš kod treba da koristi metodu `max()`.

```js
assert.match(
  code,
  /\.max/g,
  'Your code should use the <code>max()</code> method.'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];
    // Add your code below this line

    const output = undefined; // Change this line

    // Add your code above this line

    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];

    const output = d3.max(positionData, d => d[2]);

    d3.select('body').append('h2').text(output);
  </script>
</body>
```
