---
id: 587d7fa8367417b2b2512bcb
title: Naučite o SVG u D3
challengeType: 6
forumTopicId: 301489
dashedName: learn-about-svg-in-d3
---

# --description--

<dfn>SVG</dfn> inamaanisha <dfn>Scalable Vector Graphics</dfn>.

Ovde "scalable" znači da, ako povećate ili smanjite veličinu objekta, neće izgledati kao slika sa pikselima. Prilagođava se prema sistemu prikaza, bilo da je na malom ekranu telefona ili na velikom televizijskom ekranu.

SVG se koristi za kreiranje standardnih geometrijskih oblika. Pošto D3 predstavlja podatke vizuelno, koristi SVG za generisanje oblika te reprezentacije. SVG oblici za veb stranicu moraju biti unutar oznake HTML `svg`.

CSS može biti skalabilno kada stilovi koriste relativne jedinice (kao što su `vh`, `vw`, ili procenat), ali korišćenje SVG-a je lakše za kreiranje reprezentacije podataka za slike.

# --instructions--
Dodajte čvor za `svg` na `body` koristeći `append()`. Postavite svojstvo `width` postavljeno na konstantu `w` i svojstvo `height` postavljeno na konstantu `h` koristeći metode `attr()` ili `style()` za svaku od njih. Videćete ga u rezultatima jer postoji ružičasta boja `background-color` korišćena na etiketi `style`.

**Napomena:** Kada koristite atribute širine i visine za `attr()`, dimenzije nisu dostupne. Ovo je osnova za određivanje veličine - element će imati odnos širina-visina od 5:1, bez obzira na nivo skaliranja ili smanjenja veličine.

# --hints--
Vaš dokument treba da ima element 1 za `svg`.

```js
assert.lengthOf(document.querySelectorAll('svg'), 1);
```

Element za `svg` mora imati svojstvo `width` postavljeno za `500` ili sa širinom od `500px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const width = window.getComputedStyle(svg)['width'];
assert.isTrue(svg?.getAttribute('width') === '500' || width === '500px');
```

Element `svg` mora imati svojstvo `height` definisano za `100` ili biti postavljeno na dužinu od `100px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const height = window.getComputedStyle(svg)['height'];
assert.isTrue(svg?.getAttribute('height') === '100' || height === '100px');
```

# --seed--

## --seed-contents--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3.select('body');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
  </script>
</body>
```
