---
id: 587d7fa9367417b2b2512bd0
title: Pretvorite SVG elemente
challengeType: 6
forumTopicId: 301488
dashedName: invert-svg-elements
---

# --description--
Možda ste primetili linijski grafikon koji izgleda kao da je naopako, ili obrnuto. To je zbog načina na koji SVG koristi koordinate (x, y).

U SVG-u, početna koordinata se nalazi u gornjem levom uglu. Koordinata `x` sa vrednošću 0 određuje oblik na levoj ivici SVG oblasti. Koordinata `y` sa vrednošću 0 određuje oblik na gornjoj ivici SVG oblasti. Velike vrednosti `x` pomeraju pravougaonik udesno. Velike vrednosti `y` pomeraju pravougaonik nadole.

Da biste učinili linije jednakima gornjem ivici, morate promeniti način merenja koordinata `y`. Potrebno je uzeti u obzir dužinu linije i ukupnu visinu SVG područja.

Visina SVG područja je 100. Ako imate merenje podataka od 0 u skupu, želite da linija počne ispod SVG područja (a ne iznad). Da biste to uradili, koordinata `y` zahteva vrednost od 100. Ako je vrednost merenja podataka 1, započeli biste sa koordinatom `y` sa vrednošću 100 da postavite liniju ispod. Zatim morate uzeti u obzir dužinu linije od 1, pa će konačna koordinata `y` biti 99.

Poravnanje `y` koje je `y = heightOfSVG - heightOfBar` će učiniti da redovi budu poravnati sa gornjom stranom.

# --instructions--
Promenite callback funkciju svojstva `y` ili postavite redove jednake gornjoj vrednosti. Napomena da je red `height` tri puta vrednost podataka `d`.

**Napomena:** Generalno, odnos je `y = h - m * d`, gde `m` je konstanta koja meri dimenzije podataka.

# --hints--
`rect` prvi mora imati vrednost `y` od `64`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('y'),
  (h - dataset[0] * 3).toString()
);
```

`rect` drugi treba da ima vrednost od `y` od `7`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('y'),
  (h - dataset[1] * 3).toString()
);
```

`rect` treći mora imati vrednost od `y` za `34`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('y'),
  (h - dataset[2] * 3).toString()
);
```

`rect` mora imati vrednost `y` od `49`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('y'),
  (h - dataset[3] * 3).toString()
);
```

`rect` peti treba da ima vrednost od `y` od `25`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('y'),
  (h - dataset[4] * 3).toString()
);
```

`rect` šestо mora imati vrednost `y` od `46`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('y'),
  (h - dataset[5] * 3).toString()
);
```

`rect` sedmi mora imati vrednost `y` od `13`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('y'),
  (h - dataset[6] * 3).toString()
);
```

`rect` menta trebao bi da ima vrednost od `y` od `58`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('y'),
  (h - dataset[7] * 3).toString()
);
```

`rect` devetog treba da ima vrednost od `y` za `73`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('y'),
  (h - dataset[8] * 3).toString()
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => {
        // Add your code below this line


        
        // Add your code above this line
      })
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```
