---
id: 587d7fa7367417b2b2512bc8
title: Ongeza madarasa na D3
challengeType: 6
forumTopicId: 301473
dashedName: add-classes-with-d3
---

# --description--
Korišćenje više ugrađenih stilova na HTML elementima teško je za održavanje, čak i u malim aplikacijama. Mnogo je lakše dodati klasu elementima, a zatim definisati stil te klase pomoću CSS pravila. D3 omogućava da pomoću metode `attr()` dodate bilo koji HTML atribut, uključujući i naziv klase.

Metoda `attr()` radi isto kao `style()`. Prihvata vrednosti razdvojene zarezom i može koristiti funkciju povratnog poziva (callback function). Evo primera dodavanja klase `container` za izbor:

```js
selection.attr('class', 'container');
```

Zapamtite da će parametar `class` ostati isti svaki put kada želite dodati klasu, a samo parametar `container` će se menjati.

# --instructions--
Dodajte putanju za `attr()` u kod koji se nalazi u editoru i stavite klasu `bar` u elemente/komponente `div`.

# --hints--
Vaši elementi za `div` moraju imati klasu od `bar`.

```js
assert.isTrue(document.querySelector('div')?.classList.contains('bar'));
```

Tvoj kod bi trebalo da koristi njia `attr()`.

```js
assert.match(code, /\.attr/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line
  </script>
</body>
```
