---
id: 587d7fab367417b2b2512bd9
title: Dodaj oznaku za tačke na dijagramu raspršenosti
challengeType: 6
forumTopicId: 301477
dashedName: add-labels-to-scatter-plot-circles
---

# --description--
Možete dodati tekst/oznake za tačku u dijagram rasipanja.

Cilj je prikazati vrednosti napisane zarezom za prvi deo (`x`) i drugi deo (`y`) svakog elementa u `dataset`.

Čvorovi za `text` zahtevaju atribute od `x` i `y` kako bi se postavili u SVG. U ovom izazovu, vrednost od `y` (koja određuje dužinu) može koristiti istu vrednost koju `circle` koristi za svoj atribut `cy`. Vrednost od `x` može biti malo veća od vrednosti od `cx` za `circle`, kako bi labela nestala. Ovo će pomeriti labelu ka desnoj strani postavljene tačke.

# --instructions--
Označite svaku tačku na dijagramu raspršene tačke koristeći elemente `text`. Tekst oznake treba da bude dve vrednosti razdvojene zarezom i razmakom. Na primer, oznaka prve tačke je `34, 78`. Postavite svojstvo `x` da ima vrednost jedinica `5` veće od vrednosti koje ste koristili za svojstvo `cx` na `circle`. Postavite svojstvo `y` na isti način korišćen za vrednost `cy` na `circle`.

# --hints--
Tvoj kod treba da ima 10 elemenata od `text`.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

Prvi red treba da ima tekst `34, 78`, vrednost `x` od `39`, na vrednosti `y` od `422`.

```js
const labelOne = document.querySelector('text');
assert.strictEqual(labelOne?.textContent, '34, 78');
assert.strictEqual(labelOne?.getAttribute('x'), '39');
assert.strictEqual(labelOne?.getAttribute('y'), '422');
```

Drugi label treba da sadrži tekst `109, 280`, vrednost `x` od `114`, sa vrednošću `y` od `220`.

```js
const labelTwo = document.querySelectorAll('text')[1];
assert.strictEqual(labelTwo?.textContent, '109, 280');
assert.strictEqual(labelTwo?.getAttribute('x'), '114');
assert.strictEqual(labelTwo?.getAttribute('y'), '220');
```

Treći label treba da ima tekst `310, 120`, vrednost `x` od `315`, na vrednosti `y` od `380`.

```js
const labelThree = document.querySelectorAll('text')[2];
assert.strictEqual(labelThree?.textContent, '310, 120');
assert.strictEqual(labelThree?.getAttribute('x'), '315');
assert.strictEqual(labelThree?.getAttribute('y'), '380');
```

Četvrti label treba da ima tekst `79, 411`, vrednosti `x` za `84`, na vrednosti `y` od `89`.

```js
const labelFour = document.querySelectorAll('text')[3];
assert.strictEqual(labelFour?.textContent, '79, 411');
assert.strictEqual(labelFour?.getAttribute('x'), '84');
assert.strictEqual(labelFour?.getAttribute('y'), '89');
```

Peti labela treba da ima tekst od `420, 220`, vrednost od `x` od `425`, sa vrednošću `y` od `280`.

```js
const labelFive = document.querySelectorAll('text')[4];
assert.strictEqual(labelFive?.textContent, '420, 220');
assert.strictEqual(labelFive?.getAttribute('x'), '425');
assert.strictEqual(labelFive?.getAttribute('y'), '280');
```

Šesta etiketa treba da ima tekst od `233, 145`, vrednost `x` od `238`, i vrednost `y` od `355`.

```js
const labelSix = document.querySelectorAll('text')[5];
assert.strictEqual(labelSix?.textContent, '233, 145');
assert.strictEqual(labelSix?.getAttribute('x'), '238');
assert.strictEqual(labelSix?.getAttribute('y'), '355');
```

Sedma oznaka treba da ima tekst `333, 96`, vrednost od `x` za `338`, sa vrednošću od `y` za `404`.

```js
const labelSeven = document.querySelectorAll('text')[6];
assert.strictEqual(labelSeven?.textContent, '333, 96');
assert.strictEqual(labelSeven?.getAttribute('x'), '338');
assert.strictEqual(labelSeven?.getAttribute('y'), '404');
```

Osma oznaka treba da ima tekst od `222, 333`, vrednost `x` od `227`, na vrednosti `y` od `167`.

```js
const labelEight = document.querySelectorAll('text')[7];
assert.strictEqual(labelEight?.textContent, '222, 333');
assert.strictEqual(labelEight?.getAttribute('x'), '227');
assert.strictEqual(labelEight?.getAttribute('y'), '167');
```

Deveti label bi trebalo da ima tekst `78, 320`, vrednost `x` za `83`, sa vrednošću `y` od `180`.

```js
const labelNine = document.querySelectorAll('text')[8];
assert.strictEqual(labelNine?.textContent, '78, 320');
assert.strictEqual(labelNine?.getAttribute('x'), '83');
assert.strictEqual(labelNine?.getAttribute('y'), '180');
```

Deseta etiketa mora imati tekst `21, 123`, vrednost `x` od `26`, na vrednost `y` od `377`.

```js
const labelTen = document.querySelectorAll('text')[9];
assert.strictEqual(labelTen?.textContent, '21, 123');
assert.strictEqual(labelTen?.getAttribute('x'), '26');
assert.strictEqual(labelTen?.getAttribute('y'), '377');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg.selectAll('text').data(dataset).enter().append('text');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', d => d[0] + 5)
      .attr('y', d => h - d[1])
      .text(d => d[0] + ', ' + d[1]);
  </script>
</body>
```
