---
id: 587d7fa9367417b2b2512bce
title: Nateraj smer svake linije
challengeType: 6
forumTopicId: 301487
dashedName: dynamically-set-the-coordinates-for-each-bar
---

# --description--
Poslednji izazov je kreirao i dodao pravougaonik elementu `svg` za svaku tačku u `dataset` koja predstavlja liniju. Nažalost, sve je bilo postavljeno jedno na drugo.

Lokacija za postavljanje pravougaonika povezana je sa svojstvima `x` i `y`. Govori D3 gde da počne crtati oblik u području `svg`. Posledji izazov je sve postavio na 0, pa je svaka linija postavljena u gornji levi ugao.

Za linijski grafikon, sve linije bi trebalo da budu na istoj vertikalnoj nivou, što znači da vrednost `y` ostaje konstantna (na 0) za sve linije. Međutim, vrednost `x` mora da se promeni kada dodate nove linije. Imaj na umu da velike vrednosti `x` guraju stvari dalje udesno. Kada pređete preko elemenata redova podataka u `dataset`, vrednost `x` bi trebalo da se poveća.

Metoda ``attr()`` u D3 prima povratnu funkciju za eksplicitno postavljanje tog atributa. Povratna funkcija prima dva argumenta: jedan za samu tačku podataka (obično ``d``) i jedan za indeks tačke podataka u nizu podataka. Drugi argument je opcionalan. Ovo je struktura:

```js
selection.attr('property', (d, i) => {});
```

Važno je znati da NIJE potrebno pisati petlju za `for` niti koristiti `forEach()` za iteraciju po elementima u skupu podataka. Napomenuti da metoda `data()` analizira skup podataka, i da se bilo koja sledeća metoda `data()` izvršava jednom za svaki element u skupu podataka.

# --instructions--
Promenite funkciju poziva atributa `x` ili ponovite uzorak 30 puta.

**Napomena:** Sve linije imaju širinu od 25, pa povećanje vrednosti za `x` za 30 povećava razmak između redova. Bilo koja vrednost veća od 25 funkcioniše u ovom primeru.

# --hints--
Prvi `rect` mora imati vrednost `x` od `0`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('x'),
  '0'
);
```

`rect` druga mora imati vrednost `x` od `30`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('x'),
  '30'
);
```

`rect` treći treba da ima vrednost `x` od `60`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('x'),
  '60'
);
```

`rect` četvrti mora imati vrednost od `x` od `90`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('x'),
  '90'
);
```

`rect` peto treba da ima vrednost `x` od `120`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('x'),
  '120'
);
```

`rect` šesti treba da ima vrednost `x` od `150`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('x'),
  '150'
);
```

`rect` sedmi treba da ima vrednost od `x` od `180`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('x'),
  '180'
);
```

`rect` ya nane mora imati vrednost od `x` od `210`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('x'),
  '210'
);
```

`rect` deveti deo treba da ima vrednost od `x` od `240`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('x'),
  '240'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        // Add your code below this line


      
        // Add your code above this line
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        return i * 30;
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
