---
id: 587d7fa9367417b2b2512bcf
title: Promeni na silom dužinu svake linije
challengeType: 6
forumTopicId: 301486
dashedName: dynamically-change-the-height-of-each-bar
---

# --description--
Dužina svake linije može se postaviti na vrednost polja podataka u redovima podataka, kao što je prisilno postavljena vrednost `x`.

```js
selection.attr('property', (d, i) => {});
```

Ovde `d` će biti vrednost elementa podataka, i `i` će biti reprezentacija elementa podataka u redovima podataka.

# --instructions--
Promenite funkciju povratne vrednosti atributa `height` ili vratite vrednost podataka pomnoženu sa 3.

**Napomena:** Zapamtite da množenje svih elemenata podataka sa jednim brojem povećava veličinu podataka (kao dodavanje slike). Ovo pomaže u vidljivom razlikovanju vrednosti redova u ovom primeru.

# --hints--
`rect` prvi treba da ima `height` od `36`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('height'),
  '36'
);
```

Drugi `rect` treba da ima `height` od `93`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('height'),
  '93'
);
```

Treći `rect` treba da ima `height` od `66`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('height'),
  '66'
);
```

`rect` ya nne inapaswa kuwa na `height` ya `51`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('height'),
  '51'
);
```

Peti `rect` treba da ima `height` od `75`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('height'),
  '75'
);
```

Šesti `rect` treba da ima `height` od `54`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('height'),
  '54'
);
```

`rect` ya saba inapaswa kuwa na `height` ya `87`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('height'),
  '87'
);
```

`rect` ya nane inapaswa kuwa na `height` ya `42`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('height'),
  '42'
);
```

Deveti `rect` treba da ima `height` od `27`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('height'),
  '27'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        // Add your code below this line


        
        
        // Add your code above this line
      });
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        return d * 3;
      });
  </script>
</body>
```
