---
id: 587d7fac367417b2b2512bdd
title: Tumia vipimo vinavyobadilika
challengeType: 6
forumTopicId: 301495
dashedName: use-dynamic-scales
---

# --description--
Putanje za D3 `min()` i `max()` su važne jer pomažu u postavljanju standarda.

Ako imate kompleksan skup podataka, jedan prioritet je postaviti dimenziju tako da prikaz odgovara širini i visini SVG kontejnera. Želite da su svi podaci nacrtani unutar SVG-a kako bi bili vidljivi na veb stranici.

Primer ispod postavlja meru x-osu podataka raspršenog grafikona. Funkcija ``domain()`` prenosi informacije o merama o sirovim vrednostima podataka za grafikona. Funkcija ``range()`` pruža informacije o stvarnoj poziciji na veb stranici za prikaz.

Na primer, domenni opseg počinje od 0 do najveće vrednosti u skupu. Koristi metodu `max()` zajedno sa funkcijom povratne pozive koja zavisi od vrednosti x u redovima podataka. Opsegni opseg koristi širinu SVG-a (`w`), ali takođe uključuje unutrašnji prostor. Ovo ostavlja prostor između tačaka crtanja markera i ivica SVG-a.

```js
const dataset = [
  [34, 78],
  [109, 280],
  [310, 120],
  [79, 411],
  [420, 220],
  [233, 145],
  [333, 96],
  [222, 333],
  [78, 320],
  [21, 123]
];
const w = 500;
const h = 500;

const padding = 30;
const xScale = d3
  .scaleLinear()
  .domain([0, d3.max(dataset, d => d[0])])
  .range([padding, w - padding]);
```

Unutrašnji prostor može izgledati teško na početku. Zamislite X osu kao horizontalnu liniju od 0 do 500 (vrednost širine SVG). Uključivanje unutrašnjeg prostora u putanju od `range()` prisiljava crtež da počne na 30 duž te linije (umesto 0), i da se završi na 470 (umesto 500).

# --instructions--
Koristi parametar `yScale` za kreiranje dimenzije Y ose koja varira linearno. Domenski opseg treba da počne od nule i do najveće vrednosti `y` u skupu. Opseg treba da koristi visinu SVG-a (`h`) i uključuje interni prostor.

**Napomena:** Zapamtite da postavite dijagram pravilno. Kada postavljate opseg za Y koordinate, najveća vrednost (dužina minus unutrašnji prostor) je prvi argument, a najmanja vrednost je drugi argument.

# --hints--
Tekst u `h2` trebao bi da bude `30`.

```js
assert.strictEqual(output, 30);
assert.strictEqual(document.querySelector('h2')?.textContent, '30');
```

`domain()` funkcije `yScale` treba da bude jednak `[0, 411]`.

```js
assert.deepEqual(yScale.domain(), [0, 411]);
```

`range()` funkcije `yScale` treba da bude jednak `[470, 30]`.

```js
assert.deepEqual(yScale.range(), [470, 30]);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    // Padding between the SVG boundary and the plot
    const padding = 30;

    // Create an x and y scale

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    // Add your code below this line

    const yScale = undefined;


    // Add your code above this line

    const output = yScale(411); // Returns 30
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const padding = 30;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const output = yScale(411);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
