---
id: 587d7fab367417b2b2512bd7
title: Kreiraj raspršeni chat za SVG cikluse
challengeType: 6
forumTopicId: 301484
dashedName: create-a-scatterplot-with-svg-circles
---

# --description--
Dijagram raspršenih tačaka je drugi tip vizualizacije podataka. Obično koristi krugove za prikaz tačaka podataka, od kojih svaka ima dve vrednosti. Ove vrednosti se odnose na osu `x` i `y`, i koriste se za postavljanje kruga u prezentaciji.

SVG ima mogućnost kreiranja kružnog oblika putem `circle`. Funkcioniše na način sličan elementima `rect` koje ste koristili za linijski grafikon.

# --instructions--
Koristite puteve `data()`, `enter()` i `append()` za povezivanja `dataset` za nove elemente `circle` dodate u SVG.

**Napomena:** Krugovi neće biti vidljivi jer još nismo postavili njihove svojstva elemenata. To ćemo učiniti u sledećem izazovu.

# --hints--
Tvoj kod mora imati 10 komponenti od `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
  </script>
</body>
```
