---
id: 587d7fa8367417b2b2512bc9
title: Podesite dužinu elementa transformacijom
challengeType: 6
forumTopicId: 301493
dashedName: update-the-height-of-an-element-dynamically
---

# --description--
Prethovni izazovi uključivali su kako prikazati podatke iz niza podataka i kako dodati klase za CSS. Možete povezati ove lekcije da biste kreirali jednostavan linijski grafikon. Postoje dva koraka za to:

1. Generišite `div` za svaki element podataka u redu sa podacima

2. Vrati dužinu `div` transformacijom, koristeći funkciju povratnog poziva u metodi `style()` koja postavlja dužinu jednaku vrednosti podataka.

Zapamti strukturu za postavljanje stila korišćenjem funkcije povratnika (callback):

```js
selection.style('cssProperty', d => d);
```

# --instructions--
Dodajte putanju/način `style()` u kod u editoru da postavite atribut `height` za svaki element. Koristite funkciju povratne pozive (callback) da vratite vrednost elementa podataka zajedno sa podržanim nizom karaktera `px`.

# --hints--
`div` prvi treba da ima `height` od piksela `12`.

```js
assert.strictEqual(document.querySelectorAll('div')[0]?.style.height, '12px');
```

Drugi `div` treba da ima `height` od `31` piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[1]?.style.height, '31px');
```

Treći `div` treba da ima `height` od `22` piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[2]?.style.height, '22px');
```

Četvrti `div` treba da ima `height` od `17` piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[3]?.style.height, '17px');
```

Peti `div` treba da ima `height` od `25` piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[4]?.style.height, '25px');
```

Šesti `div` treba da ima `height` od `18` piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[5]?.style.height, '18px');
```

Sedmi `div` treba da ima `height` od `29` piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[6]?.style.height, '29px');
```

Osmi `div` treba da ima `height` od `14` piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[7]?.style.height, '14px');
```

Deveti `div` treba da ima `height` od `9` piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[8]?.style.height, '9px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => `${d}px`);
  </script>
</body>
```
