---
id: 587d7fa8367417b2b2512bcd
title: Kreiraj šipku za svaku tačku podataka u skupu
challengeType: 6
forumTopicId: 301482
dashedName: create-a-bar-for-each-data-point-in-the-set
---

# --description--
Posledji izazov je dodao samo jedan kvadrat na element za `svg` koji predstavlja bar. Ovde ćete kombinovati ono što ste naučili do sada o `data()`, `enter()`, i SVG oblike da biste kreirali i dodali kvadrat za svaku tačku podataka u `dataset`.

Početni izazov je pokazao obrazac kako da se kreira i doda `div` za svaki element u `dataset`:

```js
d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
```

Postoje male razlike kada koristite elemente `rect` umesto elemenata `div`. Elementi `rect` moraju se dodati na element `svg`, a ne direktno u `body`. Takođe, morate navesti gde da postavite svaki `rect` unutar oblasti `svg`. Pozicioniranje trake će biti obrađeno u sledećem zadatku.

# --instructions--
Koristite puteve `data()`, `enter()` i `append()` za kreiranja sa dodavanjem `rect` za svaki element u `dataset`. Sve trake bi trebalo da se prikazuju iznad svakog; ovo će biti popravljeno u sledećem izazovu.

# --hints--
Vaš dokument mora imati 9 elemenata od `rect`.

```js
assert.lengthOf(document.querySelectorAll('rect'), 9);
```

Vaš kod bi trebalo da koristi metodu `data()`.

```js
assert.match(code, /\.data/g);
```

Vaš kod bi trebalo da koristi `enter()`.

```js
assert.match(code, /\.enter/g);
```

Vaš kod bi trebalo da koristi metodu `append()`.

```js
assert.match(code, /\.append/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      // Add your code below this line

      // Add your code above this line
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
