---
id: 587d7faa367417b2b2512bd6
title: Dodaj privremeni placeholder za element D3
challengeType: 6
forumTopicId: 301470
dashedName: add-a-tooltip-to-a-d3-element
---

# --description--
Tooltip prikazuje više informacija o elementu na stranici kada korisnik pređe mišem preko tog elementa. Postoji nekoliko načina za dodavanje tooltip-a na prikaz slika. Ovaj izazov koristi SVG element `title`.

`title` je praćen putem `text()` dodavanja podataka pritiskom na šipke.

# --instructions--
Dodajte element ``title`` ispod svakog čvora od ``rect``. Zatim pozovite metodu ``text()`` sa callback funkcijom kako bi tekst prikazao vrednost podataka.

# --hints--
Tvoj kod mora imati 9 komponenti `title`.

```js
assert.lengthOf(document.querySelectorAll('title'), 9);
```

Prvi element `title` treba da ima tekst privremenog naslova `12`.

```js
assert.strictEqual(document.querySelectorAll('title')[0]?.textContent, '12');
```

Drugi element `title` bi trebalo da ima tekst privremenog naslova `31`.

```js
assert.strictEqual(document.querySelectorAll('title')[1]?.textContent, '31');
```

Treće polje `title` treba da ima tekst privremenog naslova `22`.

```js
assert.strictEqual(document.querySelectorAll('title')[2]?.textContent, '22');
```

Četvrti element od `title` treba da ima tekst privremenog placeholder-a za `17`.

```js
assert.strictEqual(document.querySelectorAll('title')[3]?.textContent, '17');
```

Peti element od `title` treba da ima privremeni tekst naslovnica za `25`.

```js
assert.strictEqual(document.querySelectorAll('title')[4]?.textContent, '25');
```

Šesto polje `title` treba da sadrži tekst privremenog naslova `18`.

```js
assert.strictEqual(document.querySelectorAll('title')[5]?.textContent, '18');
```

Sedma stavka `title` treba da ima privremeni tekst za `29`.

```js
assert.strictEqual(document.querySelectorAll('title')[6]?.textContent, '29');
```

Osmo polje `title` bi trebalo da ima tekst privremenog naslova `14`.

```js
assert.strictEqual(document.querySelectorAll('title')[7]?.textContent, '14');
```

Deveti element `title` treba da ima privremeni tekst za pomoć od `9`.

```js
assert.strictEqual(document.querySelectorAll('title')[8]?.textContent, '9');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar')
      .append('title')
      .text(d => d);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```
