---
id: 657efa642593c5746acc5c81
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

U telu izjave za `if`, morate konvertovati svako veliko slovo u malo slovo i dodati donje crte pre ovog malog slova.

Koristi sekvencijalni način abekasa `.lower()` da pretvori velika slova u mala. Zatim, dodaj razmak ispod tog slova. Stavi rezultat u parametar nazvan `converted_character`.

# --hints--

Trebalo bi da postavite vrednost izmenjenog karaktera u parametar nazvan `converted_character`.

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )
      
      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

Ne bi trebalo da imaš `pass` u tvojoj izjavi od `if`.

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
