---
id: 657effaa2a5e0277d71f9cbe
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Sada, promena `snake_cased_char_list` čuva listu izmenjenih slova. Da biste ove slova spojili u jedan niz znakova, možete koristiti metodu `.join()`.

Metoda `join` radi spajanjem svakog elementa liste u niz znakova, koji je razdvojen posebnim nizom znakova, poznatim kao separator.

```py
result_string = ''.join(characters)
```

Primer iznad spaja elemente liste `characters` u jedan niz znakova, gde je svaki element spojen korišćenjem praznog niza znakova kao razdvojnika.

Sada, nakon ciklusa od `for`, koristite metodu `.join()` za kombinovanje elemenata u `snake_cased_char_list` koristeći prazan niz znakova kao razdvojnik. Sačuvajte te rezultate u novu promenljivu nazvanu `snake_cased_string`.


# --hints--

Trebalo bi da spojite slova iz liste `snake_cased_char_list` u sekvencu slova koristeći metodu `.join()` i staviti tu sekvencu u promenljivu nazvanu `snake_cased_string` nakon vašeg ciklusa od `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
