---
id: 657e928716b77b2277980276
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovom projektu, kreiraćeš program koji prima niz znakova formatiran po strukturi `camelCase` ili `PascalCase` kao ulaz i konvertuje ga u niz znakova formatiran po strukturi `snake_case` koristeći dve metode. Prvo, koristićeš petlju (loop) sa for, a zatim ćeš koristiti list comprehension da postigneš isti rezultat. Videćeš kako list comprehension može učiniti vaš kod kraćim.

Počnite definišući novu funkciju nazvanu `convert_to_snake_case()` koja prima niz karaktera nazvan `pascal_or_camel_cased_string` kao ulaz. Trenutno, dodajte deklaraciju `pass` unutar te funkcije.

# --hints--

Trebalo bi da definišete novu funkciju koja se zove `convert_to_snake_case()`. Nemojte zaboraviti dvotačku na kraju i koristite `pass` za popunjavanje tela funkcije.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

Vaša funkcija treba da uzme parametar nazvan `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
