---
id: 657efce98e958b75df86b305
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Unutar tela deklaracije `if`, dodati ćete karakter koji je promenjen na listu koju ste prethodno kreirali.

Za ovo, metoda `.append()` će biti korišćena. Ova metoda dodaje nešto koje je dato na kraju korišćene liste.

Koristite putanju `.append()` na `snake_cased_char_list` da dodate `converted_character` u listu.

# --hints--

Trebalo bi da koristiš metodu `.append()` za dodavanje izmenjenog karaktera na `snake_cased_char_list`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_char_list\.append\s*\(\s*converted_character\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            converted_character = '_' + char.lower()
--fcc-editable-region--
```
