---
id: 657ef2a86d4e545cec9a85fb
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

U camel case i pascal case, velika slova označavaju početak novih reči. Da biste promenili sekvencu karaktera unosa u snake case, moraćete proveriti da li su slova u sekvenci karaktera unosa velika.

Možete koristiti metodu `.isupper()` za niz karaktera da proverite da li je slovo veliko. Ova metoda vraća `True` ako je slovo veliko i `False` ako nije.

Unutar petlje `for`, dodaj podatke `if` da proveriš da li je trenutni karakter veliko slovo. Premesti podatke `pass` unutar novih podataka `if`.

# --hints--

Treba da napišeš novi unos za if sa char.isupper() kao uslov. Zapamti da dodaš dve tačke na kraju i koristi osnovnu reč pass za popunjavanje tela unosa if.

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

Trebali biste da promenite podatke za `pass` unutar tela petlje i podatke za `if`.

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
