---
id: 657ee28cefc4945568287673
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Pošto je prazan niz spreman, sada možete početi da iterirate kroz sekvencu znakova unosa i da ga pretvorite u snake case.

Unutar funkcije, ispod liste koju ste kreirali, dodajte petlju za iteraciju kroz `pascal_or_camel_cased_string` koja koristi `for`. Uverite se da pozivate ciljni parametar `char`. Za sada, dodajte deklaraciju `pass` kao placeholder unutar tela petlje.

# --hints--

Treba da napišeš novi krug za `for` sa ciljnim parametrom koji se zove `char`. Ne zaboravi dvotačku na kraju i koristi `pass` da popuniš telo kruga.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
