---
id: 657f425dbab54e11993c80f0
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Sada, u ovom projektu koristili ste petlju `for` za iteraciju kroz vaš ulazni niz karaktera i pretvaranje ga u željeni rezultat. Sada ćete započeti prelazak od petlje `for` ka list comprehension-u.

Počnite sa komentarisanjem svih linija koda unutar funkcije `convert_to_snake_case()`. Ne brišite ih jer će vam pomoći kada implementirate logiku koristeći listu opisa.

Zapamtite da dodate glavnu vrednost ``pass`` u telo funkcije kako biste sprečili da kod ne radi tokom testiranja.

# --hints--

Trebaće ti da uneseš komentare u svih 10 redova koda unutar funkcije `convert_to_snake_case()` i ništa drugo ne radi. Dodaj `pass` da popuniš telo funkcije i izbegne grešku.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
