---
id: 657f025ec86c3d7c4177b6be
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Nakon što se funkcija završi, sada ga možete koristiti unutar druge funkcije.

Kreiraj novu funkciju nazvanu `main()` i `pass` kao telo te funkcije.

# --hints--

Treba da definišete novu funkciju koja se zove `main()`. Ne zaboravite na dve tačke na kraju i koristite `pass` da popunite telo funkcije.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(main)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--


--fcc-editable-region--
```
