---
id: 657f47b12c51e41b3149e584
title: Korak 21
challengeType: 20
dashedName: step-21
---

# --description--

Međutim, konačni rezultati i dalje nisu baš ono što želiš postići. Morate primeniti drugačiju matematičku formulaciju za filtrirane karaktere sa blokom `if`. Koristićete blok `else` za to:

```py
spam = [i * 2 if i > 0 else -1 for i in iterable]
```

Zapamtite da, za razliku od sekcije `if`, format `if`/`else` mora biti između matematičkog izraza i glavne reči `for`.

Promeni listu skraćivanja ili kada slovo nije veliko da ne bude promenljivo.

# --hints--

Treba da promenite vašu listu skraćenica kako biste procenili matematički izraz `'_' + char.lower()` ako je `char.isupper()` i drugi karakteri.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_expr().is_equivalent("'_' + char.lower() if char.isupper() else char")
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string if char.isupper()]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
