---
id: 657f0044be09db790b1eb1c5
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

U Pascalu `case`, niz znakova počinje velikim slovom. Nakon konverzije svih slova u mala i dodavanja donžjika, postoji mogućnost da imaš dodatni donžjik na početku svog niza znakova.

Najjednostavniji način za popravku ovoga je korišćenjem metode `.strip()` niza znakova, koja uklanja iz niza znakova bilo koje početno ili krajnje slovo između skupa znakova dato kao njen argument. Primer:

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

Putanj za `strip()` se koristi za `original_string`. Ovo uklanja bilo kakav linijski znak ispod početka i kraja. Rezultat gore navedenog primera biće niz slova `'example_string'`.

Deklariši novi parametar nazvan `clean_snake_cased_string` i postavi njegovu vrednost na rezultat metode `.strip()` koja se koristi za `snake_cased_string`, prosljeđujući `'_'` kao argument toj metodi.

# --hints--

Treba da postavite vrednost niza znakova koji se dobija iz `snake_cased_string.strip('_')` u parametar nazvan `clean_snake_cased_string`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
