---
id: 657f465f8e718b19c5105ae5
title: Korak 20
challengeType: 20
dashedName: step-20
---

# --description--

Opisne liste prihvaćaju uslovne klauzule, kako bi se izvršila evaluacija matematičkog izraza koji je dostupan samo ukoliko su ispunjeni određeni uslovi:

```py
spam = [i * 2 for i in iterable if i > 0]
```

Kao što vidite iz rezultata, lista karaktera generisanih iz `pascal_or_camel_cased_string` je spojena. Pošto se matematički izraz unutar liste procjenjuje za svaki karakter, rezultat je sekvenca malih slova, a svi su karakteri razdvojeni donjom crtom (underscore).

Pratite primer iznad da dodate stavku za `if` u vaš spisak opisa, kako bi matematički izraz bio izvršen samo ako je slovo veliko.

# --hints--

Trebalo bi da dodate odeljak `if` sa uslovom `char.isupper()` u vašu opisnu listu.

```js
({ 
    test: () => assert(runPython(`
    ifs = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_ifs()
    len(ifs) == 1 and ifs[0].is_equivalent("char.isupper()")
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
