---
id: 657f456223b8c1187b461987
title: Korak 19
challengeType: 20
dashedName: step-19
---

# --description--

U Python, <dfn>list comprehension</dfn> je struktura koja vam omogućava kreiranje nove liste korišćenjem matematičkog izraza za svaki element u postojećem iterabilu i opcionalno filtriranje elemenata pomoću uslova. Pored toga što su konciznije, list comprehensions često rade brže.

Osnovna list razumevanje ima matematički izraz koji sledi sa blokom `for`:

```py
spam = [i * 2 for i in iterable]
```

Ovo koristi parametar `i` za ponavljanje preko `iterable`. Svaka stavka rezultujućeg popisa dostupna je evaluacijom matematičkog izraza `i * 2` u trenutnom ciklusu.

U ovom koraku, morate popuniti prazan spisak `snake_cased_char_list` koristeći sintaksu *list comprehension*-a.

Pretvorite vašu praznu listu u list comprehension koja menja svaki karakter u `pascal_or_camel_cased_string` u mala slova i dodaje novi red na početak (prethodno komentisani kod može pomoći u pisanju matematičkog izraza). Koristite `char` za iteraciju preko `pascal_or_camel_cased_string`.

# --hints--

Morate zamijeniti `snake_cased_char_list` sa list komprehenzijom koja iterira nad `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

Vaša list comprehension mora da koristi `char` za iteraciju preko `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

Tvoja list comprehensija treba da evaluira `'_' + char.lower()` za svaki `char` u `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
