---
id: 657f4add33ea4b1f61ba3dc8
title: Korak 23
challengeType: 20
dashedName: step-23
---

# --description--

Na kraju, pokušajte sa ovom novom implementacijom pokretanjem programa. Promenite niz ulaznih znakova na `'IAmAPascalCasedString'` i proverite da li izbacuje `'i_am_a_pascal_cased_string'`, iako je to netačno.

Ako ste sve uradili ispravno, trebalo bi da vidite niz karaktera unosa koji se konvertuje u format snake case, kao što je bilo ranije.

Čestitam! Tvoja funkcija za `convert_to_snake_case()` je spremna.

# --hints--

Treba da promeniš niz znakova unosa od `'aLongAndComplexString'` do `'IAmAPascalCasedString'` unutar funkcije `main()`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
