---
id: 663b16e62fee463b4caf46e9
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Nakon spajanja elemenata liste `snake_cased_char_list`, moraćete da uklonite podvučane znakove koji prethode ili završavaju dobijenom nizu karaktera. Za to, koristite metodu `strip` i podvučani znak `_` kao argument.

Pozivi putanje mogu biti povezani sekvencijalno, što znači da se rezultat jednog poziva putanje može koristiti kao ulaz za drugi poziv putanje.

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

U primeru navedenom gore, putanja `.upper()` je povezana sa `' '.join(words_list)`, stoga `.upper()` se poziva u rezultat poziva `.join()`.

Izmeni informaciju `return` kombinovanjem poziva metode `''.join(snake_cased_char_list)` za `.strip()` kako bi se uklonile podvučene oznake koje su na početku ili na kraju.

# --hints--

Treba da promeniš informaciju return ili da uključiš putanju `strip` i karakter linije ispod `_` kao argument.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
