---
id: 69860157f7fce39f9df399a3
title: Algoritam Tekeleze za Malkija N
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

Izazov Kraljice N traži da staviš kraljicu N na šahovsku tablu od N×N tako da se dve kraljice ne sukobe (da ne dele red, kolonu ili dijagonalu).

Na primer, ako postoji daska od 4x4, jedna validna konfiguracija je:

```md
[1, 3, 0, 2]
```

Ovo znači da u redu 0, kraljica je u koloni 1; u redu 1, kraljica je u koloni 3; u redu 2, kraljica je u koloni 0; i u redu 3, kraljica je u koloni 2.

Iz vizuelne perspektive, ovaj raspored izgleda kao:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Gde `Q` predstavlja kraljicu, a `.` predstavlja prazno polje.

U ovoj laboratoriji, ćeš implementirati rešenje za izazov Kraljica N koristeći metodu dubinskog pretraživanja.

**Cilj**: Implementirajte priče korisnika ispod i prođite sve testove ili završite laboratoriju.

**Priče korisnika:**

1. Funkcija mora da ima nazvanu `dfsNQueens`.
2. Ta funkcija treba da primi samo jedan argument: kompletan broj `n`.
3. Ako je `n` manji od `1`, funkcija treba da vrati prazan niz podataka (`[]`).
4. Funkcija treba da vrati listu rešenja; svako rešenje je niz dužine `n`, gde je element na indeksu `i` indeks kolone (počevši od 0) kraljice u redu `i`.

# --hints--

Trebalo bi da imate funkciju nazvanu `dfsNQueens` koja prima jedan argument.

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

Ako `n` je manje od `1`, funkcija bi trebalo da vrati prazan red podataka.

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

Funkcija bi trebalo da vrati niz rešenja, gde je svako rešenje niz dužine `n`.

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)` treba vratiti `[[0]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)` treba da vrati `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)` trebalo bi da vrati `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)` treba da vrati `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)` treba da vrati `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length` treba biti `10`.

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length` trebao bi da bude `92`.

```js
assert.lengthOf(dfsNQueens(8), 92);
```

`dfsNQueens` Trebalo bi da vrati tačne rezultate.

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
