---
id: 6140c9d35015ae0ba0c250e8
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Dodajte `div` unutar vašeg `body` elementa i dodelite mu `class` od `wheel`.

Unutar vašeg novog `div`, dodajte šest `span` elemenata sa `class` postavljeno na `line`, i šest `div` elemenata sa `class` postavljeno na `cabin`.

# --hints--

Trebalo bi da kreirate new `div` unutar vašeg `body` elementa.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert.lengthOf(divs, 1);
```

Vaš `div` unutar vašeg `body` elementa bi trebalo da ima `class` od `wheel`.

```js
assert.isTrue(document.querySelector('body div')?.classList?.contains('wheel'));
```

Vaš `.wheel` element bi trebalo da sadrži šest `span` elemenata.

```js
assert.lengthOf(document.querySelectorAll('.wheel span'), 6);
```

Vaši šest `span` elementa unutar `.wheel` element trebalo bi da imaju `class` od `line`.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert.isTrue(spans?.every(span => span?.classList?.contains('line')));
assert.lengthOf(document.querySelectorAll('.line'), 6);
```

Vaš element `.wheel` bi trebalo da sadrži šest elemenata `div`.

```js
assert.lengthOf(document.querySelectorAll('.wheel div'), 6);
```

Vaših šest `div` elemenata unutar `.wheel` elementa treba da ima `class` od `cabin`.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert.isTrue(divs?.every(div => div?.classList?.contains('cabin')));
assert.lengthOf(document.querySelectorAll('.cabin'), 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
