---
id: 6140cdebd39d6a101e747529
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Kreiraj selektor za vaše `.line` elemente. Počnite postavljanjem `background-color` na `black`, `width` na `50%` i `height` na `2px`.

# --hints--

Trebalo bi da imate selektor `.line`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line'));
```

Vaš `.line` selektor bi trebalo da ima `background-color` svojstvo set do `black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor, "black");
```

Vaš `.line` selektor bi trebalo da ima `width` svojstvo set do `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.width, "50%");
```

Vaš `.line` selektor bi trebalo da ima `height` svojstvo set do `2px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.height, "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
