---
id: 6140cd32d018ed0f600eefce
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Postavi svojstvo `position` selektora `.wheel` na `absolute`. Postavi i `height` i `width` na `55vw`.

# --hints--

Vaš selektor `.wheel` bi trebalo da ima svojstvo `position` do set do `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.position, 'absolute');
```

Vaš `.wheel` selektor bi trebalo da ima svojstvo `height` set do `55vw`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.height, '55vw');
```

Vaš `.wheel` selector bi trebalo da ima `width` svojstvo set do `55vw`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.width, '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
