---
id: 6140c7e645d8e905819f1dd4
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Za početak, dodajte `link` element za fajl `./styles.css`.

# --hints--

Vaš kod bi trebalo da ima `link` element.

```js
assert.match(code, /<link/);
```

Trebalo bi da imate jedan `link` element.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Vaš `link` element treba da bude unutar vašeg `head` elementa.

```js
assert.exists(document.querySelector('head > link'));
```

Tvoj element `link` treba da ima atribut `rel` sa vrednošću `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Element vašeg `link` treba da ima atribut `href` sa vrednošću `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>

  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css

```
