---
id: 6140d1a351e88f159ed54fca
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Koristeći isti obrazac, kreiraj zasebni selektor za treći `.line`, četvrti `.line`, peti `.line` i šesti `.line`.

Postavi svojstvo `transform` za treći `.line` na `rotate(120deg)`, četvrti na `rotate(180deg)`, peti na `rotate(240deg)` i šesti na `rotate(300deg)`.

# --hints--

Trebalo bi da kreirate selektor `.line:nth-of-type(3)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

Vaš `.line:nth-of-type(3)` selektor bi trebalo da ima `transform` svojstvo set do `rotate(120deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform, 'rotate(120deg)');
```

Trebalo bi da kreirate selektor `.line:nth-of-type(4)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

Vaš `.line:nth-of-type(4)` selektor bi trebalo da ima `transform` svojstvo set do `rotate(180deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform, 'rotate(180deg)');
```

Trebalo bi da kreirate selektor `.line:nth-of-type(5)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

Vaš `.line:nth-of-type(5)` selektor bi trebalo da ima svojstvo `transform` svojstvo set do `rotate(240deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform, 'rotate(240deg)');
```

Trebalo bi da kreirate selektor `.line:nth-of-type(6)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

Vaš `.line:nth-of-type(6)` selektor bi trebalo da ima `transform` svojstvo set do `rotate(300deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform, 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
