---
id: 6140d10d50636e14695013b2
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Kreiraj selektor koji cilja vaš drugi element `.line`. Postavi svojstvo `transform` na `rotate(60deg)`.

Zapamtite da svojstvo `transform` omogućava manipulisanje oblikom elementa. U ovom slučaju, korišćenje vrednosti `rotate(60deg)` rotiraće element oko svoje tačke `transform-origin` za 60 stepeni u smeru kazaljke na satu.

# --hints--

Trebalo bi da kreirate selektor `.line:nth-of-type(2)`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)'));
```

Tvoj `.line:nth-of-type(2)` selektor bi trebalo da ima `transform` svojstvo set do `rotate(60deg)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)')?.transform, 'rotate(60deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
