---
id: 6140cbeec34e970dfe75e710
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Kreirajte selektor za element `.wheel`. Počnite postavljanjem `border` na `2px solid black`, `border-radius` na `50%`, i `margin-left` na `50px`.

# --hints--

Trebalo bi da imate selektor `.wheel`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

Vaš `.wheel` selektor treba da ima `border` svojstvo set do `2px solid black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.border, '2px solid black');
```

Vaš `.wheel` selektor bi trebalo da ima `border-radius` svojstvo set do `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius, '50%');
```

Tvoj `.wheel` selektor bi trebalo da ima `margin-left` svojstvo set do `50px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft, '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
