---
id: 6140cfc08ca9c5128c3e6478
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Postavi svojstvo `position` selektora `.line` na `absolute`, svojstvo `left` na `50%`, i svojstvo `top` na `50%`.

# --hints--

Vaš `.line` selektor bi trebalo da ima svojstvo `position`, set do `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.position, 'absolute');
```

Your `.line` selector should have a `left` property set to `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.left, '50%');
```

Tvoj `.line` selektor bi trebalo da ima svojstvo `top` set do `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.top, '50%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
}
--fcc-editable-region--
```
