---
id: 686b9720ee1d032bd77a480a
title: Pronađi greške na validatoru ISBN-a
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

ISBN (International Standard Book Number) je jedinstveni identifikator dodeljen komercijalnim knjigama. Može imati 10 ili 13 cifara, a poslednja cifra je kontrolna cifra koja se računa iz ostalih cifara.

Camperbot je pokušao da napravi svoj ISBN validator. Međutim, napravili su nekoliko grešaka na putu.

U ovoj laboratoriji, treba da popraviš postojeći kod i učiniš ga funkcionalnim.

Očekivano ponašanje:
Kada korisnik pokrene aplikaciju, prikazuje upozorenje `Enter ISBN and length: `. Korisnik može uneti ISBN kod koji želi da ispravi u formatu `ISBN,length`. ISBN kod ne sme imati crte, praćen dužinom (`10` ili `13`), odvojen zarezima.

Primjeri unosa:
`1530051126,10` za kodove ISBN-10.
`9781530051120,13` za kodove ISBN-13.

Kako dobiti broj inspekcije:

- Nije potrebno znati detaljnu logiku računanja u ovoj laboratoriji. Operateri `calculate_check_digit_10` i `calculate_check_digit_13` će osigurati izračunavanje i vratiti očekivanu kontrolnu cifru kao niz slova.
- Kontrolna cifra za ISBN-10 kodove može biti broj od `0` do `9` ili velika slova `X`.
- Kontrolna cifra za ISBN-13 kodove može biti broj od `0` do `9`.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratorijum.

Priče o korisniku:

1. Treba da popraviš `IndentationError` u postojećem kodu.
1. Čak i ako korisnik ne unese vrednost razdvojenu zarezom, program bi trebalo da obradi `IndexError` bez padanja.
1. Kada korisnik ne unese vrednost razdvojenu zarezom, treba da vidi poruku `Enter comma-separated values.` na konzoli, a program bi trebalo da se zaustavi.
1. Čak i ako korisnik unese nebrojnu vrednost za dužinu, program bi trebalo da obradi `ValueError` bez padanja.
1. Kada korisnik unese nebrojnu vrednost za dužinu, treba da vidi poruku `Length must be a number.` na konzoli, a program bi trebalo da se zaustavi.
1. Treba da popraviš jednu grešku u funkciji `validate_isbn`.
1. Treba da popraviš `TypeError` u postojećem kodu koji se dešava kada korisnik unese važeći ISBN kod.
1. Treba da popraviš `IndexError` u postojećem kodu kada korisnik unese važeći ISBN kod.
1. Čak i ako korisnik unese neispravan ISBN kod sa nebrojnim karakterima, program bi trebalo da obradi `ValueError` bez padanja.
1. Kada korisnik unese neispravan ISBN kod sa nebrojnim karakterima, treba da vidi poruku `Invalid character was found.` na konzoli.
1. Kada korisnik unese `1530051126,10`, treba da vidi poruku `Valid ISBN Code.`.
1. Kada korisnik unese `9781530051120,13`, treba da vidi poruku `Valid ISBN Code.`.

**Važno**: Moraćete da unesete komentar na pozivu `main()` u opštem delu kako bi mere funkcionisale pravilno.

Kada pristupite ovom projektu, korisnik bi trebalo da vidi sledeće poruke u zavisnosti od vrednosti koje unose.

| ISBN Kod | Dužina | Poruka | Primer unosa |
|-----------|--------|---------| ------------- |
| Valid | Valid | `Valid ISBN Code.` | `1530051126,10` |
| Invalid broj | Valid | `Invalid ISBN Code.` | `1530051125,10` |
| Navedena dužina je nevažeća ili je ostavljeno prazno | Valid | `ISBN-10 code should be 10 digits long.` ili `ISBN-13 code should be 13 digits long.`, u zavisnosti od unete dužine. | `9781530051120,10` ili `1530051126,13` |
| Sadrži nenumeričke znakove (osim kontrole cifre) | Valid | `Invalid character was found.` | `15-0051126,10` |
| Bilo šta | Invalid broj | `Length should be 10 or 13.` | `1530051126,9` |
| Bilo šta | Sadrži nenumeričke znakove ili je prazno | `Length must be a number.` | `1530051125,A` |
| Nije razdvojeno zarezom | Nije razdvojeno zarezom | `Enter comma-separated values.` | `1530051125` |

Možete koristiti sledeće vrednosti da ručno testirate svoj kod, ako želite.

Primeri unosa za validne ISBN-10 kodove:
`1530051126,10`
`9971502100,10`
`080442957X,10`

Primeri unosa za validne ISBN-13 kodove:
`9781530051120,13`
`9781947172104,13`

# --hints--

Trebalo bi da ostavite komentar na akciju (callback) `main` kako bi drugi parametri mogli da rade ispravno.

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

Trebalo bi da imaš funkciju `validate_isbn`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

Trebalo bi da imaš funkciju `calculate_check_digit_10`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

Trebalo bi da imaš funkciju `calculate_check_digit_13`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

Trebalo bi da imaš funkciju `main`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

Kada korisnik unese vrednost koja nije razdvojena zarezom, trebalo bi da vidite poruku `Enter comma-separated values.` u konzoli.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Kada korisnik unese nenumeričku vrednost za dužinu, trebalo bi da vidite poruku `Length must be a number.` u konzoli.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Kada korisnik unese neispravan ISBN kod sa slovima koja nisu brojevi, trebalo bi da vidi poruku `Invalid character was found.` u konzoli.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Kada korisnik unosi `1530051126,10`, treba da vidite poruku `Valid ISBN Code.` u konzoli.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Kada korisnik unese `9781530051120,13`, bi trebalo da vidi poruku `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Kada korisnik unese `1530051125,10`, trebalo bi da vidi poruku `Invalid ISBN Code.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

Kada korisnik unosi `9781530051120,10`, trebalo bi da vidi poruku `ISBN-10 code should be 10 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

Kada korisnik unese `1530051126,13`, trebalo bi da vidi poruku `ISBN-13 code should be 13 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

Kada korisnik unese `15-0051126,10`, treba da vidi poruku `Invalid character was found.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Kada korisnik unese `1530051126,9`, trebalo bi da vidi poruku `Length should be 10 or 13.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

Kada korisnik unese `1530051125,A`, treba da vidi poruku `Length must be a number.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Kada korisnik unese `1530051125`, trebalo bi da vidi poruku `Enter comma-separated values.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Kada korisnik unese `9971502100,10`, trebalo bi da vidi poruku `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Kada korisnik unosi `080442957X,10`, trebalo bi da vidi poruku `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Kada korisnik unese `9781947172104,13`, bi trebalo da vidi poruku `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
