---
id: afcc8d540bea9ea2669306b6
title: Napravi funkciju koja ponavlja niz slova
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

U ovoj laboratoriji, kreiraćete funkciju koja ponavlja sekvencu slova data u određeno vreme. Za potrebe ove laboratorije, nemojte koristiti metodu/način koji je ugrađen unutar `.repeat()`.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratorijum.

Priče o korisniku:

1. Treba da kreiraš funkciju nazvanu `repeatStringNumTimes` koja prima dva parametra: niz znakova i broj.
2. Funkcija treba da vrati niz znakova ponovljen navedeni broj puta.
3. Ako je broj manji ili jednak nuli, funkcija treba da vrati prazan niz znakova.

# --hints--

Treba da kreiraš funkciju nazvanu `repeatStringNumTimes`.

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes` Trebalo bi da primi dva parametra.

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

Funkcija `repeatStringNumTimes` treba da uvek vrati niz znakova.

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)` mora vratiti niz znakova `***`.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` trebalo bi da vrati sekvencu slova `abcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` trebalo bi da vrati sekvencu slova `abcabcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` treba da vrati niz znakova `abc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` mora vratiti niz slova `********`.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` mora vratiti prazan niz znakova (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)` treba da vrati `""`.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

Putanja koja je napravljena unutar `repeat()` ne bi trebalo koristiti.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
