---
id: 68b94f23f7b3ee14078c89c5
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, koristićete funkciju `iframe` za prikaz videa. Osnovni dizajn HTML je pripremljen za vas.

Počnite kreiranjem elementa za `h1` sa tekstom `iframe Video Display`.

# --hints--

Trebalo bi da imaš element `h1`.

```js
assert.exists(document.querySelector('h1'))
```

Vaše polje `h1` treba da ima tekst `iframe Video Display`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
