---
id: 68ba9c4f1688914d72876095
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Element `iframe` takođe uzima svojstvo elementa `src` sa vrednošću koja određuje URL ili prikazan resursni način.

Dodajte svojstvo elementa `src` sa vrednošću `https://www.youtube.com/embed/I0_951_MPE0` na vaš element `iframe`.

Trenutno bi trebalo da vidite video prikazan na stranici, ali i dalje postoje funkcije koje trebate dodati elementu.

# --hints--

Element vašeg `iframe` mora imati atribut elementa `src` postavljen na `https://www.youtube.com/embed/I0_951_MPE0`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
