---
id: 68ba9c4f1688914d72876094
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

U prvom kursu o elementu `iframe`, naučili ste da je element modifikovan kao `img`. To takođe znači da može uzeti atribute od `width` i `height` kako bi odredio svoju dužinu i širinu.

Postavi svoj element od svojstva `iframe` svojstvo `width` od `560` na svojstvo `height` od `315`.

# --hints--

Tvoje svojstvo `iframe` mora imati atribut `width` postavljen na `560`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

Element vašeg `iframe` mora imati svojstvo od `height` postavljeno na `315`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
