---
id: 68ba9c4f1688914d72876096
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Atribut elementa je `allow`. To je kao lista dozvola koja pregledaču govori koji elementi mogu da koriste `iframe`.

Ovde postoji element `iframe` sa karakteristikom elementa `allow`:

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

Dodaj svojstvo elementu `allow` sa vrednostima `accelerometer`, `autoplay`, na `clipboard-write`.

`accelerometer` omogućava `iframe` korišćenje senzora pokreta ili da detektuje stvari kao podizanje i rotiranje uređaja. `autoplay` omogućava da video počne automatski, i `clipboard-write` omogućava iframe da piše podatke na korisnikovu međupaletu.

# --hints--

Vaše svojstvo `iframe` trebalo bi da ima svojstvo elementa `allow` postavljeno na `accelerometer`.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

Svojstvo elementa `allow` vašeg elementa `iframe` mora imati `autoplay` kao jednu od svojih vrednosti.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

Svojstvo elementa `allow` vašeg elementa `iframe` mora imati `clipboard-write` kao jednu od njegovih vrednosti.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
