---
id: 68ba9c4f1688914d72876097
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Povećajte `encrypted-media`, `gyroscope`, na `web-share` sa postojećim vrednostima u svojstvu elementa/polja `allow`.

Ovo troje će omogućiti korišćenje skrivene medijske ekstenzije za zaštitu videa, pružanje pristupa senzorima pokreta i orijentacije uređaja, i dozvoljavanje deljenja sadržaja iframe-a putem nativnog dijaloga za deljenje uređaja.

# --hints--

Atribut elementa `allow` vašeg elementa `iframe` treba da ima `encrypted-media` kao jednu od njegovih vrednosti.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

Svojstvo polja ``allow`` vašeg polja ``iframe`` mora imati ``gyroscope`` kao jednu od njegovih vrednosti.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

Svojstvo elementa za `allow` vašeg elementa `iframe` mora imati `web-share` kao jednu od svojih vrednosti.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
