---
id: 68ba9c4f1688914d72876098
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Atribut sledećeg elementa koji ćete dodati je `referrerpolicy`. Ovo je pravilo koje određuje koje podatke delite kada povezujete svoju stranicu sa drugom stranicom.

Dodaj svojstvo elementa `referrerpolicy` i postavi ga na `strict-origin-when-cross-origin`. Ovo deli punu adresu na istom veb-sajtu, samo naziv sajta na drugom web-sajtu, i nema ništa na nesigurnim sajtovima.

# --hints--

Vaš element `iframe` mora imati svojstvo elemenata `referrerpolicy` postavljeno na `strict-origin-when-cross-origin`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
