---
id: 6619240f46cec8e04d77e03a
title: Osnovna vežba funkcije A
challengeType: 1
dashedName: top-basic-functions-exercise-a
---

# --description--

Kreiraj funkciju koja prima ceo broj. Ova funkcija treba da vrati `integer + 7` ako je ceo broj manji od `10`. Ako je ceo broj veći ili jednak `10`, treba da vrati `integer - 3`.

Ime funkcije treba da bude `addOrSubtract`.

# --hints--

Trebalo bi da imaš funkciju nazvanu `addOrSubtract`.

```js
assert.isFunction(addOrSubtract);
```

Tvoja funkcija bi trebalo da primi kompletan broj kao argument.

```js
assert.match(addOrSubtract.toString(), /\s*addOrSubtract\(\s*\w+\s*\)/);
```

Morate vratiti pun broj koji je dat + 7 ako je pun broj manji od 10.

```js
assert.strictEqual(addOrSubtract(5), 12);
```

Treba vratiti kompletan broj umanjen za 3 ako je kompletan broj veći ili jednak 10.

```js
assert.strictEqual(addOrSubtract(10), 7);
```




# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addOrSubtract(num) {
  if (num < 10) {
    return num + 7;
  } else {
    return num - 3;
  }
}
```
