---
id: 661e17c6068359c3ccf2f4d8
title: Vežbe osnovnih funkcija D
challengeType: 1
dashedName: top-basic-functions-exercise-d
---

# --description--

Napiši funkciju, sa imenom `lastLetter`, koja prima niz znakova kao argument i vraća poslednji karakter tog niza.

# --hints--

Trebalo bi da imaš funkciju sa imenom `lastLetter`.

```js
assert.isFunction(lastLetter);
```

Funkcija mora da primi niz slova kao parametar.

```js
assert.match(lastLetter.toString(), /\s*lastLetter\(\s*\w+\s*\)/);
```

Trebalo bi da vratite poslednje slovo te sekvence.

```js
assert.strictEqual(lastLetter('Sem'), 'm');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function lastLetter(str) {
  return str[str.length - 1];
}
```
