---
id: 661e151f068359c3ccf2f4d7
title: Osnovna vežba funkcije C
challengeType: 1
dashedName: top-basic-functions-exercise-c
---

# --description--

Napišite funkciju, pod nazivom `capitalize`, koja prima niz slova kao parametar i vraća novi niz slova čije prvo slovo je veliko.

# --hints--

Trebalo bi da imaš funkciju sa imenom `capitalize`.

```js
assert.isFunction(capitalize);
```

Funkcija treba da primi niz znakova kao parametar.

```js
assert.match(capitalize.toString(), /\s*capitalize\(\s*\w+\s*\)/);
```

Vaša funkcija treba da vrati novu sekvencu slova sa velikim početnim slovom.

```js
assert.strictEqual(capitalize('sem'), 'Sem');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function capitalize(str) {
  return str[0].toUpperCase() + str.slice(1);
}
```
