---
id: 661e131f068359c3ccf2f4d6
title: Osnovna vežba B
challengeType: 1
dashedName: top-basic-functions-exercise-b
---

# --description--

Napišite funkciju, sa imenom `multiply`, koja prima dva parametra i vraća njihov proizvod.

# --hints--

Trebalo bi da imaš funkciju sa imenom `multiply`.

```js
assert.isFunction(multiply);
```

Vaša funkcija bi trebalo da primi dva puna broja kao argumente.

```js
assert.match(multiply.toString(), /\s*multiply\(\s*\w+\s*,\s*\w+\s*\)/);
```

Trebalo bi da vratite ova dva proizvoda, sa ovim kompletnim brojevima.

```js
assert.strictEqual(multiply(10, 10), 100);
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js 
function multiply(a, b) {
  return a * b;
}
```
