---
id: 67d83df6f82eda3868dd2a84
title: Kreiraj lika za RPG
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

U ovoj laboratoriji ćeš naučiti osnove Python za izradu male aplikacije koja kreira lika za RPG avanturu.

**Cilj:** Popunite priče korisnika ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da postoji funkcija koja se zove `create_character`.
1. Funkcija treba da primi, po redosledu, ime lika praćeno sa tri statistike: snaga, inteligencija i šarm.
1. Ime lika mora biti validirano:
   - Ako ime lika nije niz slova, funkcija treba da vrati `The character name should be a string`.
   - Ako ime lika je prazan niz slova, funkcija treba da vrati `The character should have a name`.
   - Ako ime lika ima više od 10 slova, funkcija treba da vrati `The character name is too long`.
   - Ako ime lika sadrži razmak, funkcija treba da vrati `The character name should not contain spaces`.
1. Statistike takođe moraju biti validirane:
   - Ako je jedna ili više statistika ceo broj, funkcija treba da vrati `All stats should be integers`.
   - Ako je jedna ili više statistika manja od 1, funkcija treba da vrati `All stats should be no less than 1`.
   - Ako je jedna ili više statistika veća od 4, funkcija treba da vrati `All stats should be no more than 4`.
   - Ako ukupan zbir svih statistika nije jednak 7, funkcija treba da vrati `The character should start with 7 points`.
1. Ako sve vrednosti prođu validaciju, funkcija treba da vrati niz slova sa četiri reda:
   - prvi red treba da sadrži ime lika
   - redovi 2-4 moraju početi sa skraćenici statistike, `STR`, `INT` ili `CHA` (u ovom redosledu), zatim razmak, a zatim tačka sa decimalom (`●`) u skladu sa vrednošću statistike, i prazna tačka (`○`) za 10. Primer: ako je vrednost snage 3, mora da sledi tačka 3 sa decimalom praćena tačkom 7 za celo broj. Tačke su uklonjene u editoru.

Ovo je sekvenca slova koja mora biti vraćena na `create_character('ren', 4, 2, 1)`:

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

NAPOMENA: iako su `str` i `int` uobičajene statističke skraćenice, zapamtite da su ključne reči sačuvane u Python i ne bi trebalo koristiti ih kao imena parametara.

# --hints--

Trebalo bi da imaš funkciju koja se zove `create_character`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

Kada je `create_character` pozvano sa prvim argumentom koji nije niz slova, trebalo bi da vrati `The character name should be a string`.

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

Kada `create_character` je naznačeno prvim argumentom koji je niz slova, ne bi trebalo da vrati `The character name should be a string`.

```js
({
    test: () => runPython(
        `assert create_character('ren', 1, 2, 4) != 'The character name should be a string'`
    )
})
```

Kada `create_character` je pozvano sa prvim argumentom koji je prazan niz znakova, trebalo bi da vrati `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

Kada se `create_character` pozove sa prvim argumentom koji nije prazan niz znakova, ne bi trebalo da vrati `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should have a name'
`
    )
})
```

Wada `create_character` bude pozvan prvim argumentom koji ima više od 10 karaktera, treba da vrati `The character name is too long`.

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

Funkcija `create_character` ne bi trebalo da kaže da je karakter predugačak kada nije duži od 10 karaktera.

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

Kada `create_character` je navedeno prvim argumentom koji ima prostor, trebalo bi da vrati `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

Kada `create_character` bude pozvan prvim argumentom koji nema mesta, ne bi trebalo da vrati `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('chacha', 4, 2, 1) != 'The character name should not contain spaces'
`
    )
})
```

Kada `create_character` je pozvan sa drugim, trećim ili četvrtim argumentom koji nije pun broj, treba da vrati `All stats should be integers`.

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

Kada je `create_character` nazvan drugim, trećim i četvrtim argumentom koji su svi potpuni brojevi, ne bi trebalo da vrati `All stats should be integers`.

```js
({
    test: () => runPython(
        `
assert create_character('friend', 4, 2, 1) != 'All stats should be integers'
`
    )
})
```

Kada `create_character` je upućen drugim, trećim ili četvrtim stavkom koji spadaju pod `1`, treba da vrati `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

Kada `create_character` je pozvan sa drugim, trećim i četvrtim argumentom, koji nisu pod `1`, ne bi trebalo da vrati `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no less than 1'
`
    )
})
```

Kada `create_character` je pozvan sa drugim, trećim ili četvrtim argumentom koji je veći od `4`, treba da vrati `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

Kada je `create_character` navedeno drugim, trećim i četvrtim argumentom koji nisu veći od `4`, ne bi trebalo da vrati `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no more than 4'
`
    )
})
```

Kada `create_character` bude pozvan putem drugog, trećeg ili četvrtog argumenta koji ne uključuju `7`, bi trebalo da vrati `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

Kada `create_character` se pozove sa drugim, trećim i četvrtim argumentima koji kombinuju `7`, ne bi trebalo da vrati `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)` treba da vrati `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

Kada `create_character` dobije važeću vrednost, trebalo bi da vrati statistiku lika kako je potrebno.

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
