---
id: 61329d52e5010e08d9b9d66b
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Još jedna važna karakteristika `meta` za Pristupačnost i SEO je opis `description`. Vrednost atributa elementa `content` koristi se od strane pretraživača kako bi pružili detalje vaše stranice.

Dodaj karakteristiku `meta` sa svojstvom karakteristike `name` postavljenu na `description`, i sa svojstvom karakteristike `content` koje je korisno.

# --hints--

Treba da dodaš novu komponentu `meta` unutar `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

Treba da dodeliš `meta` elementu atribut `name` sa vrednošću `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

Trebati je dodeliti atribut elementu `meta` sa svojstvom `content`.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

Vrednost atributa elementa `content` ne sme da prekorači 165 karaktera. _Ovo je maksimalna dužina opisa na Google-u._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
