---
id: 613297a923965e0703b64796
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Možda ste već upoznati sa elementom `meta`; koji se koristi za opisivanje informacija o stranici, kao što su naslov, opis, ključne reči i autor.

Podesite vašu stranicu polje za `meta` sa odgovarajućom vrednošću od `charset`.

Atribut elementa `charset` opisuje kodiranje znakova stranice, a danas je `UTF-8` jedino kodiranje koje podržava većina pregledača.

# --hints--

Treba da kreiraš element `meta` unutar elementa `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Trebalo bi da mu dodelite oznaku `meta` sa svojstvom elementa `charset` čija je vrednost `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
