---
id: 61329b210dac0b08047fd6ab
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sa elementima `meta`, opis `viewport` objašnjava pretraživaču kako da prikaže stranicu. Uključivanje povećava vidljivost na mobilnim telefonima i poboljšava _SEO_ (kvalitet pretraživača).

Dodajte opis `viewport` sa atributom funkcije `content` koji definiše `width` i `initial-scale` za stranicu.

# --hints--

Trebali biste da kreirate drugi element od `meta` unutar `head`.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

Treba mu dodeliti `meta` svojstvo polja `name` sa vrednošću `viewport`.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

Trebalo bi da mu dodelite `meta` svojstvo polja `content` sa vrednošću `width=device-width, initial-scale=1`.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
