---
id: 6133d11ef548f51f876149e3
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Pretraživačka tehnologija je važan deo pristupačnosti, a čitači ekrana zavise od toga da pružite strukturu vaše stranice. Ovo se postiže korišćenjem semantičkih elemenata HTML.

Dodajte element `header` na `main` na vašu stranicu.

Funkcija ``header`` će se koristiti za inicijalizaciju stranice, kao i za pružanje navigacionog menija.

Element `main` će biti na glavnom sadržaju vaše stranice.

# --hints--

Treba da dodate element od `header` unutar `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Trebalo bi da dodaš element `main` unutar `body`.

```js
assert.exists(document.querySelector('body > main'));
```

Element `header` mora da stoji prije elementa `main`.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    
  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
