---
id: 613e2546d0594208229ada50
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

U okviru `header`, pružite kontekst o stranici tako da organizujete unutar jednog elementa koji sadrži `img`, `h1`, i `nav`.

`img` trebalo bi da prikaže `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`, sa `id` od `logo`, i sa tekstom `alt` od `freeCodeCamp`.

`h1` treba da sadrži tekst od `HTML/CSS Quiz`.

# --hints--

Treba dodati element `img` u element `header`.

```js
assert.exists(document.querySelector('header > img'));
```

Treba dodati element `h1` u element `header`.

```js
assert.exists(document.querySelector('header > h1'));
```

Trebalo bi da dodate element `nav` u element `header`.

```js
assert.exists(document.querySelector('header > nav'));
```

Trebalo bi da postavite elemente `img`, `h1`, i `nav` u tom redosledu.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

Trebalo bi mu dati atribut elementu `img` sa svojstvom elementa `src` čija je vrednost `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

Trebalo bi mu dati atribut elementa `img` sa svojstvom elementa `id` sa vrednošću `logo`.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

Trebalo bi mu dati atribut elementa `img` sa atributom elementa `alt` sa vrednošću `freeCodeCamp`.

```js
assert.equal(document.querySelector('img')?.alt, 'freeCodeCamp');
```

Trebalo je mu dodeliti atribut sa tekstom `h1` `HTML/CSS Quiz`.

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
