---
id: 6140827cff96e906bd38fc2b
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Kao što je opisano u <a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">freeCodeCamp Style Guide</a>, logo mora da zadrži omjer stranica od `35 / 4`, i imati unutrašnji prostor oko teksta.

Prvo, promenite `background-color` u `#0a0a23` ili ćete videti logo. Zatim, koristite svojstvo `aspect-ratio` da postavite željeni odnos stranica na `35 / 4`. Na kraju, dodajte unutrašnji razmak `padding` za `0.4rem` oko svega.

# --hints--

Trebalo bi mu dati unutrašnji prostor od `#logo` na `background-color` za `#0a0a23`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

Trebalo bi da koristiš atribut `aspect-ratio`.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

Ne smete koristiti atribut `height`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

Ne smete da menjate atribut `width`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

Trebalo bi mu dati `img` unutrašnji prostor od `aspect-ratio` od `35 / 4`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

Treba mu dati `img` unutrašnji prostor od `padding` za `0.4rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--
#logo {
  width: max(10rem, 18vw);

}
--fcc-editable-region--

```
