---
id: 6141fabd6f75610664e908fd
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Pošto je ovo kratak test, biće vam potrebna forma za korisnike da pošalju odgovore. Možete razdvojiti sadržaj unutar forme, odnosno koristeći elemente od `section`.

Unutar elementa ``main``, kreiraj formular sa tri polja raspoređena unutar ``section``. Zatim, napravi da se formular pošalje na ``https://freecodecamp.org/practice-project/accessibility-quiz``, koristeći ispravnu metodu.

# --hints--

Trebalo bi da ga postavite unutar elementa `form`, a taj element je unutar vašeg elementa `main`.

```js
assert.exists(document.querySelector('main > form'));
```

Treba da postavite tri elementa od `section` unutar vašeg elementa `form`.

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

Trebalo bi mu dati atribut elementa `form` sa atributom elementa `action`.

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

Treba da mu dodeliš atribut elementa `action` vrednost `https://freecodecamp.org/practice-project/accessibility-quiz`.

```js
assert.equal(document.querySelector('main > form')?.action, 'https://freecodecamp.org/practice-project/accessibility-quiz');
```

Trebalo bi mu dati atribut elementu `form`, a onaj je atribut elementa `method`.

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

Trebalo bi mu dodeliti atribut elementu `form` sa tipom atributa `method` i vrednošću `post`.

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
