---
id: 614202874ca576084fca625f
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Svaka funkcija za `region` zahteva oznaku, što pomaže korisnicima čitača ekrana da razumeju svrhu tog područja. Jedan način za dodavanje oznake je dodavanje elementa naslova sekcije unutar tog područja, a zatim referenciranje atributom `aria-labelledby`.

Dodajte sledeća svojstva za `aria-labelledby` na elemente `section`:

- `student-info`
- `html-questions`
- `css-questions`

Zatim, unutar svakog elementa `section`, dodaj jedan element sa karakteristikama `h2` koji odgovaraju karakteristiki `id` koja je povezana s `aria-labelledby`. Izvuci relevantan tekstualni sadržaj za svaki element `h2`.

# --hints--

Trebalo bi mu dati prvu instancu `section` atribut `aria-labelledby` sa vrednošću `student-info`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

Trebalo bi mu dati drugi atribut `section` svojstvo `aria-labelledby` sa vrednošću `html-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

Trebalo bi mu dati treći atribut `section` svojstvo `aria-labelledby` sa vrednošću `css-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

Trebalo bi da postavite unutar jednog elementa od `h2`, unutar prvog elementa od `section`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

Trebalo bi da postavite u prvi element od `h2`, unutar drugog elementa od `section`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

Moraš da postaviš u prvi element od `h2`, unutar trećeg elementa od `section`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

Trebalo bi mu dati prvi atribut `h2` svojstvo `id` sa vrednošću `student-info`.

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

Trebalo bi mu dati drugi atribut `h2` svojstvo `id` sa vrednošću `html-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

Trebalo bi mu dati treći atribut `h2` svojstvo `id` sa vrednošću `css-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

Prvom elementu `h2` treba da dodeliš odgovarajući tekst. _Savet: izabrao bih `Student Info`._

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[0]?.textContent?.length, 1);
```

Trebaš mu dati drugi deo sadržaja teksta za `h2`. Savet: Izabrao bih `HTML`.

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[1]?.textContent?.length, 1);
```

Trebaš mu dati treći deo sadržaja teksta za `h2`. _Savet: Izabrao bih `CSS`_

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[2]?.textContent?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
