---
id: 614ccc21ea91ef1736b9b578
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dobrodošli u prvi deo kratkog testa pristupa. Kako počinjete da budete programer sa iskustvom od HTML i CSS, započeli smo vas sa osnovnom strukturom.

Počnite sa ovom putnjom pristupačnosti pružanjem atributa elementu `lang`` za vaš element ``html``. Ovo će pomoći čitaču ekrana da detektuje jezik stranice.

# --hints--

Trebalo bi mu dati atribut elementa `html` svojstvo elementa `lang`. _Napomena: Možete koristiti vrednost `en` na engleskom._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
