---
id: 66f686b8ebdb982fa8e14330
title: Napravi program za konverziju datuma
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

U ovom laboratoriju, ćete napisati kod programa za konverziju datuma koji pretvara zadani datum u različite formate. Na primer, današnji datum ``Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)`` će biti formatiran na sledeće dva načina:

- `(MM/DD/YYYY): 9/27/2024`.
- `(Month Day, Year): September 27, 2024`.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da kreirate parametar nazvan `currentDate` i postavite mu vrednost na trenutni datum i vreme koristeći objekat od `Date`.
2. Trebalo bi da kreirate parametar nazvan `currentDateFormat` i postavite mu vrednost na niz znakova `Current Date and Time: [current date]`. Izmenite `[current date]` sa rezultatom parametra `currentDate`.
3. Trebalo bi da ispišete vrednost `currentDateFormat` u konzolu.
4. Trebalo bi da kreirate funkciju nazvanu `formatDateMMDDYYYY` koja prima datumski objekat kao argument. Možete se na ovaj parametar pozvati bilo kojim imenom koje želite.
5. Vaša funkcija `formatDateMMDDYYYY` treba da konvertuje trenutni datum u format `month/day/year` i vrati niz znakova `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`.
6. Trebalo bi da kreirate funkciju nazvanu `formatDateLong` koja prima datumski objekat kao argument.
7. Vaša funkcija `formatDateLong` treba da konvertuje trenutni datum u format `Month Day, Year` i vrati niz znakova `Formatted Date (Month Day, Year): [formatted date]`.

**Napomena:** Da bi mere bile tačne, uverite se da koristite `en-US` kao jezično područje prilikom zakazivanja datuma.

# --hints--

Trebalo bi da imaš parametar nazvan `currentDate` koji drži trenutni datum i vreme koristeći `Date`.

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

Trebalo bi da imaš parametar nazvan `currentDateFormat` koji drži trenutni datum u formatu `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>`.

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

Trebalo bi da napišete vrednost `currentDateFormat` u konzoli.

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

Trebalo bi da imate funkciju `formatDateMMDDYYYY`

```js
assert.isFunction(formatDateMMDDYYYY);
```

Funkcija `formatDateMMDDYYYY` treba da primi jedan parametar.

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

Kada element datuma sadrži `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, funkcija `formatDateMMDDYYYY` treba da vrati `Formatted Date (MM/DD/YYYY): 9/27/2024`.

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

Trebalo bi da imaš funkciju `formatDateLong`

```js
assert.isFunction(formatDateLong);
```

Funkcija `formatDateLong` mora da primi jedan parametar.

```js
assert.lengthOf(formatDateLong, 1);
```

Kada datumski objekat sadrži `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, funkcija `formatDateLong` treba da vrati `Formatted Date (Month Day, Year): September 27, 2024`.

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```
