---
id: 68e6aff5325410118e76e7be
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Na ovoj radionici, vežbaće korišćenje ARIA za izradu lako dostupnog kontrolera zvuka.

Početi, dodaj element `h1` sa tekstom `Audio Controls`.

# --hints--

Trebalo bi da imaš element `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Vaš element za `h1` bi trebalo da bude unutar vašeg elementa za `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Vaš element za `h1` treba da ima tekst `Audio Controls`.

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
