---
id: 68eac57fadcc662ce917e75e
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Nakon toga, dodaj komponentu za `button` sa svojstvom komponente `type` koje je postavljeno na vrednost `button`.

Tekst elementa za `button` treba biti `Play`.

# --hints--

Trebalo bi da imaš element `button`.

```js
assert.exists(document.querySelector('button'));
```

Vaša komponenta za `button` mora imati svojstvo komponente `type`.

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

Vaš element za `button` treba da ima svojstvo elementa `type` postavljeno za `button`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

Vaš element za `button` bi trebalo da ima tekst od `Play`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
