---
id: 68eb7f2397da251a40e04aa2
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sledeće, dodajte element za `div` ili uključite sve elemente kontrole zvuka.

Unutar elementa `div`, postavite unutar element/polja `span` sa atributom `id` koji je postavljen na `volume-label`.
Takođe, tekst elementa `span` treba da bude `Volume`.

`id` Kasnije će se vratiti sa ARIA atributima ili pomoćne tehnologije moći da prepoznaju ovaj tekst kao oznaku za kontrolu zvuka.

# --hints--

Trebalo bi da imaš element `div` ispod tvog elementa `button`.

```js
assert.exists(document.querySelector('button + div'));
```

Trebalo bi da imaš element `span` unutar svog elementa `div`.

```js
assert.exists(document.querySelector('button + div > span'));
```

Tvoj element od `span` mora imati svojstvo `id`.

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

Vaš element `span` mora imati svojstvo `id` sa vrednošću postavljenu na `volume-label`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

Vaše polje `span` mora imati tekst od `Volume`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
