---
id: 68ec0a48c9d6db3594d768c3
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Sada je vreme da dodamo još jednu funkciju od `span`.

Ovaj element sa `span` treba da ima svojstvo polja od `id` postavljeno na `volume-description`.
Takođe, tekst ovog elementa sa `span` treba da bude `Adjust the sound level`.

Kao što je slučaj sa prethodnim elementom `span`, referenca na `id` od `span` će biti prikazana kasnije sa ARIA atributima kako bi asistivne tehnologije mogle da prepoznaju ovaj tekst kao opis kontrole zvuka.

# --hints--

Trebalo bi da imaš dva elementa `span` unutar elementa `div`.

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

Tvoja druga komponenta iz `span` treba da ima svojstvo komponente `id`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

Vaš drugi element od `span` treba da ima svojstvo atributa `id` sa vrednošću `volume-description`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

Vaš drugi element za `span` mora imati tekst od `Adjust the sound level`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
