---
id: 68ed488821845990c911a223
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Kontrola glasa za pristupnost je skoro gotova, ali postoji nešto poslednje što treba dodati.

Dodajte element `button` koji ima atribut elementa `type` postavljen na `button` i sa tekstom `Mute` ispod elementa `div`.

Sa tim promenama, završili ste ovaj radionicu!

# --hints--

Trebalo bi da imaš element `button` ispod elementa `div`.

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

Tvoj element za `button` treba da ima svojstvo elementa `type`.

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

Vaš element za `button` mora imati svojstvo elementa `type` sa vrednošću `button`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

Element za `button` treba da ima tekst od `Mute`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
           aria-labelledby="volume-label volume-description">
    </div>
  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
             aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
