---
id: 68ed3a96e0ebcc36565b39e5
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Sledeće, morate dodati element `input` sa svojstvom elementa `type` postavljeno na `range` ili kreirati audio sampler. Postavite svojstvo elementa `min` na `0`, svojstvo elementa `max` na `100`, i svojstvo elementa `value` na `50` ili definišite osnovni nivo zvuka.

# --hints--

Trebalo bi da imaš element `input` ispod svog drugog elementa `span`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

Vaše polje za `input` treba da ima svojstvo elementa od `type`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

Vaš element `input` mora imati atribut elementa `type` sa vrednošću `range`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

Vaše polje `input` mora imati svojstvo elementa `min`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

Element vašeg `input` treba da ima svojstvo elementa `min` sa vrednošću `0`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

Vaš element `input` mora imati svojstvo elemenata/komponente `max`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

Vaš element `input` mora imati svojstvo elemenata `max` sa vrednošću `100`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

Vaš element za `input` mora imati svojstvo elementa `value`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

Vaš element `input` mora imati svojstvo polja `value` sa vrednošću `50`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </div>

  </body>
</html>
```
