---
id: 68ed4204ab7edc66b6cb2383
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Kao što ste zapamtili iz prethodnih lekcija, atribut elementa `aria-labelledby` koristi se kada postoji postojeći tekst na stranici koji može da služi kao oznaka. U ovom slučaju, dodali ste atribut elementa `id` vašim elementima `span` koji će sada biti korišćeni ovde.

Dodaj svojstvo elementa `aria-labelledby` za element `input`, kada je postavljeno na `volume-label volume-description`.

# --hints--

Vaš element za `input` mora imati atribut elementa `aria-labelledby`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

Element vašeg `input` mora imati svojstvo elementa `aria-labelledby` sa vrednošću `volume-label volume-description`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
