---
id: 66f53dc2c5bd6a11d6c3282f
title: Napravi program za generator lozinki
challengeType: 26
dashedName: lab-password-generator
---

# --description--

U ovom laboratoriju, ćete vežbati korišćenje funkcija za kreiranje generatora nasumičnih lozinki.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba da kreirate funkciju nazvanu `generatePassword` koja prima argument koji pokazuje dužinu generisane lozinke. Argument možete nazvati bilo čime što želite.
2. Vaša funkcija treba da vrati niz znakova koji predstavlja nasumično generisanu lozinku. Treba da koristite sledeći niz znakova i različite metode za `Math` kako biste pomogli vraćanju novog niza sa slučajnim slovima unutar njega: `ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`.
3. Treba da definišete argument nazvan `password` i postavite mu vrednost kao rezultat poziva funkcije `generatePassword` sa numeričkim argumentom koji predstavlja željenu dužinu lozinke.
4. Treba da imate `console.log` koja štampa jedan niz kombinovan sa porukom `Generated password:` i argumentom `password`, razdvojenim razmakom.

# --hints--

Trebalo bi da imaš funkciju `generatePassword` sa parametrom. Možeš navesti parametar bilo kojim imenom koje želiš.

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

Vaša funkcija `generatePassword` treba da vrati niz slova/karaktera.

```js
const result = generatePassword(5);
assert.isString(result);
```

Tvoja funkcija `generatePassword` treba da vrati novu sekvencu sa ispravnom dužinom.

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

Vaša funkcija treba da vrati lozinku generisanu nasumično sa važećim karakterima.

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

Vaša lozinka mora da vrati/sadrži nasumično generisanu lozinku koja ima više od jednog različitog karaktera.

```js
const passwordSetSize1 = new Set(generatePassword(12)).size;
const passwordSetSize2 = new Set(generatePassword(12)).size;
const passwordSetSize3 = new Set(generatePassword(12)).size;
assert.isAbove(passwordSetSize1 + passwordSetSize2 + passwordSetSize3, 3);
```

Vaša funkcija mora vratiti novu sekvencu nasumičnih brojeva svaki put kada je pozovete.

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

Trebalo bi da imate kriterijum za `password`.

```js
assert.isDefined(password);
```

Vaš parametar za `password` mora biti niz slova.

```js
assert.isString(password);
```

Trebalo bi da pozoveš funkciju `generatePassword` sa argumentom broja i sačuvaš vraćenu lozinku u parametar `password`.

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

Trebalo bi da odštampate sekvencu koja spaja `Generated password:` i `password`, razdvojenu jednim razmakom, koristeći `+` ili template literal.

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
