---
id: 66eaddd04a9e533fba689001
title: Napravi objavu za blog
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

U ovoj radionici, ćeš se naviknuti kako da stilizuješ pozadinske teme i ivice kreiranjem šablona za blog objavu.

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Treba da imaš `div` element sa klasom `blog-post-card` koji će sadržati sve elemente tvoje kartice.
2. Unutar elementa `.blog-post-card` treba da imaš element `img` sa važećim atributom `alt` koji sadrži tekst i klasom `post-img`. Možeš koristiti `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` kao vrednost atributa `src`.
3. Treba da imaš `div` sa klasom `post-content` unutar elementa `.blog-post-card`, koji sadrži sledeće:
- Element `h2` sa klasom `post-title` i tekstom naslova tvoje blog objave.
- Element `p` sa klasom `post-excerpt` i tekstom sažetka tvoje blog objave.
- Element `a` sa klasom `read-more` i tekstom `Read More`.
4. Treba da postaviš boju pozadine elementa `body` na vrednost koja nije bela.
5. Primeni sledeće stilove na element `.blog-post-card`:
    - Mandhari ya nyuma nyeupe.
    - Pembe zilizopindika.
    - Upana wa chaguo lako.
    - Ulinganifu wa maandishi wa chaguo lako.
6. Element `.post-img` treba stilizovati tako da slika popuni celu širinu kartice i ima donju ivicu.
7. Element `.post-content` treba stilizovati tako da ima unutrašnje razmake.
8. Elementi `.post-title` i `.post-excerpt` treba da imaju boju teksta različitu od podrazumevane i razmake sa svih strana.
9. Element `.read-more` treba stilizovati kao dugme i treba da ima:
- Boju teksta različitu od podrazumevane.
    - Rangi ya mandhari ya nyuma.
    - Ukingo upande wote.
    - Nafasi ya ndani upande wote.
    - Pembe zilizopindika.
- Svojstvo `display` postavljeno na `inline-block`.
    - Athari ya kupaa inayobadilisha rangi ya mandhari ya nyuma.

**Napomena:** Uverite se da ste povezali svoj stilski list na HTML i koristili CSS.

# --hints--

Trebalo bi da imaš `div` sa klasom objekta `blog-post-card`.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

Treba da ima karakteristiku `img` sa `class` od `post-img`. Proverite da vaša slika ima atribut karakteristike `alt` sa tekstom i atribut karakteristike `src` sa vrednošću.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

Trebalo bi da imaš `div` sa klasom objekata `post-content`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

Treba da imaš element `h2` sa klasom `post-title`. Uveri se da sadrži tekst naslova tvoje blog objave.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

Treba imati ``p`` sa klasom objekata ``post-excerpt``. Osigurajte da ima tekstualni sadržaj za sažetak vašeg blog posta.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

Trebalo bi da imaš element sa `a` koji ima `class` od `read-more`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

Vaše polje `.read-more` trebalo bi da sadrži tekst od `Read More`.

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

Vaše polje za `.blog-post-card` mora imati svojstvo `border-radius` sa vrednošću (ne sme biti `0` ili negativna vrednost).

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

Tvoj element `.blog-post-card` mora imati belu pozadinu.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

Vaš element za `body` bi trebalo da ima pozadinssku boju koja nije bela.

```js
const bodyBg = getComputedStyle(document.body).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.isFalse(whiteColors.includes(bodyBg.toLowerCase()));
```

Trebalo bi da ciljate na `.blog-post-card` da biste postavili njegovu osobinu `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

Trebalo bi da ciljate na `.blog-post-card` da postavite njegov/njezin atribut `text-align`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

Trebalo bi da ciljate na `.post-content` i postavite mu atribut `padding`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

Vaša funkcija za `.read-more` bi trebalo da ima rastući efekat.

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

Treba da ciljate `.read-more` i postavite njegovo svojstvo/atribut `color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

Trebalo bi da ciljate na `.read-more` da postavite njegovu osobinu `background-color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

Trebalo bi da ciljate `.read-more` kako biste postavili njegov atribut `margin`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

Trebalo bi da ciljate na `.read-more` da biste postavili njegovu osobinu `display`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

Trebalo bi da ciljate na `.read-more` i postavite njegovo svojstvo od `border-radius`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

Treba ti da ciljaš na `.read-more` kako bi postavio njegovu osobinu `padding`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

Vaše polje `.post-img` mora da popuni `width` sa karte.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

Vaše polje za `.post-img` treba da ima vrednost `border-bottom`.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

Elementi `.post-title` i `.post-excerpt` moraju imati granice i neobične boje teksta.

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
