---
id: 587d7fb0367417b2b2512bef
title: Servirajte HTML fajl
challengeType: 2
forumTopicId: 301516
dashedName: serve-an-html-file
---

# --description--

Možete odgovoriti na zahteve serviranjem fajlova pomoću `res.sendFile(path)`. Možete ga postaviti unutar handlera rute `app.get('/', ...)`. Iza kulisa, ova metoda će postaviti odgovarajuća zaglavlja koja obaveštavaju vaš pregledač kako da obradi fajl koji želite da pošaljete, u zavisnosti od njegovog tipa. Zatim će pročitati i poslati fajl. Ova metoda zahteva apsolutnu putanju fajla. Preporučujemo da koristite globalnu Node promenljivu `__dirname` da izračunate putanju kao što je ova:

```js
absolutePath = __dirname + '/relativePath/file.ext'
```

# --instructions--

Pošaljite fajl `/views/index.html` kao odgovor na GET zahtev na ruti `/`. Ako pogledate trenutni program, trebali biste da vidite veliki HTML naslov (sa formom koju ćemo koristiti kasnije…), bez ikakvog stila.

**Napomena:** Možete urediti rešenje prethodnog izazova ili kreirati novo rešenje. Ako kreirate novo rešenje, imajte u vidu da Express prolazi kroz rute odozgo nadole, i obrađuje prvi podudarni handler. Morate da komentarišete kod prethodnog rešenja, inače će server nastaviti da odgovara nizom karaktera.

# --hints--

Vaš program treba da servira fajl views/index.html

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.match(
    data,
    /<h1>.*<\/h1>/,
    'Your app does not serve the expected HTML'
  );
```
