---
id: 587d7fb2367417b2b2512bf5
title: Dobijte parametar rute od klijenta
challengeType: 2
forumTopicId: 301513
dashedName: get-route-parameter-input-from-the-client
---

# --description--

Kada gradite API, morate dozvoliti korisnicima da komuniciraju sa vama o onome što žele da dobiju iz vašeg servisa. Na primer, ako klijent zahteva informacije o korisniku koji je sačuvan u bazi podataka, oni trebaju način da vam kažu kog korisnika žele da vide. Jedan mogući način da postignete ovo je korišćenjem parametara rute. Parametri rute su imenovani delovi URL-a, razdvojeni slash znakovima (/). Svaki deo hvata vrednost dela URL-a koja odgovara njegovom položaju. Uhvaćene vrednosti se mogu pronaći u `req.params` objektu.

<blockquote>route_path: '/user/:userId/book/:bookId'<br>actual_request_URL: '/user/546/book/6754' <br>req.params: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

Izgradite echo server, postavljen na ruti `GET /:word/echo`. Odgovorite sa JSON objektom, koristeći format `{echo: word}`. Možete dobiti echo reč iz `req.params.word`. Možete testirati vašu rutu iz adresne trake pregledača, posetivši nekoliko odgovarajućih ruta, na primer `your-app-rootpath/freecodecamp/echo`.

# --hints--

Test 1 : Echo server treba da vraća reči tačno

```js
  const response = await fetch(code + '/eChOtEsT/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'eChOtEsT',
    'Test 1: the echo server is not working as expected'
  );
```

Test 2 : Echo server treba da vraća reči tačno

```js
  const response = await fetch(code + '/ech0-t3st/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'ech0-t3st',
    'Test 2: the echo server is not working as expected'
  );
```
