---
id: 587d7fb2367417b2b2512bf6
title: Dobijte query parametar od klijenta
challengeType: 2
forumTopicId: 301512
dashedName: get-query-parameter-input-from-the-client
---

# --description--

Još jedan uobičajeni način za dobijanje ulaza od klijenta je stavljanje stringa podataka nakon rute, koristeći query string. Query string počinje sa znakom pitanja (?), i sadrži parove polje=vrednost. Svaki par je razdvojen znakom ampersand (&). Express može da parsira podatke iz query stringa, i popuni `req.query` objekat. Neki karakteri, kao što je znak procenta (%), ne mogu da postoje u URL-u i moraju biti kodirani u drugačijem formatu pre slanja. Ako koristite API iz JavaScript-a, možete koristiti posebne metode za enkodiranje/dekodiranje ovih karaktera.

<blockquote>route_path: '/library'<br>actual_request_URL: '/library?userId=546&#x26;bookId=6754' <br>req.query: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

Izgradite API endpoint, postavljen na `GET /name`. Odgovorite sa JSON objektom, koristeći format `{ name: 'firstname lastname'}`. Parametri za ime i prezime treba da budu kodirani u query string, na primer `?first=firstname&last=lastname`.

**Napomena:** U sledećem zadatku ćete primiti podatke od POST zahteva, na istoj ruti `/name`. Ako želite, možete koristiti `app.route(path).get(handler).post(handler)`. Ova sintaksa vam omogućava da ulančate handlere za različite HTTP metode na istoj ruti. To može uštedeti mnogo pisanja, i učiniti kod čistijim.

# --hints--

Test 1 : Vaš API endpoint treba da odgovori sa `{ "name": "Mick Jagger" }` kada se `/name` pozove sa `?first=Mick&last=Jagger`

```js
  const response = await fetch(code + '/name?first=Mick&last=Jagger');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Mick Jagger',
    'Test 1: "GET /name" route does not behave as expected'
  );
```

Test 2 : Vaš API endpoint treba da odgovori sa `{ "name": "Keith Richards" }` kada se `/name` pozove sa `?first=Keith&last=Richards`

```js
  const response = await fetch(code + '/name?last=Richards&first=Keith');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Keith Richards',
    'Test 2: "GET /name" route does not behave as expected'
  );
```
