---
id: 66c07238b01053abaf812065
title: Napravi kalkulator za faktore
challengeType: 26
dashedName: build-a-factorial-calculator
---

# --description--

Faktorijela je proizvod celog broja i svih celih brojeva ispod njega. Na primer, faktorijela za `5` je `5 * 4 * 3 * 2 * 1 = 120`. U ovom laboratoriju, kreiraćete kalkulator faktorijela koji uzima broj od korisnika i izračunava faktorijelu tog broja.

**Cilj:** Popunite priče korisnika ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da deklarišete parametar `num` i dodelite mu željenu brojčanu vrednost. Dodeljeni broj mora biti između `1` i `20`, uključivo.
1. Napravite funkciju nazvanu `factorialCalculator` koja prima broj kao argument i vraća njegov faktorijel.
1. Unutar funkcije, deklarišite parametar `result` i dodelite mu vrednost `1`. Koristeći petlju, iterirajte kroz sve brojeve od `1` do ulaznog broja (uključivo) i za svaki broj, pomnožite parametar `result` sa trenutnim brojem i sačuvajte rezultat u parametar `result`. Možete izabrati da koristite petlju `for`, `while` ili `do...while` ovde.
1. Trebalo bi da pozovete funkciju `factorialCalculator` sa `num` kao argumentom i sačuvate rezultat u parametar `factorial`.
1. Trebalo bi da konačni rezultat sačuvate u formatu `Factorial of [num] is [factorial]` i dodelite ga parametru `resultMsg`.
1. Trebalo bi da ispišete vrednost `resultMsg` na konzolu.

# --hints--

Trebalo bi da imaš kriterijum `num`.

```js
assert.isDefined(num);
```

Trebalo bi da postavite vrednost na parametar `num`

```js
assert.isNotNull(num);
```

Vrednost `num` bi trebalo da bude između `1` i `20`.

```js
assert.isAtLeast(num, 1);
assert.isAtMost(num, 20);
```

Vrednost `num` mora biti broj.

```js
assert.isNumber(num);
```

Trebalo bi da imaš funkciju nazvanu `factorialCalculator`.

```js
assert.isFunction(factorialCalculator);
```

Funkcija `factorialCalculator` mora da primi broj kao argument.

```js
assert.match(factorialCalculator.toString(), /\s*factorialCalculator\(\s*\w+\s*\)/);
```

Faktoria za `5` treba da vrati `120`.

```js
assert.strictEqual(factorialCalculator(5), 120);
```

Fabrika za `7` treba da vrati `5040`.

```js
assert.strictEqual(factorialCalculator(7), 5040);
```

Trebalo bi da pozoveš svoju funkciju `factorialCalculator` sa parametrom `num` kao argumentom.

```js
assert.match(__helpers.removeJSComments(code), /factorialCalculator\(\s*num\s*\)\s*;?\s*$/m);
```

Trebalo bi da objaviš parametar `factorial` i staviš rezultat funkcije `factorialCalculator` unutar njega.

```js
assert.isDefined(factorial);
assert.strictEqual(factorial, factorialCalculator(num));
```

`factorialCalculator` Vaše treba da pruži tačne rezultate.

```js
assert.strictEqual(factorialCalculator(6), 720);
assert.strictEqual(factorialCalculator(9), 362880);
assert.strictEqual(factorialCalculator(11), 39916800);
```

`resultMsg` treba da ima niz karaktera u formatu `Factorial of [num] is [factorial]`.

```js
assert.strictEqual(resultMsg, `Factorial of ${num} is ${factorialCalculator(num)}`);
assert.notMatch(__helpers.removeJSComments(code), /Factorial of \d{1,2} is \d+/);
```

Trebalo bi da unesete vrednost `resultMsg` u konzolu.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(resultMsg\)\s*;?\s?$/m);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5;

function factorialCalculator(n) {
  let result = 1;

  for (let i = 2; i <= n; i++) {
    result *= i;
  }

  return result;
}

const factorial = factorialCalculator(num);
const resultMsg = `Factorial of ${num} is ${factorial}`;
console.log(resultMsg);
```
