---
id: 587d78af367417b2b2512b00
title: Koristi svojstvo `align-self`
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvzfv'
forumTopicId: 301107
dashedName: use-the-align-self-property
---

# --description--
Konačno svojstvo za fleksibilne delove je `align-self`. Ovo svojstvo vam omogućava da podesite poravnanje svakog pojedinačnog dela, umesto da ih postavljate sve zajedno. Ovo je važno jer druge uobičajene metode podešavanja koje koriste svojstva kao što su CSS `float`, `clear` i `vertical-align` ne funkcionišu za fleksibilne delove.

`align-self` prihvata vrednost jednaku `align-items` i bilo koju vrednost postavljenu atributom `align-items`.

# --instructions--
Povećaj svojstvo CSS `align-self` za `#box-1` na `#box-2`. Postavi `#box-1` na vrednost `center` i postavi `#box-2` na vrednost `flex-end`.

# --hints--
Element za `#box-1` mora imati svojstvo `align-self` ako ima vrednost `center`.

```js
const boxOne = document.querySelector('#box-1');
const alignment = window.getComputedStyle(boxOne)['align-self'];
assert.strictEqual(alignment, 'center');
```

Element `#box-2` bi trebalo da ima svojstvo `align-self` sa vrednošću `flex-end`.

```js
const boxTwo = document.querySelector('#box-2');
const alignment = window.getComputedStyle(boxTwo)['align-self'];
assert.strictEqual(alignment, 'flex-end');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    align-self: center;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    align-self: flex-end;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
