---
id: 587d78ad367417b2b2512afb
title: Koristi svojstvo `flex-shrink` da smanjiš veličinu elemenata
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PBfr'
forumTopicId: 301113
dashedName: use-the-flex-shrink-property-to-shrink-items
---

# --description--
Sada, sve osobine u izazovu se odnose na varijablni alat/mehanizam (roditelj varijablih delova). Međutim, postoje nekoliko važnih svojstava za varijabilne delove.

Prvi je atribut za `flex-shrink`. Kada se koristi, omogućava komadu da smanji veličinu ako je kontejner previše mali. Komadi smanjuju veličinu kada je širina roditeljskog kontejnera manja od ukupne širine svih izmenljivih komada unutar njega.

Svojstvo `flex-shrink` uzima brojeve kao vrednost. Veći broj, to će komad smanjiti veličinu u poređenju sa drugim komadi unutar kontejnera. Na primer, ako jedan komad ima vrednost `flex-shrink` od `1` i drugi ima vrednost `flex-shrink` od `3`, on sa vrednošću `3` će se smanjiti za tri puta više nego drugi.

# --instructions--
Dodaj svojstvo CSS `flex-shrink` za `#box-1` i `#box-2`. Postavi `#box-1` na vrednost `1` i `#box-2` na vrednost `2`.

# --hints--
Polje `#box-1` trebalo bi da ima svojstvo `flex-shrink` sa vrednošću od `1`.

```js
const boxOne = document.querySelector('#box-1');
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
assert.equal(flexShrink, '1');
```

Element za `#box-2` mora imati atribut od `flex-shrink` sa vrednošću `2`.

```js
const boxTwo = document.querySelector('#box-2');
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
assert.equal(flexShrink, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;
    flex-shrink: 1;
  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;
    flex-shrink: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
