---
id: 587d78ae367417b2b2512aff
title: Koristi atribut 'order' za ponovno uređenje delova
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvNAG'
forumTopicId: 301116
dashedName: use-the-order-property-to-rearrange-items
---

# --description--
Svojstvo `order` se koristi za određivanje CSS redos u kojem izgledaju promenljivi delovi u promenljivom kontejneru. Podrazumevano, delovi će se pojaviti u istom redosledu kakav je bio u izvoru od HTML. Ovo svojstvo prima broj kao vrednost, i mogu se koristiti negativni brojevi.

# --instructions--
Povećajte atribut/svojstvo CSS `order` za `#box-1` na `#box-2`. Postavi `#box-1` vrednost za `2` i `#box-2` vrednost za `1`.

# --hints--
Dinamički deo `#box-1` mora imati svojstvo `order` ako ima vrednost od `2`.

```js
const boxOne = document.querySelector('#box-1');
const order = window.getComputedStyle(boxOne)['order'];
assert.strictEqual(order, '2');
```

Dinamički deo `#box-2` treba da ima svojstvo `order` sa vrednošću `1`.

```js
const boxTwo = document.querySelector('#box-2');
const order = window.getComputedStyle(boxTwo)['order'];
assert.strictEqual(order, '1');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    order: 2;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    order: 1;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
