---
id: 587d78ae367417b2b2512afc
title: Koristi svojstvo `flex-grow` da proširiš element
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c2p78cg'
forumTopicId: 301111
dashedName: use-the-flex-grow-property-to-expand-items
---

# --description--
Suprotno od `flex-shrink` je svojstvo `flex-grow`. Zapamtite da `flex-shrink` kontroliše veličinu objekata kada kontejner raste. Svojstvo `flex-grow` kontroliše veličinu objekata kada roditeljski kontejner se širi.

Koristeći primer kao onaj iz prošlog izazova, ako jedan element ima vrednost `flex-grow` od `1` a drugi ima vrednost `flex-grow` od `3`, onaj sa vrednošću `3` će biti tri puta veći od drugog.

# --instructions--
Povećajte svojstvo CSS `flex-grow` za `#box-1` na `#box-2`. Postavite `#box-1` na vrednost `1` i `#box-2` na vrednost `2`.

# --hints--
Element za `#box-1` bi trebalo da ima svojstvo `flex-grow` sa vrednošću `1`.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
assert.equal(flexGrow, '1');
```

Element `#box-2` mora imati svojstvo `flex-grow` sa vrednošću `2`.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
assert.equal(flexGrow, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-grow: 1;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-grow: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
