---
id: 587d78ae367417b2b2512afe
title: Koristi skraćenicu flex
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cbpW2tE'
forumTopicId: 301112
dashedName: use-the-flex-shorthand-property
---

# --description--Postoji prečica dostupna za postavljanje više fleks atributa zajedno. Atributi `flex-grow`, `flex-shrink` i `flex-basis` mogu se postaviti zajedno koristeći atribut `flex`.

Na primer, `flex: 1 0 10px;` će postaviti element da ima `flex-grow: 1;`, `flex-shrink: 0;`, i `flex-basis: 10px;`.

Podešavanja atributa su obično `flex: 0 1 auto;`.

# --instructions--
Dodajte svojstvo CSS `flex` za `#box-1` na `#box-2`. Postavite `#box-1` na vrednost da `flex-grow` bude `2`, `flex-shrink` bude `2`, i `flex-basis` bude `150px`. Postavite `#box-2` na vrednost da `flex-grow` bude `1`, `flex-shrink` bude `1`, i `flex-basis` bude `150px`.

Ove vrednosti će uzrokovati da `#box-1` raste i popuni dodatni prostor dvostruko od `#box-2` kada je kontejner veći od 300px, i da se smanji za dvostruko od `#box-2` kada je kontejner manji od 300px. 300px je ukupna širina vrednosti `flex-basis` za dva boksa.

# --hints--
Element za `#box-1` treba imati svojstvo `flex` postavljeno na vrednost `2 2 150px`.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];

assert.equal(flexGrow, '2');
assert.equal(flexShrink, '2');
assert.equal(flexBasis, '150px');
```

Element za `#box-2` mora imati svojstvo `flex` postavljeno na vrednost `1 1 150px`.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];

assert.equal(flexGrow, '1');
assert.equal(flexShrink, '1');
assert.equal(flexBasis, '150px');
```

Vaš kod bi trebalo da koristi svojstvo `flex` za `#box-1` i `#box-2`.

```js
assert.lengthOf(code.match(/flex:\s*?\d\s+?\d\s+?150px;/g), 2);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    flex: 2 2 150px;
    height: 200px;
  }

  #box-2 {
    background-color: orangered;
    flex: 1 1 150px;
    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
