---
id: 587d78ad367417b2b2512afa
title: Koristi svojstvo `flex-wrap` da pokrije red ili kolonu
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cQv9ZtG'
forumTopicId: 301114
dashedName: use-the-flex-wrap-property-to-wrap-a-row-or-column
---

# --description--
CSS flexbox ima funkciju za podelu fleksibilnog kontejnera u više redova (ili kolona). Po podrazumevanju, fleksibilni kontejner će sve fleksibilne elemente postaviti zajedno. Na primer, ceo red će biti na jednoj liniji.

Međutim, korišćenje atributa ``flex-wrap`` sprečava da `CSS` pokrije delove. To znači da su dodatni delovi premešteni u novi red ili kolonu. Tačka za zaustavljanje gde se pokrivanje dešava zavisi od veličine delova i veličine kontejnera.

CSS Takođe ima opcije za smer pokrivanja:

<ul><li><code>nowrap</code>: ovo su podrazumevani podešavanja, i ne uvija elemente.</li><li><code>wrap</code>: uvija elemente u više redova od vrha do dna ako je red, i od leve ka desno ako je kolona.</li><li><code>wrap-reverse</code>: uvija elemente u više redova od dna do vrha ako je red, i od desne ka levo ako je kolona.</li></ul>

# --instructions--
Trenutna postavka ima previše kutija na jedan red. Dodajte atribut CSS `flex-wrap` na element `#box-container` i postavite mu vrednost `wrap`.

# --hints--
Element za `#box-container` treba da ima svojstvo od `flex-wrap` ako ima vrednost `wrap`.

```js
const boxContainer = document.querySelector('#box-container');
const flexWrap = window.getComputedStyle(boxContainer)['flex-wrap'];
assert.strictEqual(flexWrap, 'wrap');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;
    flex-wrap: wrap;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```
