---
id: 587d78ae367417b2b2512afd
title: Koristi svojstvo `flex-basis` za postavljanje početne veličine elementa
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c3d9nCa'
forumTopicId: 301108
dashedName: use-the-flex-basis-property-to-set-the-initial-size-of-an-item
---

# --description--
Atribut `flex-basis` opisuje početnu veličinu elementa pre nego što CSS izvrši izmene na `flex-shrink` ili `flex-grow`.

Dimenzije koje se koriste sa svojstvom `flex-basis` su jednake drugim svojstvima veličine (`px`, `em`, `%`, itd.). Vrednost `auto` određuje veličinu komponenti u zavisnosti od sadržaja.

# --instructions--
Postavite početnu veličinu kutija koristeći `flex-basis`. Dodajte atribut CSS `flex-basis` za `#box-1` na `#box-2`. Navedite `#box-1` sa vrednošću `10em` i `#box-2` sa vrednošću `20em`.

# --hints--
Komponenta `#box-1` mora imati svojstvo `flex-basis`.

```js
const boxOne = document.querySelector('#box-1');
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

Polje za `#box-1` mora imati vrednost `flex-basis` od `10em`.

```js
assert.match(code, /#box-1\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?10em;/g);
```

Element od `#box-2` treba da ima svojstvo `flex-basis`.

```js
const boxTwo = document.querySelector('#box-2');
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

Element `#box-2` treba da ima vrednost `flex-basis` od `20em`.

```js
assert.match(code, /#box-2\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?20em;/g);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-basis: 10em;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-basis: 20em;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
