---
id: 587d78ab367417b2b2512af2
title: Koristi svojstvo `flex-direction` za kreiranje reda
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cBEkbfJ'
forumTopicId: 301110
dashedName: use-the-flex-direction-property-to-make-a-row
---

# --description--
Dodavanje `display: flex` na element ga pretvara u fleksibilni kontejner. Ovo omogućava porještanje bilo koje dete tog elementa u red ili kolonu. To radite dodavanjem svojstva `flex-direction` na roditeljski element i postavljanjem kao red ili kolonu. Kreiranje reda porješće dete horizontalno, a kreiranje kolone porješće dete vertikalno.

Chaguzi nyingine za `flex-direction` ni `row-reverse` na `column-reverse`.

**Napomena:** Podrazumevana vrednost atributa `flex-direction` je `row`.

# --instructions--
Postavi atribut za CSS `flex-direction` na elementu `#box-container`, i dodeli mu vrednost od `row-reverse`.

# --hints--
Element `#box-container` mora imati svojstvo `flex-direction` postavljeno na `row-reverse`.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'row-reverse');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: row-reverse;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
