---
id: 587d78ac367417b2b2512af4
title: Koristi svojstvo `flex-direction` da napraviš vertikalni red (ili kolonu).
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cZmWeA4'
forumTopicId: 301109
dashedName: use-the-flex-direction-property-to-make-a-column
---

# --description--
Poslednja dva izazova koristila su skup svojstava `flex-direction` koje je postavljen na `row`. Ovo svojstvo takođe može kreirati redove raspoređivanjem dece vertikalno prilagodljivog kontejnera.

# --instructions--
Dodajte atribut CSS `flex-direction` na element `#box-container` i postavite mu vrednost `column`.

# --hints--
Element za `#box-container` mora imati atribut `flex-direction`, postavljen na `column`.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: column;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
