---
id: 65f0694fb2296f3caadf8347
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Magnituda vektora $\mathbf{a}$, ili norma, obično se prikazuje kao $\\| \mathbf{a} \\|$. Može se izračunati kao kvadratni koren zbira kvadrata njegovih komponenti:
\\[ \\| \mathbf{a} \\| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2} \\]

Izračunaj normu vektora i vrati rezultat iz vaše metode `norm`.

# --hints--

Vaša funkcija za `norm` treba da koristi dati formulu kako bi izračunala i vratila broj koji predstavlja normu postojećeg vektora. Ne zaokružujte vrednost.

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
