---
id: 65f055b9190fc41ca35549b8
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Python pruža različite putanje koje imaju dva podznaknika (underscore) na početku i kraju svojih imena. Možda već znate neke, kao što su `__init__` i `__str__`. Ove putanje, koje prate strukturu imena `__<name>__`, nazivaju se specijalne metode, magične metode ili *dunder* metode (što označava dvostruki podznaknik).

Definisati specijalne metode u klasi objekata utiče na ponašanje te klase. Poziva se u posebnim situacijama (primer `__init__` prilikom kreiranja instanci, `__str__` kada je objekat štampan ili prosleđen na `str()`). U ovom projektu, naučićete o nekim često korišćenim specijalnim metodama.

Trenutno, postavi vrednost `x` za atribut elementa `x` objekta `Vector`.

# --hints--

Treba da postaviš vrednost `x` za `self.x`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
