---
id: 65f058e06f34fd1f0ee6e55d
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

U okviru klase `Vector`, kreiraj praznu metodu za `norm` i dodeli joj parametar od `self`.

# --hints--

Vaša klasa za `Vector` treba da ima metodu od `norm`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

Tvoja putanja ``norm`` mora imati atribut ``self``.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
