---
id: 65f01969115f933073b6be03
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Vektor je veličina koja ima dužinu (ili magnitudu) i smer, i ne može se predstaviti samo jednim brojem. U fizici, vektor se često koristi za predstavljanje sile, brzine, ubrzanja, i bilo koje druge veličine.

U ovom projektu ćete kreirati vektorsko prostranstvo, što je skup gde je definisan niz operacija između elemenata sadržanih u tom skupu. Sve detalje ćete uskoro naučiti.

Trenutno, započnite projekat objavljivanjem praznog klase sa imenom `Vector`.

# --hints--

Trebalo bi da objaviš kurs praznih stvari pod nazivom `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
