---
id: 65f07c9b1ffb814d856dcffc
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Pokušajte da se potvrdi da li `norm` radi kako je očekivano tako što će se ispisati vrednost koju ta metoda vraća.

# --hints--

Treba da pozoveš `norm()` na `v1` i ispišeš rezultate.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1.norm())")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
--fcc-editable-region--
```
