---
id: 65f06a8e5a57673d700c79c3
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Izvan klase `Vector`, kreiraj instancu od `Vector` prosljeđujući pune brojeve `2` i `3` kao vrednosti za `x` i `y` i postavi njegovu vrednost u parametar nazvan `v1`.

# --hints--

Treba da deklarišete parametar `v1` i dodelite mu vrednost `Vector(2, 3)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
