---
id: 65f03d9f92eac9183a4d3281
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Vektor se može definisati sa dve koordinate, `x` i `y`, u Euklidskom prostoru. Rastojanje između porekla koordinata i tačke `(x, y)` biće njegova dužina, ili norma. A smer vektora će biti usmeren ka `(x, y)`.

<img alt="vector ya vipimo viwili (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Unutar klase objekata `Vector`, kreiraj metodu `__init__` i sa dodeli mu tri parametra, `self`, `x`, i `y`.

# --hints--

Trebalo bi definisati metodu `__init__` unutar klase objekata `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

Tvoja putanja za `__init__` mora da primi tri parametra: `self`, `x`, i `y`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
