---
id: 65f3fa60a93b84110b3f1708
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Želite da vratite izraz smislen sekvence karaktera. Na primer, u slučaju `v1`, želite da vratite sekvencu karaktera koja prikazuje tupel sa vrednostima komponenti vektora: `(2, 3)`.

Iz puta `__str__`, vrati niz karaktera koji predstavlja vektor kao tuple sa komponentama vektora u redosledu. Zatim, iz klase `Vector` i odštapi `v1` da proveriš rezultate.

# --hints--

Tvoja putanja za `__str__` treba da vrati niz znakova koji predstavlja vektor kao tupel sa komponentama vektora u redosledu.

```js
({ test: () => runPython(`
v1 = Vector(2, 3)
v2 = Vector(0, 0)
v3 = Vector(-2, -3.5)
assert str(v1) == '(2, 3)'
assert str(v2) == '(0, 0)'
assert str(v3) == '(-2, -3.5)'
`) })
```

Trebalo bi da odštampate `v1`.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
