---
id: 65f3fa097f26b510db6c710b
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Dobro, `norm` radi kako je očekivano. Sada, kada pokušate da štampate `v1`, dobićete reprezentaciju sekvence karaktera elementa (nešto poput `<__main__.Vector object at 0x11eb778>`).

U klasi objekata `Vector`, deklarisati praznu metodu za `__str__` radi implementacije predstavljanja čitljive sekvence karaktera. Zapamtite da mu date parametar `self`.

Pažljivo nemoj zalepiti `v1` do `__str__`, vrati sekvencu slova, inače ćeš dobiti `TypeError`, jer `__str__` mora uvek da vrati sekvencu slova.

# --hints--

Trebalo bi da objaviš putanju `__str__` unutar klase objekata `Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

Tvoj put/putanja ``__str__`` mora imati atribut ``self``.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
