---
id: 65f4520e363e2642f8112e33
title: Korak 19
challengeType: 20
dashedName: step-19
---

# --description--

Kao što možete da vidite, nešto nije u redu uopšte. Putanje za `norm` i `__str__` dobijene iz `R2Vector` ne mogu se promeniti sa vektoru tri dimenzije. Njegova implementacija zahteva veći učinak.

Svaki objekat u Python ima poseban atribut nazvan `__dict__`, koji je rečnik koji skladišti atribute elementa tog objekta.

Uklonite postojeći poziv za `print`. Zatim, ispišite svojstvo `__dict__` vaših vektora `v1` na `v2` ili vidite kako izgledaju.

# --hints--

Nemoj imati `print(v2)` na `print(v2.norm())` u svom kodu.

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

Trebalo bi da ispišete atribut `__dict__` od `v1` na `v2`.

```js
({ test: () => assert(runPython(`
(_Node(_code).has_call("print(v1.__dict__)") and _Node(_code).has_call("print(v2.__dict__)")) or _Node(_code).has_call("print(v1.__dict__, v2.__dict__)")
`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v2.norm())
print(v2)
--fcc-editable-region--
```
