---
id: 65f4535bbdb28d436ff3ddc9
title: Korak 20
challengeType: 20
dashedName: step-20
---

# --description--

Kao što vidite iz rezultata, `__dict__` ima vrednosti svojstava vašeg elementa. Umesto da eksplicitno sabirate kvadrat od `self.x` i `self.y`, ćete iterirati preko vrednosti sačuvanih u `__dict__` kako biste izračunali normu.

Unutar tela puta `norm`, promenite sadržaj unutar običnih zagrada i izraz koji podiže svaku vrednost `val` u `self.__dict__.values()` na snagu `2`. Takođe, prosledite taj generativni izraz kao argument funkciji `sum`, kako bi se svi kvadrati grupisali zajedno pre računanja kvadratnog korena.

# --hints--

Treba ti vratiti `sum(val**2 for val in self.__dict__.values())**0.5` putem `norm()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in self.__dict__.values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.__dict__)
print(v2.__dict__)

```
