---
id: 65f412c208c3791fee305acf
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

U Python, možete primorati korišćenje samo string argumenata dodavanjem `*` kao dodatnog argumenta u potpisu funkcije ili metode. Izmenite sve metode za `__init__` dodavanjem `*` kao drugog parametra (nakon `self`). Bilo koji parametar postavljen nakon toga zahteva korišćenje string argumenta u pozivu funkcije/metode.

Ovo znači da takođe treba da promeniš poziv za `super().__init__(x, y)`. To učini dajući `x` vrednost `x`, i `y` vrednost `y`.

Konačno, promenite inicijalizaciju `v1` i `v2` koristeći string argumente.

# --hints--

Treba da obavezujete argumente reči dodavanjem `*` kao drugi parametar u sve vaše metode za `__init__`.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

Trebalo bi da promeniš svoj poziv `super.__init__(x, y)` ili koristi argumente sa rečima.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

Trebalo bi da promeniš raspon vrednosti od `v1` do `R2Vector(x=2, y=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

Trebalo bi da promeniš raspodelu vrednosti od `v2` u `R3Vector(x=2, y=2, z=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
