---
id: 65f40051d6b09a139f253e8e
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Nasleđivanje vam omogućava da definišete klasu objekata iz već postojećih. Nova klasa, nazvana podklasa, nasleđuje sve metode i atribute postojeće klase, nazvane roditeljska klasa.

```py
class Tree:
    def sprout(self):
        print('Making new leaves!')

class Oak(Tree):
    pass
    
Oak().sprout() # Output: Making new leaves!
```

U gornjem primeru, podklasa `Oak` nasleđuje od `Tree` i nasleđuje metodu `sprout` iz roditeljske klase `Tree`.

Kreiraj novu klasu nazvanu `R3Vector` i prati primer iznad da bi naslijedio od klase `R2Vector`.

# --hints--

Treba da definišete novu klasu `R3Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R3Vector")
    `))
})
```

Tvoja nova klasa `R3Vector` treba da nasledi od `R2Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").inherits_from("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--

--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
