---
id: 65f8709620b2ce1a62608f5a
title: Korak 23
challengeType: 20
dashedName: step-23
---

# --description--

Ako trebate dobiti određeni atribut putem transformacije počevši od unosa sekvence slova, funkcija koja se nalazi unutar ``getattr()`` je ono što vam treba. Uzima nešto kao svoj prvi argument, i sekvencu slova koja prikazuje ime atributa elementa kao svoj drugi argument.

Počnite da modifikujete putanju ``__str__`` promenom niza karaktera koji vraća ``__str__()`` i generativnim izrazom koji se ponavlja kroz svojstva objekta i poziva ``getattr()`` za svako svojstvo ``i``.

# --hints--

Treba da vratite ponavljajući izraz generatora na `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

Treba da ponovite sintaksu generatora koji koristi `i` kao parametar za ponavljanje.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

Morate vratiti generativnu sintaksu koja poziva `getattr(self, i)` za svaki element `i` koji se nalazi u `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
