---
id: 65f9bc10a9fb1612e066e9e8
title: Korak 42
challengeType: 20
dashedName: step-42
---

# --description--

Kada sabirate dva vektora, svaka komponenta jednog vektora dodaje se odgovarajućoj komponenti drugog vektora. Na primer, sabiranje `(1, 2)` i `(2, 4)` stvara treći vektor `(3, 6)`, gde je `3` zbir `1` i `2`, a `6` je zbir `2` i `4`.

Rečnik `kwargs` će imati parove ključ-vrednost potrebne za inicijalizovanje novog vektora klase jednakog dva vektora dodata zajedno.

Konvertujte prazan rečnik u razumevanje rečnika korišćenjem `vars(self)` i za svaki ključ (`i`) kreirajte par ključ-vrednost, gde je ključ trenutni iterativni ključ, a njegova vrednost je zbir vrednosti atributa `i` za dva posrednika.

# --hints--

Trebalo bi da promeniš prazan rečnik koji je inicijalizovan vrednošću `kwargs` u kreiranje rečnika koji se ponavlja na osnovu `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

Generisanje rečnika je dobilo vrednost `kwargs`, treba da koristi parametar `i` za ponavljanje nad `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

Generisanje rečnika dobilo je vrednost `kwargs` i treba da koristi `i` kao ključ.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

Generisanje rečnika je dobilo vrednost `kwargs` i treba da postavi sumu atributa elementa `i` od `self` na atribut elementa `i` od `other` za ključ `i` za svaki `i` u `vars(self)`. Koristite funkciju `getattr()` za to.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
