---
id: 65f9b17a638f0d0dcce8c354
title: Korak 38
challengeType: 20
dashedName: step-38
---

# --description--

Da biste kreirali vektor prostor, potrebno je definisati kako vektori treba da se ponašaju u nekoliko slučajeva. Vektori mogu biti sabrani, stvarajući nove vektore.

Poseban način `__add__` može se izvršiti anuliranjem onoga što se dešava podrazumevano kada se dva objekta dodaju zajedno koristeći parametar `+`.

Trenutno, pokušaj dodavanja dve instance za `R2Vector` ili `R3Vector` izaziva grešku. Kreirajte praznu metodu za `__add__` unutar klase `R2Vector` i date joj dva parametra: `self`, i `other`.

# --hints--

Trebalo bi da definišete putanju `__add__` unutar klase `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__add__")`)) })
```

Tvoja putanja ``__add__`` mora imati dva parametra, ``self`` i ``other``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
