---
id: 65f93d4647ce2670dc6f095c
title: Korak 31
challengeType: 20
dashedName: step-31
---

# --description--

Putanja za `__repr__` je tajno nazvana od strane funkcije `repr`. Osigurajte da `__repr__` radi ispravno dodavanjem `f'\nrepr = {repr(v1)}'` i `f'\nrepr = {repr(v2)}'` kao drugi argument za vaš prvi i drugi poziv od `print`, respektivno.

# --hints--

Trebalo bi da dodaš `f'\nrepr = {repr(v1)}'` kao drugi argument svom pozivu `print(f'v1 = {v1}')`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v1 = {v1}', f'\\\\nrepr = {repr(v1)}')")`)) })
```

Trebalo bi da dodate `f'\nrepr = {repr(v2)}'` kao drugi argument za vaš poziv `print(f'v2 = {v2}')`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v2 = {v2}', f'\\\\nrepr = {repr(v2)}')")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
--fcc-editable-region--
```
