---
id: 65f93acbb514857003510e79
title: Korak 29
challengeType: 20
dashedName: step-29
---

# --description--

Objavi kriterijum `args` na način da bude onaj koji se vraća spajanjem elemenata koji su u `arg_list` koristeći niz karaktera `', '`.

# --hints--

Treba da deklarišete parametar `args` i postavite mu vrednost na onu koja se vraća spajanjem elemenata koji su locirani `arg_list` koristeći niz znakova `', '`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("args").is_equivalent("args = ', '.join(arg_list)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]

--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
