---
id: 65f9cb678070ca1668898c70
title: Korak 44
challengeType: 20
dashedName: step-44
---

# --description--

Na isti način kako se ``__add__`` poziva ispod poklopca kada se dve stvari dodaju zajedno, putanja ``__sub__`` se poziva indirektno u stanju izlaska.

Sada, definiši praznu funkciju/metodu za `__sub__` i ukloni dva parametra: `self`, i `other`. Unutar svoje nove metode, kreiraj deklaraciju za `if` da proveriš da li `self` i `other` nisu u istoj klasi objekata i vrati `NotImplemented`, kao što si ranije uradio.

# --hints--

Treba da definišete putanju `__sub__` unutar klase objekta `R2Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__sub__")
    `))
})
```

Tvoj put/način za `__sub__` treba da ima dva parametra, `self` i `other`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").has_args("self, other")
    `))
})
```

Trebalo bi da kreirate informaciju za `if` koja proverava da li `self` i `other` pripadaju istoj klasi objekata.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
    `))
})
```

Treba ti da vratiš `NotImplemented` iznutra tela `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
