---
id: 65f9dd6e5a08af19c196c2df
title: Korak 45
challengeType: 20
dashedName: step-45
---

# --description--

Vektor dobijen oduzimanjem jednog vektora od drugog nalazi se iz računanja razlike svake njihove komponente. Na primer, oduzimanje `(2, 4)` od `(7, 3)` stvara treći vektor `(5, -1)`, gde je `5` razlika između `7` i `2`, a `-1` je razlika između `3` i `4`.

Kao što ste prethodno učinili unutar putanje `__add__`, deklarišite varijablu `kwargs` nakon deklaracije `if`. Postavite vrednost *dictionary comprehension*-a koji iterira kroz atribute elementa objekta i kreira par ključ-vrednost za svaki ključ `i`, gde je ključ isti ključ trenutne petlje, a njegova vrednost je razlika između atributnih vrednosti dva borca iz `i`.

# --hints--

Treba da deklarišete varijablu `kwargs` i dodelite joj vrednost dictionary comprehension koja iterira kroz atribute objekta.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

Komprehensija rečnika koja je data `kwargs` mora da koristi varijablu `i` za iteraciju preko svojstava atributa objekta.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

Komprehensija rečnika koji je dat `kwargs` treba da koristi `i` kao ključ.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

Diktionaralna kompozicija koja je dodeljena `kwargs` treba da postavi razliku između atributa `i` od `self` i atributa `i` od `other` za ključ `i` za svaki `i`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
