---
id: 666827a6fd0dbaafe8330ea6
title: Korak 35
challengeType: 20
dashedName: step-35
---

# --description--

`__getattr__` je drugi specijalni način koji igra ulogu pri pristupu svojstvima elementa.

Standardna implementacija `__getattribute__` uzrokuje `AttributeError` kada svojstvo zahtevanog elementa nije svojstvo modela ili ne postoji u hijerarhiji klasa objekata.

U tom slučaju, `__getattr__` se naziva ako je definisan klasom objekata. Možete ga smatriti kao alternativnu metodu kada standardne procedure pristupa svojstvima elementa ne uspeju.

Promeni putanju `__getattribute__` u `__getattr__` i promeni niz karaktera koji vraća ta putanja u `'calling __getattr__'`.

# --hints--

Treba da definišete metod nazvan `__getattr__` sa dva parametra, `self` i `attr`, unutar klase objekata `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_args("self, attr")`)) })
```

Tvoja putanja ``__getattr__`` bi trebalo da vrati niz karaktera ``'calling __getattr__'``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_return("'calling __getattr__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __getattribute__(self, attr):                
        return 'calling __getattribute__'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

print(v1.x)
print(getattr(v1, 'x'))

```
