---
id: 66682150151af29efec9727d
title: Korak 33
challengeType: 20
dashedName: step-33
---

# --description--

Metoda za `__getattribute__` se poziva unutar opsega svaki put kada pokušate dohvatiti svojstvo model elementa. Ako svojstvo elementa nije pronađeno na nivou modela, ta metoda će pretraživati na nivou klase objekta, a na kraju u roditeljskim klasama.

Unutar vaše klase objekata, definisać `__getattribute__` metod sa dva parametra, `self` i `attr`, i naterati ga da vrati niz karaktera `'calling __getattribute__'`. Uskoro ćete promeniti standardnu implementaciju kako biste videli kako funkcioniše pristup svojstvu elementa.

# --hints--

Definišite metodu pod nazivom `__getattribute__` koja ima dva parametra, `self` i `attr`, unutar klase `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_args("self, attr")`)) })
```

Tvoja putanja ``__getattribute__`` bi trebalo da vrati niz karaktera ``'calling __getattribute__'``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_return("'calling __getattribute__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

```
