---
id: 699b64096301fd75e3efeb1b
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Kada ponovo otvorite konzolu, videćete drugu poruku o grešci:

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

Kako ste se setili iz prethodnih lekcija, možete učiniti svojstvo opcionim koristeći oznaku ``?`` pored imena svojstva kao ovde:

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

U ovom primeru, `numberOfAlbums` je postavljeno kao opcionalno, pa TypeScript neće baciti grešku u ovom slučaju.

Ispravite svoj atribut od `bio` ili ga postavite na opcionalno.

# --hints--

Tvoja karakteristika `bio` bi trebalo da bude postavljena kao opcionalna.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
