---
id: 699b68f94a0d154c2a7a5ce0
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Sada nakon što je `profile` završeno, vreme je da kreiramo nešto novo što predstavlja uloge korisnika.

Kreiraj novi parametar nazvan `userRoles` i postavi vrednost objekta. Vaš objekat mora imati sledeća svojstva i vrednosti:

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

Trebalo bi da imaš parametar nazvan `userRoles`.

```js
assert.isDefined(userRoles);
```

Vaš parametar `userRoles` treba da bude nešto.

```js
assert.isObject(userRoles);
```

Tvoj objekat sa `userRoles` treba da ima svojstvo od `admin`.

```js
assert.property(userRoles, "admin");
```

Vaš atribut za `admin` mora imati vrednost `"full-access"`.

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

Tvoj deo `userRoles` mora imati karakteristiku `editor`.

```js
assert.property(userRoles, "editor");
```

Tvoja vrednost za `editor` bi trebalo da bude `"limited-access"`.

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

Vaš predmet za `userRoles` mora imati svojstvo `viewer`.

```js
assert.property(userRoles, "viewer");
```

Tvoja vrednost za `viewer` mora biti `"read-only"`.

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
