---
id: 699b61ca8c9cb9bc1968364e
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Kada otvorite konzolu, trebalo bi da vidite sledeću poruku o grešci:

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

U standardnom JavaScript, imao je dozvolu da doda bilo koju osobinu i vrednost koje želiš na objekt `profile`. Ali u TypeScript, kompajler proverava da li objekt ima samo navedene osobine (`username`, `age` i `isLoggedIn` u ovom slučaju). Pošto `mood` nije definisano u tipu, TypeScript ga tretira kao nevažeću osobinu i generiše grešku.

Uklonite `mood: null` iz svog koda.

# --hints--

Ne bi trebalo da imaš `mood: null` u svom kodu.

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
