---
id: 699b4e284291adcfcb90df47
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, ćе se naviknuti na rad sa osnovnim tipovima zajedno sa tipovima objekata i nizovima podataka za izgradnju sigurnog profila korisnika tipa.

Počnite kreiranjem parametra nazvanog `profile`, a zatim postavite vrednost objekta. Vaš objekat mora imati ova tri svojstva i njihove vrednosti:

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

Trebalo bi da imate parametar nazvan `profile`.

```js
assert.exists(profile);
```

Vaš parametar za `profile` treba da bude nešto.

```js
assert.isObject(profile);
```

Trebalo bi da imaš atribut `username` na tvom objektu `profile`.

```js
assert.property(profile, "username");
```

Vaše svojstvo za `username` treba da ima vrednost od `"codeLearner"`.

```js
assert.propertyVal(profile, "username", "codeLearner");
```

Vaš `profile` mora imati svojstvo od `age`.

```js
assert.property(profile, "age");
```

Tvoja vrednost `age` bi trebalo da bude `25`.

```js
assert.propertyVal(profile, "age", 25);
```

Vaš predmet sa `profile` mora imati svojstvo `isLoggedIn`.

```js
assert.property(profile, "isLoggedIn");
```

Tvoja vrednost za `isLoggedIn` bi trebalo da bude `false`.

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
