---
id: 699b5b7cb68b6d35afca61b7
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Iako je moguće nastaviti dodavanje svojstava kao ovih objektu `profile`, bilo bi dobro ograničiti vrste dozvoljenih svojstava. Bolje bi bilo definisati strukturu objekta `profile` prije toga.

U prethodnim lekcijama, naučili ste kako da radite sa vrstama predmeta kao ovakve:

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

U ovom primeru, `composer` je objekat tipa element unutar stringa. Ovaj objekat ima tri važna svojstva: `name`, `instrument` i `isActive`. Prva dva svojstva moraju biti tipa `string` i `isActive`, a treće mora biti tipa `boolean`.

Ispravite vaš postojeći element za `profile` ili uključite tip elementa u red. `username` mora imati tip `string`, `age` mora imati tip `number` i `isLoggedIn` mora imati tip `boolean`.

# --hints--

Vaš objekat za `profile` mora imati tip objekta unutar linije sa karakteristikom `username`, postavljen na `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Vaš objekat za `profile` mora imati tip objekta unutar linije koja ima karakteristiku `age` postavljenu na `number`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Vaš predmet od `profile` mora imati tip predmeta unutar linije sa karakteristikom `isLoggedIn`, postavljen na `boolean`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
--fcc-editable-region--
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}

console.log(profile);
```
