---
id: 699b71d98c820ac7ab9eea2a
title: Korak 13
challengeType: 1
dashedName: step-13
---

# --description--

Ako ponovo otvoriš konzolu, videćeš novu poruku o grešci:

```md
Type 'number' is not assignable to type 'string'.
```

Ovo se dešava zato što pokušavate dodati svojstvo `guest` koje ima numeričku vrednost.

Kako bi se uklonio kvar sa TypeScript, promenite vrednost atributa `guest` na `"read-only"`.

Nakon ovih poslednjih promena, radionica vašeg profila korisnika će biti završena!

# --hints--

Vaše svojstvo od `guest` za predmet `userRoles` treba da ima vrednost `"read-only"`.

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
