---
id: 699b5370a43e299b4673bcaa
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Biće korisno da vidite stavke stvari iz `profile` u konzoli dok je gradite.

Dodaj `console.log` koji loguje `profile` na konzoli.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da registrujete `profile` u konzoli.

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
