---
id: 699b70b10ae344c802d106b0
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

Sada je vreme da dodamo još nekoliko atributa na objektu `userRoles`.

Povećaj atribut za `moderator` sa vrednošću `"medium-access"`.

Zatim ispod toga, dodaj atribut za `guest` sa vrednošću `3`.

# --hints--

Vaš `userRoles` mora imati svojstvo od `moderator`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

Tvoja vrednost `moderator` bi trebalo da bude `"medium-access"`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));
```

Vaš `userRoles` mora imati svojstvo od `guest`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.guest);
```

Tvoja vrednost `guest` mora biti numerička vrednost `3`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.guest.value.matches("3"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
--fcc-editable-region--
  viewer: "read-only"
  
--fcc-editable-region--
};

console.log(userRoles);
```
