---
id: 699b65caf17a8e8d340af27d
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Sada je vreme da dodate poslednji atribut na objekat `profile`.

U prethodnim lekcijama, naučili ste kako da radite sa tipovima nizova podataka. Naučili ste o različitim tipovima nizova podataka kao što su `string[]`, ili `number[]`:

```ts
const scores: number[] = [95, 87, 100];
```

U ovom primeru, polje `scores` prihvata samo brojeve. Ako pokušate dodati drugi tip kao `string`, ili `null`, onda TypeScript će vratiti grešku.

Izmeni svoj objekat od `profile` ili dodaj svojstvo `programmingLanguages`, koje je tipa `string[]`. Unutar tog objekta, postavi ovo svojstvo na `["JavaScript", "Python", "C++"]`.

# --hints--

Trebalo bi da imaš atribut `programmingLanguages` unutar svog objekta `profile`.

```js
assert.property(profile, "programmingLanguages");
```

Tvoja vrednost za `programmingLanguages` bi trebalo da bude `["JavaScript", "Python", "C++"]`.

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

Tvoja karakteristika `programmingLanguages` mora imati tip `string[]`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
--fcc-editable-region--
  
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
  
--fcc-editable-region--
};

console.log(profile);
```
