---
id: 699b6dd1a2961bd38491ef2e
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Element ``profile`` je imao ključeve i strogo tipiziranje. Ako pokušate dodati nedefinisan atribut ili postaviti vrednost nepodudarnog tipa, `TypeScript` će nam dati grešku.

Biće dobro da element `userRoles` ima manje strogih i više prilagodljivih funkcija.

Kao što se sećate iz prethodnih lekcija, možete postići ove rezultate koristeći tip `Record`:

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

U ovom primeru od `studentGrades`, TypeScript neće morati da zna sve predmete unapred. Ovo vam omogućava da dodate više predmeta kao `Art` ili `Music` kasnije. TypeScript ga samo zanima da su vrednosti svakog ključa brojevi.

Izmeni svoj element `userRoles` dodavanjem tipa `Record` unutar linije. Postavi ga na `<string, string>` tako da svaki ključ bude niz bilo kojih karaktera, a svaka vrednost niz karaktera.

# --hints--

Tvoj komponenta za `userRoles` mora imati tip od `Record` unutar linije postavljene na `Record<string, string>`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
