---
id: 699b62c99ec8c63470313365
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Sada je vreme da dodate još jedan atribut na vaš objekat od `profile`.

Počnite sa ispravljanjem tipa objekta unutar linije kako bi se uključilo svojstvo `bio` koje ima tip postavljen na `string`.

# --hints--

Trebalo bi da ispraviš tip objekta unutar linije kako bi uključio svojstvo `bio` postavljeno na `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
