---
id: 615f4172e9eec061d6456221
title: Korak 19
challengeType: 0
dashedName: step-19
---

# --description--

Prekomerno korišćenje podebljanog teksta ili privlačenje pažnje na važne informacije. Umesto ciljanja svakom elementu koji zahteva podebljanje, bolje je koristiti klasu elemenata za primenu stila podebljanog teksta na svaki element.

Navedite atribut elementa `p` sa vrednošću `class` za vaše drugo polje od `bold`.

# --hints--

Tvoj drugi element od `p` mora imati svojstvo elementa `class` sa vrednošću `bold`.

```js
assert(document.querySelectorAll('p')?.[1]?.classList?.contains('bold'));
```

Prvi atribut vašeg elementa `p` ne sme imati atribut elementa `class` sa vrednošću `bold`.

```js
assert(!document.querySelector('p')?.classList?.contains('bold'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
--fcc-editable-region--
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
--fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px;
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}
```
