---
id: 615f34ecc1091b4fd5a8a484
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Unutar vašeg elementa `head`, dodajte element `link` sa svojstvom elementa `rel` postavljeno na `stylesheet` i svojstvom elementa `href` postavljeno na `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Ovo je import porodična grupa fontova od `Open Sans`, sa vrednostima težine fonta `400`, `700`, i `800`.

Takođe dodajte funkciju `link` za povezivanje vašeg fajla `styles.css`.

# --hints--

Vaš kod treba da ima dva elementa od `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

Vaši elementi za `link` treba da budu void elementi, ne bi smele imati zatvarajući tag `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Vaša dva elementa od `link` treba da budu unutar elementa/komponente `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Svi vaši elementi sa `link` moraju imati svojstvo elementa `rel` postavljeno na `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Jedna komponenta vaših `link` treba da ima atribut komponente `href` postavljen na `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

Jedan od vaših komponenti za `link` treba da ima svojstvo elementa `href` postavljeno na `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
