---
id: 615f41c979787462e76dab90
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Vaša nova klasa još uvek nije dobila nikakav stil. Kreirajte selektor za `.bold` i postavite atribut `font-weight` na `800` ili učinite tekst podebljanim.

Nastavi i ukloni funkciju `font-weight` iz tvog selektora `h1` takođe.

# --hints--

Trebalo bi da imaš birač za `.bold`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold'));
```

Vaš element od `.bold` mora imati svojstvo `font-weight` postavljeno na `800`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold')?.fontWeight === '800');
```

Tvoj kandidat za `h1` ne bi trebalo da ima atribut `font-weight`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === "");
```

Ne smete da uklonite vaš izbor za `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p class="bold">Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--
h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px;
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

--fcc-editable-region--
```
