---
id: 615f405b89a7ec5f8e2d11f4
title: Korak 17
challengeType: 0
dashedName: step-17
---

# --description--

Napravi selektor za `.divider` novi i postavi svojstvo `border-bottom` na `1px solid #888989`. Takođe mu dodaj gornju i donju ivicu od `2px`. Ne sme imati levu niti desnu ivicu.

# --hints--

Trebalo bi da napravite selektor za `.divider`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.divider'));
```

Vaš selektor za `.divider` mora imati svojstvo `border-bottom` postavljeno na `1px solid #888989`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.divider')?.borderBottom === '1px solid rgb(136, 137, 137)');
```

Vaš selektor za `.divider` mora imati svojstvo `margin` postavljeno na `2px 0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.divider')?.margin === '2px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
}

p {
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
