---
id: 615f2d4150fe0d4cbd0f2628
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Ispod tvog elementa ``h1``, dodaj element ``p`` sa tekstom ``8 servings per container``.

# --hints--

Treba da dodaš novu komponentu/funkcionalnost za `p`.

```js
assert.exists(document.querySelector('p'));
```

Vaša komponenta `p` mora biti unutar vaše komponente `body`.

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

Tvoj element `p` mora da sledi nakon tvog elementa `h1`.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

Vaš element `p` mora imati tekst od `8 servings per container`.

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
    
</body>
</html>
--fcc-editable-region--
```

```css

```
