---
id: 615f34948891834dd77655a6
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Dodaj drugi element od `p` koji sadrži tekst `Serving size 2/3 cup (55g)`.

# --hints--

Trebalo bi da ima drugi deo `p`.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

Vaš drugi element za `p` treba da bude unutar vašeg elementa za `body`.

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

Tvoj drugi element za `p` treba da dođe nakon tvog postojećeg elementa za `p`.

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

Vaš drugi element od `p` mora imati tekst `Serving size 2/3 cup (55g)`.

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
    
</body>
</html>
--fcc-editable-region--
```

```css

```
