---
id: 615f38cabc64e3556f98cc1a
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Granice se mogu koristiti za povezivanje i davanju prioriteta sadržaju.

Napravi ključ za `.label` na kipa sa atributom `border` da bi bio `2px solid black`.

# --hints--

Trebalo bi da imaš birač za `.label`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.label'));
```

Vaš selektor za `.label` mora imati atribut `border` postavljen na `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.border === '2px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--

--fcc-editable-region--
```
