---
id: 615f2abbe7d18d49a1e0e1c8
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dao ti je osnovni model za HTML.

Kreiraj element sa `h1` unutar tvog elementa za `body` i daj mu tekst od `Nutrition Facts`.

# --hints--

Trebalo bi da dodaš novu komponentu `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Tvoj element `h1` mora biti unutar tvog elementa `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

Vaš element od `h1` treba da ima tekst `Nutrition Facts`.

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
