---
id: 615f38279e5c3d53692ea441
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Umotajte vaše elemente od `h1` i `p` unutar komponente za `div`. Dodajte `div` to svojstvo komponente za `class` postavljeno na `label`.

# --hints--

Trebalo bi da kreirate novu komponentu za `div`.

```js
assert(document.querySelector('div'));
```

Vaš novi element za `div` treba da ima svojstvo komponente `class` postavljeno na `label`.

```js
assert(document.querySelector('div')?.classList?.contains('label'));
```

Vaše komponente za `h1` i `p` moraju biti unutar vaše nove komponente za `.label`.

```js
const children = [...document.querySelectorAll('h1'), ...document.querySelectorAll('p')];
assert(children?.every(child => child?.parentElement?.classList?.contains('label')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

--fcc-editable-region--
<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
--fcc-editable-region--
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}
```
