---
id: 615f357957e370510f21ea16
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Kreiraj selektor za `body` sa svojstvom `font-family` da bude `Open Sans` ukoliko ima alternativni put/način `sans-serif`.

Zapamtite da fontovi koji imaju razmak u imenu moraju biti obuhvaćeni navodnicima za CSS.

# --hints--

Trebalo bi da imaš birač za `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

Tvoj selektor za `body` mora imati svojstvo od `font-family` postavljeno na `"Open Sans", sans-serif`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
