---
id: a97fd23d9b809dac9921074f
title: Izgradnja generalne funkcionalnosti za opcione argumente
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

U ovom okruženju ćete napraviti funkciju koja prihvata do dva argumenta i sabira ih, ali ako postoji samo jedan argument, funkcija će vratiti drugu funkciju koja čeka drugi broj za sabiranje.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratorijum.

Priče o korisniku:

1. Kreiraj funkciju `addTogether`.
2. Ako funkcija prima dva argumenta, funkcija treba da vrati zbir ova dva argumenta.
3. Ako je dostavljen samo jedan argument, `addTogether` treba da vrati funkciju.
   - Funkcija koja je vraćena kada se pozove sa jednim argumentom, treba da vrati zbir.

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. Ako bilo koji argument nije važeći broj, vrati `undefined`.

# --hints--

Trebalo bi da imaš funkciju `addTogether`.

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)` trebalo bi da vrati 5.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` treba da vrati 53.4.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` trebalo bi vratiti `undefined`.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` treba da vrati `undefined`.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` trebalo bi da vrati `undefined`.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` Trebalo bi da vrati funkciju.

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)` Trebalo bi da vrati 12.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` trebalo bi vratiti `undefined`.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` trebalo bi da vrati `undefined`.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
