---
id: 695774002591bbc5f8cf3e53
title: Funkcionalni kit za izradu sa popustom
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

U ovom radnom listu ćeš napisati funkciju koja računa konačnu cenu proizvoda nakon primene procentualnog popusta.

Na primer, ako je cena proizvoda `50` i popust od `20` se primeni, iznos popusta je `10`, a krajnja cena je `40`.

**Cilj:** Popunite korisničke priče navedene ispod da biste dovršili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da definišete funkciju pod nazivom `apply_discount`.

2. Funkcija `apply_discount` treba da prihvata samo dva parametra: `price` i `discount`.

3. Ako `price` nije broj (`int` ili `float`), funkcija bi trebalo da vrati niz znakova `The price should be a number`.

4. Ako `discount` nije broj (`int` ili `float`), funkcija bi trebalo da vrati niz slova `The discount should be a number`.

5. Ako `price` je manji ili jednak `0`, funkcija treba da vrati niz znakova `The price should be greater than 0`.

6. Ako `discount` je manje od `0` ili veće od `100`, funkcija bi trebalo da vrati sekvencu karaktera `The discount should be between 0 and 100`.

7. Ako su oba unosa validna, funkcija bi trebalo da izračuna popust kao procenat od cene.

8. Funkcija mora vratiti konačnu cenu nakon primene popusta.

# --hints--

Trebalo bi da imaš funkciju nazvanu `apply_discount`.

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

Tvoja funkcija za `apply_discount` mora da primi dva parametra: `price` i `discount`.

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

Kada `apply_discount` bude upućeno sa `price` (prvo polje) koje nije broj (`int` ili `float`), treba da vrati `The price should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

Kada `apply_discount` bude pozvan od `discount` (druga funkcija) koja nije broj (`int` ili `float`), treba da vrati `The discount should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

Kada `apply_discount` bude pozvano od strane `price` koje je manje ili jednako `0`, trebalo bi da vrati `The price should be greater than 0`.

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

Kada `apply_discount` je definisano preko `discount`, koje je manje od `0` ili veće od `100`, treba da vrati `The discount should be between 0 and 100`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)` bi trebalo da vrati `80`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)` treba da vrati `100`.

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)` trebalo bi da vrati `50`.

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

Kada `apply_discount` je primenjeno sa popustom od `100`, treba da vrati `0`.

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)` treba da vrati `59.6`.

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
