---
id: 5d822fd413a79914d39e98d0
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Kreiraj element `div` unutar `body` koji ima klasu objekata `background-buildings`. Ovo će biti kontejner za grupu zgrada.

# --hints--

Trebalo bi da kreirate komponentu `div`.

```js
assert.exists(document.querySelector('div'));
```

Vaš element `div` mora biti unutar `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Tvoj element za `div` treba da ima klasu objekata `background-buildings`

```js
assert([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```
