---
id: 5d822fd413a79914d39e98cf
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Možete videti `body` (što je kutija duboko unutar vaše stranice); okružna kutija je element od `html`. Napravite da `body` popuni celokupno vidljivo područje dajući `height` od `100vh`. Uklonite podrazumevano iz `margin` sa `body` postavljanjem `margin` na `0`. Na kraju, postavite atribut `overflow` na `hidden` da sakrije trake za skrolovanje vidljive kada nešto prelazi izvan vidljivog područja.

# --hints--

Trebalo bi da koristiš ključnik `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` vaš treba da imate `height` od `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

`body` vaše treba da ima `margin` od/sa `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

`body` mora imati svojstvo `overflow` postavljeno na `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```
