---
id: 5d822fd413a79914d39e98d4
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Unesite dijelove vaše zgrade `width` na `height` vigezo i ove vrijednosti: `70%` na `10%` do `.bb1a`, `80%` na `10%` do `.bb1b`, `90%` na `10%` do `.bb1c`, i `100%` na `70%` do `.bb1d`. Zapamtite da su ovi procenati u odnosu na roditeljski element i proverite da visina iznosi do 100% - popunjavanje kontejnera vertikalno.

# --hints--

Trebalo bi da koristiš alat za klasifikaciju objekata da napraviš model `.bb1a`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

Trebalo bi mu dati `.bb1a` `width` od `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

Trebalo bi mu dati `.bb1a` `height` od `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

Trebalo bi da koristiš alat za klasifikaciju objekata da napraviš model `.bb1b`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

Trebalo bi mu dati `.bb1b` `width` od `80%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

Treba da mu daš `.bb1b` `height` od `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

Treba da koristiš alat za klasifikaciju entiteta da napraviš model `.bb1c`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

Trebalo bi mu dati `.bb1c` `width` od `90%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

Trebalo bi mu dati `.bb1c` `height` od `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

Trebalo bi da koristite alat za klasifikaciju objekata da biste kreirali model `.bb1d`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

Trebalo bi mu dati `.bb1d` `width` od `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

Trebalo bi mu dati `.bb1d` `height` od `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```
