---
id: 5d822fd413a79914d39e98d1
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Ukloni svoj element od `.background-buildings` `width` na `height` iz `100%` ili mu dati punu širinu i visinu roditeljskog elementa, `body`.

# --hints--

Trebalo bi da koristiš klasu objekata `background-buildings` da izabereš odgovarajući atribut.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

Vaša funkcija od `.background-buildings` treba da ima `width` za `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

Vaša karakteristika za `.background-buildings` treba da ima `height` od `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--
    
```
