---
id: 5d822fd413a79914d39e98cb
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sledeće, dodajte tagove za otvaranje i zatvaranje `head` i `body` unutar elementa `html`.

# --hints--

Trebalo bi da imate etiketu za otvaranje `head`.

```js
assert(code.match(/<head\s*>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Trebalo bi da imaš oznaku za otvaranje `body`.

```js
assert(code.match(/<body\s*>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Elementi za `head` i `body` bi trebalo da budu sestrinski elementi.

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

Element `head` treba da bude unutar elementa `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

Element za `body` treba da bude unutar elementa `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">
  
</html>
--fcc-editable-region--

```
