---
id: 5d822fd413a79914d39e98d7
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Za korišćenje varijable, unesite ime varijablama unutar običnih zagrada sa `var` ispred, kao što je ovo: `var(--variable-name)`. Bilo koja vrednost koju dodelite varijabli koristiće se za bilo koji atribut koji koristite.

Dodaj varijablu `--building-color1` koju ste kreirali u prethodnom koraku kao vrednost svojstva `background-color` klase objekata `.bb1a`.

# --hints--

Atribut `background-color` elementa `.bb1a` treba biti postavljen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a')?.backgroundColor)
```

Trebalo bi da koristiš `var(--building-color1)` za postavljanje atributa `background-color` elementu `.bb1a`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.getPropVal('background-color', true), 'var(--building-color1)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #999;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}
--fcc-editable-region--
.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
    
```
