---
id: 5d822fd413a79914d39e98ca
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Dodajte otvarajući i zatvarajući tag `html` ispod `DOCTYPE` ili se nalazite na mestu gde želite da počnete unositi kod. Proverite da ste jezik podesili na engleski.

# --hints--

Tvoje saopštenje za `DOCTYPE` trebalo bi biti na početku tvog HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Tvoj element za `html` mora imati oznaku za otvaranje, i svojstvo elementa `lang` sa vrednošću `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Element vašeg `html` treba da ima zatvarajući tag.

```js
assert(code.match(/<\/html\s*>/));
```

Etikete za `html` moraju biti u ispravnom redosledu.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Trebalo bi da ima samo jedan element `html`.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
