---
id: 5d822fd413a79914d39e98d6
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Sada imate nešto što liči na strukturu. Vi ste spremni da kreirate svoju prvu varijablu. Deklaracije varijabli počinju sa dve crte (dashes) (`-`) i daju se ime i vrednost kao ovo: `--variable-name: value;`. U pravilima klase `bb1`, kreirajte varijablu nazvanu `--building-color1` i dodelite joj vrednost `#999`.

# --hints--

Trebalo bi da kreirate novu varijablu pod nazivom `--building-color1`.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('--building-color1'));
```

Treba da definišete varijablu ``--building-color1`` unutar ``.bb1``.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
```

Trebalo bi mu dati vrednost `--building-color1` od `#999`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1').trim(),'#999');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}

.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
    
```
