---
id: 5d822fd413a79914d39e98d2
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Ubaciti u `div` sa klasom objekta `bb1` u kontejner pozadnjih zgrada. Otvori svoj fajl `styles.css`, i dodeli mu klasu objekta `.bb1` od `width` od `10%` i `height` od `70%`. "bb" znači "background building", ovo će biti tvoja prva zgrada.

# --hints--

Trebalo bi da kreirate novu komponentu za `div`.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

Trebalo bi mu dodeliti novu osobinu klase objekta `div` `bb1`.

```js
assert.exists(document.querySelector('div.bb1'));
```

Trebalo bi da koristiš klasifikator objekata `.bb1` kako bi mu dao taj stil elementa.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

Trebalo bi mu dati atribut klase entiteta `.bb1` od `width` od `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

Treba da mu date atribut od `.bb1` klase objekata `height` iz `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
    
```
