---
id: 5d822fd413a79914d39e98cd
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

U CSS, možete ciljati sve koristeći zvezde. Povećajte opseg za sve koristeći selektor `*`, i dobijete svojstvo `border` sa vrednošću `1px solid black`. Ovo je metoda koja pomaže da se vidi gde su stvari i njihov veličina. Bićete ga dobiti kasnije.

# --hints--

Trebalo bi da koristiš alat `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Trebalo bi da koristiš atribut `border` za kreiranje stila za sve elemente.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Svi elementi moraju imati granicu od `1px solid black`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
