---
id: 5d822fd413a79914d39e98cc
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Unutar `head`, postavite unutar elementa `meta` sa svojstvom elementa `charset` sa vrednošću `UTF-8`, elementa `title` sa zaglavljem `City Skyline`, i elementa `link` koji povezuje vaš fajl za `styles.css`.

# --hints--

Treba da kreiraš element `meta` unutar elementa `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('meta'));
```

Trebalo bi mu dati oznaku `meta` sa atributom elementa `charset` sa vrednošću `UTF-8`.

```js
const linkElement = document.querySelector('meta');
const charset = linkElement?.getAttribute("charset").toLowerCase();
assert.strictEqual(charset, 'utf-8');
```

Vaš kod mora imati karakteristiku `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Element ``title`` treba da bude unutar elementa ``head``.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

Vaš projekat bi trebalo da ima zaglavlje `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

Zapamtite, velika i mala slova zajedno sa pravopisom su važne za naslov.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

Vaš kod mora imati element `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Ne smete da menjate vaše postojeće oznake za `head`. Proverite da niste izbrisali oznaku za zatvaranje.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Trebalo bi da imate jedan element sa `link` bez oznake za zatvaranje.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Tvoj element `link` treba da bude unutar tvog elementa `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Vaš element od `link` treba da ima svojstvo atributa `rel` sa vrednošću `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Vaš element za `link` mora imati atribut elementa `href` sa vrednošću `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
