---
id: 66ec4c8e9878d8441956516f
title: Napravi tabelu popisa knjiga
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba da kreirate element `table` koji prikazuje informacije o knjigama.
1. Vaša tabela mora imati zaglavlje tabele sa jednom redom unutra.
1. Red u zaglavlju tabele mora imati četiri elementa zaglavlja tabele, sa tekstom `Title`, `Author`, `Genre` i `Publication Year`, u tom redosledu.
1. Vaša tabela mora imati telo tabele sa najmanje pet redova unutra.
1. Svaki red u telu tabele mora imati četiri elementa podataka tabele koji prikazuju Naslov knjige, Autor, Žanr i Godinu izdavanja.
1. Vaša tabela mora imati podnožje tabele sa jednom redom unutra.
1. Red u podnožju tabele mora imati element podatka tabele koji zauzima četiri kolone i ima tekst `Total Books: N`, gde se `N` treba zameniti brojem knjiga u vašoj tabeli.

# --hints--

Trebalo bi da ima jedan element `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Trebalo bi da imaš jedan element `thead` unutar tvog elementa `table`.

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

Trebalo bi da imaš jedan element `tr` unutar svog elementa `thead`.

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

Trebalo bi da imaš četiri elementa `th` unutar sekvence tvog elementa `thead`.

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

Vaše četiri elementa za `th` treba da imaju tekst od `Title`, `Author`, `Genre`, i `Publication Year`, u tom redosledu.

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent.trim(), 'Title');
assert.equal(ths[1]?.textContent.trim(), 'Author');
assert.equal(ths[2]?.textContent.trim(), 'Genre');
assert.equal(ths[3]?.textContent.trim(), 'Publication Year');
```

Treba da imaš jedan element od `tbody` unutar tvog elementa `table`.

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

Vaš element za `tbody` mora da ima najmanje pet redova.

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

Svaki red u vašeg dela `tbody` treba da ima četiri elementa `td` kao decu.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

Svaki svaki element `td` u tabelarnom telu mora imati tekst sa informacijama o knjizi.

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.trim().length, 1);
});
```

Trebalo bi da imaš jedan element `tfoot` unutar tvog elementa `table`.

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

Trebalo bi da imaš jedan element `tr` u tvom elementu `tfoot`.

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

Polje `td` u redu vašeg polja `tfoot` treba da ima atribut `colspan` postavljen na `4`.

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

Polje `td` u redu vašeg polja `tfoot` trebalo bi da ima tekst/sadržaj `Total Books: N` gde je `N` broj knjige u vašoj tabeli.

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent.trim(), `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
