---
id: 60a3e3396c7b40068ad6996a
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Trenutno bi trebalo da se naviknete na osnovne karakteristike koje stranica HTML treba da ima.

Generišite svoj kod sa deklaracijom ``DOCTYPE``, elementom ``html`` čiji je jezik postavljen na engleski, elementom ``head`` i elementom ``body``.

# --hints--

Tvoj kod mora imati prefiks `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Trebalo bi da imaš etiketu pakovanja `<html>` sa karakteristikom elementa `lang` sa vrednošću `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Vaš kod bi trebalo da sadrži element `head` unutar elementa `html`.

```js
assert(document.querySelectorAll('head').length === 1);
```

Tvoj kod mora imati element `body` unutar elementa `html`.

```js
assert(document.querySelectorAll('body').length === 1);
```

Vaša komponenta `head` mora da stoji prije komponente `body`.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
