---
id: 60a3e3396c7b40068ad6997e
title: Korak 21
challengeType: 0
dashedName: step-21
---

# --description--

Sada je `.one` savršeno centriran, ali njegov gornji rub se proteže izvan platna za crtanje i granice okvira, pomerajući celo platno nadole za 20 piksela.

Dodajte `padding` od `1px` na element `.canvas` kako biste funkciji `.one` pružili čvrst mehanizam za gađenje.

# --hints--

Trebalo bi da postaviš vrednost za `padding` na `1px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

Vaše polje `.canvas` mora imati vrednost `padding` za `1px`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--

--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
