---
id: 60a3e3396c7b40068ad6997d
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Koristi vizetizer da staviš element `.one` na platno za crtanje.

Dodaj kratku osobinu od `margin` sa vertikalnom granicom od `20px` i sa horizontalnom granicom od `auto`.

# --hints--

Trebalo bi da postaviš vrednost za `margin` na `20px auto`.

```js
const marginFilter = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.style.margin === '20px auto');
assert(marginFilter.length === 2);
```

Vaše polje za `.one` trebalo bi imati vrednost `margin` od `20px auto`.

```js
const oneMargin = new __helpers.CSSHelp(document).getStyle('.one')?.getPropertyValue('margin');
assert(oneMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
