---
id: 60a3e3396c7b40068ad6996f
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Dodajte element `div` unutar `body`.

Postavi svojstvo elementa za `class` jednako `canvas`.

Ovo će se koristiti kao platno za crtanje/slikanje vašeg umetničkog dela.

# --hints--

Vaš kod treba da ima element `div`.

```js
assert(document.querySelectorAll('div').length === 1)
```

Element `div` mora imati svojstvo elementa `class` sa vrednošću `canvas`.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
