---
id: 60a3e3396c7b40068ad69977
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Koristi `padding` da popraviš prostor unutar elementa.

U `.frame`, koristite svojstvo padding-a od `padding` da dodate razmak između elemenata `.frame` i `.canvas` za `50px`. Svojstvo padding-a će dodati razmak gore, dole, levo i desno od ivice i platna za crtež unutar elementa.

# --hints--

Trebalo bi da postaviš vrednost `padding` na `50px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

Vaše polje za `.frame` trebalo bi imati vrednost `padding` od `50px`.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
