---
id: 60a3e3396c7b40068ad6996e
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Rub je oblast van kutije, i može se koristiti za regulisanje prostora između drugih kutija ili komponenti.

Ovde donji element ima veliki gornji rub, što ga dodatno spušta na stranici.

Sada, pošto ste razumeli primer kutije CSS, počnimo sa slikom Rothka.

Uklonite element `<img>`.

# --hints--

Ne smete imati element `img` u vašem kodu.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
