---
id: 60a3e3396c7b40068ad69970
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Pre nego što počnete sa dekorisanjem ``div`` koji ste dodali, morate da povežete vaš `CSS` sa vašim `HTML`.

Dodaj element `link` da poveže tvoju datoteku od `styles.css`. Postavi `href` na `styles.css`, i zapamti da postaviš atribut elementa `rel` na `stylesheet`.

# --hints--

Vaš kod mora imati element `link`.

```js
assert(/<link/.test(code))
```

Vaš element `link` mora imati svojstvo elementa `rel` sa vrednošću `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

Vaš element `link` mora imati atribut elementa `href` sa vrednošću `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
