---
id: 60a3e3396c7b40068ad6997f
title: Korak 22
challengeType: 0
dashedName: step-22
---

# --description--

Dodavanje jednog piksela prostora na vrh, dno, lijevo i desno platna za crtanje promijenilo je njegove dimenzije u 502 x 602 piksela.

Promeni atribut `padding` u postavljanje `overflow` na `hidden` - razvlači platno za crtanje po njegovim početnim dimenzijama.

# --hints--

Treba da ukloniš atribut `padding` sa selektora `.canvas`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

Trebalo bi da postaviš atribut `overflow` na `hidden`.

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

Vaše polje za `.canvas` mora imati vrednost od `overflow` od `hidden`.

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
