---
id: 60a3e3396c7b40068ad6996b
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Unutar polja `head`, dodaj labelu za `meta` koja postavlja `charset` na `UTF-8`, i polja `title` sa vrednošću `Rothko Painting`.

Unutar elementa `body`, dodaj element `img` koji ima `src` od `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.# --hints--

Vaš kod mora imati oznaku `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Etiketa za `meta` treba da postavi `charset` na `UTF-8`.

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Vaš kod mora imati karakteristiku `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` treba da bude `Rothko Painting`.

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

Tvoj kod bi trebalo da ima karakteristiku `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

Funkcionalnost za `img` treba da ima `src` od `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
