---
id: 60a3e3396c7b40068ad69979
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Dodaj novu funkciju `div` unutar svoje funkcije `.canvas`.

Napravi novi element od `div` svojstvo elementa `class` sa vrednošću `one`. Ovo će biti vaš prvi pravougaonik.

# --hints--

Trebalo bi da kreirate novu komponentu za `div`.

```js
assert(document.querySelectorAll('div').length === 3);
```

Moraš da postaviš unutar novog elementa `div`, unutar tvog elementa `.canvas`.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

`div` novi mora imati svojstvo polja `class` sa vrednošću `one`.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
