---
id: 60a3e3396c7b40068ad6996c
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

U primeru kutije od CSS, svaki element iz HTML tretira se kao kutija sa četiri područja.

Zamislite da dobijate kutiju od vašeg omiljenog onlajn prodavca -- sadržaj su proizvodi unutar kutije, ili u našem slučaju, naslov, pasus ili element slike.

Promeni atribut `src` u `<img>` iz `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` u `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

Karakteristika `img` mora imati svojstvo `src` od `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```
