---
id: 60a3e3396c7b40068ad69975
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Napiši novi zakon koristeći alat za razvrstavanje predmeta `.frame`.

Koristite kratki deo `border` da obuhvatite element `.frame` do čvrste, crne linije širine `50px`.

# --hints--

Vaš kod treba da ima selektor za `.frame`.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

Trebalo bi da postaviš atribut `border` na `50px solid black`.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

Vaša funkcija za `.frame` treba da ima `50px solid black` `border`.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
