---
id: 60a3e3396c7b40068ad69974
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Svaki crtež zahteva ram.

Umetni element `.canvas` u drugi element `div`. Dodeli/Postavi `div` klasu `frame`.# --hints--

Treba da dodaš novu komponentu/funkcionalnost za `div`.

```js
assert(document.querySelectorAll('div').length === 2)
```

Vaša komponenta `.canvas` treba biti smeštena unutar nove komponente `div`.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

`div` nov mora imati `class` sa vredno od `frame`.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

`div` novi treba da bude unutar vašeg elementa za `body`.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
