---
id: 69446f18fdcb11aa516f7c5c
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovom radionici, kreiraćete funkciju za `countup` koja stvara niz podataka brojeva počevši od 1 do zadatnog broja korišćenjem petlje.

Počnite kreiranjem prazne funkcije nazvane `countup` koja prima jedan parametar.

# --hints--

Trebalo bi da imaš funkciju pod nazivom `countup`.

```js
assert.isFunction(countup);
```

Tvoja funkcija za `countup` treba da primi jedan parametar.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
