---
id: 69447054012a22f38c944ff0
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Dodaj izjavu za `if` unutar svoje funkcije `countup`.

Proveri da li je `number` manji od `1` i vrati prazan red podataka.

Ovo će se koristiti kao bazni slučaj za vaš ponavljajući ciklus.

# --hints--

Kada je `number` manje od `1`, funkcija treba da vrati prazan red podataka.

```js
assert.deepStrictEqual(countup(0), []);
```

Ako `number` nije manje od `1`, funkcija ne bi trebalo da vrati ništa.

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
