---
id: 6948705160b66cf93a916e29
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Prikaži red sa praznim podacima nazvan `countArray` unutar funkcije `countup`, neposredno iznad vaše deklaracije `if`.

Ovaj red podataka će sačuvati vrednost `number` pri svakom pozivu funkcije.

# --hints--

Trebalo bi da objavite prazan niz podataka pod nazivom `countArray`.

```js
assert.match(countup.toString(), /(let|const|var)\s+countArray\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
