---
id: 694870bcf3c6874112b6f951
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Za sve brojeve veće ili jednake od `1`, kreiraj blok `else` i unutar njega, pozovi funkciju `countup` za `number - 1` i postavi vrednost rezultata na `countArray`.

Zatim koristi `.push()` da dodaš `number` za `countArray`, i konačno vrati `countArray`.

Ovo osigurava da rekurzivni ciklus (recursion) konačno dosegne bazni slučaj i terminira.

# --hints--

Morate postaviti vrednost rezultata poziva za `countup` za `number - 1` za `countArray`.

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

Trebalo bi da guraš `number` unutar `countArray`.

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

Moraš vratiti `countArray`.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
