---
id: 68f1f993460f06a64c033061
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Atribut `poster` je minijaturna slika videa. Razmislite o videima koje gledate na YouTube-u. Prikazuje se dok se video preuzima. Ako se atribut izostavi, prikazan je prvi kadar videa tokom faze preuzimanja.

Sada dodajte atribut `poster` sa vrednošću `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` vašem elementu `video`.

# --hints--

Your `video` element should have the `poster` attribute with the value `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--
    
--fcc-editable-region--
  >
  </video>
</body>
</html>
```
