---
id: 68f1f57723debf9383c75947
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Atribut `loop` ponovo će pokrenuti video kada se reprodukcija završi. Zamislite internet mem koji ponavlja reprodukciju. Izostavljanje atributa `loop` učiniće da se video reprodukuje samo jednom.

Atribut `loop` je booleanski atribut i ne zahteva vrednost.

Dodajte atribut `loop` elementu `video`.

# --hints--

Your `video` element should have the `loop` attribute.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
