---
id: 68f1f6f97af9749a42260a43
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Atribut `controls` pruža kontrole reprodukcije uključujući reprodukciju, povratak i kontrolu jačine za element `video`.

Atribut `controls` je booleanski atribut i ne zahteva vrednost.

Dodajte atribut `controls` na element `video`.

Sada bi trebalo da vidite element `video` prikazan na stranici.

# --hints--

Your `video` element should have the `controls` attribute.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
