---
id: 68f1f8c248306da204655117
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Atribut `muted` će utišati zvuk pri prvom pokretanju. Ako je omogućen `controls`, korisnik će moći ponovo da uključi zvuk. Izostavljanje atributa `muted` će pustiti zvuk pri prvom pokretanju.

Atribut `muted` je booleanski atribut i ne zahteva vrednost.

Dodajte atribut `muted` elementu `video`.

# --hints--

Your `video` element should have the `muted` attribute.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
