---
id: 68f1f371fefd448b4f70b898
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

U prethodnoj lekciji ste naučili o različitim atributima dostupnim za element `video`. Atribut `width` određuje širinu videa u pikselima.

Dodajte atribut `width` elementu `video` sa vrednošću `640`.

# --hints--

Your `video` element should have the attribute `width` with the value `640`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
