---
id: 68f1fac259a4afabb478b960
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Možda ste primetili da niste povezali na stvarni video. To ćete uraditi u sledećoj fazi. Kada je reč o tipovima video fajlova, postoje razlike u podršci pregledača. Da bi se to podržalo, možete koristiti elemente `source` unutar elementa `video`, a pregledač će izabrati prvi kompatibilan `source`.

Evo primera elementa `source`:

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

Element `source` je prazan element (void element), pa nema zatvarajući tag.

Dodajte element `source` unutar vašeg elementa `video`.

# --hints--

You should have a `source` element inside of your `video` element.

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
