---
id: 68f26a5f79e97c5032c10642
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Koristili ste video fajl sa ekstenzijom `mp4`, i morate reći pretraživaču da bi znao kako da pročita taj fajl.

Koristićete atribut `type` za specifikovanje MIME tipa `video/mp4`.

MIME (Multipurpose Internet Mail Extensions) je standard za opisivanje dokumenata u oblicima različitim od ASCII teksta, na primer, audio, video i slike.

MP4, zvanično poznat kao MPEG-4 Deo 14, digitalni multimedijski kontejnerski format je. Široko se koristi za skladištenje videa i zvuka, ali može da uključuje i druge tipove podataka poput titlova i statičnih slika. MP4 fajlovi su dizajnirani za striming preko Interneta i kompatibilni su sa mnogim uređajima i platformama.

Sada dodajte atribut `type` i vrednost `video/mp4`.

# --hints--

Your `source` element should have a `type` attribute.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

Your `source` element should have a `type` attribute with the value `video/mp4`.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
--fcc-editable-region--
      
--fcc-editable-region--
    >
  </video>
</body>
</html>
```
