---
id: 68f26d52b13e0f5b6e2d2e09
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Drugi uobičajeni MIME tip je MIME tip `video/webm`.

WebM je open-source audiovizuelni format medijskih fajlova koji je razvila kompanija Google, prvenstveno dizajniran za web sadržaj. Podržava video kodeke kao što su VP8, VP9 i AV1, kao i audio kodeke kao što su Vorbis i Opus, što ga čini popularnim izborom za HTML5 video i audio elemente.

Ispod vašeg prvog `source` elementa, dodajte još jedan `source` element i dodelite mu atribut `src` sa vrednošću `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` i atribut `type` sa vrednošću `video/webm`.

# --hints--

You should have a second `source` element.

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

Your second `source` element should have a `src` attribute.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

Your second `source` element should have a `src` attribute with a value of `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

Your second `source` element should have a `type` attribute.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

Your second `source` element should have a `type` attribute with a value of `video/webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
