---
id: 6853e58d52ec64349a5760ca
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovoj radionici, ćete implementirati algoritma najkraćeg puta. Napisaćete funkciju Python koja računa najkraći put između čvorova u grafu, i takođe vraća pređeni put.

Na primer, uzimajući u obzir graf gde su gradovi povezani putevima različitih udaljenosti, algoritam će pronaći najkraći put od jednog grada do drugog. Ako želite putovati od Grada A do Grada D, algoritam može otkriti da je putovanje A ⇨ B ⇨ C ⇨ D (ukupno: km 15) kraći nego direktan put A ⇨ D (km 20).

Da počneš, definiši parametar nazvan `INF` i postavi mu vrednost `float('inf')`, koja predstavlja beskonačnu pozitivnu granicu. Kasnije, koristićeš ga da prikaže beskonačnu udaljenost između dve čvorove.

# --hints--

Treba da definišete parametar `INF` za vaš fajl.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

Treba da postaviš vrednost `float('inf')` za parametar `INF`.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
