---
id: 686251fb49ebc489ba60ccb5
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Biće vam potrebno da kreirate dvodimenzionalni niz ili predstavite matricu susednosti grafa. Ova matrica će se koristiti za reprezentovanje težine ivica između čvorova u grafu.

Kreiraj parametar nazvan `adj_matrix` i postavi mu vrednost u vidu 2D niza koji predstavlja sledeću grafičku strukturu sa težinama:

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Trebalo bi da kreirate parametar nazvan `adj_matrix`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix` treba da bude 2D niz sa navedenim vrednostima.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
