---
id: 686253f1b051998ad4904e3e
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Sada ćete kreirati glavnu funkciju koja prima tri parametra: matricu susednosti, početni čvor i opcioni ciljni čvor.

Kreiraj funkciju nazvanu `shortest_path` koja prima tri parametra: `matrix`, `start_node`, i `target_node`. Postavi `None` kao podrazumevanu vrednost za `target_node`.

Parametar ``target_node`` je postavljen na ``None`` po podrazumevanju, što ukazuje da ako nije navedjen ciljni čvor, funkcija treba da izračuna najkraći put od početnog čvora do svih ostalih čvorova u grafu.

Povećaj deklaraciju za `pass` unutar tela funkcije za sada.

# --hints--

Trebalo bi da kreirate funkciju nazvanu `shortest_path`.

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

Funkcija mora da primi tri parametra: `matrix`, `start_node`, na `target_node` dok `target_node` ima podrazumevanu vrednost od `None`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
