---
id: 686b73007fc582db0132c361
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Potrebno je izračunati najkraću poznatu udaljenost od početnog čvora do svakog drugog čvora u grafu.

Da biste to uradili, kreirajte parametar nazvan `distances` i počnite postavljanjem kao niz sa samo jednim elementom: `INF`.


# --hints--

Trebalo bi da kreirate parametar nazvan `distances`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("distances")`
        )
      );
    }
});
```

Lista za `distances` treba biti inicijalizovana na `[INF]`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF]")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)

--fcc-editable-region--
```
