---
id: 686b78bc13d518e1267448e4
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Osim praćenja udaljenosti, takođe je potrebno pratiti stvarne staze koje su korišćene za dolazak do svakog čvora.

Kreiraće se lista gde svaki element čuva putanju do tog čvora. Inicijalno, putanja za svaki čvor će sadržati samo taj čvor.

List comprehensions pružaju kratak način za kreiranje liste. Na primer:

```py
[x * 2 for x in range(3)]
```

Kreiraj parametar nazvan `paths` i počni korišćenjem opisne liste koja stvara listu sa `[node_no]` za svaki `node_no` u `range(n)`.

# --hints--

Trebalo bi da imaš parametar koji se zove `paths`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

Treba vam da inicijalizujete `paths` koristeći enumeraciju.

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

Vaš alat za opisivanje mora da koristi `node_no` koji prolazi preko `range(n)`.

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

Vaš popis opisa treba da proceni `[node_no]` za svaki `node_no` u `range(n)`.

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0

--fcc-editable-region--
```
