---
id: 68760464eb9c2e79cc913fac
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

U ovom koraku, dodati ćete ciklus koji će se pokrenuti za svaki čvor u grafu. Ovaj ciklus će omogućiti algoritmu da koriguje rastojanje i putanju kroz više ciklusa.

Napravite petlju za `for` koja se izvršava `n` puta. Koristi `_` kao parametar ciklusa jer nema potrebe da koristiš ponovnu vrednost.

Unutar ciklusa, morate se pripremiti za izabiranje sledećeg procesnog čvora kreiranjem dva parametra:

- održavanje najmanjeg rastojanja pronađenog do sada u tekućoj iteraciji
- i drugi, za skladištenje reprezentacije čvora sa tom najmanjom udaljenošću.

Definiši parametre `min_distance` i `current`, a postavi ih na `INF` i `-1`, respektivno.

# --hints--

Trebalo bi da imaš ciklus `for` unutar funkcije `shortest_path`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

Tvoj ciklus za `for` treba da se ponovi preko `range(n)`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

Vaš ciklus for mora da koristi `_` kao ponavljajući parametar.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

Treba da kreiraš parametar nazvan `min_distance` i popuniš ga sa `INF` unutar ciklusa for.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

Trebalo bi da kreirate parametar nazvan `current` i popunite ga sa `-1` unutar ciklusa for.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
--fcc-editable-region--

--fcc-editable-region--
```
