---
id: 68777816e95d4936f129819f
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Pre nego što pokušate da izračunate udaljenost do suseda, morate potvrditi da je taj sused dostupan i da nije posetjen. Zatim ćete izračunati ukupnu udaljenost koja bi stigla do tog suseda preko trenutnog čvora.

Unutar ciklusa od `for`, dodaj izjavu za `if` koja proverava:

- Ako je `distance` jednak `INF` (što znači da postoji most između ovih čvorova)
- Susjed `node_no` još uvek nije poseta/posećen

Unutar bloka `if`, kreiraj parametar nazvan `new_distance` i postavi mu vrednost zbir `distances[current]` (najkraća udaljenost do trenutnog čvora) i `distance` (udaljenost od trenutnog čvora do suseda).

# --hints--

Trebalo bi da imaš deklaraciju `if` koja proverava da li `distance` nije jednako `INF` i da li susedni čvor nije posetađen.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

Treba da napraviš parametar koji se zove `new_distance` unutar bloka `if`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

Trebalo bi da postaviš vrednost `new_distance` kao zbir `distances[current]` i `distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
              min_distance = distances[node_no]
              current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):  
            distance = matrix[current][node_no]

--fcc-editable-region--
```
