---
id: 68e4c0ff9edcc92771d7549a
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Trenutno, lista `distances` ima samo jedan element. Međutim, potreban vam je jedna vrednost udaljenosti za svaki čvor u grafu.

U Python, možete proširiti listu punim brojem da je ponovite. Na primer, `[0] * 3` kreira `[0, 0, 0]`.

Potvrdite spisak `distances` ili počnite sa člankom `n` za `INF` uz dodavanje `[INF]` za `n`.

# --hints--

Trebali biste ispraviti `distances` ili početi sa člankom `n` za `INF`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF]

--fcc-editable-region--
```
