---
id: 6894c1b715b8d18a18ef3203
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovoj radionici ćete se vežbati sa Python klasama i objektima gradeći inventar muzičkih instrumenata.

Da biste definisali klasu u Pythonu, koristite ključnu reč `class` praćenu imenom klase i dvotačkom.

```python
class ClassName:
    # class contents go here
```

Zapamtite da u Pythonu imena klasa obično koriste PascalCase, što znači da svaka reč počinje velikim slovom i nema razmaka između reči.

Počnite kreiranjem klase pod nazivom `MusicalInstrument`. Dodajte `pass` kao telo klase za sada.

# --hints--

You should create a class called `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
