---
id: 6895e0bdb0348bf46952ab67
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Sada je vreme da kreirate instance vaše klase `MusicalInstrument`.

Da biste kreirali instancu klase, pozivate ime klase kao funkciju, prosleđujući potrebne argumente za metod `__init__` (osim `self`, koji je automatski obrađen od strane Pythona).

```python
my_instance = ClassName(arg1, arg2)
```

Kada kreirate ovu instancu, Python poziva vaš metod `__init__` i skladišti vrednosti kao instance atribute koji pripadaju samo ovoj specifičnoj instanci.

Izvan definicije klase, kreirajte instancu pod imenom `instrument_1` za `Oboe` koji je `woodwind` instrument.

# --hints--

You should create an instance of the `MusicalInstrument` class named `instrument_1`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

You should pass `Oboe` and `woodwind` as arguments when creating `instrument_1`.

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--

--fcc-editable-region--
```
