---
id: 6895e235935106f88281b648
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Možete pristupiti atributima vaših instanci koristeći notaciju tačke:

```py
instance_name.attribute_name
```

Ovo vam omogućava da povratite vrednosti koje su bile sačuvane tokom inicijalizacije.

Pokušajte da ispišete atribute `name` i `instrument_type` za obe vaše instance. Bićete u mogućnosti da vidite vrednosti koje ste dodelili kada ste ih kreirali.

# --hints--

You should print the `name` attribute of `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

You should print the `instrument_type` attribute of `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

You should print the `name` attribute of `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

You should print the `instrument_type` attribute of `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--

--fcc-editable-region--
```
