---
id: 6895e3f91c069e0068a42582
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Unutar vaše metode ``play``, koristite ``print()`` da prikažete poruku o igranju instrumenta.

Koristite f-string da uključite naziv instrumenta u poruku: ``The [instrument name] is fun to play!``. ``[instrument name]`` treba zameniti stvarnim nazivom instrumenta.

# --hints--

Tvoja metoda `play` treba da koristi f-string za prikaz poruke `The [instrument name] is fun to play!`. Koristi `self.name` za pristup nazivu instrumenta.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
--fcc-editable-region--
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')
```
