---
id: 6895e4940e39950288e406c2
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Zatim, dodajte još jednu metodu u vašu klasu nazvanu `get_fact` koja vraća činjenicu o instrumentu.

Ova metoda bi trebalo da ima samo parametar `self` i da vrati f-string koji kaže `The [instrument name] is part of the [instrument type] family of instruments.`.

Za razliku od metode `play` koja direktno ispisuje poruku, ova metoda bi trebalo da vrati string koji se može koristiti negde drugde u vašem kodu.

# --hints--

You should define a method named `get_fact` with a `self` parameter.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

Tvoja metoda `get_fact` treba da vrati f-string koji sadrži naziv i tip instrumenta.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')
```
