---
id: 6895e0452b3370f26f365a46
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

U metodi ``__init__``, morate dodeliti parametre instancijalnim atributima. Instancijalni atributi su varijable koje pripadaju specifičnoj instanci klase.

Unutar vaše metode ``__init__``, dodelite parametar ``name`` na ``self.name`` i parametar ``instrument_type`` na ``self.instrument_type``.

# --hints--

Parametar `name` treba da dodeliš promenljivoj `self.name` unutar metode `__init__`.

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

You should assign the `instrument_type` parameter to `self.instrument_type` inside the `__init__` method.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
--fcc-editable-region--
        pass
--fcc-editable-region--
```
