---
id: 6895e197180c9af6988e1e19
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Kreirajte drugu instancu nazvanu `instrument_2` za `Trumpet` koji je `brass` instrument.

Primećujete kako možete kreirati više instanci iz istog klasnog šablona. Svaka instanca skladišti svoje atribute nezavisno od drugih instanci.

# --hints--

You should create an instance named `instrument_2` of the `MusicalInstrument` class.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

You should pass `Trumpet` and `brass` as arguments when creating `instrument_2`.

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
--fcc-editable-region--

--fcc-editable-region--
```
