---
id: 6895dffc64e97df0e62e9ef9
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Metode su funkcije definisane unutar klase koje mogu izvršavati akcije koristeći podatke te klase. Evo kako kreirati metode:

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

U Pythonu, metoda ``__init__`` je specijalna metoda poznata kao inicijalizator metoda, koja se poziva kada kreirate novu instancu klase.

Unutar vaše klase ``MusicalInstrument``, dodajte metodu ``__init__`` sa tri parametra: ``self``, ``name``, i ``instrument_type``.

Parametar ``self`` je standardan način za referisanje na instancu klase i zahteva se kao prvi parametar za sve metode instance u klasi.

# --hints--

You should define an `__init__` method inside your `MusicalInstrument` class.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

Your `__init__` method should have three parameters: `self`, `name`, and `instrument_type`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
--fcc-editable-region--
    
    pass
--fcc-editable-region--
```
