---
id: 6895e4e023e07603bcef46f9
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Sada koristite metode koje ste kreirali.

Pozovite metodu `play()` na vašoj instanci `instrument_1`. Ovo će ispisati poruku o sviranju oboa.

Da biste pozvali metodu na instanci, koristite notaciju tačke:

```python
instance_name.method_name()
```

# --hints--

You should call the `play()` method on the `instrument_1` instance.

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_1.play()")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--

--fcc-editable-region--
```
