---
id: 66faac4139dbbd5dd632c860
title: Popuni obrazac za prijavu za posao
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Trebalo bi da imate `div` element sa klasom `container`.
1. Unutar `div` elementa trebalo bi da imate `form` element.
1. Forma treba da sadrži `input` element tipa `text` sa atributom `name` za unos punog imena korisnika.
1. Trebalo bi da imate još jedan `input` element tipa `email` sa atributom `email` za unos adrese elektronske pošte korisnika.
1. Forma treba da uključuje `select` element sa atributom `position` koji omogućava korisniku da izabere radnu poziciju.
1. Trebalo bi da imate `fieldset` element sa klasom `radio-group`.
1. Unutar `.radio-group` trebalo bi da imate grupu `input` elemenata tipa `radio` sa odgovarajućim oznakama za izbor opcija dostupnosti (na primer, Puno radno vreme, Nepuno radno vreme). Atribut `name` grupe treba da bude `availability`.
1. Trebalo bi da imate `textarea` element sa atributom `id="message"` za unos poruke.
1. Trebalo bi da povežete svaki `input`, `select` i `textarea` element sa odgovarajućim `label` elementom.
1. Trebalo bi da imate `button` element tipa `submit` za slanje forme.
1. Dodajte pseudo-klasu `:focus` na `input` i `textarea` elemente kako biste promenili boju ivice i uklonili podrazumevani okvir kada su aktivni.
1. `input`, `select` i `textarea` elementi treba da imaju pseudo-klasu `:invalid` koja menja boju ivice u crvenu kada se pronađe nevažeći unos.
1. `input`, `select` i `textarea` elementi treba da imaju pseudo-klasu `:valid` koja menja boju ivice u zelenu kada se unese ispravan podatak.
1. `button` element treba da ima pseudo-klasu `:hover` koja menja boju pozadine kada se pređe mišem preko njega.
1. Koristite pseudo-klasu `:checked` na `.radio-group input[type="radio"]` kako biste dodali boju ivice, boju pozadine i senku kada je radio dugme izabrano.
1. Koristite pseudo-klasu `:checked` na radio dugmadima kako biste promenili boju teksta povezanog `label` elementa kada je opcija izabrana.
1. Dodajte pseudo-klasu `:first-of-type` na `input` element kako biste drugačije stilizovali prvi deo unosa (na primer, zaobljene uglove).

**Napomena:** Uverite se da ste povezali svoj stilski list na HTML i koristili CSS.

# --hints--

Trebalo bi da imaš element sa klasom objekta `container`, koji je `div`.

```js
assert.exists(document.querySelector("div.container"));
```

Unutar elementa `div`, trebalo bi da ima element `form`.

```js
assert.exists(document.querySelector("div.container > form"));
```

Forma mora imati element `input` tipa `text` i identifikatorom `name` za unos punog imena korisnika.

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

Trebalo bi da imate još jedan `input` element tipa `email` sa atributom `email` za unos adrese elektronske pošte korisnika.

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

Form treba da sadrži element sa identifikatorom `select` koji ima elemente `position` sa komponentama `option`.

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

Trebalo bi da imaš element sa klasom objekta `.radio-group`, odnosno `fieldset` ili `section`.

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

Unutar `.radio-group` bi trebalo da postoji grupa elemenata tipa `input` sa tipom `radio` za odabir opcija dostupnosti. Grupa `name` bi trebalo da bude `availability`.

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

Trebalo bi da imaš element sa ID-om `textarea` koji ima identifikator `message` za unos poruke.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

Morate povezati svaki element od `input`, `select`, i `textarea` na elementu `label`.

```js
const els = document.querySelectorAll('input, select, textarea');
assert.isNotEmpty(els);
els.forEach(el => {
  const label = document.querySelector(`label[for="${el.id}"]`) || el.parentElement;
  assert.exists(label); 
  assert.equal(label.tagName, "LABEL");
})
```

Trebalo bi da imate `button` element tipa `submit` za slanje forme.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

Treba da dodate pseudo klasu `:focus` za elemente `input` i `textarea` da promenite boju ivice tokom fokusa. Koristite padajući meni za navedeni redosled.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'));
```

Trebalo bi da koristite selektor lažne klase za `:focus` za `input` i `textarea`, a takođe uklonite običnu liniju tokom pažnje.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('outline'), 'none');
```

Elementi `input`, `select` i `textarea` bi trebalo da imaju pseudo-klasu `:invalid` koja menja boju ivice u crvenu kada se detektuje nevažeći unos. Koristite birač liste za navedeni redosled.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

Elementi `input`, `select` i `textarea` treba da imaju pseudo-klasu `:valid` koja menja boju ivice u zelenu kada se unese validni unos. Koristite padajući izbornik za navedeni redosled.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

Element ``button`` bi trebalo da ima pseudoklasu ``:hover`` koja menja boju pozadine kada se pređe preko.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

Treba da koristiš klasu `:checked` na `.radio-group input[type="radio"]` da dodaš boju ivice kada je odabrano radio dugme.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
```

Trebalo bi da koristiš pseudo-klasu `:checked` na `.radio-group input[type="radio"]` da dodaš pozadinsku boju kada je odabran radio taster.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('background-color'));
```

Trebalo bi da koristiš lažni klasa `:checked` na `.radio-group input[type="radio"]` da dodaš senku kutije kada je odabrano radio dugme.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

Trebalo bi da koristiš lažni klasa `:checked` na radio dugmićima za promenu boje teksta `label` kada je izabrana opcija.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

Dodajte pseudo klasu za `:first-of-type` za elemente `input` da stilizuje prvi deo unosa na različit način.

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```
