---
id: 6720b3191bbafd269bae0ae8
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Kreiraj `div` element koji ima `id` sa `greeting-card` i `class` sa `card`.

Unutar `div` elementa, dodaj `h1` sa tekstom `Happy Birthday!`. Zatim dodaj paragraf element sa `class` nazvanim `message` i tekstom `Wishing you all the happiness and joy on your special day!`.

# --hints--

Trebalo bi da imate element `div`.

```js
assert.exists(document.querySelector("div"));
```

Element `div` treba da ima `id` od `greeting-card`.

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

Element `div` bi trebalo da ima `class` od `card`.

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

Trebalo bi da imate element `h1` unutar elementa `div`.

```js
assert.exists(document.querySelector("div > h1"));
```

Element `h1` treba da sadrži tekst `Happy Birthday!`.

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

Trebalo bi da imate element `p` unutar `div`.

```js
assert.exists(document.querySelector("div > p"));
```

Element `p` bi trebalo da ima `class` od `message`.

```js
assert.exists(document.querySelector("div > p.message"));
```

Element `p` treba da sadrži tekst `Wishing you all the happiness and joy on your special day!`.

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
