---
id: 6720b7f3155a9278d6b6a8b8
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sada je vreme da stilizujete svoju čestitku.

Dodajte selektor za element `body`, zatim:

- promenite `font-family` tako da bude `Arial` praćen opštim `sans-serif`,
- dodajte padding sa svih strana za `40px`,
- podesite svojstvo `text-align` na `center`.

# --hints--

Trebalo bi da imate selektor `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Trebalo bi da set `font-family` u `"Arial", sans-serif` selektoru, unutar `body`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

Trebalo bi da set `padding` do `40px` u `body` selektor.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

You should set `text-align` to `center` in the `body` selector.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
