---
id: 672b4434dd6d6a1edc405419
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Dodajte dva elementa ``a`` unutar elementa ``.card-links``.

Prvi treba da ima tekst ``Send Card``, ``class`` sa vrednošću ``send-link`` i ``href`` sa vrednošću ``#send``.

Drugi treba da ima tekst `Share on Social Media``, ``class`` sa vrednošću ``share-link`` i ``href`` sa vrednošću ``#share`.

# --hints--

Prvi element `a` unutar `.card-links` treba da ima tekst `Send Card`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

Prvi element `a` treba da ima `class` sa vrednošću `send-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

Prvi element `a` treba da ima `href` sa vrednošću `#send`.

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

Drugi element `a` unutar `.card-links` treba da ima tekst `Share on Social Media`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

Drugi `a` element bi trebalo da ima `class` sa vrednošću od `share-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

Drugi `a` element treba da ima `href` sa vrednošću `#share`.

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
