---
id: 67332db0f793bf08f2a18528
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Vreme je da popunimo drugi `section`!

Dodajte element `h2` u element `#share` koji sadrži tekst `Sharing your card...`, a zatim dodajte element `p` sa tekstom `Your card was shared on social media!`.

# --hints--

Element `#share` trebao bi da sadrži element `h2`.

```js
assert.exists(document.querySelector("#share > h2"));
```

Element `h2` bi trebalo da sadrži tekst `Sharing your card...`.

```js
assert.strictEqual(document.querySelector("#share > h2").innerText.trim(), "Sharing your card...");
```

Drugi element unutar `#share` treba da bude `p` element.

```js
assert.strictEqual(document.querySelector("#share").children[1].tagName, "P");
```

Element `p` treba da ima tekst `Your card was shared on social media!`.

```js
assert.strictEqual(document.querySelector("#share").children[1].innerText.trim(), "Your card was shared on social media!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
