---
id: 69a00f673f171e29159e3ad8
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Ovaj čvor bi zahtevao dve osobine: element za dodavanje i pokazivač na sledeći čvor.

Unutar funkcije `add`, kreirajte objekat nazvan ``node``; trebalo bi da ima osobinu nazvanu ``element`` i osobinu ``next`` postavljenu na ``null``, koja je pokazivač na sledeći element liste.

# --hints--

You should create an object named `node` inside your `add` function.

```js
assert.match(__helpers.removeJSComments(add.toString()), /const|var|let\s+node\s*=\s*{/);
```

Your `node` object should have a property named `element` set to `element`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /element\s*:\s*element/);
```

Your `node` object should have a property named `next` set to `null`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /next\s*:\s*null/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
