---
id: 69a03aaa1e6fe8152788ce6a
title: Korak 14
challengeType: 1
dashedName: step-14
---

# --description--

Unutar funkcije `remove`, koristite `let` da kreirate varijablu `previous` postavljenu na `null`, i ponovo `let` da kreirate varijablu `current` postavljenu na `list.head`. Ovo su varijable koje ćete koristiti u petlji koja će prolaziti kroz listu.

`current` prati čvor koji se trenutno proverava u listi, počevši od glave i pomerajući se napred po jedan čvor. `previous`, s druge strane, prati čvor odmah pre trenutnog.

Ovo je važno prilikom uklanjanja čvora jer morate ažurirati vrednost `next` prethodnog čvora tako da pokazuje na čvor posle onog koji se uklanja, kako bi lista ostala povezana.

# --hints--

You should use `let` to create a `previous` variable set to `null` inside your `remove` function.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

You should use `let` to create a `current` variable set to `list.head` inside your `remove` function.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
