---
id: 69a013a86c97dac3784b2636
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Sada, log bi trebalo da prikaže nešto poput `{ head: { element: 42, next: null }, length: 1 }` u konzoli. 

Međutim, ovo nije način na koji biste rukovali slučajem u kojem je dodeljeni čvor prvi element u listi, to jest, glava (head).

Da biste to pravilno uradili, koristite `if` izjavu da proverite da li je lista prazna pozivom funkcije `isEmpty` sa `list`. Ako je to tačno, onda premestite svoju liniju `list.head = node` u telo (body) izjave `if`.

# --hints--

You should have an `if` statement with the condition `isEmpty(list)`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*/);
```

You should move your `list.head = node` line into the `if` statement.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*list\.head\s*=\s*node;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
