---
id: 69a00cd4ec513345d40e95ec
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Sledeći zadatak na kojem ćete raditi je funkcija koja vam omogućava da proverite da li je lista prazna.

Da biste to uradili, kreirajte funkciju pod nazivom `isEmpty` koja prima parametar `list`. Unutar funkcije, vratite proveru da vidite da li je svojstvo `length` objekta `list` striktno jednako `0`.

# --hints--

You should create an `isEmpty` function.

```js
assert.isFunction(isEmpty)
```

Your `isEmpty` function should have a `list` parameter.

```js
const isEmptyParams = __helpers.getFunctionParams(isEmpty.toString());
assert.deepEqual(isEmptyParams[0].name, "list");
```

Your `isEmpty` function should return `list.length === 0`.

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
