---
id: 69a03ddce82b0626f329b001
title: Korak 15
challengeType: 1
dashedName: step-15
---

# --description--

Sada, kreiraj `while` petlju koja proverava da li `current` nije strogo `null` i da `current.element` nije striktno jednak elementu.

To znači da petlja treba da nastavi iteriranje po listi sve dok se ne dođe do kraja liste ili pronađe element za uklanjanje.

# --hints--

You should create a `while` loop with the condition `current !== null && current.element !== element`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /while\s+\(current\s*!==\s*null\s*&&\s*current\.element\s*!==\s*element\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    --fcc-editable-region--
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
