---
id: 69a014bd9b1730e94136279f
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Ako lista nije prazna, želite da pronađete njen kraj, dodate novi element tamo i usmerite ga na `null`. To znači da morate iterirati kroz listu od `head`.

Da biste to počeli, otvorite `else` klauzulu. Unutar nje, koristite `let` da kreirate varijablu `current` i postavite je na `list.head`.

# --hints--

You should have an `else` clause.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{/);
```

You should create a `current` variable set to `list.head` inside the `else` clause.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{\s*let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
