---
id: 69a00a051b9a18694d00c16b
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovom radionici, implementiraćete spajanu listu.

Da biste počeli, kreirajte funkciju pod nazivom `initList`. Unutar nje, vratite objekat sa svojstvom `head` postavljeno na `null`, i svojstvom `length` od `0`.

Ovo je funkcija koju ćete pozvati za inicijalizaciju spajane liste. Više kod `head` predstavlja prvi element u listi; inicijalno je to `null` jer nema elemenata u listi. Više kod `length` prati broj čvorova (elemenata) u listi, koji ćete inkrementirati ili dekrementirati prilikom dodavanja ili uklanjanja čvora iz liste.

# --hints--

You should create an `initList` function.

```js
assert.isFunction(initList)
```

Your `initList` function should return an object.

```js
assert.isObject(initList());
```

Your `initList` function should return an object with a `head` set to `null`.

```js
assert.propertyVal(initList(), "head", null);
```

Your `initList` function should return an object with a `length` set to `0`.

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
