---
id: 69a045fd236d0d973a59f103
title: Korak 19
challengeType: 1
dashedName: step-19
---

# --description--

Ako je `previous` jednako `null`, to znači da je element koji treba ukloniti na samom početku liste, pa biste trebali ažurirati `head` da pokazuje na sledeći čvor kako biste uklonili prvi element.

Da biste to uradili, dodajte `else` klauzulu. Unutar nje, postavite `list.head` na `current.next`.

# --hints--

You should add an `else` clause to your `if` statement.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*/);
```

You should set `list.head` to `current.next` inside the `else` clause.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*list\.head\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    --fcc-editable-region--
    } 
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
