---
id: 69a012c6224fac85832247dd
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Primetite li rezultate `true`, `{ head: null, length: 0 }` i `true` u konzoli? Ovo se dešava zato što niste ništa dodelili listi `myList`, i niste inkrementirali svojstvo `length`.

Dodelite `node` na `list.head`, a ispod toga, inkrementirajte svojstvo `length` od `list`.

# --hints--

You should assign `node` to `list.head`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.head\s*=\s*node;?/);
```

You should increment the `length` property of `list`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.length\+\+;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
