---
id: 69a0102e761a245bb1cbd8e7
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

U ovom trenutku, trebalo bi da testirate spajanu listu radi boljeg razumevanja.

Kreirajte varijablu `myList` i postavite je na `initList()`. Ispod nje, logujte poziv `isEmpty()` koristeći `myList` kao parametar u konzolu. Zatim, nakon toga, pozovite `add` sa `myList` i brojem po izboru, zatim logujte `myList` u konzolu. Ispod toga, logujte još jedan poziv `isEmpty()` u konzolu koristeći `myList` kao parametar ponovo.

# --hints--

You should create a `myList` variable set to `initList()`.

```js
assert.deepEqual(myList, initList());
```

You should log a call to `isEmpty(myList)` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\);?/)
```

You should call `add` with `myList` and a number of your choice as the parameters.

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

You should log your `myList` variable to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?/)
```

You should log `isEmpty(myList)` to the console again.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?\s*console\.log\(isEmpty\(myList\)\);?/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
