---
id: 69a0179a85b00b6127e0af45
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Unutar petlje `while`, ažurirajte `current` na `current.next` kako bi petlja napredovala kroz listu po jedan čvor, i zaustavila se kada dođe do čvora sa svojstvom `next` od `null`, tj. kraj liste.

# --hints--

You should update the value of your `current` variable to `current.next` inside the `while` loop.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
