---
id: 587d78a9367417b2b2512ae8
title: Nauci kako Bezier krive funkcionisu
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bDrs8'
forumTopicId: 301058
dashedName: learn-how-bezier-curves-work
---

# --description--

Poslednji izazov je predstavio `animation-timing-function` svojstvo i nekoliko kljucnih reci koje menjaju brzinu animacije tokom trajanja. CSS nudi i druge opcije osim kljucnih reci koje daju finiju kontrolu nad tim kako se animacija izvodi, kroz upotrebu Bezier krivih.

U CSS animacijama, Bezier krive se koriste zajedno sa `cubic-bezier` funkcijom. Oblik krive predstavlja kako ce se animacija odvijati. Kriva se nalazi u koordinatnom sistemu 1x1. X osa tog sistema predstavlja vreme animacije, a Y osa predstavlja promenu u animaciji.

`cubic-bezier` funkcija ima cetiri glavne tacke u ovom 1x1 koordinatnom sistemu: `p0`, `p1`, `p2` i `p3`. `p0` i `p3` su fiksne tacke pocetka i kraja, uvek na koordinatama (0, 0) i (1, 1). Vi postavljate x i y vrednosti za preostale dve tacke, a njihov polozaj odredjuje oblik krive koji animacija prati. U CSS-u se to radi navodenjem x i y vrednosti `p1` i `p2` anchor tacaka ovako: `(x1, y1, x2, y2)`. Kada sve spojimo, evo primera Bezier krive u CSS kodu:

```css
animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
```

U primeru iznad, x i y vrednosti su iste za svaku tacku (x1 = 0.25 = y1 i x2 = 0.75 = y2), sto, ako se secate geometrije, daje pravu liniju od pocetka do tacke (1, 1). Ova animacija je linearna promena elementa kroz vreme i ekvivalentna je `linear` kljucnoj reci. Drugim recima, menja se konstantnom brzinom.

# --instructions--

Za element sa id `ball1`, promenite `animation-timing-function` vrednost sa `linear` na njen ekvivalent u `cubic-bezier` funkciji. Koristite vrednosti tacaka date u prethodnom primeru.

# --hints--

`animation-timing-function` vrednost za element sa id `ball1` treba da bude linearni ekvivalent `cubic-bezier` funkcije.

```js
const ballOne = document.querySelector('#ball1');
const ballOneStyle = window.getComputedStyle(ballOne); 
assert.equal(
  ballOneStyle?.animationTimingFunction,
    'cubic-bezier(0.25, 0.25, 0.75, 0.75)'
);
```

`animation-timing-function` vrednost za element sa id `ball2` ne treba da se menja.

```js
const ballTwoElement = document.querySelector('#ball2');
const ballTwoStyle = window.getComputedStyle(ballTwoElement); 
const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(
  ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
