---
id: 587d78a5367417b2b2512ad9
title: Koristite CSS Transform scale svojstvo za promenu veličine elementa
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZVSg'
forumTopicId: 301076
dashedName: use-the-css-transform-scale-property-to-change-the-size-of-an-element
---

# --description--

Da biste promenili veličinu elementa, CSS ima `transform` svojstvo, zajedno sa svojom `scale()` funkcijom. Sledeći primer koda udvostručuje veličinu svih paragraf elemenata na stranici:

```css
p {
  transform: scale(2);
}
```

# --instructions--

Povećajte veličinu elementa sa identifikatorom `ball2` na 1.5 puta njegove originalne veličine.

# --hints--

`transform` svojstvo za `#ball2` treba da bude postavljeno da ga skalira na 1.5 puta njegove veličine.

```js
assert.match(code,
    /#ball2\s*?{\s*?left:\s*?65%;\s*?transform:\s*?scale\(1\.5\);\s*?}|#ball2\s*?{\s*?transform:\s*?scale\(1\.5\);\s*?left:\s*?65%;\s*?}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .ball {
    width: 40px;
    height: 40px;
    margin: 50 auto;
    position: fixed;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    border-radius: 50%;
  }
  #ball1 {
    left: 20%;
  }
  #ball2 {
    left: 65%;

  }


</style>

<div class="ball" id= "ball1"></div>
<div class="ball" id= "ball2"></div>
```

# --solutions--

```html
<style>
  .ball {
    width: 40px;
    height: 40px;
    margin: 50 auto;
    position: fixed;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    border-radius: 50%;
  }
  #ball1 {
    left: 20%;
  }
  #ball2 {
    left: 65%;
    transform: scale(1.5);
  }
</style>
<div class="ball" id= "ball1"></div>
<div class="ball" id= "ball2"></div>