---
id: 587d781e367417b2b2512acb
title: Zaključajte element za roditelja koristeći apsolutno pozicioniranje
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJ7c3'
forumTopicId: 301060
dashedName: lock-an-element-to-its-parent-with-absolute-positioning
---

# --description--

Sledeća opcija za CSS `position` svojstvo je `absolute`, koja zaključava element na mesto u odnosu na njegov roditeljski kontejner. Za razliku od `relative` pozicioniranja, ovo uklanja element iz normalnog toka dokumenta, tako da okolne stavke ne reaguju na njega. CSS offset svojstva (top ili bottom i left ili right) se koriste za podešavanje pozicije.

Jedna jedinstvena stvar kod apsolutnog pozicioniranja je da će biti zaključano u odnosu na najbližeg pretka koji ima postavljenu *poziciju*. Ako zaboravite da dodate pravilo pozicije roditeljskom elementu (ovo se obično radi pomoću `position: relative;`), pretraživač će nastaviti da traži više u lancu i na kraju će koristiti `body` tag kao podrazumevanu opciju.

# --instructions--

Zaključajte `#searchbar` element gore-desno u njegovom roditelju `section` tako što ćete deklarisati njegovo `position` svojstvo kao `absolute`. Dajte mu offset od `top` i `right` od po 50 piksela.

# --hints--

`#searchbar` element treba da ima `position` svojstvo postavljeno na `absolute`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.position, 'absolute');
```

Vaš kod treba da koristi `top` offset od 50 piksela na `#searchbar` elementu.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.top, '50px');
```

Vaš kod treba da koristi `right` offset od 50 piksela na `#searchbar` elementu.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.right, '50px');
```

# --seed--

## --seed-contents--

```html
<style>
  #searchbar {



  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```

# --solutions--

```html
<style>
  #searchbar {
    position: absolute;
    top: 50px;
    right: 50px;
  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>