---
id: 587d78a8367417b2b2512ae4
title: Napravi CSS otkucaj srca koristeci beskonacan broj ponavljanja animacije
challengeType: 0
videoUrl: 'https://scrimba.com/c/cDZpDUr'
forumTopicId: 301062
dashedName: make-a-css-heartbeat-using-an-infinite-animation-count
---

# --description--

Evo jos jednog primera kontinuirane animacije koja koristi `animation-iteration-count` svojstvo, i koristi srce napravljeno u prethodnom izazovu.

Animacija srca koja traje jednu sekundu sastoji se iz dva dela sa animacijom. `heart` elementi (ukljucujuci `:before` i `:after` delove) animiraju promenu velicine pomocu `transform` svojstva, a pozadinska boja `div` elementa menja se pomocu `background` svojstva.

# --instructions--

Nastavite da srce kuca tako sto cete dodati `animation-iteration-count` svojstvo klasama `back` i `heart` i postaviti vrednost na `infinite`. Selektori `heart:before` i `heart:after` ne zahtevaju dodatna animation svojstva.

# --hints--

`animation-iteration-count` svojstvo za klasu `heart` treba da ima vrednost `infinite`.

```js
const heartElement = document.querySelector('.heart');
 const heartStyle = window.getComputedStyle(heartElement);
 assert.equal(heartStyle?.animationIterationCount, 'infinite');
```

`animation-iteration-count` svojstvo za klasu `back` treba da ima vrednost `infinite`.

```js
const backElement = document.querySelector('.back');
 const backStyle = window.getComputedStyle(backElement);
 assert.equal(backStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>
  .back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: white;
    animation-name: backdiv;
    animation-duration: 1s;

  }

  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
    animation-name: beat;
    animation-duration: 1s;

  }
  .heart:after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart:before {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }

  @keyframes backdiv {
    50% {
      background: #ffe6f2;
    }
  }

  @keyframes beat {
    0% {
      transform: scale(1) rotate(-45deg);
    }
    50% {
      transform: scale(0.6) rotate(-45deg);
    }
  }

</style>
<div class="back"></div>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: white;
    animation-name: backdiv;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
    animation-name: beat;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }
  .heart:after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart:before {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }

  @keyframes backdiv {
    50% {
      background: #ffe6f2;
    }
  }

  @keyframes beat {
    0% {
      transform: scale(1) rotate(-45deg);
    }
    50% {
      transform: scale(0.6) rotate(-45deg);
    }
  }
</style>
<div class="back"></div>
<div class="heart"></div>
```
