---
id: 587d78a9367417b2b2512ae9
title: Koristi Bezier krivu za pomeranje grafike
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bnRCK'
forumTopicId: 301071
dashedName: use-a-bezier-curve-to-move-a-graphic
---

# --description--

Prethodni izazov je obradio `ease-out` kljucnu rec, koja opisuje animaciju koja ubrzava na pocetku, a usporava na kraju. Levo je prikazana razlika izmedu `ease-out` (za plavi element) i `linear` (za crveni element). Sekvenca animacije slicna `ease-out` kljucnoj reci moze se postici prilagodenom kubnom Bezier funkcijom.

Uopsteno, promena anchor tacaka `p1` i `p2` pravi razlicite Bezier krive koje kontrolisu kako se animacija krece kroz vreme. Evo primera Bezier krive koja imitira `ease-out` stil:

```css
animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
```

Imajte na umu da sve `cubic-bezier` funkcije pocinju u `p0` tacki (0, 0) i zavrsavaju u `p3` tacki (1, 1). U ovom primeru, kriva se brze krece duz Y ose (krece od 0, ide do `p1` y vrednosti 0, zatim do `p2` y vrednosti 1) nego duz X ose (od 0, zatim 0 za `p1`, pa do 0.58 za `p2`). Zbog toga se promena kod animiranog elementa desava brze u tom delu od samog vremena animacije. Pred kraj krive, odnos promene x i y vrednosti se menja: y ide od 1 do 1 (bez promene), dok x ide od 0.58 do 1, pa se promena animacije odvija sporije u odnosu na vreme animacije.

# --instructions--

Da biste videli efekat ove Bezier krive u akciji, promenite `animation-timing-function` elementa sa id `red` u `cubic-bezier` funkciju sa x1, y1, x2, y2 vrednostima `0, 0, 0.58, 1`. Time ce oba elementa prolaziti kroz animaciju na isti nacin.

# --hints--

`animation-timing-function` svojstvo elementa sa id `red` treba da bude `cubic-bezier` funkcija sa x1, y1, x2, y2 vrednostima redom 0, 0, 0.58, 1.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.equal(
  redStyle?.animationTimingFunction, 'cubic-bezier(0, 0, 0.58, 1)'
);
```

Element sa id `red` vise ne treba da ima `animation-timing-function` svojstvo sa vrednoscu `linear`.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.notEqual(redStyle?.animationTimingFunction, 'linear');
```

Vrednost `animation-timing-function` svojstva za element sa id `blue` ne treba da se menja.

```js
const blueElement = document.querySelector('#blue');
const blueStyle = window.getComputedStyle( blueElement);
const blueBallAnimation = __helpers.removeWhiteSpace(
  blueStyle?.animationTimingFunction
);
assert.isTrue(
  blueBallAnimation == 'ease-out' ||
    blueBallAnimation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```

# --solutions--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```
