---
id: 587d78a7367417b2b2512ae1
title: Napravi kretanje koristeci CSS animaciju
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7amZfW'
forumTopicId: 301051
dashedName: create-movement-using-css-animation
---

# --description--

Kada elementi imaju definisan `position`, kao sto su `fixed` ili `relative`, CSS offset svojstva `right`, `left`, `top` i `bottom` mogu da se koriste u pravilima animacije za kreiranje kretanja.

Kao sto je prikazano u primeru ispod, element mozete pomeriti nadole pa nagore tako sto postavite `top` svojstvo keyframe-a na `50%` na 50px, dok ostaje 0px u prvom (`0%`) i poslednjem (`100%`) frejmu.

```css
@keyframes rainbow {
  0% {
    background-color: blue;
    top: 0px;
  }
  50% {
    background-color: green;
    top: 50px;
  }
  100% {
    background-color: yellow;
    top: 0px;
  }
}
```

# --instructions--

Dodajte horizontalno kretanje `div` animaciji. Koristeci `left` offset svojstvo, dopunite `@keyframes` pravilo tako da rainbow pocinje na 0 piksela u `0%`, ide na 25 piksela u `50%` i zavrsava na -25 piksela u `100%`. Ne menjajte `top` svojstvo u editoru, animacija treba da ima i vertikalno i horizontalno kretanje.

# --hints--

`@keyframes` pravilo za `0%` treba da koristi `left` offset od 0px.

```js
assert.match(code,/[^50]0%\s*?{[\s\S]*?left:\s*?0px(;[\s\S]*?|\s*?)}/gi);
```

`@keyframes` pravilo za `50%` treba da koristi `left` offset od 25px.

```js
assert.match(code,/50%\s*?{[\s\S]*?left:\s*?25px(;[\s\S]*?|\s*?)}/gi);
```

`@keyframes` pravilo za `100%` treba da koristi `left` offset od -25px.

```js
assert.match(code,/100%\s*?{[\s\S]*?left:\s*?-25px(;[\s\S]*?|\s*?)}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;

    }
    50% {
      background-color: green;
      top: 50px;

    }
    100% {
      background-color: yellow;
      top: 0px;

    }
  }
</style>

<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;
      left: 0px;
    }
    50% {
      background-color: green;
      top: 50px;
      left: 25px;
    }
    100% {
      background-color: yellow;
      top: 0px;
      left: -25px;
    }
  }
</style>
<div id="rect"></div>
```
