---
id: 587d78a3367417b2b2512acf
title: Promenite poziciju preklapajućih elemenata koristeći z-index svojstvo
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM94aHk'
forumTopicId: 301046
dashedName: change-the-position-of-overlapping-elements-with-the-z-index-property
---

# --description--

Kada su elementi postavljeni na preklapajući način (tj. koristeći `position: absolute | relative | fixed | sticky`), element koji se pojavi kasnije u HTML oznakama će po defaultu biti iznad drugih elemenata. Međutim, `z-index` svojstvo može specificirati redosled kako se elementi slažu jedan iznad drugog. Mora biti ceo broj (tj. broj bez decimalnog dela), i veće vrednosti `z-index` svojstva elementa uzimaju višu poziciju u redosledu od onih sa manjim vrednostima.

# --instructions--

Dodajte `z-index` svojstvo elementu sa klasom `first` (crveni pravougaonik) i postavite ga na 2 kako bi bio iznad drugog elementa (plavi pravougaonik).

# --hints--

Element sa klasom `first` treba da ima `z-index` vrednost od 2.

```js
const firstElement = document.querySelector('.first');
const firstStyle = window.getComputedStyle(firstElement); 
assert.equal(firstStyle?.zIndex, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;

  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>

<div class="first"></div>
<div class="second"></div>
```

# --solutions--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;
    z-index: 2;
  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>

<div class="first"></div>
<div class="second"></div>