---
id: 587d781e367417b2b2512acc
title: Zakljucaj element za prozor pregledaca koristeci fixed pozicioniranje
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDNUR'
forumTopicId: 301061
dashedName: lock-an-element-to-the-browser-window-with-fixed-positioning
---

# --description--

Sledeca tehnika rasporeda koju nudi CSS je `fixed` pozicioniranje, sto je vrsta apsolutnog pozicioniranja koja vezuje element za prozor pregledaca. Kao i apsolutno pozicioniranje, koristi se zajedno sa CSS offset svojstvima i uklanja element iz normalnog toka dokumenta. Ostali elementi ga vise ne uzimaju u obzir, pa ponekad moze biti potrebna dodatna korekcija rasporeda.

Glavna razlika izmedu `fixed` i `absolute` pozicioniranja je sto se element sa `fixed` pozicijom ne pomera kada korisnik skroluje stranicu.

# --instructions--

Navigaciona traka u kodu ima id `navbar`. Promenite njen `position` u `fixed` i postavite offset na 0 piksela od `top` i 0 piksela od `left`. Nakon dodavanja koda, skrolujte prozor prikaza da vidite kako navigacija ostaje na mestu.

# --hints--

`#navbar` element treba da ima `position` postavljen na `fixed`.

```js
const navbarElement = document.querySelector('#navbar');
const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.position, 'fixed');
```

Vas kod treba da koristi `top` offset od 0 piksela na `#navbar` elementu.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
assert.equal(navbarStyle?.top, '0px');
```

Vas kod treba da koristi `left` offset od 0 piksela na `#navbar` elementu.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.left, '0px');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {



    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```

# --solutions--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {
    position: fixed;
    top: 0;
    left: 0;
    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```
