---
id: 587d78a9367417b2b2512aea
title: Ucini kretanje prirodnijim pomocu Bezier krive
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7akWUv'
forumTopicId: 301063
dashedName: make-motion-more-natural-using-a-bezier-curve
---

# --description--

Ovaj izazov animira element tako da imitira kretanje bacene lopte. Prethodni izazovi su ukljucivali `linear` i `ease-out` kubne Bezier krive, ali nijedna ne prikazuje realistican let lopte. Zato je potrebno prilagoditi Bezier krivu.

`animation-timing-function` se automatski ponavlja za svaki keyframe kada je `animation-iteration-count` postavljen na beskonacno. Posto postoji keyframe pravilo na sredini trajanja animacije (na `50%`), nastaje sekvenca koja se ponavlja dva puta u kretanju lopte nagore i nadole.

Sledeca kubna Bezier kriva imitira kretanje bacene lopte:

```css
cubic-bezier(0.3, 0.4, 0.5, 1.6);
```

Primetite da je vrednost y2 veca od 1. Iako je kubna Bezier kriva postavljena u koordinatni sistem 1x1 i x vrednosti mogu biti samo od 0 do 1, y vrednost moze biti i veca od 1. To proizvodi efekat odskoka, sto je odlicno za simulaciju lopte koja se baca.

# --instructions--

Promenite `animation-timing-function` vrednost elementa sa id `green` u `cubic-bezier` funkciju sa x1, y1, x2, y2 vrednostima 0.311, 0.441, 0.444, 1.649.

# --hints--

`animation-timing-function` vrednost za element sa id `green` treba da bude `cubic-bezier` funkcija sa navedenim x1, y1, x2, y2 vrednostima.

```js
const greenElement = document.querySelector('#green');
 const greenStyle = window.getComputedStyle(greenElement);
assert.equal(greenStyle?.animationTimingFunction, 'cubic-bezier(0.311, 0.441, 0.444, 1.649)');
```

# --seed--

## --seed-contents--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.69, 0.1, 1, 0.1);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.311, 0.441, 0.444, 1.649);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```
