---
id: 587d78a8367417b2b2512ae7
title: Promeni tajming animacije pomocu kljucnih reci
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKvwCM'
forumTopicId: 301045
dashedName: change-animation-timing-with-keywords
---

# --description--

U CSS animacijama, `animation-timing-function` svojstvo kontrolise brzinu kojom se animirani element menja tokom trajanja animacije. Ako je animacija auto koji putuje od tacke A do tacke B za odredeno vreme (vas `animation-duration`), `animation-timing-function` odreduje kako taj auto ubrzava i usporava tokom puta.

Postoji vise unapred definisanih kljucnih reci za uobicajene opcije. Na primer, podrazumevana vrednost je `ease`, koja pocinje sporo, ubrzava u sredini, pa ponovo usporava na kraju. Druge opcije su `ease-out` (pocinje brzo pa usporava), `ease-in` (pocinje sporo pa ubrzava pred kraj) i `linear` (konstantna brzina tokom cele animacije).

# --instructions--

Za elemente sa id vrednostima `ball1` i `ball2`, dodajte `animation-timing-function` svakom i postavite `#ball1` na `linear`, a `#ball2` na `ease-out`. Primetite razliku u nacinu kretanja elemenata tokom animacije, iako zavrsavaju istovremeno jer dele isti `animation-duration` od 2 sekunde.

# --hints--

`animation-timing-function` vrednost za element sa id `ball1` treba da bude `linear`.

```js
const ballOne =document.querySelector('#ball1'); 
const ballOneStyle = window.getComputedStyle(ballOne); 

const ball1Animation = __helpers.removeWhiteSpace(ballOneStyle?.animationTimingFunction);
assert.isTrue(ball1Animation == 'linear' || ball1Animation == 'cubic-bezier(0,0,1,1)');
```

`animation-timing-function` vrednost za element sa id `ball2` treba da bude `ease-out`.

```js
const ballTwo = document.querySelector('#ball2'); 
const ballTwoStyle = window.getComputedStyle(ballTwo); 

const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)');
```

# --seed--

## --seed-contents--

```html
<style>

  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;

  }
  #ball2 {
    left:56%;

  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left:56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
