---
id: 587d78a6367417b2b2512adc
title: Koristite CSS Transform skewY svojstvo za iskošenje elementa duž Y ose
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZ2uB'
forumTopicId: 301075
dashedName: use-the-css-transform-property-skewy-to-skew-an-element-along-the-y-axis
---

# --description--

Pošto `skewX()` funkcija iskošava izabrani element duž X ose za dati stepen, nije iznenađenje da `skewY()` svojstvo iskošava element duž Y (vertikalne) ose.

# --instructions--

Iskošite element sa identifikatorom `top` za -10 stepeni duž Y ose koristeći `transform` svojstvo.

# --hints--

Element sa identifikatorom `top` treba da bude iskošen za -10 stepeni duž njegove Y ose.

```js
assert.match(code, /#top\s*?{\s*?.*?\s*?transform:\s*?skewY\(-10deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;

  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;
    transform: skewY(-10deg);
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>