---
id: 587d781e367417b2b2512aca
title: Pomerite relativno pozicioniran element pomocu CSS offseta
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bQEA4'
forumTopicId: 301065
dashedName: move-a-relatively-positioned-element-with-css-offsets
---

# --description--

`top` ili `bottom`, kao i `left` ili `right` offseti govore pregledacu koliko element treba da bude pomeren u odnosu na mesto gde bi se nalazio u normalnom toku dokumenta. Pomerate element od referentne strane, pa se on zapravo pomera u suprotnom smeru. Kao sto ste videli u prethodnom izazovu, `top` offset je pomerio `h2` nadole. Slicno tome, `left` offset pomera element udesno.

# --instructions--

Koristite CSS offsete da pomerite `h2` za 15 piksela udesno i 10 piksela nagore.

# --hints--

Vas kod treba da koristi CSS offsete da pozicionira `h2` na odredjeno mesto pomeranjem 10px nagore. Drugim recima, pomerite ga 10px od `bottom` pozicije gde bi inace bio.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.bottom, '10px');
```

Vas kod treba da koristi CSS offsete da pozicionira `h2` na odredjeno mesto pomeranjem 15px udesno. Drugim recima, pomerite ga 15px od `left` pozicije gde bi inace bio.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.left , '15px');
```

# --seed--

## --seed-contents--

```html
<head>
<style>
  h2 {
    position: relative;


  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```

# --solutions--

```html
<head>
<style>
  h2 {
    position: relative;
    left: 15px;
    bottom: 10px;
  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```
