---
id: 587d78a8367417b2b2512ae3
title: Animirajte elemente neprekidno koristeći beskonačan broj animacija
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDVfq'
forumTopicId: 301041
dashedName: animate-elements-continually-using-an-infinite-animation-count
---

# --description--

Prethodni izazovi su pokrili kako koristiti neka svojstva animacije i `@keyframes` pravilo. Drugo svojstvo animacije je `animation-iteration-count`, koje vam omogućava da kontrolišete koliko puta želite da ponovite animaciju. Evo primera:

```css
animation-iteration-count: 3;
```

U ovom slučaju animacija će se zaustaviti nakon što se pokrene 3 puta, ali je moguće učiniti da se animacija nastavi zauvek postavljanjem te vrednosti na `infinite`.

# --instructions--

Da bi lopta nastavila da skače udesno u beskonačnom ciklusu, promenite `animation-iteration-count` svojstvo na `infinite`.

# --hints--

`animation-iteration-count` svojstvo treba da ima vrednost `infinite`.

```js
const ballElement = document.querySelector('#ball');
const ballStyle = window.getComputedStyle(ballElement); 
assert.equal(ballStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: 3;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>