---
id: 587d78a5367417b2b2512ad8
title: Napravi teksturu dodavanjem diskretnog paterna kao pozadinske slike
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQdwJC8'
forumTopicId: 301052
dashedName: create-texture-by-adding-a-subtle-pattern-as-a-background-image
---

# --description--

Jedan nacin da dodate teksturu i vizuelni interes pozadinskoj slici je dodavanje suptilnog paterna. Vazna je ravnoteza, jer ne zelite da pozadina bude preupadljiva i odvlaci paznju sa prednjeg sadrzaja. Svojstvo `background` podrzava `url()` funkciju za povezivanje slike teksture ili paterna. URL adresa se stavlja pod navodnike unutar obicnih zagrada.

# --instructions--

Koristeci URL `https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png`, postavite `background` za celu stranicu preko `body` selektora.

# --hints--

Vas `body` element treba da ima `background` svojstvo postavljeno na `url()` sa datim linkom.

```js
assert.match(code,/background(-image)?:\s*?url\(\s*("|'|)https:\/\/cdn-media-1\.freecodecamp\.org\/imgr\/MJAkxbh\.png\2\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {

  }
</style>
```

# --solutions--

```html
<style>
  body {
    background: url("https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png");
  }
</style>
```
