---
id: 587d781b367417b2b2512abc
title: Podesite background-color svojstvo teksta
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDqwA6'
forumTopicId: 301032
dashedName: adjust-the-background-color-property-of-text
---

# --description--

Umesto podešavanja cele pozadine ili boje teksta da bi prednji plan bio lakše čitljiv, možete dodati `background-color` elementu koji sadrži tekst koji želite da istaknete. Ovaj izazov koristi `rgba()` umesto uobičajenih `hex` ili `rgb()`.

<blockquote>rgba označava:<br>  r = crvena<br>  g = zelena<br>  b = plava<br>  a = alfa/nivo providnosti</blockquote>

RGB vrednosti mogu biti od 0 do 255. Alfa vrednost može biti od 1, što je puna boja ili potpuno neprovidna, do 0, što je potpuno providna ili čista. `rgba()` je odlična za korišćenje u ovom slučaju, jer vam omogućava podešavanje providnosti. To znači da ne morate potpuno da prekrijete pozadinu.

Koristićete `background-color: rgba(45, 45, 45, 0.1)` za ovaj izazov. On proizvodi tamno sivu boju koja je skoro nevidljiva zbog niske vrednosti providnosti od 0.1.

# --instructions--

Da bi tekst bio vidljiviji, podesite `background-color` elementa `h4` na datu `rgba()` vrednost.

Takođe za `h4`, uklonite `height` svojstvo i dodajte `padding` od 10px.

# --hints--

Vaš kod treba da doda `background-color` svojstvo elementu `h4` i postavi ga na `rgba(45, 45, 45, 0.1)`.

```js
assert(
  /(background-color|background):rgba\(45,45,45,0?\.1\)(;?}|;)/gi.test(
    code.replace(/\s/g, '')
  )
);
```

Vaš kod treba da doda `padding` svojstvo elementu `h4` i postavi ga na 10 piksela.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.equal(h4style?.paddingTop, "10px");
assert.equal(h4style?.paddingRight, "10px");
assert.equal(h4style?.paddingBottom, "10px");
assert.equal(h4style?.paddingLeft, "10px");
```

`height` svojstvo na `h4` elementu treba da bude uklonjeno.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.notEqual(h4style?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;


  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    padding: 10px;
    background-color: rgba(45, 45, 45, 0.1);
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>