---
id: 587d78a4367417b2b2512ad4
title: Prilagodite nijansu boje
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp38TZ'
forumTopicId: 301036
dashedName: adjust-the-hue-of-a-color
---

# --description--

Boje imaju nekoliko svojstava uključujući hue, saturation, i lightness. CSS3 je predstavio `hsl()` funkciju kao alternativni način za izbor boja izričito navođenjem ovih svojstava.

**Hue** je ono što ljudi obično misle kao 'boja'. Ako razmišljate o spektru boja koji počinje crvenom sa leve strane, prelazi kroz zelenu u sredini, i plavu sa desne strane, hue je gde se boja nalazi na toj liniji. U `hsl()`, hue koristi koncept kruga boja umesto linije, gde je ugao boje na krugu dat kao vrednost između 0 i 360.

**Saturation** je količina sivosti u boji. Boja sa punom saturacijom nema nimalo sive, a boja sa niskom saturacijom je blizu sive. Ovo se daje kao procenat gde je 100% puna saturacija.

**Lightness** je količina bele ili crne u boji. Procenti se daju od 0% (crna) do 100% (bela), gde je 50% normalna boja.

Evo nekoliko primera korišćenja `hsl()` sa bojama sa punom saturacijom i normalnom lightness:

<table><thead><tr><th>Boja</th><th>HSL</th></tr></thead><tbody><tr><td>crvena</td><td>hsl(0, 100%, 50%)</td></tr><tr><td>žuta</td><td>hsl(60, 100%, 50%)</td></tr><tr><td>zelena</td><td>hsl(120, 100%, 50%)</td></tr><tr><td>cijan</td><td>hsl(180, 100%, 50%)</td></tr><tr><td>plava</td><td>hsl(240, 100%, 50%)</td></tr><tr><td>magenta</td><td>hsl(300, 100%, 50%)</td></tr></tbody></table>

# --instructions--

Promenite `background-color` svakog `div` elementa prema nazivima klasa (`green`, `cyan`, ili `blue`) koristeći `hsl()`. Sve tri treba da imaju punu saturaciju i normalnu lightness.

# --hints--

Vaš kod treba da koristi `hsl()` funkciju za deklarisanje zelene boje.

```js
assert.match(code,/\.green\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Vaš kod treba da koristi `hsl()` funkciju za deklarisanje cijan boje.

```js
assert.match(code,/\.cyan\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Vaš kod treba da koristi `hsl()` funkciju za deklarisanje plave boje.

```js
assert.match(code,/\.blue\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

`div` element sa klasom `green` treba da ima `background-color` zelene boje.

```js
const greenElement = document.querySelector(".green");
const greenStyle = window.getComputedStyle(greenElement); 
assert.equal(greenStyle?.backgroundColor, 'rgb(0, 255, 0)');
```

`div` element sa klasom `cyan` treba da ima `background-color` cijan boje.

```js
const cyanElement = document.querySelector(".cyan");
const cyanStyle = window.getComputedStyle(cyanElement); 
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

`div` element sa klasom `blue` treba da ima `background-color` plave boje.

```js
const blueElement = document.querySelector(".blue");
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .blue {
    background-color: #000000;
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>

<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: hsl(120, 100%, 50%);
  }

  .cyan {
    background-color:   hsl(180, 100%, 50%);
  }

  .blue {
    background-color: hsl(240, 100%, 50%);
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>