---
id: 587d78a4367417b2b2512ad5
title: Prilagodite ton boje
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDJvT7'
forumTopicId: 301038
dashedName: adjust-the-tone-of-a-color
---

# --description--

`hsl()` opcija u CSS takođe čini lakše za promenu tona boje. Mešanje bele sa čistom bojom stvara tintu te boje, a dodavanje crne stvara senku. U suprotnom, ton se stvara dodavanjem sive ili mešanjem tinte i senke. Zapamtite da 's' i 'l' u `hsl()` označavaju saturation i lightness, redom. Procenat saturation menja količinu sive, a procenat lightness određuje koliko bele ili crne je u boji. Ovo je korisno kada imate osnovnu boju koja vam se sviđa, ali trebate različite varijacije te boje.

# --instructions--

Svi elementi imaju `background-color` postavljen na `transparent` po defaultu. Naš `nav` element trenutno izgleda da ima `cyan` pozadinsku boju, jer element iza njega ima `background-color` postavljen na `cyan`. Dodajte `background-color` za `nav` element koristeći istu cyan boju, ali sa `80%` saturation i `25%` lightness da promenite njegov ton i senku.

# --hints--

`nav` element treba da ima `background-color` sa promenjenim tonom cijan boje korišćenjem `hsl()`.

```js
// Computed style of hsl(180, 80%, 25%) results in rgb(13,115,115)
assert.equal(
  new __helpers.CSSHelp(document).getStyle('nav').getPropVal('background-color', true), 
  'rgb(13,115,115)'
)
```

# --seed--

## --seed-contents--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {

  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>

<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```

# --solutions--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {
    background-color: hsl(180, 80%, 25%);
  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>

<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>