---
id: 587d78a5367417b2b2512ad6
title: Napravite postepen CSS `linear-gradient()`
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4dpt9'
forumTopicId: 301047
dashedName: create-a-gradual-css-linear-gradient
---

# --description--

Postavljanje boje na HTML elemente nije ograničeno samo na jednu solidnu boju. CSS pruža mogućnost korišćenja prelaza boja, poznatih i kao gradienti, na elementima. Ovo se postiže kroz `background` svojstvo i `linear-gradient()` funkciju. Evo opšte sintakse:

```css
background: linear-gradient(gradient_direction, color 1, color 2, color 3, ...);
```

Prvi argument definiše smer prelaza odakle počinje - može biti definisan u stepenima, gde `90deg` kreira vertikalni prelaz (sa leva na desno) i `45deg` kreira dijagonalni prelaz (od dole levo ka gore desno). Sledeći argumenti definišu redosled boja koje se koriste u prelazu.

Primer:

```css
background: linear-gradient(90deg, red, yellow, rgb(204, 204, 255));
```

# --instructions--

Koristite `linear-gradient()` za `background` svojstvo `div` elementa, sa smerom od 35 stepeni da promenite boju od `#CCFFFF` do `#FFCCCC`.

# --hints--

`div` element treba da ima `linear-gradient` `background` sa navedenim smerom i bojama.

```js
const divElement = document.querySelector('div');
const divStyle = window.getComputedStyle(divElement); 
assert.match(divStyle?.background, /linear-gradient\(35deg, rgb\(204, 255, 255\), rgb\(255, 204, 204\)\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;

  }

</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;
    background: linear-gradient(35deg, #CCFFFF, #FFCCCC);
  }
</style>
<div></div>