---
id: 587d78a7367417b2b2512ae2
title: Kreirajte vizuelni pravac bledenjem elementa s leva na desno
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJqqAE'
forumTopicId: 301054
dashedName: create-visual-direction-by-fading-an-element-from-left-to-right
---

# --description--

Za ovaj izazov, promenićete `opacity` elementa sa animacijom tako da element postepeno nestaje kako stiže do desne strane ekrana.

U prikazanoj animaciji, okrugli element sa gradijent pozadinom se pomera na desnu stranu na 50% animacije prema `@keyframes` pravilu.

# --instructions--

Ciljajte element sa identifikatorom `ball` i dodajte `opacity` svojstvo postavljeno na 0.1 na `50%`, tako da element izbledi kako se pomera udesno.

# --hints--

`keyframes` pravilo za bledenje treba da postavi `opacity` svojstvo na 0.1 na 50%.

```js
assert.match(code,/@keyframes fade\s*?{\s*?50%\s*?{\s*?(?:left:\s*?60%;\s*?opacity:\s*?0?\.1;|opacity:\s*?0?\.1;\s*?left:\s*?60%;)/gi
);
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;

    }
  }

</style>

<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;
      opacity: 0.1;
    }
  }
</style>
<div id="ball"></div>