---
id: 587d78a7367417b2b2512ae0
title: Koristite CSS animaciju za promenu hover stanja dugmeta
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4vZAa'
forumTopicId: 301073
dashedName: use-css-animation-to-change-the-hover-state-of-a-button
---

# --description--

Možete koristiti CSS `@keyframes` animaciju da promenite boju dugmeta u njegovom hover stanju.

Evo primera promene širine slike prilikom hover-a:

```html
<style>
  img {
    width: 30px;
  }
  img:hover {
    animation-name: width;
    animation-duration: 500ms;
  }

  @keyframes width {
    100% {
      width: 40px;
    }
  }
</style>

<img src="https://cdn.freecodecamp.org/curriculum/applied-visual-design/google-logo.png" alt="Google's Logo" />
```

# --instructions--

Zapamtite da `ms` označava milisekunde, gde je 1000ms jednako 1s.

Koristite CSS `@keyframes` animaciju da promenite `background-color` elementa `button` u `#4791d0` kada korisnik pređe mišem preko njega. `@keyframes` pravilo treba da ima samo jedan unos na `100%`.

# --hints--

`@keyframes` pravilo treba da koristi boju pozadine za `animation-name`.

```js
assert.match(code, /@keyframes\s+?background-color\s*?{/g);
```

Trebalo bi da postoji jedno pravilo ispod `@keyframes` koje menja `background-color` u `#4791d0` na 100%.

```js
assert.match(code, /100%\s*?{\s*?background-color:\s*?#4791d0;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }


</style>

<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }

  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>