---
id: 587d78a6367417b2b2512ade
title: Napravite složeniji oblik koristeći CSS i HTML
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpz4fr'
forumTopicId: 301050
dashedName: create-a-more-complex-shape-using-css-and-html
---

# --description--

Jedan od najpoznatijih oblika na svetu je oblik srca, i u ovom izazovu ćete napraviti jedan koristeći čisti CSS. Ali prvo, morate razumeti pseudo-elemente `::before` i `::after`. `::before` stvara pseudo-element koji je prvo dete izabranog elementa; `::after` stvara pseudo-element koji je poslednje dete izabranog elementa. U sledećem primeru, pseudo-element `::before` se koristi za dodavanje pravougaonika elementu sa klasom `heart`:

```css
.heart::before {
  content: "";
  background-color: yellow;
  border-radius: 25%;
  position: absolute;
  height: 50px;
  width: 70px;
  top: -50px;
  left: 5px;
}
```

Da bi pseudo-elementi `::before` i `::after` funkcionisali ispravno, moraju imati definisano `content` svojstvo. Ovo svojstvo se obično koristi za dodavanje stvari poput slika ili teksta na izabrani element. Kada se `::before` i `::after` pseudo-elementi koriste za pravljenje oblika, `content` svojstvo je i dalje potrebno, ali je postavljeno na prazan string. U prethodnom primeru, element sa klasom `heart` ima pseudo-element `::before` koji proizvodi žuti pravougaonik visine `50px` i širine `70px`. Ovaj pravougaonik ima zaobljene uglove zbog `border-radius` od 25% i apsolutno je pozicioniran `5px` od leve i `50px` iznad vrha elementa.

# --instructions--

Promenite element na ekranu u srce. U selektoru `.heart::after`, promenite `background-color` u `pink` i `border-radius` na 50%.

Zatim, ciljajte element sa klasom `heart` (samo `heart`) i popunite `transform` svojstvo. Koristite `rotate()` funkciju sa -45 stepeni.

Na kraju, u selektoru `.heart::before`, postavite njegovo `content` svojstvo na prazan string.

# --hints--

`background-color` selektora `.heart::after` treba da bude `pink`.

```js
const heartAfter = code.match(/\.heart::after\s*{[\s\S]+?[^\}]}/g)?.[0];
assert.match(heartAfter, /({|;)\s*background-color\s*:\s*pink\s*(;|})/g);
```

`border-radius` selektora `.heart::after` treba da bude 50%.

```js
assert.lengthOf(code.match(/border-radius\s*?:\s*?50%/gi),2);
```

`transform` svojstvo za klasu `heart` treba da koristi `rotate()` funkciju postavljenu na -45 stepeni.

```js
assert.match(code,/transform\s*?:\s*?rotate\(\s*-45deg\s*?\)/gi);
```

`content` selektora `.heart::before` treba da bude prazan string.

```js
assert.match(code,/\.heart::before\s*?{\s*?content\s*?:\s*?("|')\1\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: ;
  }
  .heart::after {
    background-color: blue;
    content: "";
    border-radius: 25%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: ;
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
  }
  .heart::after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: "";
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>