---
id: 587d78a6367417b2b2512adb
title: Koristite CSS Transform skewX svojstvo za iskošenje elementa duž X ose
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLP8Sr'
forumTopicId: 301074
dashedName: use-the-css-transform-property-skewx-to-skew-an-element-along-the-x-axis
---

# --description--

Sledeća funkcija `transform` svojstva je `skewX()`, koja iskošava izabrani element duž njegove X (horizontalne) ose za dati stepen.

Sledeći kod iskošava paragraf element za -32 stepena duž X ose.

```css
p {
  transform: skewX(-32deg);
}
```

# --instructions--

Iskošite element sa identifikatorom `bottom` za 24 stepena duž X ose koristeći `transform` svojstvo.

# --hints--

Element sa identifikatorom `bottom` treba da bude iskošen za 24 stepena duž njegove X ose.

```js
assert.match(code, /#bottom\s*?{\s*?.*?\s*?transform:\s*?skewX\(24deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;

  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>