---
id: 587d78a3367417b2b2512ad0
title: Centrirajte element horizontalno koristeći margin svojstvo
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJqU4'
forumTopicId: 301043
dashedName: center-an-element-horizontally-using-the-margin-property
---

# --description--

Još jedan način za centriranje elementa je postavljanje margine na auto. Jedan od načina da se to uradi je postavljanje `margin` svojstva na `auto`.

Ovaj način takođe funkcioniše za slike. Slike su inline elementi po defaultu, ali mogu biti pretvoreni u block elemente postavljanjem `display` svojstva na `block`.

# --instructions--

Centrirajte `div` na stranici dodavanjem `margin` svojstva sa vrednošću `auto`.

# --hints--

`div` treba da ima `margin` postavljen na `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('div')?.margin ,'auto');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;

  }
</style>
<div></div>
```

# --solutions--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;
    margin: auto;
  }
</style>
<div></div>