---
id: 587d78a5367417b2b2512ada
title: Koristi CSS transform scale svojstvo za uvecanje elementa pri hoveru
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLPJuM'
forumTopicId: 301077
dashedName: use-the-css-transform-scale-property-to-scale-an-element-on-hover
---

# --description--

`transform` svojstvo ima vise funkcija koje vam omogucavaju da uvecate, pomerite, rotirate, izoblicite itd. elemente. Kada se koristi zajedno sa pseudo-klasama kao sto je `:hover`, koje opisuju odredeno stanje elementa, `transform` moze poboljsati interaktivnost elemenata.

Evo primera uvecanja pasus elemenata na 2.1 puta originalne velicine kada korisnik predje misem preko njih:

```css
p:hover {
  transform: scale(2.1);
}
```

**Napomena:** Primena transformacije na `div` element uticace i na sve njegove child elemente.

# --instructions--

Dodajte CSS pravilo za `div` `hover` stanje i koristite `transform` svojstvo da uvecate `div` element na 1.1 puta originalne velicine kada korisnik predje misem preko njega.

# --hints--

Velicina `div` elementa treba da se uveca za 1.1 puta kada korisnik predje misem preko njega.

```js
assert.match(code, /div:hover\s*?{\s*?transform:\s*?scale\(1\.1\);/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }



</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }
  div:hover {
    transform: scale(1.1);
  }
</style>
<div></div>
```
