---
id: 587d7791367417b2b2512ab3
title: Kreirajte vizuelnu ravnotežu koristeći text-align svojstvo
challengeType: 0
videoUrl: 'https://scrimba.com/c/c3b4EAp'
forumTopicId: 301053
dashedName: create-visual-balance-using-the-text-align-property
---

# --description--

Ovaj deo kurikuluma se fokusira na vizuelni dizajn. Prvi niz izazova gradi na osnovama postavljanja kartice koje su pokazane da ilustruju neke osnovne principe.

Tekst je često veliki deo web sadržaja. CSS ima nekoliko opcija za njegovo poravnanje koristeći `text-align` svojstvo.

`text-align: justify;` rasprostire razmake između reči tako da svaki red ima istu širinu.

`text-align: center;` centrira tekst

`text-align: right;` poravnava tekst desno

A `text-align: left;` (podrazumevano) poravnava tekst levo.

# --instructions--

Poravnajte tekst `h4` elementa koji kaže "Google" na centar. Zatim postavite red paragrafa koji opisuje kako je Google osnovan da bude poravnat na opravdanje.

# --hints--

Vaš kod treba da koristi `text-align` svojstvo na `h4` elementu da ga postavi na `center`.

```js
const h4Element =document.querySelector('h4')
const h4Style = window.getComputedStyle(h4Element);
assert.equal(h4Style?.textAlign, 'center');
```

Vaš kod treba da koristi `text-align` svojstvo na `p` elementu da ga postavi na `justify`.

```js
const pElement =document.querySelector('p')
const pStyle = window.getComputedStyle(pElement);
assert.equal(pStyle?.textAlign, 'justify');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {

  }
  p {

  }
  .links {
    margin-right: 20px;

  }
  .fullCard {
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
  }
  p {
    text-align: justify;
  }
  .links {
    margin-right: 20px;

  }
  .fullCard {
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at Stanford University.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>