---
id: 587d78a7367417b2b2512adf
title: Nauci kako rade CSS @keyframes i animation svojstva
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakprhv'
forumTopicId: 301059
dashedName: learn-how-the-css-keyframes-and-animation-properties-work
---

# --description--

Da biste animirali element, potrebno je da razumete animation svojstva i `@keyframes` pravilo. Animation svojstva odreduju kako animacija treba da se ponasa, a `@keyframes` pravilo odreduje sta se desava tokom same animacije. Ukupno postoji osam animation svojstava. Ovaj izazov ce prvo pokriti dva osnovna:

`animation-name` postavlja ime animacije, koje `@keyframes` kasnije koristi da poveze pravila sa tom animacijom.

`animation-duration` postavlja trajanje animacije.

`@keyframes` omogucava da precizno definisete sta se desava tokom cele animacije. To se radi tako sto se CSS svojstva dodeljuju odredenim "frejmovima" animacije, u procentima od 0% do 100%. Ako to uporedimo sa filmom, CSS svojstva na 0% prikazuju kako element izgleda na pocetku, a svojstva na 100% kako izgleda na kraju. CSS zatim interpolira promene kroz zadato vreme trajanja. Evo primera koji prikazuje upotrebu `@keyframes` i animation svojstava:

```css
#anim {
  animation-name: colorful;
  animation-duration: 3s;
}

@keyframes colorful {
  0% {
    background-color: blue;
  }
  100% {
    background-color: yellow;
  }
}
```

Za element sa id `anim`, gornji kod postavlja `animation-name` na `colorful` i `animation-duration` na 3 sekunde. `@keyframes` pravilo je povezano sa istim imenom `colorful`. Ono postavlja plavu boju na pocetku animacije (0%), koja se menja u zutu na kraju (100%). Niste ograniceni samo na pocetak i kraj, mozete definisati svojstva za bilo koji procenat izmedu 0% i 100%.

# --instructions--

Napravite animaciju za element sa id `rect` tako sto cete postaviti `animation-name` na `rainbow` i `animation-duration` na 4 sekunde. Zatim deklarisite `@keyframes` pravilo i postavite `background-color` na pocetku animacije (`0%`) na plavu, u sredini (`50%`) na zelenu, a na kraju (`100%`) na zutu.

# --hints--

Element sa id `rect` treba da ima `animation-name` svojstvo sa vrednoscu `rainbow`.

```js
const rectElement = document.querySelector('#rect');
const rectStyle = window.getComputedStyle(rectElement);
assert.equal(rectStyle?.animationName, 'rainbow');
```

Element sa id `rect` treba da ima `animation-duration` svojstvo sa vrednoscu `4s`.

```js
const rectElement = document.querySelector('#rect');
const rectStyle = window.getComputedStyle(rectElement);
assert.equal(rectStyle?.animationDuration, '4s');
```

`@keyframes` pravilo treba da koristi ime animacije `rainbow`.

```js
assert.match(code ,/@keyframes\s+?rainbow\s*?{/g);
```

`@keyframes` pravilo za `rainbow` treba da koristi `background-color` vrednost `blue` na 0%.

```js
assert.match(code,/0%\s*?{\s*?background-color:\s*?blue;\s*?}/gi);
```

`@keyframes` pravilo za `rainbow` treba da koristi `background-color` vrednost `green` na 50%.

```js
assert.match(code,/50%\s*?{\s*?background-color:\s*?green;\s*?}/gi);
```

`@keyframes` pravilo za `rainbow` treba da koristi `background-color` vrednost `yellow` na 100%.

```js
assert.match(code,/100%\s*?{\s*?background-color:\s*?yellow;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
  }

  #rect {


  }




</style>
<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
    }
    50% {
      background-color: green;
    }
    100% {
      background-color: yellow;
    }
  }
</style>
<div id="rect"></div>
```
