---
id: 587d78a3367417b2b2512ace
title: Pomerite elemente levo ili desno koristeci float svojstvo
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDqu2'
forumTopicId: 301066
dashedName: push-elements-left-or-right-with-the-float-property
---

# --description--

Sledeca tehnika za raspored ne koristi `position`, vec postavlja `float` svojstvo elementa. Elementi koji su float-ovani uklanjaju se iz normalnog toka dokumenta i pomeraju na `left` ili `right` stranu roditeljskog elementa koji ih sadrzi. Ovo se cesto koristi zajedno sa `width` svojstvom da bi se definisalo koliko horizontalnog prostora je elementu potrebno.

# --instructions--

Dati markup bi bolje radio kao raspored u dve kolone, gde su `section` i `aside` elementi jedan pored drugog. Dodelite elementu `#left` `float` svojstvo sa vrednoscu `left`, a elementu `#right` `float` svojstvo sa vrednoscu `right`.

# --hints--

Element sa id vrednoscu `left` treba da ima `float` vrednost `left`.

```js
const leftElement = document.querySelector('#left');
const leftStyle = window.getComputedStyle(leftElement);
assert.equal(leftStyle?.float, 'left');
```

Element sa id vrednoscu `right` treba da ima `float` vrednost `right`.

```js
const rightElement = document.querySelector('#right');
const rightStyle = window.getComputedStyle(rightElement);
assert.equal(rightStyle?.float, 'right');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
    #left {

      width: 50%;
    }
    #right {

      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```

# --solutions--

```html
<head>
  <style>
    #left {
      float: left;
      width: 50%;
    }
    #right {
      float: right;
      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>