---
id: 587d78a6367417b2b2512add
title: Napravite grafiku koristeci CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDWPs6'
forumTopicId: 301048
dashedName: create-a-graphic-using-css
---

# --description--

Menjanjem razlicitih selektora i svojstava mozete napraviti zanimljive oblike. Jedan od najjednostavnijih za probu je oblik polumeseca. U ovom izazovu koristite `box-shadow` svojstvo koje dodaje senku elementu, kao i `border-radius` svojstvo koje kontrolise zaobljenje uglova elementa.

Napravicete kruzni element, providan, sa ostrijom senkom blago pomerenom u stranu, a ta senka ce zapravo formirati oblik polumeseca.

Da biste napravili kruzni element, `border-radius` treba postaviti na vrednost 50%.

Možda se sećate iz prethodnog izazova da `box-shadow` prima vrednosti `offset-x`, `offset-y`, `blur-radius`, `spread-radius` i `color` tim redosledom. Vrednosti `blur-radius` i `spread-radius` su opcione.

# --instructions--

Izmenite kvadratni element u editoru kako biste napravili oblik meseca. Prvo promenite `background-color` u `transparent`, zatim postavite `border-radius` na 50% da biste dobili kružni oblik. Na kraju promenite `box-shadow` tako da `offset-x` bude 25px, `offset-y` 10px, `blur-radius` 0, `spread-radius` 0 i `color` bude `blue`.

# --hints--

Vrednost `background-color` svojstva treba da bude `transparent`.

```js
assert.match(code,/background-color:\s*?transparent;/gi);
```

Vrednost `border-radius` svojstva treba da bude `50%`.

```js
assert.match(code,/border-radius:\s*?50%;/gi);
```

Vrednost `box-shadow` svojstva treba da bude 25px za `offset-x`, 10px za `offset-y`, 0 za `blur-radius`, 0 za `spread-radius` i na kraju `blue` za `color`.

```js
assert.match(code,/box-shadow:\s*?25px\s+?10px\s+?0(px)?\s+?0(px)?\s+?blue\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .center {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 0px;
    box-shadow: 25px 10px 10px 10px green;
  }

</style>
<div class="center"></div>
```

# --solutions--

```html
<style>
  .center {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    width: 100px;
    height: 100px;
    background-color: transparent;
    border-radius: 50%;
    box-shadow: 25px 10px 0 0 blue;
  }
</style>
<div class="center"></div>
```
