---
id: 587d781d367417b2b2512ac8
title: Prilagodi hover stanje anchor taga
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakRGcm'
forumTopicId: 301035
dashedName: adjust-the-hover-state-of-an-anchor-tag
---

# --description--

Ovaj izazov pokriva upotrebu pseudo-klasa. Pseudo-klasa je kljucna rec koja se moze dodati selektoru da bi se ciljalo posebno stanje elementa.

Na primer, stil anchor taga moze da se promeni u hover stanju koriscenjem pseudo-klasnog selektora `:hover`. Ovo je CSS kod koji menja `color` anchor taga u crvenu boju tokom hover stanja:

```css
a:hover {
  color: red;
}
```

# --instructions--

U editoru vec postoji CSS pravilo koje postavlja sve `a` tagove na crnu boju. Dodajte pravilo tako da kada korisnik predje misem preko `a` taga, `color` bude plava.

# --hints--

Boja anchor taga treba da ostane crna; dodajte samo CSS pravilo za `:hover` stanje.

```js
const anchorElement = document.querySelector("a"); 
const anchorStyle = window.getComputedStyle(anchorElement);
assert.equal(anchorStyle?.color, 'rgb(0, 0, 0)');
```

Anchor tag treba da ima plavu `color` vrednost tokom hover stanja.

```js
assert.match(code,
    /a:hover\s*?{\s*?color:\s*?(blue|rgba\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?,\s*?1\s*?\)|#00F|rgb\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?\))\s*?;\s*?}/gi
  );
```

# --seed--

## --seed-contents--

```html
<style>
  a {
    color: #000;
  }



</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```

# --solutions--

```html
<style>
  a {
    color: #000;
  }
  a:hover {
    color: rgba(0,0,255,1);
  }
</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```
