---
id: 587d781c367417b2b2512ac2
title: Podesite velicinu fonta za vise heading elemenata
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpQNT3'
forumTopicId: 301067
dashedName: set-the-font-size-for-multiple-heading-elements
---

# --description--

`font-size` svojstvo koristi se za odredjivanje velicine teksta u elementu. Ovo pravilo moze se primeniti na vise elemenata kako bi se dobio konzistentan izgled teksta na stranici. U ovom izazovu postavicete vrednosti za sve tagove od `h1` do `h6` kako biste uskladili velicine heading elemenata.

# --instructions--

U `style` tagovima postavite `font-size` za:

- `h1` tag na 68px.
- `h2` tag na 52px.
- `h3` tag na 40px.
- `h4` tag na 32px.
- `h5` tag na 21px.
- `h6` tag na 14px.

# --hints--

Vas kod treba da postavi `font-size` svojstvo za `h1` tag na 68 piksela.

```js
 const fontSizeOfh1 = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh1 ,'68px');
```

Vas kod treba da postavi `font-size` svojstvo za `h2` tag na 52 piksela.

```js
 const fontSizeOfh2 = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh2 ,'52px');
```

Vas kod treba da postavi `font-size` svojstvo za `h3` tag na 40 piksela.

```js
 const fontSizeOfh3 = new __helpers.CSSHelp(document).getStyle('h3')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh3 ,'40px');
```

Vas kod treba da postavi `font-size` svojstvo za `h4` tag na 32 piksela.

```js
 const fontSizeOfh4 = new __helpers.CSSHelp(document).getStyle('h4')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh4 , '32px');
```

Vas kod treba da postavi `font-size` svojstvo za `h5` tag na 21 piksel.

```js
 const fontSizeOfh5 = new __helpers.CSSHelp(document).getStyle('h5')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh5 ,'21px');
```

Vas kod treba da postavi `font-size` svojstvo za `h6` tag na 14 piksela.

```js
 const fontSizeOfh6 = new __helpers.CSSHelp(document).getStyle('h6')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh6 , '14px');
```

# --seed--

## --seed-contents--

```html
<style>






</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```

# --solutions--

```html
<style>
  h1 {
    font-size: 68px;
  }
  h2 {
    font-size: 52px;
  }
  h3 {
    font-size: 40px;
  }
  h4 {
    font-size: 32px;
  }
  h5 {
    font-size: 21px;
  }
  h6 {
    font-size: 14px;
  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```
