---
id: 58a7a6ebf9a6318348e2d5aa
title: Izmenite režim popunjavanja animacije
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDmcE'
forumTopicId: 301064
dashedName: modify-fill-mode-of-an-animation
---

# --description--

To je dobro, ali još uvek ne radi ispravno. Obratite pažnju kako se animacija vraća u prvobitno stanje nakon što prođe `500ms`, uzrokujući da dugme vrati na prvobitnu boju. Želite da dugme zadrži navedenu boju.

Ovo se može postići postavljanjem `animation-fill-mode` svojstva na `forwards`. `animation-fill-mode` određuje stil koji se primenjuje na element kada je animacija završena. Možete ga postaviti ovako:

```css
animation-fill-mode: forwards;
```

# --instructions--

Postavite `animation-fill-mode` svojstvo za `button:hover` na `forwards` kako bi dugme zadržalo boju kada korisnik pređe mišem preko njega.

# --hints--

`button:hover` treba da ima `animation-fill-mode` svojstvo sa vrednošću `forwards`.

```js
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-fill-mode\s*?:\s*?forwards\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-name\s*?:\s*?background-color\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-duration\s*?:\s*?500ms\s*?;[\s\S]*}/gi );
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    /* Only change code below this line */

    /* Only change code above this line */
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    animation-fill-mode: forwards;
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>