---
id: 6626a060c4006f793e10cb33
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Sada nastavniku je potrebna tvoja pomoć da promeni ocenu učenika u slovo/graduši.

Implementiraj funkciju `getGrade` koja prima broj `score` kao argument. Tvoja funkcija treba da vrati niz slova koji predstavlja veličinu slova u zavisnosti od znaka.

Ovde su znaci i njihovi mostovi za odgovarajuća slova:

| Područje oznaka | Klasa |
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"` |
| `60 - 69`   | `"D"` |
| `0 - 59`    | `"F"` |

**Napomene**

- Zapamtite se da ste naučili o uslovnim izjavama (`if`, `else if`, na `else`).
- Zapamtite se da ste naučili o operatorima poređenja (`>`, `<`, `>=`, `<=`, `===`).
- Zapamtite se da ste naučili o logičkim operatorima AND (`&&`).

# --hints--

Tvoja funkcija za `getGrade` treba da vrati `"A++"` ako je uslov `100`.

```js
assert.strictEqual(getGrade(100), "A++");
```

Vaša funkcija za `getGrade` treba da vrati `"A"` ako je uslov `94`.

```js
assert.strictEqual(getGrade(94), "A");
```

Tvoja funkcija za `getGrade` treba da vrati `"B"` ako je vrednost između `80` i `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Tvoja funkcija za `getGrade` treba da vrati `"C"` ako je status `78`.

```js
assert.strictEqual(getGrade(75), "C");
```

Vaša funkcija od `getGrade` treba da vrati `"D"` ako je vrednost između `60` i `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Tvoja funkcija za `getGrade` treba da vrati `"F"` ako je uslov `57`.

```js
assert.strictEqual(getGrade(57), "F");
```

Vaša funkcija za `getGrade` treba da vrati `"F"` ako je vrednost između `0` i `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
