---
id: 662693f82c91a66be46c881b
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

Nastavnik je završio ocenjivanje ispita svojih učenika i treba mu vaša pomoć da izračuna prosečne ocene klase.

Implementirajte funkciju `getAverage` koja prima red podataka o ocenama sa ispita i vraća prosečnu vrednost tih ocena.

Prosečna vrednost se računa sabiranjem svih rezultata, a zatim deljenjem sa ukupnim brojem dostupnih rezultata.

```md
average = sum of all scores / total number of scores
```

Nekoliko funkcija je pripremljeno da biste mogli da testirate svoj kod.

**Napomene**

- Možete koristiti petlju za iteriranje nad nizom podataka `scores` i dodavanje svih vrednosti.
- Možete koristiti svojstvo `length` za dobijanje ukupnog zbira postojećih vrednosti.

# --hints--

Tvoja funkcija za `getAverage` treba da vrati broj.

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` trebalo bi da vrati `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` treba da vrati `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` treba da vrati `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

Tvoja funkcija za `getAverage` treba da vrati prosečne ocene sa ispita.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
