---
id: 6626b4c58c027d86478ff5eb
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Učitelj je veoma srećan sa programom koji je kreirao do sada. Ali sada žele jednostavan način da provere da li učenik ima ocenu za prolazak. Ocena za prolazak je bilo šta što nije `"F"`.

Implementiraj funkciju `hasPassingGrade` koja prima ocenu učenika kao parametar. Tvoja funkcija bi trebalo da vrati `true` ako učenik ima pološnu ocenu i `false` ako nema.

**Napomene**

- Koristite funkciju `getGrade` da dobijete ocenu učenika. Zatim proverite da li je ocena prolazna ili ne.

# --hints--

Vaša funkcija `hasPassingGrade` treba da vrati vrednost tačno ili netačno.

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

Vaša funkcija za `hasPassingGrade` treba da vrati `true` ako je vrednost `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

Tvoja funkcija za `hasPassingGrade` treba da vrati `false` ako je uslov `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

Tvoja funkcija od `hasPassingGrade` treba da vrati `true` za sve prolazne ocene.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {
  
}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
