---
id: 68de40f53fcaf106f51cc20f
title: Kreiraj reproduktor za audio i video pomoću HTML-a
challengeType: 25
dashedName: napravi-html-audio-i-video-plejer
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Trebalo bi da ima element `h1` za glavni naslov stranice.
2. Trebalo bi da ima dva elementa `section`.
3. Unutarima elementa `section`, trebalo bi da ima element `h2` za naslov reprodukovanog videa.
4. Ispod elementa `h2`, trebalo bi da ima element `video` sa atributima `controls` i `width`. Atribut `width` mora biti postavljen na `640`.
5. Unutarima elementa `video`, trebalo bi da ima element `source` sa atributom `src` koji referencira video fajl i atributom `type`.
   - Možete koristiti `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.
6. Unutarima elementa `section`, trebalo bi da ima element `h2` za naslov reprodukovane pesme.
7. Ispod elementa `h2`, trebalo bi da ima element `audio` sa atributima `controls` i `loop`, i atributom `src` koji referencira audio fajl.
   - Možete koristiti `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`.
   - Ili `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`.

# --hints--

Trebalo bi da imaš element `h1` za glavni naslov stranice.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

`h1` Ne sme biti prazno.

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

Trebalo bi da ima samo jedan element `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Trebalo bi da ima dva elementa od `section`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

Unutar prvog elementa od `section`, treba da postoji polje za `h2` za naslov reprodukovanog videa.

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Ispod elementa `h2`, bi trebalo da postoji element `video`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

Element za `video` treba da ima atribut `controls`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

Element `video` treba da ima atribut `width` postavljen na `640`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

Unutar elementa za `video`, treba da ima element za `source`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

Funkcija za `source` mora da ima atribut `src` koji se odnosi na video fajl.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

Funkcija za `source` treba da ima svojstvo `type`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

Svojstvo `type` mora da odgovara medijskom tipu fajla svojstva `src`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

Unutar drugog elementa `section`, trebalo bi da postoji polje `h2` za naslov trenutno reprodukovane pesme.

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Svi elementi `h2` moraju imati neki tekstualni sadržaj.

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

Ispod elementa `h2` bi trebalo da postoji element `audio`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

Element `audio` mora imati svojstvo `controls`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

Element za `audio` mora imati svojstvo od `loop`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

Element za `audio` treba da ima svojstvo `src` koje se odnosi na audio fajl.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
