---
id: 680a3d4c065557260de2ff33
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovoj radionici, kreiraćeš program nazvan Caesarov šifar koji je metoda kodiranja koja menja slova u alfabetu kako bi sakrila poruku.

Počnite deklarisanjem parametra nazvanog `shift` i postavite vrednost `5` za vaš novi parametar.

# --hints--

Treba da objaviš parametar koji se zove `shift`. Obrati pažnju na postavljanje imena parametra na početak linije.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

Trebalo bi da postavite vrednost `5` za vaš parametar `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
