---
id: 680a82806c4c6bc3567f2708
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Objavi novi kriterijum nazvan `text` i postavi mu vrednost koja je niz slova `hello world`. Ovo će biti skrivena poruka.

# --hints--

Trebalo bi da objaviš parametar nazvan `text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

Trebalo bi da postavite vrednost sekvence karaktera `hello world` za `text`. Zapamtite da formatirate sekvencu karaktera sa pojedinačnim ili dvostrukim navodnicima i obratite pažnju na velika i mala slova.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
