---
id: 680a5e7f3d47715196b12ca0
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Sada nakon što ste prebacili sav vaš alfabet, uklonite poziv za `print(shifted_alphabet)`.

# --hints--

Treba da uklonite `print(shifted_alphabet)` iz svog koda.

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
print(shifted_alphabet)
--fcc-editable-region--
```
