---
id: 680a7e200400089967b8b1f4
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Funkcija ``str.maketrans()`` prima dva niza karaktera iste dužine i vraća tabelu prevoda koja prikazuje svaki karakter iz prvog niza odgovarajućim karakterom u drugom nizu. Svaki karakter u tabeli prevoda pohranjen je kao Unicode broj, broj koji jedinstveno identifikuje taj karakter.

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

Naučićete više o vrsti obrasca koji vraća `str.maketrans()` kasnije u kurikulumu. Trenutno, kreirajte tabelu prevoda koja prikazuje karakter iz `alphabet` za karakter iz `shifted_alphabet` i unesite je u parametar nazvan `translation_table`.

# --hints--

Trebalo bi da deklarišete parametar nazvan `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

Trebalo bi da pozoveš `str.maketrans()` prosljeđujući `alphabet` i `shifted_alphabet` kao argumente, i stavi rezultat u `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
