---
id: 680a431692a5f863f726c438
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

U ovoj radionici, kreiraćeš Caesarov šifar. Ovo je jedna od najjednostavnijih metoda enkripcije teksta, koja uključuje zamenu svakog slova običnog teksta slovom koje se nalazi na određenoj poziciji u alfabetu. Na primer, sa pomeranjem od 5, `a` će biti promenjeno u `f`, `b` u `g` i tako dalje.

Da biste implementirali ovaj šifar, biće vam potrebno da kreirate novu verziju vašeg alfabeta koja počinje na poziciji koju ukazuje promena. Kao što ste naučili u prethodnoj lekciji, možete izvući deo sekvence slova koristeći se string slicing:

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

Kreiraj varijablu nazvanu `shifted_alphabet` i koristi sintaksu za sečenje niza znakova da postavi deo `alphabet` koji dolazi iz primera (ili modela) `shift`. Zatim, pozovi ugrađenu funkciju (built-in) `print()` da ispiše `shifted_alphabet` na terminal i proveri rezultat.

# --hints--

Trebalo bi da objavite novi parametar nazvan `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

Trebalo bi da postaviš vrednost `alphabet[shift:]` za `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

Treba da pozoveš funkciju `print()` prosljeđujući `shifted_alphabet` kao argument.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
