---
id: 680a41ef5bfdc059cfd9d4b5
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Deklariši još jedan parametar nazvan `alphabet` i postavi vrednost za niz znakova `abcdefghijklmnopqrstuvwxyz` za ovaj parametar.

# --hints--

Trebalo bi da objavite parametar koji se zove `alphabet`. Obratite pažnju na postavljanje imena parametra na početak reda.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

Trebalo bi da postavite vrednost niz znakova `abcdefghijklmnopqrstuvwxyz` za svoj parametar `alphabet`. Zapamtite da koristite pojedinačne ili dvostruke navodne znakove za zatvaranje niza znakova.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
