---
id: 680a836235461fce1ff53b3f
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Funkcija `translate()` prima kao argument tabelu prevoda kreiranu od strane `str.maketrans()`. Poziva se na niz znakova i vraća kopiju originalnog niza gde su znaci promenjeni prema tabeli prevoda:

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

Izvrši poziv metode `translate()` na `text` prosleđujući `translation_table` kao argument i postavi rezultat u parametar nazvan `encrypted_text`.

# --hints--

Trebalo bi da ima parametar nazvan `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Treba da postaviš `text.translate(translation_table)` na `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
