---
id: 680a549f60f5a3fedfc0edd9
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Kao što vidite iz rezultata, pomaknuti alfabet počinje sa slovom `f` jer `shift` ima vrednost `5`. Ali sada prvih pet slova alfabeta – `a`, `b`, `c`, `d` i `e` – nedostaje u pomaknutom alfabetu, pa ćete morati da ih dodate na kraj pomaknutog alfabeta. Možete koristiti operater `+` da spojite nedostajući deo alfabeta.

Promeni postojeću dodelu vrednosti za `shifted_alphabet`: koristi sintaksu sečenja (slicing) da izvadi nedostajući prvi deo od `alphabet` i doda ga u `alphabet[shift:]`.

Kao podsetnik, `sentence[start:stop]` obnavlja znakove za `sentence` počevši od pozicije `start` (uključujući) do `stop` (isključujući).

# --hints--

Ne smete direktno menjati distribuciju vrednosti `alphabet`.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

Trebalo bi da postaviš vrednost spajanja `alphabet[shift:]` na prvi nedostajući deo `alphabet` u raspodeli vrednosti `shifted_alphabet` na istom redu kada je objavljeno. Koristi `shift` da prikažeš granicu rezanja.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
