---
id: 6810891bd6de8a87a833df42
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Kako vidite iz rezultata, poruka je sakrivena/enkodirana. Sledeći korak biće da učinite vaš kod ponovo upotrebljivim ukoliko želite da enkodirate drugu poruku.

Za to, biće vam potrebno kreirati funkciju. Kao podsetnik, ovako kreirate funkciju koja se zove `spam` i koja štampa `Spam!` u terminal:

```py
def spam():
    print('Spam!')
```

Kreiraj funkciju nazvanu `caesar`. Postavi sav kod koji si napisao unutar tela te funkcije. Obazi pažnju da sve linije unutar tela funkcije imaju isti nivo uvlačenja (indentacije).

# --hints--

Trebalo bi da ima funkciju nazvanu `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("caesar")`)) })
```

Treba da premestite sav kod koji ste do sada napisali unutar tela funkcije `caesar`. Uverite se da postavite isti nivo uvlačenja za sve linije unutar tela funkcije...

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_body().is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'\\nshift = 5\\nshifted_alphabet = alphabet[shift:] + alphabet[:shift]\\ntranslation_table = str.maketrans(alphabet, shifted_alphabet)\\ntext = 'hello world'\\nencrypted_text = text.translate(translation_table)\\nprint(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
print(encrypted_text)
--fcc-editable-region--
```
