---
id: 6818b80519e6c6c47a8bb6e8
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Kao što ste naučili u prethodnom kursu, funkcije mogu imati parametre, koji su parametri koje se mogu pozvati unutar funkcije. Evo funkcije `sum` sa dve parametre, `num1` i `num2`.

```py
def sum(num1, num2):
    print(num1 + num2)
```

Logika za skrivanje i izmene još uvek je direktno napisana unutar vaše funkcije. Uklonite dva parametra iz vaše funkcije koja se nazivaju `text` i `shift`, i obrišite deklaraciju parametara `text` i `shift` iz tela funkcije `caesar`.

# --hints--

Tvoja funkcija za `caesar` mora imati dva parametra, `text` i `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

Ne bi trebalo da ima parametar za `text` deklarisan unutar tela funkcije `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

Ne sme imati parametar za `shift` koji je deklarisan unutar tela funkcije `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
