---
id: 6818e9c7c8dc5694e8f64fc3
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Sada, pokušajte funkciju `caesar` pozivanjem sa nizom karaktera `freeCodeCamp` i brojem `3` kao argumentima. Umestite poziv te funkcije u parametar nazvan `encrypted_text`.

# --hints--

Trebalo bi da imaš parametar nazvan `encrypted_text` u globalnom opsegu.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Trebalo bi da postaviš vrednost `caesar('freeCodeCamp', 3)` u parametar `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
