---
id: 6818eb3f98c0bca5af7c8d03
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Vaš kod sada može biti ponovo korišćen. Međutim, funkcija `caesar` ispisuje poruku na terminal i vraća `None` podrazumevano. Da biste to potvrdili, ispišite `encrypted_text` na kraju svog koda.

# --hints--

Trebalo bi da ispišete `encrypted_text` izvan tela funkcije.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--

--fcc-editable-region--
```
