---
id: 6818ed1adc249bb91d373729
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Prethodnim lekcijama, naučili ste da deklaracija `return` se koristi za vraćanje vrednosti iz funkcije, kako biste mogli da ga koristite negde drugde u svom kodu:

```py
def spam():
    return 'Spam!'
```

Obriši poziv za `print(encrypted_text)` iz vaše funkcije. Zatim, obriši deklaraciju parametra za `encrypted_text` i direktno vrati `text.translate(translation_table)` iz vaše funkcije.

# --hints--

Ne smijete imati `print(encrypted_text)` unutar funkcije `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

Ne smete imati parametar `encrypted_text` unutar funkcije `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

Vaša funkcija za `caesar` treba da vrati `text.translate(translation_table)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
