---
id: 6818fcb250f34e62cab32c39
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Možda ste primetili da, iako je poruka skrivena, velika slova nisu dotaknuta. To se dešava zato što tabela prevoda ne uključuje velika slova.

Da biste ovo popravili, biće vam potrebno da promenite niz znakova koji se šalje za vaš poziv od `str.maketrans()` koristeći metodu `upper()`:

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

Ispravite vaš poziv za `str.maketrans()` povezujući ga sa svaku veliku slovo verziju tog samog argumenta.

# --hints--

Morate promeniti argumente vašeg poziva `str.maketrans()` povezivanjem na veliko slovo svakog samog argumenta.

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
