---
id: 6819ca464db219ab86f37dc4
title: Korak 20
challengeType: 20
dashedName: step-20
---

# --description--

Kao što ste već potvrdili, promene korišćene za skrivanje teksta moraju biti ceo pozitivan broj. Međutim, ne sme da prelaze `25` (poslednji primer od `alphabet`).

Dodaj drugi uslov u izraz za `if` koji proverava da li `shift` je veće od `25`. Zapamtite da logička operacija ILI u Python se dešava putem operatera `or`.

Pia, ažuriraj poruku koja se vraća na `Shift must be an integer between 1 and 25.`

# --hints--

Morate koristiti operater `or` da dodate drugi uslov nakon `shift < 1` koji proverava da li je `shift` veće od `25`. Ne smete menjati postojeći uslov.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

Tvoj drugi deo od `if` bi trebalo da vrati `Shift must be an integer between 1 and 25.`

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
