---
id: 6819c5c9faf2548e556163d5
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Kama zapamtite iz prethodnih lekcija, operater `not` je unarni logički operator koji odbacuje matematičke izraze:

```py
print(not True) # False
print(not False) # True
```

Koristite operater `not` da ispravite stanje deklaracije `if` ili funkcije koja vraća `Shift must be an integer value.` kada `shift` nije pun broj.

# --hints--

Trebalo bi da koristite operater `not` za negiranje `isinstance(shift, int)` u vašem `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
