---
id: 6819bc0637b80256a33adccc
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Sada ste implementirali osnovnu funkcionalnost enkodiranja, došlo je vreme da dodate autentifikaciju. Za to će vam trebati deklaracija `if`.

Na početku tela vaše funkcije, kreiraj deklaraciju za `if`. Trenutno, koristi `True` kao uslov, a unutar tela izjave `if` vrati niz karaktera `Shift must be an integer value.`.

# --hints--

Trebalo bi da imaš deklaraciju `if` unutar svoje funkcije `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

Tvoja izjava o `if` treba da koristi `True` kao svoj uslov.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Trebalo bi da vratiš niz slova `Shift must be an integer value.` iz svog izraza `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
