---
id: 6819cf38880d8cc9b9f6af7a
title: Korak 21
challengeType: 20
dashedName: step-21
---

# --description--

Python omogućava vam da postavite podrazumevanu vrednost za parametre u funkciji, kreirajući funkciju koja se može pozvati sa nekoliko argumenata ili bez ikakvih argumenata. Evo kako kreirati funkciju sa parametrom ``name`` koji ima podrazumevanu vrednost:

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

Dodaj treći parametar nazvan `encrypt` u svoju funkciju i sa podrazumevanom vrednošću `True`.

# --hints--

Trebalo bi da dodate treći parametar nazvan `encrypt` sa podrazumevanom vrednošću `True` u vašu funkciju. Ne menjajte postojeće parametre.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
