---
id: 6819be9055828d65c813eec5
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Funkcija za `isinstance()` vraća `True` ako je prvi argument primer drugog argumenta, i `False` inače:

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

Promeni trenutnu uslovnu deklaraciju `if` u poziv `isinstance()`. Dodaj `shift` kao prvi argument, i `int` kao drugi argument.

# --hints--

Izjava za `if` mora koristiti `isinstance(shift, int)` kao uslov.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if True:
--fcc-editable-region--
        return 'Shift must be an integer value.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
