---
id: 681a1d60e037127e7c8e0230
title: Korak 22
challengeType: 20
dashedName: step-22
---

# --description--

Koristi parametar dodan u prethodnom koraku da odluči da li treba funkcija da sakrije tekst poslat joj (standardna karakteristika, `encrypt=True`), ili da otvori poruku koja je ranije sačuvana.

Kreiraj deklaraciju za `if` koja proverava da li je `encrypt` netačno. Unutar nove deklaracije za `if`, postavi `shift` na `- shift`. Ovo je važno kako bi se omogućile promene u drugačijem smeru u vezi sa procesom skrivanja.

# --hints--

Trebalo bi postaviti `shift` na `- shift` kada `encrypt` nije tačno.

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
