---
id: 681a2a6cf95e82ca30866364
title: Korak 23
challengeType: 20
dashedName: step-23
---

# --description--

Objavi dva funkcionera koja se zovu `encrypt` i `decrypt`, oba uključujući parametre `text` i `shift`.

Koristi `encrypt` za proces kodiranja i `decrypt` za proces dekodiranja, izvršavajući oba ta dela pod opisnim imenom.

Pozovi funkciju `caesar` prosljeđujući `text` i `shift` iz obe vaše nove funkcije, ali obavezno prosljeđuj i `False` kao treći argument za funkciju `decrypt`.

# --hints--

Trebalo bi da imaš funkciju koja se zove `encrypt` sa dve parametre, `text` i `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

Tvoja funkcija za `encrypt` treba da vrati `caesar(text, shift)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

Trebalo bi da imaš funkciju koja se zove `decrypt` sa dva parametra, `text` i `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

Vaša funkcija za `decrypt` bi trebalo da vrati `caesar(text, shift, False)`.

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
