---
id: 66c8ba975ee7230e29f6c4af
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Ako je godišnji prihod podnosioca zahteva veći ili jednak `minIncomeForCondo`, I ako je njihov kreditni rejting veći ili jednak `minCreditScoreForCondo`, onda su kvalifikovani za kredit za kondo i auto.

Proverite da li je to tačno u funkciji `getLoanMessage`. Ako jeste, vratite string `"You qualify for a condo and car loan."`

# --hints--

Your `getLoanMessage` function should return a string if the applicant qualifies for a condo and car loan.

```js
assert.isString(getLoanMessage(45000, 680));
```

Your `getLoanMessage` function should return the string `"You qualify for a condo and car loan."`.

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

Your `getLoanMessage` function should return `undefined` if the applicant's annual income and credit score do not meet the requirements for a condo loan.

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
