---
id: 66c8ba41a77db20a93f9d7a1
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, pregledaćete JavaScript uslovne izraze kreiranjem aplikacije za proveru kvalifikacije za kredit.

Aplikacija će proveriti da li je korisnik kvalifikovan za kredit za dupleks, automobil i kondominijum na osnovu godišnjeg prihoda i kreditičkog rejtinga. Kreditički rejting je broj koji predstavlja koliko ste dobri u upravljanju pozajmljenim novcem.

Da biste počeli, kreirajte sledeće varijable i vrednosti.

| Variable Name | Value   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

You should create a `minIncomeForDuplex` variable set to `60000`.

```js
assert.equal(minIncomeForDuplex, 60000);
```

You should create a `minCreditScoreForDuplex` set to `700`.

```js
assert.equal(minCreditScoreForDuplex, 700);
```

You should create a `minIncomeForCar` variable set to `30000`.

```js
assert.equal(minIncomeForCar, 30000);
```

You should create a `minCreditScoreForCar` variable set to `650`.

```js
assert.equal(minCreditScoreForCar, 650);
```

You should create a `minIncomeForCondo` variable set to `45000`.

```js
assert.equal(minIncomeForCondo, 45000);
```

You should create a `minCreditScoreForCondo` variable set to `680`.

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
