---
id: 66c8ba975ee7230e29f6c4ad
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Da biste proverili za koji kredit je korisnik kvalifikovan na osnovu `annualIncome` i `creditScore`, morate koristiti izjavu `if/else if` ili ternarni operator direktno unutar funkcije `getLoanMessage`. Zatim ćete vratiti odgovarajuću poruku u bloku svakog uslova.

Počevši od kredita za dupleks stanovanje, proverite da li `annualIncome` je veće ili jednako od `minIncomeForDuplex` I ako je `creditScore` veće ili jednako od `minCreditScoreForDuplex`.

Ako je taj uslov tačan, onda je podnosilac prijave kvalifikovan za kredit za dupleks stanovanje i ostale kredite. Dakle, unutar provere, vratite string `"You qualify for a duplex, condo, and car loan."`

# --hints--

Your `getLoanMessage` function should return a string.

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

Your `getLoanMessage` function should return the string `"You qualify for a duplex, condo, and car loan."`.

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

Your `getLoanMessage` function should return `undefined` if the applicant's annual income and credit score do not meet the requirements for a duplex loan.

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```
 
# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
