---
id: 66c8ba975ee7230e29f6c4ac
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Kada je korisnik kvalifikovan za kredit, želite da mu prikažete poruku u konzoli.

Za to, kreiraćete funkciju unutar koje ćete imati neke provere koje će vratiti za koji kredit je podnosilac zahteva kvalifikovan.

Kreirajte praznu funkciju `getLoanMessage` sa parametrima `annualIncome` i `creditScore`.

# --hints--

You should create a `getLoanMessage` function

```js
assert.isFunction(getLoanMessage)
```

Your `getLoanMessage` function should have an `annualIncome` and `creditScore` as parameters.

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage` should be an empty function.

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
