---
id: 65f569725359e10d345bc52a
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Razvijaćete program koji računa i crta trajektoriju projektila na osnovu ugla, brzine i visine paljenja/lansiranja.

Počnite sa import `math`, koristićete ćete ga mnogo u ovom projektu jer ima važne metode kao što su `math.radians`, `math.cos`, `math.sin` i druge.

Takođe definiši ove varijable da imaju vrednost ubrzanja gravitacije Zemlje i neke posebne konstante koje će vam kasnije pomoći (kopiraj i zalepi za ove).

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

Koristi `import math` za import modul od `math`.

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

Trebalo bi da imaš varijablu `GRAVITATIONAL_ACCELERATION` sa vrednošću `9.81`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

Trebalo bi da imaš varijablu `PROJECTILE` sa vrednošću `"∙"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

Trebalo bi da imaš varijablu `x_axis_tick` sa vrednošću `"T"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

Trebalo bi da imaš varijablu `y_axis_tick` sa vrednošću `"⊣"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
