---
id: 65fd4bd84561a14d3e720061
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Atribut klase `__slots__` ima posebnu upotrebu u klasama Python. Deklarisanje `__slots__` i postavljanje njegove vrednosti na sekvencijalni niz karaktera sprečava kreiranje atributa članice da budu samo oni koji su prisutni u tom nizu. Takođe, sprečava kreiranje specifičnog atributa člana `__dict__` i omogućava pristup atributima članice efikasnije.

Trebalo bi da koristite parametar ``__slots__`` unutar klase kako biste definisali koje su svojstva element te klase: postavite vrednost za tuple ``__slots__`` koja sadrži niz od 3 karaktera, svaki jednak jednom od imena svojstava elemenata deklarisanih u ``__init__``.

# --hints--

Treba da postaviš vrednost za `__slots__` tuple koji uključuje `'__height'`, `'__speed'`, i `'__angle'`.

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
