---
id: 65fd4b3bedee044b4b957d33
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Kreiraj klasu za `Projectile` sa metodom `__init__` koja postavlja tu klasu koristeći tri argumenta: početnu brzinu, početnu visinu i početni ugao paljenja krila, u ovom redosledu.

Unutar puta ``__init__``, postavite ove tri vrednosti argumenata u svojstva tri privatna polja: ``__speed``, ``__height``, i ``__angle``. Omjer koji se pruža će biti u stepenima; međutim, treba ga sačuvati unutar u radianima. Da biste ovo postigli, koristite funkciju ``math.radians()`` za konverziju ugla iz stepeni u radijane kada ga postavljate u ``__angle``.

Korišćenje dve podvučnice ispred imena članice klase uzrokuje promenu imena (name mangling) u Python. To znači da ta svojstva nisu dostupna direktno izvan klase koristeći deklarisana imena, već moraju biti dostupna pod promenjenim imenima kao `ball._Projectile__height` ako su potrebna izvana, što ukazuje da su namenjena samo za interno korišćenje.

# --hints--

Trebalo bi da koristiš `class Projectile:` za kreiranje klase.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

Trebalo bi da imaš putanju `__init__` unutar klase `Projectile`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

Nova karakteristika klase `Projectile` treba da ima tri pravilno postavljena atributa.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
