---
id: 660400f02cf3a76795b82d4e
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Prvi faza inicijalizacije je put koji računa pravac projekta, što je ravna površina pređena od bacanja do mesta gde se projekat dodiruje zemlje.

Kreiraj putanju `__calculate_displacement`, koja ima samo `self` kao parametar, i vrati projektnu perspektivu.

Koristite sledeću formulu za izračunavanje vektora projekta:
\[ d = \frac{v \cdot \cos\theta \cdot \left(v\cdot\sin\theta + \sqrt{v^2 \cdot \sin^2\theta + 2 \cdot g \cdot h}\right)}{g} \]

Gde je $d$ smer, $v$ je početna brzina, $\theta$ je ugao, a $h$ je početna visina projekcije.
Za $g$, možete koristiti varijaciju `GRAVITATIONAL_ACCELERATION`.

Treba koristiti metode `math.cos()` i `math.sin()` za trigonometrijske funkcije i `math.sqrt()` za izračunavanje kvadratnog korena. Takođe treba znati da se $x^y$ piše kao `x ** y` u Python-u. Takođe, $\sin^2\theta$ znači da je vrednost dobijena iz sinusa zatim kvadrirana.

Zapamtite da za ime manglinga morate pozvati metodu kao `_Projectile__calculate_displacement` ako želite da testirate, ili ga koristiti izvan klase:

```py
ball = Projectile(10, 3, 45)
displacement_of_ball = ball._Projectile__calculate_displacement() # 12.6173996009878
```

# --hints--

Trebalo bi da objaviš put nazvan `__calculate_displacement` i `def __calculate_displacement(self):`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").has_function("__calculate_displacement")`))})
```

Putanja za `__calculate_displacement` treba da ima samo argument `self`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_args('self')`))})
```

Putanja `__calculate_displacement` treba da vrati tačnu vrednost.

```js
({test: () => assert(runPython(`
p = Projectile(20, 21, 22)
disp = p._Projectile__calculate_displacement()
round(disp, 2) == 55.06 and round(disp, 2) != disp
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    
--fcc-editable-region--
```
