---
id: 660fcf3e1b9bb056b2edb567
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Trenutno ste spremni da kreirate reprezentaciju sekvence karaktera.
Počnite kreiranjem instance klase `Projectile`. Unutar definicije klase, kreirajte parametar `ball` i postavite mu vrednost za poziv `Projectile` koristeći `10, 3, 45` kao argument.

Definišite metodu `__str__` unutar klase, kako bi vaš primer pružio ispravnu reprezentaciju niza karaktera, i da ga učinite tako da vrati niz sa sledećom strukturom:

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

Trebalo bi da počinje i završi sa karakterom novog reda, ugao mora biti napisan kao ceo broj u stepenima, a pomeraj treba da se ispiše sa jeddecimalnu tačku.

Pronaći ćete da je važno konvertovati uglove iz radijana u stepene. `math.degrees`

Kada budete spremni, možete `print(ball)` da testirate svoju funkciju.

# --hints--

Trebalo bi da imaš kriterijum za `ball`.

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

Parametar za `ball` mora imati vrednost `Projectile(10, 3, 45)`.

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

Representacija sekvence znakova za `Projectile(45, 45, 45)` treba da bude tačna.

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

Predstavljanje sekvence slova takođe mora biti tačno za druge primere.

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

Trebalo bi da imaš poziv za `print(ball)`.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
--fcc-editable-region--

--fcc-editable-region--
        
```
