---
id: 6616ec0bf9ee8548ce6b0f08
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Sada je vreme da se izračunaju koordinate putanje pravca! Kreiraj putanju klase `Projectile` nazvanu `__calculate_y_coordinate`, koja mora imati, osim `self`, parametar `x`.

\\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \\]

Ovo je formula za izračunavanje vertikalne pozicije $y$ za bilo koju horizontalnu poziciju $x$, uz početni ugao $\theta$, brzinu $v_0$ i visinu $y_0$.

Potrebno će vam je da koristite `math.tan()` i `math.cos()`, i zapamtite da `x ** y` je način za pisanje $x^y$, a vrednost g se nalazi u parametru `GRAVITATIONAL_ACCELERATION`.

Diferenciraj tu funkciju da bi vratila koordinatu $y$.

Kada želite da probate tu metodu, van klase, možete pozvati `ball._Projectile__calculate_y_coordinate()` i broj kao argument nakon što je `ball` objavljen.

# --hints--

Vaša klasa za `Projectile` mora da ima metodu pod nazivom `__calculate_y_coordinate`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

Putanja za `__calculate_y_coordinate` mora da ima `self, x` kao argument.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

Metoda bi trebalo da vrati tačnu vrednost. Za `Projectile(12, 12, 12)`, ako su dati četiri koordinate za $x$ za 0, 1, 2 i 3, koordinata y bi trebalo da bude približno 12, 12.18, 12.28, 12.32.

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

Metoda bi trebalo da vrati tačnu vrednost. Za `Projectile(45, 12, 22)`, ako su date x koordinata 6, 9, 10 i 12, y koordinata bi trebalo da bude otprilike 14.32, 15.41, 15.76, 16.44.

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

Putanja za `__calculate_y_coordinate` ne bi trebalo da obavija rezultate.

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
