---
id: 66335223f3e2a31d62d84367
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Vratite se na kalkulator projektnog pravca, sada ćete kreirati novu klasu objekata koja prima niz koordinata i generiše dijagram pravca.

Kreiraj novu klasu `Graph`, koja mora biti inicijalizovana sa privatnim atributom `__coordinates` gde se čuva lista koordinata. Deklariši klasni atribut `__slots__` i definiši specifične metode za `__init__` i `__repr__`.

Takođe, kreiraj primer od `Graph` prosleđujući vrednost od `coordinates` i stavi tu vrednost u novi parametar za `graph`.

# --hints--

Treba da kreiraš novu klasu za `Graph`.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_class('Graph')`
    ))
})
```

Trebalo bi da imaš atribut `__slots__` unutar klase.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_variable('__slots__')`
    ))
})
```

Kigezo cha `__slots__` mora imati vrednost od `('__coordinates')` ili `['__coordinates']`.

```js
({
    test: () => (runPython(
        `
        slots = _Node(_code).find_class('Graph').find_variable('__slots__')

        assert any(slots.is_equivalent(str) for str in ['__slots__ = ("__coordinates")', '__slots__ = ["__coordinates"]'])`
    ))
})
```

Trebalo bi da imaš putanju `__init__`.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_function('__init__')`
    ))
})
```

Putanja `__init__` treba da inicijalizuje svojstvo `__coordinates` sa prosleđenom vrednošću.

```js
({
    test: () => assert(runPython(
        `
        l = [(3, 4), (4, 1), (3, 9)]
        g = Graph(l)
        g._Graph__coordinates == l
        `
    ))
})
```

Trebalo bi da imate putanju `__repr__`.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_function('__repr__')`
    ))
})
```

Putanja za `__repr__` treba da vrati ispravnu sekvencu karaktera.

```js
({
    test: () => assert(runPython(
        `
        ball = Projectile(10, 3, 45)
        coords = ball.calculate_all_coordinates()
        repr(Graph(coords)) == f'Graph({coords})'
        `
    ))
})
```

Trebalo bi da kreirate primer `Graph` prosleđujući parametar `coordinates` i postavljajući ga u parametar `graph`.

```js
({
    test: () => {
        runPython(`

        assert _Node(_code).has_variable('graph'), "variable graph is missing"

        assert _Node(_code).find_variable('graph').is_equivalent('graph = Graph(coordinates)'), "assignment to graph is wrong"
        `)
    }
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.speed} m/s
height: {self.height} m
angle: {self.angle}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed

    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s

    def __repr__(self):
        return f"{self.__class__.__name__}({self.speed}, {self.height}, {self.angle})"

--fcc-editable-region--


ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()


--fcc-editable-region-- 
```
