---
id: 65577791ad8c26a7705e2919
title: Korak 34
challengeType: 20
dashedName: step-34
---

# --description--

Korišćenjem skraćenica rečnika, možete kreirati reчник baziran na postojećem rečniku:

```py
{key: val for key in dict}
```

U gornjem primeru, `val` je vrednost koja `key` će imati u novom rečniku, i `dict` je postojeći reчник.

Želite da pratite putanju između početnog čvora i svakog drugog čvora.

Nakon promene `distances`, kreirajte promenu `paths` i postavite mu vrednost na rečnik sa svim ključevima iz `graph`. Postavite prazan spisak za svaki ključ i koristite komprehenziju rečnika da izgradite svoj rečnik.

# --hints--

Trebalo bi da imaš promenu sa imenom `paths`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

Tvoje izmene za `paths` bi trebalo da koriste sintaksu skraćivanja reči kako bi postavile prazan niz za svaki čvor u grafu.

```js
({ test: () =>  {
    const {function_body} = __helpers.python.getDef(code, "shortest_path");       
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
