---
id: 65578d4fc3afc3b8f554c882
title: Korak 41
challengeType: 20
dashedName: step-41
---

# --description--

`min()` takođe uzima samo argument reči. Prosleđivanjem funkcije kao dodatni argument za `min()`, možete promeniti kako se stavke liste porede.

Rezultat linije koju ste crtali u prethodnom koraku je sledeći čvor po abecednom redosledu. Umesto toga, želite da izaberete neposetašnji čvor sa najmanjom udaljenošću od početnog čvora.

Koristi `key=distances.get` kao drugi argument za tvoj poziv `min()`. Na ovaj način, poređenje će biti izvršeno na osnovu vrednosti koju svaki element liste `unvisited` poseduje unutar rečnika `distances`.

# --hints--

Trebalo bi da proslediš `key=distances.get` kao drugi argument svom pozivu `min()`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
