---
id: 65576ff7888f9e96f52a4be1
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Mapa skladišti podatke u obliku parova *ključ*-*vrednost*. Ključ je odvojen od odgovarajuće vrednosti dvotačkom. A svaki par ključ-vrednost je odvojen od sledećeg para zarezom:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

Dodajte novu paru ključ-vrednost u svoj rečnik. Koristite niz znakova `'species'` kao ključ, i niz znakova `'guinea pig'` kao vrednost.

# --hints--

Trebalo bi da imaš ključe `'species'` sa vrednošću `'guinea pig'` unutar svog rečnika `copper`.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

Vaša mapa `copper` treba da ima samo jedan par ključ-vrednost.

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
