---
id: 65578cb031cd93b77a285db2
title: Korak 38
challengeType: 20
dashedName: step-38
---

# --description--

Vaša funkcija će pregledati sve čvorove povezane sa početnim čvorom. Izračunaće najkraći put za sve njih. Zatim, ukloniće početni čvor iz neposeteđenih čvorova.

Sledeći najbliži susedni čvor će biti poset procesom, koji će se ponavljati sve dok ne budu poseti svi čvorovi.

Od sada, radićeš na glavnoj petlji koja prolazi kroz čvorove u grafu. Da bi se izbegao problem pokretanja beskonačne petlje tokom razvoja algoritma, promeni poziv svoje funkcije u komentar.

# --hints--

Treba da promeniš poziv funkcije u komentar.

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
