---
id: 6557716aadbd2d9c42c0e69a
title: Korak 19
challengeType: 20
dashedName: step-19
---

# --description--

Poslednji čvor, `'D'`, koji je povezan sa `'A'` i `'C'`.

Izmenite svoj rečnik ili koristite ovaj format. Ponovo koristite listu za prikazivanje više konekcija.

# --hints--

Vaš rečnik mora imati četiri ključa pod nazivom `'A'`, `'B'`, `'C'`, na `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` Trebalo bi da bude lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["A"]) is list
  `))
})
```

`my_graph['A']` treba da bude spisak koji uključuje `'B'` i `'D'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["A"]) == 2 and "B" in my_graph["A"] and "D" in my_graph["A"]
  `))
})
```

`my_graph['B']` To bi trebalo da bude lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']` trebalo bi da bude lista koja uključuje `'A'` i `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']` Trebalo bi da bude lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["C"]) is list
  `))
})
```

`my_graph['C']` treba da bude lista koja uključuje `'B'` i `'D'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["C"]) == 2 and "B" in my_graph["C"] and "D" in my_graph["C"]
  `))
})
```

`my_graph['D']` Trebalo bi da bude lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["D"]) is list
  `))
})
```

`my_graph['D']` treba da bude spisak koji uključuje `'A'` i `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["D"]) == 2 and "A" in my_graph["D"] and "C" in my_graph["D"]
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': ['A', 'C'],
    'C': 'B'
}
--fcc-editable-region--
```
