---
id: 655771d889132f9ccd341060
title: Korak 20
challengeType: 20
dashedName: step-20
---

# --description--

Graf se naziva *graf sa težinama* kada su njegove ivice povezane sa težinom, koja predstavlja rastojanje, vreme ili neku drugu numeričku vrednost.

U vašem slučaju, ova težina će biti rastojanje između svakog čvora ili tačka u prostoru. Da biste predstavili graf sa težinama, možete promeniti svoj rečnik, koristeći listu tupola za svaku vrednost.

Prvi element u tuplu će biti povezani čvor, a drugi element će biti cijeli broj koji određuje udaljenost.

Transformišite `my_graph['A']` u listu tupela, uzimajući u obzir da je udaljenost od `A-B` `3` i da je udaljenost od `A-D` `1`.

# --hints--

`my_graph["A"]` bi trebalo da bude lista koja sadrži tuple `('B', 3)` i `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph` trebalo bi da ima 4 ključa sa imenima `'A'`, `'B'`, `'C'`, na `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
