---
id: 6557712d77ce2d9bd7e63afd
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Dodajte drugi čvor koji je povezan sa `B` na vaš graf i nazovite ga `C`.

Promenite svoj postojeći rečnik ili koristite ovu konfiguraciju: dodajte drugi ključ `'C'` do `my_graph` i dodelite mu vrednost niza karaktera `'B'`.

Pia, promeni vrednost postojećeg ključa `'B'` u listu `['A', 'C']` kako bi primio više konekcija čvorova `'B'`.

# --hints--

Vaš rečnik bi trebalo da ima 3 ključa — `'A'`, `'B'`, i na `'C'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

Vrednost `my_graph['A']` treba da bude niz karaktera `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` Trebalo bi da bude lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

Vrednost `my_graph['B']` bi trebalo da bude lista koja uključuje `'A'` i `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

Vrednost `my_graph['C']` bi trebalo da bude niz karaktera `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
