---
id: 65579228c669fcbebffd01d5
title: Korak 52
challengeType: 20
dashedName: step-52
---

# --description--

Druga greška je manja. Kada se pronađe najkraći put do susjednog čvora, `paths[current]` dobija vrijednost puta do susjednog čvora, `paths[node]`.

Ovo znači da oba kriterijuma ukazuju na istu listu. Pošto je lista dinamička, kada dodate susedni čvor na njen put, `paths[node]` i `paths[current]` se menjaju jer je to ista lista. Ovo uzrokuje netačne puteve (rute), iako je rastojanje tačno.

Možete popraviti tu grešku postavljanjem kopije `paths[current]` u putanju susjednog čvora. Stoga možete koristiti sintaksu za sečenje (slice):

```py
my_list[:]
```

Gde je `my_list` lista koju želiš da kopiraš. Promijeni raspored `paths[node] = paths[current]` koji se nalazi unutar tvog bloka `if` uklanjanjem `paths[current]`.

# --hints--

Trebalo bi da postaviš vrednost `paths[current][:]` na susednom čvoru (node).

```js
({ test: () => assert.match(code, /^(\s*)if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\s{4}paths\s*\[\s*node\s*\]\s*=\s*paths\s*\[\s*current\s*\]\s*\[\s*::?\s*\]/ms) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                    
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
