---
id: 655774955b097ea14897db12
title: Korak 28
challengeType: 20
dashedName: step-28
---

# --description--

Distanca od početnog čvora je nula, jer algoritam počinje svoju evaluaciju direktno odatle.

Nakon dodavanja `node` za `unvisited` u vašu kola, generišite izraz `if` koji se aktivira ako je čvor jednak početnom čvoru. Zatim postavite vrednost `0` za taj čvor unutar rečnika `distances`.

# --hints--

Trebalo bi da napravite izraz `if` koji se izvršava kada je `node` jednak `start`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

Unutar vaše nove deklaracije `if` morate postaviti vrednost `0` za čvor u rečniku `distances`.

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
