---
id: 655773f8b8b5db9fc6d0ae76
title: Korak 25
challengeType: 20
dashedName: step-25
---

# --description--

Za praćenje posetjenih čvorova, potrebna vam je lista svih čvorova u grafu. Kada se čvor obradi ili poseti, on se briše iz te liste.

Sada, promenite glavnu reč za `pass` i postavite na vrednost prazne liste sa imenom `unvisited`.

# --hints--

Trebalo bi da imaš parametar nazvan `unvisited` unutar funkcije `shortest_path`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("unvisited")`)) })
```

Trebalo bi da postaviš praznu vrednost za svoj parametar `unvisited`. Zapamti da obrišeš `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_body().is_equivalent("unvisited = []")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
