---
id: 6557709b0aee699a6a00528c
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Na primer, graf se može koristiti za predstavljanje dve tačke u prostoru, `A` i `B`, povezane putanjom. Takav graf je sastavljen od dva čvora povezana ivicom.

Promeni postojeći ključ `'species'` u niz karaktera `'A'`. Zatim, promeni vrednost koja odgovara nizu karaktera `'B'` ili predstavlja konekciju između čvorova `'A'` i `'B'`.

# --hints--

Vaš rečnik mora imati ključ `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Tvoj ključ za `'A'` treba da ima `'B'` kao vrednost.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Tvoj rečnik treba da ima dva ključa.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
