---
id: 655791e6cf5e03be3de73451
title: Korak 50
challengeType: 20
dashedName: step-50
---

# --description--

Ako pokušaš da ukloniš komentare iz poziva funkcije, neće raditi. Ima nekoliko grešaka koje treba ispraviti. Prva greška se javlja zato što u strukturiranim elementima unutar `if` pokušavaš da dobiješ element koji možda ne postoji u listi `paths[node]`. Stoga, moraš osigurati da `paths[node]` nije prazan pre nego što pristupiš `paths[node][-1]`.

Dodajte dodatni uslov u vaš izraz za `if`, da bi se osiguralo da `paths[node]` nije prazno pre dobijanja `paths[node][-1]`.

# --hints--

Morate dodati `paths[node]` kao prvi uslov u vašu izjavu od `if` koja je postavljena unutar. Koristite komparator `and` za povezivanje vaših uslova.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                   
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
