---
id: 6557743527cb92a06417ea97
title: Korak 26
challengeType: 20
dashedName: step-26
---

# --description--

Kreirujte ciklus od `for` koji se ponavlja preko vaše grafe, i koristite metodu `.append()` za dodavanje svakog čvora na kraj liste `unvisited`.

# --hints--

Treba da kreirate petlju za `for` koja iterira preko `graph` unutar funkcije `shortest_path`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "shortest_path");    
    assert(function_body.match(/^\s*for\s+\w+\s+in\s+graph\s*:/m));
  }
})
```

Trebalo bi da dodate svaki čvor na `unvisited` unutar vašeg kruga od `for`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const block_regex = /for\s+(\w+)\s+in\s+graph\s*/;
    const {block_body} = __helpers.python.getBlock(commentless_code, block_regex);
    const loop_condition = commentless_code.match(block_regex);
    const regex = new RegExp(`^\\s+unvisited\\.append\\s*\\(\\s*${loop_condition[1]}\\s*\\)`, "m");
    assert(block_body.match(regex));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
--fcc-editable-region--
```
