---
id: 655774d01daeeaa1978b99d5
title: Korak 29
challengeType: 20
dashedName: step-29
---

# --description--

U početku, svi ostali čvorovi u grafu se smatraju da su na beskonačnoj udaljenosti od izvornog čvora, jer još uvek nije definisana udaljenost.

Kreiraj polje `else` i postavi neograničenu vrednost za čvor u rečniku `distances`. Za to, koristi funkciju `float()` i niz karaktera `'inf'` kao argument radi generisanja decimalnog broja koji predstavlja pozitivan maksimum.

# --hints--

Trebalo bi da imate deo `else`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1else\s*:/ms));
  }
})
```

Treba da staviš `float('inf')` u `distances[node]` unutar svog novog bloka/sekcije `else`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, "else");
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*float\s*\(\s*("|')inf\1\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
--fcc-editable-region--
```
