---
id: 65577a17564ce8a8e06c1460
title: Korak 35
challengeType: 20
dashedName: step-35
---

# --description--

Leksikalno razumevanje takođe podržava sintaksu uslova za `if`/`else`:

```py
{key: val_1 if condition else val_2 for key in dict}
```

U gornjem primeru, `dict` je postojeći rečnik. Kada se `condition` proceni da je `True`, `key` će imati vrednost `val_1`, inače `val_2`.

Koristi rečnički kontekst da kreiraš rječnik zasnovan na `graph` i postavi njegovu vrijednost u varijablu `distances`. Vrijednost nuli za ključ ako je čvor jednak početnom čvoru, inače beskonačnost. Koristi `float('inf')` da to postigneš na kraju.

# --hints--

Trebalo bi da koristite sintaksu za razumevanje rečnika kako biste dodelili vrednost vašim promenama od `distances`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}distances\s*=\s*\{\s*(\w+)\s*:\s*0\s+if\s+(?:\1\s*==\s*start|start\s*==\s*\1)\s+else\s+float\s*\(\s*("|')inf\2\s*\)\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    paths = {node: [] for node in graph}
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
