---
id: 65577236b056379d5dbc7000
title: Korak 23
challengeType: 20
dashedName: step-23
---

# --description--

Sada ćete početi da razvijate algoritam za izračunavanje najkraćeg puta između svakog čvora u vašem novom grafu.

Objavi praznu funkciju koja se zove `shortest_path`. Popuni telo funkcije glavnom rečju `pass`.

# --hints--

Trebalo bi da objaviš funkciju nazvanu `shortest_path`.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
