---
id: 655773b0591c5f9f4045883e
title: Korak 24
challengeType: 20
dashedName: step-24
---

# --description--

Algoritam će početi na pomenutom čvoru. Zatim će pretraživati graf kako bi pronašao najkraći put između početnog čvora, ili *izvora*, i svih ostalih čvorova.

Dakle, vaša funkcija zahteva dva parametra: `graph`, i `start`. Uključite ih u deklaraciju vaše funkcije.

# --hints--

Tvoja funkcija mora da primi `graph` i `start` kao parametre, u ovom redosledu.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
