---
id: 65577739f57ecca6c39bb4e9
title: Korak 33
challengeType: 20
dashedName: step-33
---

# --description--

Komponenta tipa `list()` vam omogućava da napravite listu iz iterabilnog objekta.

Promeni raspodelu vrednosti tvojih promena `unvisited` koristeći `list()`, i ignoriši `graph` kao nešto što se može ponoviti.

# --hints--

Treba da koristiš `list()` za kreiranje liste iz rečnika `graph`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/list\s*\(\s*graph\s*\)/));
  }
})
```

Trebalo bi da postaviš vrednost `list(graph)` za `unvisited`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\s*=\s*list\s*\(\s*graph\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
