---
id: 65661b72d6745ebec6a96923
title: Korak 22
challengeType: 20
dashedName: step-22
---

# --description--

Na isti način, promeni dve preostale liste uzimajući u obzir da udaljenost za `C-D` je `7`.

# --hints--

`my_graph['C']` treba da bude lista koja sadrži tuplee od `('B', 4)` i `('D', 7)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 4), ("D", 7)]
    len(my_graph["C"]) == 2 and all(t in my_graph["C"] for t in tuples) 
  `))
})
```

`my_graph['D']` trebalo bi da bude lista koja sadrži tuplee iz `('A', 1)` i `('C', 7)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 1), ("C", 7)]
    len(my_graph["D"]) == 2 and all(t in my_graph["D"] for t in tuples)
  `))
})
```

`my_graph` trebalo bi da ima 4 ključa sa imenima `'A'`, `'B'`, `'C'`, na `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` mora biti lista koja sadrži tuplee sa `('B', 3)` i `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph['B']` treba da bude lista koja sadrži tuple od `('A', 3)` i `('C', 4)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}
--fcc-editable-region--
```
