---
id: 6566195b0a021bb660b2b4b1
title: Korak 21
challengeType: 20
dashedName: step-21
---

# --description--

Sada promeni `my_graph['B']` u listu tupola, gde je prvi element upotola povezani čvor, a drugi element je rastojanje. Rastojanje za `B-C` je `4`.

# --hints--

`my_graph['B']` treba da bude lista koja sadrži tuple od `('A', 3)` i `('C', 4)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

`my_graph` mora imati 4 ključa pod nazivima `'A'`, `'B'`, `'C'`, na `'D'`.

```js
({ test: () => assert(runPython(`    
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
