---
id: 6578b13757611e2825beb8a5
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Ključevi moraju biti jedinstveni unutar rečnika i mogu biti samo imutabilni tipovi podataka. To znači da ne možete koristiti listu ili drugi rečnik kao ključeve.

Dodajte još jedan ključ `'age'` u svoj rečnik i dodelite mu kompletan broj `2` kao vrijednost.

# --hints--

Morate imati ključ `'age'` sa vrednošću `2` odmah nakon `'species': 'guinea pig'`, unutar vašeg rečnika za `copper`. Ne zaboravi zarez.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Trebalo bi da imaš par od dve ključno-vrednost u tvom rječniku `copper`.

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
