---
id: 6578b57361f2f132a02e2a18
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Možete dobiti pohranjene podatke u rječniku putem njegovih ključeva:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

Nakon vašeg rečnika, pratite primer gore kako biste dobili vrednost ključeva `'species'` za `copper` i ispišite rezultate.

# --hints--

Ne smete menjati distribuciju vrednosti vašeg rečnika.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Trebalo bi da koristiš `copper['species']` da dobiješ vrednost ključa `'species'`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

Trebalo bi da pozoveš `print()` kada prosleđuješ `copper['species']` kao argument.

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
