---
id: 65789506b30453080f77470c
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Sad, već si naučio različite vrste podataka:

- Tipovi podataka koji su nepromenljivi, kao što su integeri, stringovi, tuple i booleanske vrednosti.
- Tipovi podataka koji su promenljivi, kao što su liste i rečnici.

Identifikuje se par uglastih zagrada, `{}`.

Počnite kreiranjem parametra nazvanog `copper` i postavite mu vrednost praznog rečnika korišćenjem vitičastih zagrada, na isti način na koji kreirate prazan list koristeći uglaste zagrade.

# --hints--

Trebalo bi da imaš parametar koji se zove `copper`.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

Tvoj parametar `copper` mora imati praznu vrednost rečnika. Koristi par vitičastih zagrada za to.

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
