---
id: 6579717f0920131304286804
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Da biste dodali novi par ključ-vrednost nakon definisanja rečnika, možete pristupiti ključu na isti način kao što dohvatate postojeći ključ i postaviti vrednost novog ključa koristeći operator dodeljivanja:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

Izbriši svoj poziv za `print()`. Zatim, nakon deklaracije `copper`, dodaj ključ `'food'` u svoj rečnik i postavi mu vrednost na `'hay'`.

# --hints--

Ne bi trebalo da imaš `print(copper['age'])` u svom kodu.

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

Treba da dodate ključ `'food'` na `copper` nakon izgovora rečnika.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

Treba da staviš `copper['food']` na mesto `'hay'` nakon izgovaranja rečnika.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
