---
id: 6579dd49fa8a8e1fd06b85a9
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Grafovi su strukture podataka koje predstavljaju odnose između parova elemenata.
Ovi elementi, nazvani *čvorići* (*nodes*), mogu biti stvarni predmeti iz života, kontejneri, tačke u prostoru ili nešto drugo. Povezivanje između čvorova naziva se *preseci* (*edges*).

Ovde je vizuelna reprezentacija crteža:

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Promeni ime rečnika `copper` u `my_graph`. Ovo će predstavljati crtež za testiranje vašeg algoritma.

# --hints--

Trebalo bi da promeniš ime rečnika `copper` u `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

Tvoj parametar za `my_graph` treba da bude rečnik.

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
