---
id: 6579cbab9825b8170974c69a
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Možete obrisati par ključ-vrednost iz rečnika koristeći glavnu reč `del`:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

Pre vašeg ciklusa od `for`, koristite ključnu reč od `del` da obrišete ključ za `'age'` zajedno sa njegovom vrednošću iz `copper`.

# --hints--

Trebalo bi da koristiš glavnu reč `del` za brisanje `copper['age']` pre ciklusa `for`.

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
