---
id: 65796fac81f983127558f3f4
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Sada, izmijenite vaš postojeći poziv za `print()` ili unesite vrijednost ključa za `age`.


# --hints--

Trebalo bi da koristite `copper['age']` da biste dobili pristup vrednosti ključa `age`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')age\1\s*\]/) })
```

Trebalo bi da pozoveš `print()` sa `copper['age']` kao argumentom.

```js
({ test: () => assert.match(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['species'])
--fcc-editable-region--
```
