---
id: 6579ca0923cfa7162089d2f0
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Ista sintaksa se može koristiti za promenu vrednosti postojeće ključeve.

Neposredno pre poziva `print()`, dobiti pristup ključu za `'species'` i postaviti mu vrednost na `'Cavia porcellus'` ponovo.

# --hints--

Trebalo bi da postavite vrednost `copper['species']` kao niz karaktera `'Cavia porcellus'` pre poziva `print()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
