---
id: 65797670e0c0d016f17e7660
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Sada, odmah ispod vašeg koda, ubacite `copper`.

# --hints--

Treba da zalepiš `copper` na sam kraj svog koda.

```js
({ test: () =>
  {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /print\s*\(\s*copper\s*\)\s*$/)
  }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

--fcc-editable-region--
```
