---
id: 65affe602a2a828b0382a058
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Promenite ključ `'age'` na niz znakova `'B'` i postavite mu vrednost na niz znakova `'A'` kako bi se predstavila konekcija između oba čvora u oba smera.

# --hints--

Vaš rečnik mora da ima ključ `'B'`.

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

Vaš ključ za `'B'` bi trebalo da ima `'A'` kao vrednost.

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

Vaš rečnik treba da ima ključ za `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Vaš ključ za `'A'` treba da ima `'B'` kao vrednost.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Tvoj rečnik treba da ima dva ključa.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
