---
id: 65b7d3319ebec69b983fb91b
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Kao što možete videti iz rezultata, `.items()` stvara strukturu podataka koja čuva svaki par ključ-vrednost u zasebni tuple. Da biste iterirali kroz elemente tih tuple, možete dodati drugi parametar petlje:

```py
for i, j in dict.items():
    print(i, j)
```

Promenite svoj ciklus `for` ili uzmite dva parametra ciklusa i ispišite sve unutar tela ciklusa.

# --hints--

Trebalo bi da promeniš svoj ciklus od `for` koristeći dva parametra.

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

Treba da objavite sve parametre kruga unutar vašeg kruga od `for`.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
