---
id: 65b7cd2b7bd9a684ccf1dc16
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Da biste iterirali po ključevima rečnika, možete direktno staviti rečnik u petlju `for`. Kod ispod će ispisati svaki ključ iz rečnika `dict`:

```py
for i in dict:
   print(i)
```

Promeni poziv `print()` u petlju `for` koja ponavlja preko `copper` i štampa svaki ključ.

# --hints--

Ne bi trebalo da imaš `print(copper)` u svom kodu.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(copper)")`)) })
```

Trebalo bi da kreirate ciklus `for` ponavljanja preko rečnika `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper")`)) })
```

Morate ispisati svaki ključ za `copper` unutar vašeg ciklusa od `for`.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print({var})")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
