---
id: 65b7d25992879698180e6a71
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Na kraju, ako želite da imate mogućnost iteriranja kroz parove ključ-vrednost, možete koristiti metod `.items()`.

Promeni svoj ciklus od `for` da se ponavlja preko `copper.items()` umesto `copper.values()`.

# --hints--

Treba da promeniš svoj ciklus od `for` na ponavljanje preko `copper.items()` umesto `copper.values()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.items()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.values():
    print(i)
--fcc-editable-region--
```
