---
id: 65b7cf140d34058d7ea8935f
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Ako želite da ponovite vrednosti ključa rečnika, jedan način je korišćenje `.values()`.

Promeni svoj ciklus od `for` tako da ponavlja preko `copper.values()` umesto `copper` da vidi rezultate.

# --hints--

Trebalo bi da promeniš svoj ciklus `for` da se ponavlja preko `copper.values()` umesto `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
