---
id: 68de7360e1832ea180577a8a
title: Implementiraj kalkulator za N-ti Fibonacci broj
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba kreirati funkciju sa imenom `fibonacci`.
2. Treba definisati listu sa imenom `sequence` unutar funkcije `fibonacci`, i ona mora biti inicijalizovana vrednostima `[0, 1]`.
3. Funkcija `fibonacci` mora primiti jedan parametar, neujemni ceo broj `n`.
4. Pozivanje `fibonacci(n)` mora koristiti metodu dinamičkog programiranja (dynamic programming) za izračunavanje i vraćanje `n`-og broja iz Fibonacijove sekvence, gde je svaki broj zbir prethodnih dva broja.
5. Svaki izračunati Fibonacijov broj mora biti dodan u listu `sequence`.

# --hints--

Trebao bi da ima funkciju sa imenom `fibonacci`.

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

Tvoja funkcija `fibonacci` treba da primi jedan parametar.

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

Trebalo bi da imate listu sa imenom `sequence` unutar funkcije `fibonacci` koja je pokrenuta za `[0, 1]`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

`fibonacci(0)` trebalo bi da vrati `0`.

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

`fibonacci(1)` treba vratiti `1`.

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

`fibonacci(2)` treba da vrati `1`.

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

`fibonacci(3)` treba da vrati `2`.

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

`fibonacci(5)` treba da vrati `5`.

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

`fibonacci(10)` treba vratiti `55`.

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

`fibonacci(15)` treba da vrati `610`.

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

Nije dozvoljeno koristiti recursion u vašem kodu.

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]
```
