---
id: 65c4f0b82db314d3fcc0b8e1
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Unutar klase `BinarySearchTree`, promenite `pass` koristeći `__init__` i dodajte parametar `self` na ovaj način.

# --hints--

Trebalo bi da uklonite reč `pass` iz klase `BinarySearchTree`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```

Morate definisati putanju `__init__` i dodati parametar `self` na ovaj način. Zapamtite da koristite termin/reč `pass` unutar tela putanje.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
        
--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
