---
id: 65c4f07da204f4d2f325fbd3
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Unutar putanje `__init__`, inicijalizujte svojstva elemenata za `left` i `right` za čvor da budu `None`. To je zato što čvor kada se prvi put kreira, nema levog niti desnog potomka. Zapamtite da koristite glavnu reč/ključ za `self`.

# --hints--

Treba da postavite vrednost `None` za atribut elementa `left` čvora koristeći `self.left`.

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

Treba da postavite vrednost `None` za atribut elementa `right` čvora koristeći `self.right`.

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
