---
id: 65c4f0d32ca17ad4b1636b0e
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Unutar putanji za `__init__`, ukloni `pass` i inicijalizuj `root` sa vrednošću `None`.

Atribut polja `root` predstavlja koren čvor binarnog stabla pretrage. Pošto je ovo komponenta kada se kreira novi objekat `BinarySearchTree`, on počinje sa praznim stablom, pa je atribut polja `root` konfigurisan da bude `None`.

# --hints--

Treba da uklonite deklaraciju `pass` iz puta `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Treba da inicijalizujete svojstvo elementa `root` u `None` koristeći `self.root`.

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
