---
id: 65c4f2b7178afed8e88f782f
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Ukoliko `key` < `node.key` se potvrdi da `True`, onda novi čvor bi trebalo postaviti u levi pod-stabak.

Izbriši `pass` i pozovi metodu/funkciju `_insert` za rekurzivno izvršavanje koristeći levog potomka kao prvi argument, a `key` kao drugi argument. Postavi vrednost rezultata u atribut elementa `left` trenutnog čvora.

# --hints--

Moraš da obrišeš glavnu reč `pass` iz bloka `if`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

Trebalo bi da pozoveš putanju `self._insert` kada prosleđuješ `node.left` i `key` kao argumente.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

Trebalo bi da postaviš vrednost rezultata tvog poziva `self._insert()` u atribut elementa `left` trenutnog čvora.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
