---
id: 65c4f2851645e8d84d14f5d1
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Sada je potrebno preći kroz binarno stablo rekurzivno i umetnuti vrednosti koristeći pravila binarnih stabala:

- Vrednost manja od ključa se postavlja u levo poddrvo
- Vrednost veća od ključa se postavlja u desno poddrvo

Nakon vašeg izveštaja o if koji postoji, napiši drugi izveštaj za `if` da bi proverio/da vidiš da li je `key` manje od `node.key`.

# --hints--

Trebali biste napisati drugi izveštaj za if ili proveriti da li `key` je manje od `node.key`. Zapamtite da koristite glavni termin za `pass`.

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```
