---
id: 65c4f37ba6cbcfdc77e24165
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Putanja za `insert` će biti nazvana od strane korisnika. Pored parametra `self`, takođe će biti potreban i parametar `key`. Ovaj parametar će biti vrednost ključa za umetanje u binarno stablo pretrage.

Ongeza `key` kao drugi parametar u definiciji funkcije.

# --hints--

Putanja za `insert` bi trebalo da ima parametar `self`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

Putanja za `insert` treba da ima parametar za `key`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
