---
id: 65c4ef8463f869d0eaf87c67
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

U klasi `TreeNode`, promeni `pass` u metodu `__init__` ili da može inicijalizovati svojstva objekta. Ne dodaj parametre za sada.

# --hints--

Treba da uklonite glavni element `pass` iz klase `TreeNode` i premestite ga unutar metode `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

Treba da definišete putanju `__init__` unutar vaše klase `TreeNode`. Setite se da koristite glavnu reč `pass` unutra.

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
