---
id: 65c4f300da28d8da361bfa93
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Unutar tvog bloka `elif`, promeni `pass` u poziv metode `_insert` sa desnim potomkom trenutnog čvora kao prvim argumentom i `key` kao drugim argumentom. Postavi vrednost rezultata na atribut elementa `right` trenutnog čvora.

# --hints--

Trebalo bi da uklonite glavnu reč `pass` iz bloka `elif`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Trebalo bi da pozoveš putanju `self._insert()` kada prosleđuješ `node.right` i `key` kao argumente.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

Treba da postavite vrednost rezultata vašeg poziva `self._insert()` u odgovarajuće svojstvo elementa (`node.right`) trenutnog čvora.

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
