---
id: 65c4f09e074dd8d37830ea00
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Prazna klasa za druge stvari se zove `BinarySearchTree`, koja predstavlja binarno stabljasto pretraživanje.

# --hints--

Trebalo bi da imaš klasu stvari koja se zove `BinarySearchTree`. Zapamti da koristiš ključnu reč `pass`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
