---
id: 65c4f013851cefd1a4fe4c96
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Putanja `__init__` uzima dva parametra: `self` (koji predstavlja instancu kreiranog objekta) i `key` (vrednost sačuvana u čvoru). Dodaj ova dva parametra na putanju `__init__()`.

# --hints--

Trebalo bi da dodaš parametar `self` na putanju.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

Trebalo bi da dodaš parametar `key` u putanju.

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
