---
id: 65c4eb814cc977c95cd7df0e
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovom projektu, kreiraćete binarno stablo pretrage (BST). BST je struktura podataka gde svaki čvor ima najmanje dva potomka: levi potomek mora imati vrednost manju od roditeljskog čvora, a desni potomek mora imati vrednost veću od roditeljskog čvora. Ovo omogućava efikasnu performansu za pretraživanje i sortiranje.

Ovo je prikaz binarnog staba za pretragu:

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Počnite sa deklarisanjem prazne klase objekata `TreeNode`. Klasa objekata `TreeNode` predstavlja čvor u binarnom stablu pretrage. Koristite ključnu reč `pass` da popunite telo klase i izbegnete grešku.

# --hints--

Trebalo bi da koristiš ključnu reč `class` za deklarisanje praznog objekta sa imenom `TreeNode`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
