---
id: 65c4f33bc3c3d8db7f732787
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Sada, da izvrši stvarnu injekciju, definišite praznu metodu za `insert` unutar klase objekata `BinarySearchTree` i dodelite joj parametar `self`.

# --hints--

Treba da definišete metodu za `insert` sa parametrom `self` unutar klase za `BinarySearchTree`. Zapamtite ključnu reč `pass`.

```js
({ test: () => assert.match(code, /^\s{4}def\s+insert\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node
--fcc-editable-region--    

--fcc-editable-region--

```
