---
id: 65c4f195de7e2ad5932be717
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Sledeće, morate definisati metodu ubacivanja čvora u drvo. Za to, morate definisati putanju `_insert`, koja je pomoćna funkcija i koja će biti korišćena od strane stvarnog puta za `insert` kasnije.

Ovaj put je rekurzivan, što znači da se sam poziva kako bi prošao kroz drvo sve dok ne pronađe odgovarajuće mesto za novi čvor.

Definiši putanju `_insert` sa parametrima `self`, `node` na `key`.

# --hints--

Trebalo bi da definišete putanju `_insert` unutar klase objekata `BinarySearchTree`. Setite se da koristite `pass`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

Vaša putanja za `_insert` treba da primi tri parametra: `self`, `node` i `key`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
