---
id: 65c4f02dbd3557d23f12db62
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Unutar putanje `__init__`, obriši `pass` i postavi vrednost parametra `key` na svojstvo elementa čvora `key` koristeći `self.key`.

To znači da će svojstvo polja `key` od primjera `TreeNode` imati vrijednost koja je proslijeđena prilikom kreiranja tog objekta.

# --hints--

Morate obrisati deklaraciju `pass` sa putanje `__init__`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

Trebalo bi da postaviš vrednost parametra `key` na atribut elementa `key` čvora koristeći `self.key`.

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
