---
id: 65c4f22498d22ed775ef8efb
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Sada treba proveriti da li atribut za `node` je `None`. Ako je to tačno, to znači da staza stiže do listnog čvora ili prazne tačke u stablu gde treba umetnuti novi čvor.

U telu putanje ``_insert``, promenite ``pass`` sa izrazom ``if`` koji proverava da li ``node is None``.

Unutar novog bloka od `if`, vratite `TreeNode(key)` ili kreirajte novu instancu `TreeNode` koristeći navedeni ključ. Ovo će biti novi listni čvor, i efikasno će umetnuti ključ u drvo.

# --hints--

Ne bi trebalo da imaš `pass` na tvom putu ka `_insert`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").has_pass()`)) })

```

Trebalo bi da napišete izraz `if` ili proverite da li je `node is None`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}if\s+node\s+is\s+None\s*:/m));
  }
})
```

Treba ti vratiti `TreeNode(key)` iz bloka `if`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^(\s{8})if\s+node\s+is\s+None\s*:\s*^\1\s{4}return\s+TreeNode\s*\(\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        pass
--fcc-editable-region--
```
