---
id: 65c63df529bd15a24c187c62
title: Korak 21
challengeType: 20
dashedName: step-21
---

# --description--

Napiši još jednu deklaraciju za `if` koja proverava da li je ciljni ključ manji od ključa trenutnog čvora.

Unutar bloka `if`, vrati rezultat pozivanja metode `_search` za levi potomak trenutnog čvora i `key` kao argument.

# --hints--

Morate napisati još jednu deklaraciju za `if` ili proverite da li `key` < `node.key`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

Treba da vratiš `self._search(node.left, key)` iz svog novog bloka `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_bodies()[0].is_equivalent("return self._search(node.left, key)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node

--fcc-editable-region--
```
