---
id: 65c6409418806da8d0636ffc
title: Korak 26
challengeType: 20
dashedName: step-26
---

# --description--

Zadaci za `insert` na `search` su završeni, vreme je da ih isprobate.

Zapamtite da čvorovi trenutno nisu usklađeni i samo su uneti. Funkciju usklađivanja ćete obavljati koristeći in-order pretragu kasnije.

Možete kreirati primer klase objekata u Python, kao što sledi:

```python
object_name = ClassName()
```

Izvan definicije klase, kreiraj instancu klase `BinarySearchTree` i postavi njenu vrednost u promenljivu `bst`.


# --hints--

Treba da kreiraš primer klase `BinarySearchTree`.

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

Trebalo bi da postavite novu vrednost za `BinarySearchTree` u izmene `bst`.

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
