---
id: 65c63da2ed6769a10e141341
title: Korak 20
challengeType: 20
dashedName: step-20
---

# --description--

Sada ćete definisati osnovni slučaj za rekurzivno pretraživanje. Uklonite postojeći `pass` i napišite izraz `if` koji proverava dva uslova:

- Ako `node` je `None`: Ovo ukazuje da pretraga stigla do kraja grane bez pronalaženja ključeva.
- Ako `node.key == key`: To znači da je ključ pronađen u trenutnom čvoru.

Povežite ta dva uslova koristeći operator `or` i vratite trenutni čvor unutar bloka `if`.

# --hints--

Treba da uklonite glavnu reč `pass` iz puta `_search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

Treba da napišete izraz koji koristi `if` i proverava da li `node` je `None` ili ako `node.key` je jednak `key`.

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

Trebalo bi da vratiš `node` iz bloka `if`.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
