---
id: 65c645b838e7deb080fc25e0
title: Korak 29
challengeType: 20
dashedName: step-29
---

# --description--

Sada su uneti čvorovi. Da biste proverili da li su tačno uneti, možete potražiti čvorove u drvetu.

Izvan ciklusa for, pronađi čvor `80` u primeru `bst` i dodaj ga u poziv `print`. Takođe, prvi argument funkcije `print` bi trebalo da bude `'Search for 80:'`.

# --hints--

Trebalo bi da ispišete rezultate poziva `bst.search(80)` i vaša deklaracija od `print` treba da ima prvi argument `'Search for 80:'`.

```js
assert.match(code, /^print\(('|")Search for 80:('|"),\s*bst\.search\(80\)/gm);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

--fcc-editable-region--
```
