---
id: 65c63a4c4da62e9ae18e321a
title: Korak 19
challengeType: 20
dashedName: step-19
---

# --description--

Vreme je da radimo na funkcionalnosti pretrage. Ako ste kreirali pomoćnu putanju `_insert` putem `insert`, morate kreirati pomoćnu putanju `_search` putem `search`.

Objavi put za `_search` sa tri parametra, tj. `self`, `node` i `key`.

# --hints--

Treba da objaviš putanju `_search`. Zapamti donji red i zapamti koristiti `pass`.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

Tvoj put ``_search`` bi trebalo da uzme tri parametra: ``self``, ``node``, i ``key``.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
