---
id: 65c63e6962159fa372ecdbec
title: Korak 22
challengeType: 20
dashedName: step-22
---

# --description--

Ako druga deklaracija za `if` nije `True`, to znači da je ciljni ključ veći ili jednak ključu trenutnog čvora.

U binarnom stablu za pretragu, ako je ciljni ključ veći od ključa trenutnog čvora, pretraga se nastavlja u desnom podstablu.

Nakon komada `if`, vratite rezultat poziva metode `_search` za desnog potomka trenutnog čvora i ključ kao argument.

# --hints--

Trebalo bi da pozoveš putanju `_search` za ponavljanje korišćenjem `self._search()`.

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

Morate proslediti dva argumenta za poziv metode `_search`: `node.right` i `key`.

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

Treba da vratiš rezultat pozivanja metode `_search`.

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
