---
id: 65c63fd9b387f0a6c5bf6a72
title: Korak 24
challengeType: 20
dashedName: step-24
---

# --description--

Unutar rute `search`, izbriši `pass` i pozovi pomoću asistenta `_search` sa sledećim parametrima.

- `self.root`: Ovo je koren binarnog stabla pretrage. Pretraga počinje od korena.
- `key`: Ovo je vrednost koju korisnik želi da pronađe u binarnom stablu pretrage.

Unutar, `search` predaje stvarni logiku pretrage putem asistenta `_search` koja izvodi rekurzivnu pretragu unutar binarnog stabla pretrage.

# --hints--

Trebalo bi da uklonite reč `pass` iz puta `search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

Trebalo bi da pozoveš putanju `_search` unutar putanje `search`.

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

Trebalo bi da pozoveš putanju `_search` kada prosleđuješ `self.root` i `key` kao argumente.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
