---
id: 65ca06864129380054b55dc0
title: Korak 36
challengeType: 20
dashedName: step-36
---

# --description--

Unutar bloka `if`, zameni `pass` pozivom metode `_delete`, prosljeđujući levi potomak trenutnog čvora i `key` kao argumente. Postavi poziv funkcije na levi čvor.

# --hints--

Trebalo bi da uklonite glavnu reč od `pass` iz izjave `if`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_delete").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

Trebalo bi da pozoveš putanju `_delete` sa `node.left` i `key` kao argumente.

```js
assert.match(code, /self\._delete\(\s*node\.left\s*,\s*key\s*\)/);
```

Trebalo bi postaviti vrednost rezultata poziva ``_delete()`` za levi potomak (``node.left``) trenutnog čvora.

```js
assert.match(code, /node\.left\s*=\s*self\._delete\(\s*node\.left\s*,\s*key\s*\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            pass

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
