---
id: 65d8a6fcb15a3a239ba35dfd
title: Korak 46
challengeType: 20
dashedName: step-46
---

# --description--

Nakon pronalaženja minimalne vrednosti, moraćete da obrišete čvor sa minimalnom vrednošću iz desnog poddrva rekurzivno.

Ovaj korak osigurava da se čvor sa minimalnom vrednošću briše iz stabla uz održavanje svojstva binarnog stabla pretrage (BST).

Izvrši poziv puta `_delete` za ponavljanje korišćenjem `node.right` i `node.key` kao argumenata. Postavi vrednost return poziva `_delete()` za desetić (desni) čvor trenutnog čvora.

# --hints--

Treba ti da pozoveš putanju `_delete` za ponavljanje korišćenjem `node.right` i `node.key` kao argumenata.

```js
assert.match(code, /self\._delete\(\s*node\.right\s*,\s*node\.key\s*/);
```

Trebalo bi da postaviš vrednost return poziva `_delete()` za desetić trenutnog čvora.

```js
assert.match(code, /node\.right\s*=\s*self\._delete\(\s*node\.right\s*,\s*node\.key/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left

            node.key = self._min_value(node.right) 
            
--fcc-editable-region--
    def _min_value(self, node):
        while node.left is not None:
            node = node.left

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
