---
id: 686386e76197458b7dfd52ef
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Sada ćete kreirati dugme za CTA. Poziv na akciju (CTA) je dugme ili link dizajniran da podstakne korisnika da preduzme određenu akciju, u ovom slučaju, da počne sa slušanjem.

Na spoljnoj strani najdublje unutrašnjosti `div`, kreiraj anker element sa `href` od `#` i tekstom `Start listening`.

Ovo završava HTML prve cenovne kartice.

# --hints--

Trebalo bi da kreiraš anker element.

```js
assert.exists(document.querySelector("a"))
```

Tvoj anker element treba da ima `href` podešen na `#`.

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

Tvoj anker element treba da ima tekst `Start listening`.

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
