---
id: 686383abac6c625a84458ea9
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Na ovoj radionici, ćete naučiti više o Tailwind CSS koristeći njegove pomoćne klase za sistem mreže kako biste izgradili sekciju sa cenama.

Osnovni dizajn za HTML je pripremljen za vas. Počnite kreiranjem elementa `h1` sa tekstom `Choose your listening plan`.

# --hints--

Trebalo bi da kreirate element `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Element vašeg `h1` treba da sadrži tekst `Choose your listening plan`.

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </main>
  </body>
</html>
```
