---
id: 686386e76197458b7dfd52eb
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Kreiraj element sa `div` ispod `h1`. Unutar `div`, kreiraj još jedan element sa `div`, i poređaj unutar drugog elementa sa `div` unutar njega.

I unutar `div` u najdubljem delu, kreiraj komponentu za `h2` sa tekstom od `Listener`. Ovo će predstavljati prvi cenovni plan.

# --hints--

Trebalo bi da kreirate element za `div`.

```js
assert.exists(document.querySelector("div"));
```

Treba da kreirate još jedan element za `div` unutar eksternog `div`.

```js
assert.exists(document.querySelector("div > div"));
```

Treba da napraviš još jedan element `div` unutar unutrašnjeg elementa `div`.

```js
assert.exists(document.querySelector("div > div > div"));
```

Moraš da kreiraš `h2` sa tekstom `Listener` unutar elementa `div` najdubljeg.

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
