---
id: 686386e76197458b7dfd52ee
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Unutar `ul`, kreiraj dve stavke liste. Unutar prve stavke liste, kreiraj element od `span` i dodaj atribut/label HTML `&#10003;` kao njegov tekst. Odmah nakon `span` ali unutar stavke liste, dodaj tekst `Ad-supported streaming`.

Unutar drugog elementa liste, kreiraj element `span` i dodaj atribut `&#10003;` kao njegov tekst. Odmah nakon `span`, ali unutar stavke liste, dodaj tekst `Curated playlists`.

# --hints--

Treba vam dve stavke liste unutar vašeg elementa `ul`.

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

Tvoj prvi element od `li` mora da ima element `span` sa atributom karaktera HTML `&#10003;` unutar sebe.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

Prvi element vašeg `li` treba da sadrži tekst `Ad-supported streaming` odmah nakon elementa `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

Tvoj drugi element od `li` mora imati element `span` sa karakter atributom HTML `&#10003;` unutar sebe.

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

Vaš drugi element za `li` mora da sadrži tekst `Curated playlists` odmah nakon elementa `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
