---
id: 674ef2d357676e50e4691659
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Unutar elementa `div`, kreiraj element `h2` i unesi parametre za `name` kao njegov tekst.

Takođe, unutar `div`, kreirajte paragraf sa oznakom `className` od `card-title` da unesete parametre za `title`, i drugi paragraf da unesete parametre za `bio`.

# --hints--

Trebalo bi da napravite element `h2` unutar vašeg elementa `div`.

```js
  const testElem = await __helpers.prepTestComponent(window.index.Card);
  const h2Elem = testElem.querySelector('div > h2');
  assert.exists(h2Elem);
```

Element vašeg `h2` trebalo bi da ima `{name}` kao njegov tekstualni sadržaj.

```js
  const testElem = await __helpers.prepTestComponent(window.index.Card, { name: 'nameVal' });
  const h2Elem = testElem.querySelector('div > h2');
  // trimming because we don't need to be picky about whitespace
  assert.equal(h2Elem.textContent.trim(), 'nameVal');
```

Morate kreirati etiketu za `p` sa `className` od `card-title`, unutar vašeg elementa `div`.

```js
  const testElem = await __helpers.prepTestComponent(window.index.Card);
  const pElem = testElem.querySelector('div > p.card-title');
  assert.exists(pElem);
```

Vaš element `p` sa `className` od `card-title` mora imati `{title}` kao svoj tekstualni sadržaj.

```js
  const testElem = await __helpers.prepTestComponent(window.index.Card, { title: 'titleVal' });
  const pElem = testElem.querySelector('div > p.card-title');
  // trimming because we don't need to be picky about whitespace
  assert.equal(pElem.textContent.trim(), 'titleVal');
```

Trebalo bi da napraviš još jedan element od `p` sa `{bio}` kao tekstom.

```js
  const testElem = await __helpers.prepTestComponent(window.index.Card, {title: 'titleVal', bio: 'bioVal' });
  // trimming because we don't need to be picky about whitespace
  const texts = [...testElem.querySelectorAll('div > p')].map((p) => p.textContent.trim());
  const bioText = texts.find((text) => text.includes('bioVal'));
  assert.exists(bioText);
  assert.equal(bioText, 'bioVal');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  )
}
```
