---
id: 61fd719788899952e67692b9
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Elementi `thead` i `tbody` se koriste za prikazivanje koji deo vaše tabele je zaglavlje, a koji deo sadrži glavne podatke ili sadržaj.

Dodajte `thead` na `tbody` u `table` vaš prvi, ispod odeljka za `caption`.

# --hints--

Vaš prvi element od `table` mora imati komponentu sa `thead`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('thead'));
```

Tvoj prvi element `table` treba da sadrži element `tbody`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('tbody'));
```

Tvoj element `thead` treba da se nalazi neposredno ispod elementa `caption`.

```js
assert(document.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

Element vašeg `tbody` treba da bude direktno ispod elementa vašeg `thead`.

```js
assert(document.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
