---
id: 61fd6cc9475a784b7776233e
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Ispod vašeg elementa `div`, dodajte novi element `div` sa `class` podešenim na `table-wrap`. Ovo će biti kontejner za vaše tabele.

Unutar toga, dodajte tri elementa od `table`. Koristićete CSS da ih napravite u jednu tabelu, ali korišćenje različitih tabela će pomoći čitaču ekrana da razume tok dokumenta.

# --hints--

Trebalo bi da kreirate novu komponentu za `div`.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

Tvoja nova karakteristika `div` treba da ima `class` postavljenu na `table-wrap`.

```js
assert(document.querySelector('.table-wrap')?.localName === 'div');
```

Element `.table-wrap` treba da dođe nakon vašeg postojećeg `div`.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap'));
```

Tvoja komponenta od `.table-wrap` treba da ima tri elementa sa `table`.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert(children?.length === 3);
children.forEach(child => assert(child?.localName === 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
