---
id: 61fd71d596e8f253b9408b39
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Element ``tr`` se koristi za prikaz redova tabele. Dodajte element ``tr`` unutar vaše komponente ``thead``. U vašoj novoj komponenti ``tr``, dodajte element ``td``, praćen tri elementa ``th``.

Element `td` prikazuje podatnu ćeliju, dok element `th` prikazuje zaglavnu ćeliju.

# --hints--

Vaš element za `thead` treba da ima element za `tr`.

```js
assert(document.querySelector('thead')?.children?.[0]?.localName === 'tr');
```

Tvoj element `tr` treba da ima element `td` kao prvo dete.

```js
assert(document.querySelector('tr')?.children?.[0]?.localName === 'td');
```

Tvoj element `tr` treba da ima tri elementa od `th`, nakon elementa `td`.

```js
assert(document.querySelector('tr')?.children?.[1]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[2]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[3]?.localName === 'th');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
