---
id: 61fd6ab779390f49148773bb
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Ispod vašeg elementa `h1`, kreirajte element `div`. Dodelite mu atribut elementa `id` postavljen na `years`. Želite da je ovaj element sakriven od čitača ekrana, pa mu dodelite atribut elementa `aria-hidden` postavljen na `true`.

# --hints--

Treba da napraviš novi element `div` nakon tvog elementa `h1`.

```js
assert(document.querySelector('h1')?.nextElementSibling?.localName === 'div');
```

Tvoj novi element `div` treba da ima svojstvo elemenata `id` postavljeno na `years`.

```js
assert(document.querySelector('div')?.getAttribute('id') === 'years');
```

Tvoje novo polje `div` treba da ima svojstvo polja `aria-hidden` postavljeno na `true`.

```js
assert(document.querySelector('div')?.getAttribute('aria-hidden') === 'true');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
