---
id: 61fd75ea7f663457612dba02
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Unutar svakog vašeg novog elementa sa `th`, postavite unutar element/komponente `span` koja ima svojstvo elementa `class` postavljeno na `sr-only year`. Unesite sledeći tekst (u redosledu): `2019`, `2020`, i `2021`.

Daj svoje treće polje od `th`, atribut polja `class`, koje je postavljeno na `current`.

Ostavite element `td` prazan. Ovaj element služi samo za osiguravanje da vaša tabela ima strukturu od četiri reda i povezuje odgovarajuće zaglavlja i redove.

# --hints--

Svaki element vašeg `th` mora imati element `span`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.length === 1);
  assert(th?.children?.[0]?.localName === 'span');
});
```

Svako vaše novo polje `span` mora imati atribut polja `class` postavljeno na `sr-only year`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.[0]?.classList?.contains('sr-only'));
  assert(th?.children?.[0]?.classList?.contains('year'));
});
```

Prvi element vašeg `span` treba da ima tekst `2019`.

```js
assert(document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent === '2019');
```

Vaš drugi element za `span` treba da ima tekst `2020`.

```js
assert(document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent === '2020');
```

Tvoje treće polje za `span` bi trebalo da ima tekst `2021`.

```js
assert(document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent === '2021');
```

Tvoj treći element od `th` mora imati `class` podešeno na `current`.

```js
assert(document.querySelector('table')?.querySelectorAll('th')?.[2]?.classList?.contains('current'));
```

Vaše polje za `td` treba da bude prazno.

```js
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.textContent === '');
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
