---
id: 61fd6343fb42013d99bcd7f3
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Unutar vašeg elementa `body`, postavite element `section` unutar elementa `main`.

# --hints--

Prvi element unutar elementa `body` treba da bude vaš element `main`.

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

Trebalo bi da imaš element `section` unutar svog elementa `main`.

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
