---
id: 61fd778081276b59d59abad6
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

U vašem prvom `tr`, dodajte element `th` koji sadrži tekst `Cash This is the cash we currently have on hand.`. Umotajte sav taj tekst, osim `Cash `, unutar elementa `span` sa svojstvom elementa `class` postavljeno na `description`.

Nakon toga, dodaj tri elementa `td` sa sledećim tekstom (u redosledu): `$25`, `$30`, `$28`. Postavi atribut elementa `td` sa vrednošću `class` za treći element od `current`.

# --hints--

`tr` Prvi treba da ima karakteristiku od `th`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th'));
```

Vaše polje za `th` trebalo bi imati tekst `Cash This is the cash we currently have on hand.`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th')?.innerText === 'Cash This is the cash we currently have on hand.');
```

Trebalo bi da obuhvatite tekst `This is the cash we currently have on hand.` unutar elementa `span`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.textContent === 'This is the cash we currently have on hand.');
```

Tvoj element `span` mora imati svojstvo elementa `class`, postavljeno na `description`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.classList?.contains('description'));
```

Trebalo bi da imaš tri elementa od `td`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td').length === 3);
```

Tvoj prvi element za `td` mora imati tekst `$25`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[0]?.textContent === '$25');
```

Tvoja druga komponenta za `td` mora imati tekst `$30`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[1]?.textContent === '$30');
```

Tvoj treći element od `td` treba da ima tekst `$28`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.textContent === '$28');
```

Treći vaš element od `td` mora imati `class` postavljen na `current`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
