---
id: 61fd5a93fd62bb35968adeab
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Podesite HTML sa `DOCTYPE`, elementom `html` koji označava da je dokument na engleskom jeziku, elementom `head` i elementom `body`.

Pružite svoj element sa `head` prikladnim parametrima za `meta` za `charset` i `viewport`, element/polje `title` sa odgovarajućim naslovom, i element/polje `link` za vašu stilsku listu.

# --hints--

Tvoj kod mora imati prefiks `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Vaš kod treba da ima element `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Vaš element za `html` mora imati oznaku za otvaranje sa atributom elementa `lang` sa vrednošću `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Element vašeg `html` treba da ima zatvarajući tag.

```js
assert(code.match(/<\/html\s*>/));
```

Vaš kod bi trebalo da sadrži element `head` unutar elementa `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Tvoj kod mora imati element `body` unutar elementa `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Vaša komponenta `head` mora da stoji prije komponente `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Trebalo bi da ima dva elementa za `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Jedan element od `meta` treba da ima atribut elementa `name` postavljen na `viewport`, i `content` ako je `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Drugi element od `meta` mora imati svojstvo elementa `charset` postavljeno na `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Vaš kod mora imati karakteristiku `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` Vaše treba da ima tekst.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Vaš kod mora imati element `link`.

```js
assert(/<link/.test(code))
```

Tvoj element `link` treba da bude unutar tvog elementa `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

Vaš element `link` mora imati svojstvo elementa `rel` sa vrednošću `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

Vaš element `link` mora imati atribut elementa `href` sa vrednošću `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
