---
id: 694c9c186c089f7bdff0f3c6
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Uopšteno, dobra je praksa validirati vaš unos unutar funkcije. Koristite petlju `for` da iterirate kroz niz podataka/array `crew` i proverite da li bilo koji član već ima `id` jednak unosu od `astronaut`. Ako se zapis pronađe, ispišite poruku o grešci `console.log("Duplicate ID: " + astronaut.id)`, a zatim izvršite poziv `return` da napustite funkciju ranije.

Ovdje je primjer ove metode koja se često koristi:

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

Tvoja funkcija mora koristiti rutinu `for` da proveri postojeće članove od `crew` (bilo koja validna forma `for (...)` je u redu).

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

Unutar ciklusa `for`, trebalo bi da proverite kopiju `id`. Ako je dostupno, ukucajte `"Duplicate ID: " + astronaut.id"` u konzolu, zatim pozovite `return` da izađete iz funkcije.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Kada primi validan unos, tvoja funkcija koja `addCrewMember` ne bi trebalo da poziva poruku greške od `console.log()` niti rano izlazi za `return`.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
