---
id: 694c84330cbe604821b9deef
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Komunikacioni centar zahteva način za dodavanje novih članova na spisak. Kreiraj praznu funkciju pod nazivom `addCrewMember` koja prima dva parametra, `crew` i `astronaut`.

# --hints--

Treba da kreiraš funkciju nazvanu `addCrewMember`.

```js
assert.isFunction(addCrewMember);
```

Vaša funkcija od `addCrewMember` treba da ima parametre `crew` i `astronaut`.

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(addCrewMember.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
