---
id: 694c84106e9023f289c03b84
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Vaš spisak počinje sa samo jednim odraslim kosmičkim čovekom. Kreiraj nešto nazvano `firstAstronaut` sa sledećim karakteristikama:

| Key | Value |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

Trebalo bi da kreirate razliku od `firstAstronaut`.

```js
assert.isDefined(firstAstronaut);
```

`firstAstronaut` Tvoje treba da bude nešto.

```js
assert.isObject(firstAstronaut);
```

Kao objekat ``firstAstronaut``, treba da sadrži tačne vrednosti atributa.

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
