---
id: 694c40d86fe6a78d20b6559c
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

Na ovoj radionici, koristićete JavaScript da napravite listu astronauta za svemirsku misiju.

Započnite, kreirajte prazan red podataka za vašu listu pod nazivom `squad`.

# --hints--

Treba da definiš varijablu za `squad`.

```js
assert.isDefined(squad);
```

Vaša varijabla `squad` treba da bude red podataka.

```js
assert.isArray(squad);
```

Vaš red za `squad` mora biti prazan.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
