---
id: 698c5229e47d57c6dfaa91bd
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Atribut `controls` pruža kontrole reprodukcije uključujući pauziranje/nastavljanje reprodukcije, prebacivanje i kontrolu jačine zvuka za element `audio`.

Atribut `controls` je booleanski atribut i ne zahteva vrednost.

Dodajte atribut `controls` na element `audio`.

Sada bi trebalo da vidite element `audio` prikazan na stranici.

# --hints--

Your `audio` element should have the `controls` attribute.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
