---
id: 587d8259367417b2b2512c86
title: Izvrši podešavanje prilikom unosa
challengeType: 26
dashedName: implement-insertion-sort
---

# --description--

U ovom laboratorijskom radu, implementiraćete algoritam za umetanje sortiranja (Insertion Sort). Ovaj metod funkcioniše tako što postepeno gradi sortirani niz podataka na početku liste. Počinje sa prvim elementom. Zatim ispituje sledeći element i pomera ga unazad kroz sortirani niz dok ne dođe do odgovarajuće pozicije. Nastavlja da prolazi kroz listu, pomerajući nove elemente unazad u sortirljeni deo sve dok ne stigne do kraja. Ovaj algoritam ima vremensku složenost $O(n^2)$ u prosečnom i najgorem slučaju.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba da kreirate funkciju za `insertionSort`.
1. Funkcija `insertionSort` treba da primi red podataka punih brojeva i vrati red podataka punih brojeva sortiranih od najmanjeg do najvećeg.

# --hints--

`insertionSort` Trebalo bi da bude funkcija.

```js
assert.isFunction(insertionSort);
```

`insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` Trebalo bi da vrati red nepromenljivih podataka, osim po redu/redosledu.

```js
assert.sameMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort` Trebalo bi da vrati sortiran red podataka (od najmanjeg do najvećeg).

```js

assert.sameOrderedMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`insertionSort([5, 4, 33, 2, 8])` bi trebalo da vrati `[2, 4, 5, 8, 33]`.

```js
assert.sameOrderedMembers(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` ne sme koristiti metodu koja je ugrađena unutar `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
