---
id: 63b60aaaa65f8922bfce6b7e
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

U tvom elementu `span`, kreiraj element `label` za `entry-dropdown` sa tekstualnim sadržajem `Add food or exercise:`. Zatim kreiraj element `select` sa atributom `id` postavljenim na `entry-dropdown` i atributom `name` postavljenim na `options`. Ispod njega, dodaj element `button` sa atributom `id` postavljenim na `add-entry` i tekstualnim sadržajem `Add Entry`.

Vaš element sa `button` svojstvo elementa `type`, postavljeno na `button` ili sprečava automatsko slanje forme.

# --hints--

Moraš dodati element `label` u svoj element `span`.

```js
assert.exists(document.querySelector('.controls > span > label'));
```

Tvoj novi element za `label` treba da ima svojstvo elementa `for` postavljeno na `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

Vaša nova funkcija od `label` mora imati tekst `Add food or exercise:`.

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

Trebalo bi da dodaš element `select` u svoj element `span`.

```js
assert.exists(document.querySelector('.controls > span > select'));
```

Vaš element `select` treba da dođe nakon vašeg elementa `label`.

```js
assert(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName === 'LABEL');
```

Tvoj novi element `select` treba da ima svojstvo elementa `id` postavljeno na `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

Tvoja nova karakteristika `select` treba da ima svojstvo karakteristike `name` postavljeno na `options`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

Vaš deo/atribut od `select` mora biti prazan i imati zatvarajući label za `</select>`.

```js
assert.equal(document.querySelector('.controls > span > select')?.innerHTML?.trim(), '');
```

Trebalo bi da dodaš element `button` u svoj element `span`.

```js
assert.exists(document.querySelector('.controls > span > button'));
```

Tvoj element `button` mora da dođe nakon tvog elementa `select`.

```js
assert(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName === 'SELECT');
```

Tvoja nova komponenta `button` treba da ima svojstvo komponente `id` postavljeno na `add-entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

Vaša nova funkcija za `button` mora imati tekst od `Add Entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
