---
id: 63b608ebf40c871d960fc004
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Potreban vam je još dva kod fragmenta za `fieldset` – jedan za `snacks` i drugi za `exercise`.

# --hints--

Treba da kreirate četvrti element za `fieldset` u vašem `form`.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]);
```

Tvoj četvrti element za `fieldset` treba da sledi nakon tvog trećeg elementa za `fieldset`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.previousElementSibling?.tagName, "FIELDSET");
```

Četvrti element za `fieldset` treba da ima `id` postavljeno na `snacks`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.id, "snacks");
```

Vaš četvrti element ``fieldset`` mora imati element ``legend`` unutar sebe.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]?.querySelector('legend'));
```

Vaše četvrto polje od `legend` treba da ima tekst od `Snacks`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.querySelector('legend')?.innerText, "Snacks");
```

Tvoj četvrti element od `fieldset` mora da sadrži element `div` unutar sebe.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]?.querySelector('div'));
```

Vaš četvrti element od `div` treba da ima `class` podešen na `input-container`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.querySelector('div')?.className, "input-container");
```

Trebaće ti da kreiraš peto polje u tvom `fieldset` unutar `form`.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]);
```

Tvoja petasta sekcija `fieldset` treba da sledi četvrtu sekciju `fieldset`.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.previousElementSibling?.tagName, "FIELDSET");
```

Vaše peto polje za `fieldset` mora imati `id` postavljeno na `exercise`.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.id, "exercise");
```

Tvoj peti element `fieldset` mora da ima element `legend` unutra.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]?.querySelector('legend'));
```

Vaše peto polje za `legend` treba da ima tekst od `Exercise`.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.querySelector('legend')?.innerText, "Exercise");
```

Tvoj peti element za `fieldset` mora imati element `div` unutar sebe.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]?.querySelector('div'));
```

Vaš peti element od `div` mora imati `class` podešen na `input-container`.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.querySelector('div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
