---
id: 5a858944d96184f06fd60d61
title: Napravi svoj prvi CSS mrežni kontejner (grid)
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cqwREC4'
forumTopicId: 301129
dashedName: create-your-first-css-grid
---

# --description--
Izmenite bilo koji element HTML u alat za mrežu postavljanjem njegovog atributa `display` na `grid`. Ovo vam omogućava da koristite sve druge atribute povezane sa CSS Grid.

**Napomena:** U CSS Grid, roditeljski element se zove <dfn>kontena</dfn> a njegovi potomci se zovu <dfn>vitu</dfn>.

# --instructions--
Promeni prikaz klase div sa `container` u `grid`.

# --hints--
Klasa od `container` trebalo bi da ima svojstvo `display` sa vrednošću `grid`.

```js
assert.match(code, /.container\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    display: grid;
  }
</style>
```
