---
id: 5a9036ee38fddaf9a66b5d37
title: Povećajte razmake što je moguće više pomoću grid-gap
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/ca2qVtv'
forumTopicId: 301118
dashedName: add-gaps-faster-with-grid-gap
---

# --description--
`grid-gap` je svojstvo skraćivanja za `grid-row-gap` i `grid-column-gap` iz dve prethodne funkcije koje su najlakše za korišćenje. Ako `grid-gap` ima jednu vrednost, stvoriće razmak između svih redova i svih kolona. Međutim, ako postoji dve vrednosti, koristiće prvu vrednost da postavi razmak između redova i drugu vrednost za kolone.

# --instructions--
Koristi `grid-gap` da kreiraš razmak `10px` između redova i razmak `20px` između kolona.

# --hints--
Klasa koja `container` treba imati svojstvo `grid-gap` koje stvara razmak od `10px` između redova i razmak od `20px` između kolona.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-gap\s*?:\s*?10px\s+?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>
<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-gap: 10px 20px;
  }
</style>
```
