---
id: 5a90376038fddaf9a66b5d3c
title: Poravnajte sve elemente sa jednakim razmakom koristeći justify-items
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpECn'
forumTopicId: 301120
dashedName: align-all-items-horizontally-using-justify-items
---

# --description--
Ponekad želite da svi elementi u vašoj mreži CSS dele jedno poklapanje. Možete koristiti karakteristike koje ste ranije naučili i porediti ih jednu po jednu, ili možete sve zajedno jednako poreći koristeći `justify-items` u vašem alatu za mrežu. Ova karakteristika može primiti sve vrednosti koje ste naučili u poslednja dva izazova, drugačije ako će pomeriti **sve elemente** u našoj mreži sa željenim poklapanjem.

# --instructions--
Koristi ovu osobinu da centrirate sve ravnomerno.

# --hints--
Klasa za `container` mora imati svojstvo `justify-items` sa vrednošću `center`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*justify-items\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    justify-items: center;
  }
</style>
```
