---
id: 5a9036d038fddaf9a66b5d32
title: Dodaj red za `grid-template-columns`
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c7NzDHv'
forumTopicId: 301117
dashedName: add-columns-with-grid-template-columns
---

# --description--
Samo kreiranje grid elementa nije dovoljno. Takođe morate definisati strukturu grida. Da biste dodali red u grid, koristite atribut `grid-template-columns` na grid elementu, kao što je prikazano ispod:

```css
.container {
  display: grid;
  grid-template-columns: 50px 50px;
}
```

Ovo će kreirati dva reda u vašoj mreži koje su svaki široki 50px. Broj parametara definisan za atribut `grid-template-columns` pokazuje broj redova u mreži, a vrednost svakog parametra pokazuje širinu svakog reda.

# --instructions--
Tri reda u mrežnom kontejneru od kojih svaki je širine `100px`.

# --hints--
Klasa `container` treba da ima svojstvo `grid-template-columns` sa tri jedinice od `100px`.

```js
const templateColumns = new __helpers.CSSHelp(document).getStyle(
  '.container'
)?.gridTemplateColumns;
assert.strictEqual(templateColumns, '100px 100px 100px');
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 100px 100px 100px;
  }
</style>
```
