---
id: 5a90374338fddaf9a66b5d3a
title: Poravnajte komadi ravnomjerno korišćenjem justify-self
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpKHq'
forumTopicId: 301122
dashedName: align-an-item-horizontally-using-justify-self
---

# --description--
U CSS Grid, sadržaj svake ćelije nalazi se unutar kutije nazvane <dfn>seli</dfn>. Možete poravnati prostor sadržaja unutar svoje ćelije koristeći svojstvo `justify-self` na elementu mrežice. Podrazumevano, ovo svojstvo ima vrednost `stretch`, što će učiniti da sadržaj popuni celu širinu ćelije. Ovo svojstvo CSS Grid takođe prihvata druge vrednosti:

`start`: inapanga maudhui upande wa kushoto wa seli,

`center`: inapanga maudhui katikati ya seli,

`end`: inapanga maudhui upande wa kulia wa seli.

# --instructions--
Koristi atribut `justify-self` da sastavi deo sa klasom `item2` na sredini.

# --hints--
Klasa `item2` treba da ima svojstvo `justify-self` sa vrednošću `center`.

```js
assert.match(
  code,
  /.item2\s*?{[\s\S]*justify-self\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }

  .item2 {
    background: LightSalmon;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item2 {
    justify-self: center;
  }
</style>
```
