---
id: 5a9036ee38fddaf9a66b5d34
title: Koristi CSS Grid jedinice za menjanje veličine redova i kolona
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cvE8phd'
forumTopicId: 301134
dashedName: use-css-grid-units-to-change-the-size-of-columns-and-rows
---

# --description--
Možete koristiti fiksne dimenzije i povezane dimenzije kao što su `px` i `em` u CSS Grid za definisanje veličine redova i kolona. Takođe možete koristiti ovo:

`fr`: postavlja kolonu ili red kao deo postojećeg prostora,

`auto`: Postavlja kolonu ili red da bude širina ili visina njegovog sadržaja direktno,

`%`: Postavlja kolonu ili red na procenat širine njegovog kontejnera.

Ovo je kod koji generiše rezultate na početnom prikazu:

```css
grid-template-columns: auto 50px 10% 2fr 1fr;
```

Ovaj deo stvara pet kolona. Prva kolona je široka koliko joj sadržaj, druga kolona je 50px, treća kolona je 10% njenog kontejnera, a poslednje dve kolone; preostali prostor je podeljen na tri dela, dva su dodeljena četvrtoj koloni, a jedno petoj.

# --instructions--

Napravi mrežu sa tri kolone sledećih širina: `1fr`, `100px` i `2fr`.

# --hints--
Vaša klasa za `container` treba da ima svojstvo `grid-template-columns` koje postavlja tri kolone sa sledećim širinama: `1fr`, `100px`, i `2fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?1fr\s*?100px\s*?2fr\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: auto 50px 10% 2fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 50px 50px;
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 1fr 100px 2fr;
  }
</style>
```
