---
id: 5a94fe1369fb03452672e45d
title: Postavite elemente u mrežne zone koristeći svojstvo grid-area
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cRrqmtV'
forumTopicId: 301132
dashedName: place-items-in-grid-areas-using-the-grid-area-property
---

# --description--
Nakon kreiranja šablona područja za vaš alat u mrežici, kao što je prikazano u prethodnom izazovu, možete postaviti element u vaše specifično područje referenciranjem imena koje ste dodelili. Da biste to uradili, koristite atribut `grid-area` na elementu na ovaj način:

```css
.item1 {
  grid-area: header;
}
```

Ovo omogućava mreži da zna da želiš klasu objekata `item1` da bude u području koje ima ime `header`. U ovom slučaju, element će zauzeti ceo gornji red jer je taj ceo red dobio naziv područja `header`.

# --instructions--
Postavi element sa klasom `item5` u područje `footer` koristeći atribut `grid-area`.

# --hints--
Klasa objekata `item5` bi trebalo da ima svojstvo `grid-area` sa vrednošću `footer`.

```js
assert.match(
  __helpers.removeCssComments(code),
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?footer\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: footer;
  }
</style>
```
