---
id: 5a94fe3669fb03452672e45f
title: Smanji ponavljanje korišćenjem funkcije za ponavljanje
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cQvqyHR'
forumTopicId: 301133
dashedName: reduce-repetition-using-the-repeat-function
---

# --description--
Koristili ste `grid-template-columns` i `grid-template-rows` za kreiranje mrežne strukture, uneli ste vrednosti za svaki red ili stub koji ste kreirali.

Recimo da želite mrežu sa 100 redova jednakih visina. Nije baš lako unositi vrednost 100 jednu po jednu. Srećom, postoji bolji način – korišćenjem funkcije `repeat` koja pokazuje koliko puta želite da se ponovi red ili kolona, praćen zarezom i vrednošću koju želite da ponovite.

Evo primera koji bi kreirao mrežu od 100 redova, pri čemu svaki red ima visinu od 50px.

```css
grid-template-rows: repeat(100, 50px);
```

Takođe možete ponoviti više vrednosti korišćenjem funkcije „repeat“ i ubaciti tu funkciju između drugih vrednosti prilikom kreiranja mrežnog uzorka. To izgleda ovako:

```css
grid-template-columns: repeat(2, 1fr 50px) 20px;
```

Hii inatafsiriwa kama:

```css
grid-template-columns: 1fr 50px 1fr 50px 20px;
```

**Napomena:** `1fr 50px` se ponavlja dva puta, a zatim sledi `20px`.

# --instructions--
Koristi `repeat` da ukloni duplikate iz atributa `grid-template-columns`.

# --hints--
Klasa `container` mora imati svojstvo `grid-template-columns` koje je podešeno za ponavljanje u redovima 3 sa širinom od `1fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?1fr\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: 1fr 1fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, 1fr);
  }
</style>
```
