---
id: 5a94fe0569fb03452672e45c
title: Gawanya gridi kuwa eneo la template
challengeType: 0
forumTopicId: 301130
dashedName: divide-the-grid-into-an-area-template
---

# --description--
Možete spojiti ćelije svoje mreže zajedno u <dfn>oblast </dfn> i dati toj oblasti jedinstveno ime. To učinite korišćenjem `grid-template-areas` na kontejneru, kao što sledi:

```css
grid-template-areas:
  'header header header'
  'advert content content'
  'advert footer footer';
```

Gornji kod povezuje ćelije mreže u četiri područja: `header`, `advert`, `content`, na `footer`. Svaka reč predstavlja jednu ćeliju, a svaki par navodnika predstavlja jedan red.

# --instructions--
Promenite šablon tako da oblast `footer` pokrije ceo donji red. Napomena: Ove oblasti trenutno ne mogu izazvati vidljiv efekat. Kasnije, koristićete funkciju kako biste videli kako funkcioniše korišćenjem te oblasti.

# --hints--
Klasa za `container` bi trebalo da ima svojstvo `grid-template-areas` kao primer, ali sa oblastima/regionom `footer` koji pokriva ceo donji red.

```js
assert.match(
  __helpers.removeCssComments(code),
  /.container\s*?{[\s\S]*grid-template-areas\s*?:\s*?["|']\s*?header\s*?header\s*?header\s*?["|']\s*?["|']\s*?advert\s*?content\s*?content\s*?["|']\s*?["|']\s*?footer\s*?footer\s*?footer\s*?["|']\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
    /* Only change code below this line */
      'header header header'
      'advert content content'
      'advert footer footer';
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;

    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```
