---
id: 579e2a2c335b9d72dd32e05c
title: Implementirati algoritma za sečenje i spajanje
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

U ovoj laboratoriji, biće vam potrebno da kreirate algoritam za spajanje dva reda podataka.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Kreiraj funkciju `frankenSplice` koja prima dva reda podataka i šablon.
1. Kopiraj svaki element prvog reda u drugi red, po redu, počevši od dato šablona, i vrati rezultujući red.
1. Ulazni redovi podataka ostaju isti nakon izvršavanja funkcije.

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)` trebalo bi da vrati `[4, 1, 2, 3, 5]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)` treba da vrati `["a", 1, 2, "b"]`.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)` trebalo bi da vrati `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`.

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

Svi elementi iz prvog reda moraju biti dodati u drugi red po njihovom originalnom redosledu. `frankenSplice([1, 2, 3, 4], [], 0)` treba da vrati `[1, 2, 3, 4]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

Prvi red mora ostati nepromenjen nakon izvršavanja funkcije.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

Razina treba da ostane nepromenjena nakon spuštanja toaleta.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
