---
id: 655cd899f8de09431eabb40c
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Merge sort je algoritam za sortiranje koji koristi princip razdvajanja i osvajanja (divide-and-conquer) za sortiranje kolekcija podataka. To jest, on "razdvaja" kolekciju na manje poddelove i "osvaja" te poddelove sortiranjem nezavisno, a zatim spaja sortirane poddelove.

Počnite kreiranjem funkcije nazvane `merge_sort` koja prima parametar `array`. Ova funkcija će obraditi zadatak sortiranja liste brojeva. Koristite ključnu reč `pass` u telu funkcije.

# --hints--

You should declare a function named `merge_sort` with a single parameter: `array`. Don't forget to use the `pass` keyword in the function body.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_args("array")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
