---
id: 6564aee9c077774ea49c3faf
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Algoritam spajanja (merge sort) izvodi tri akcije:

- Podeliti neuređenu sekvencu stavki na poddelove
- Sortirati stavke u poddelovima
- Spojiti sortirane poddelove

Ovo se dešava rekursivno sve dok se poddelovi ne spoje u kompletnu sortiranu sekvencu. Počnite sa podele sekvence.

Prvo, zamenite ključnu reč `pass` varijablom `middle_point`. Koristite operator celobrojnog deljenja (`//`) da podelite dužinu liste `array` sa 2 i dodelite rezultat vašoj novoj varijabli `middle_point`. Ne zaboravite da uvuku kod.

# --hints--

You should not have `pass` in the body of the function.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
})
```

You should declare a `middle_point` variable and assign `len(array) // 2` to it.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
