---
id: 656639771fed09b5c0e8fe71
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Sada kada ste podelili listu `array` na dve odvojene liste, nastavićete da delite svaku listu sve dok svaki element ne bude sam u svojoj listi. Lista sa jednim brojem je uvek sortirana.

Da biste to uradili, rekurzivno pozovite `merge_sort` unutar vaše funkcije i prosledite `left_part` kao argument poziva.

# --hints--

You should call `merge_sort` with the argument `left_part` at the bottom of your function body.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
