---
id: 656709e50ed928da35260276
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

U prethodnom koraku, dodelili ste element iz `left_part` spisku `array` jer je bio manji. Ali to neće uvek biti slučaj. U nekim slučajevima poređenja, element sa desne strane može biti manji.

Kreirajte klauzulu `else` koja se izvršava kada element u `left_part` nije manji od elementa u `right_part`.

Unutar bloka `else`, dodelite `right_part[right_array_index]` u `array[sorted_index]`.

# --hints--

You should have `array[sorted_index] = right_part[right_array_index]` in the body of your `else` clause.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_bodies()[1].find_body()[0].is_equivalent("array[sorted_index] = right_part[right_array_index]")
`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
