---
id: 656707e11e671ed54c96f7ec
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Kada uslov `if` evaluira na `True`, to znači da je element u listi `left_part` manji od elementa sa kojim se poredi u listi `right_part`.

U tom slučaju, možete staviti element pronađen na `left_array_index` u `left_part` unutar sortirane niz.

Unutar bloka `if`, uklonite `pass` i dodelite `left_part[left_array_index]` na `array[sorted_index]`.

# --hints--

You should replace `pass` with `array[sorted_index] = left_part[left_array_index]` in the body of the `if` statement.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
