---
id: 6569b4e0bd29d17d4b53b16c
title: Korak 19
challengeType: 20
dashedName: step-19
---

# --description--

Sada ćete replicirati istu logiku petlje `while` za `right_part`.

Kreirajte `while` petlju koja radi kada je `right_array_index` manje od `len(right_part)`. Unutar `while` petlje, dodelite `right_part[right_array_index]` na `array[sorted_index]`. Na kraju, inkrementirajte `right_array_index` i `sorted_index` sa `1`.

# --hints--

Your `while` loop should check if `right_array_index` is less than `len(right_part)`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

You should assign `right_part[right_array_index]` to `array[sorted_index]` in the body of the loop. 

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

You should increment `right_array_index` by `1` within the `while` loop. Note that the increment should be done after assigning the right part element to the sorted array.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

You should use the increment the current value of `sorted_index` by `1` within the `while` loop. Note that the increment should be done after assigning the right part element to the sorted array.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
