---
id: 6569b743630ee592a65a7e2f
title: Korak 20
challengeType: 20
dashedName: step-20
---

# --description--

Pre testiranja funkcije `merge_sort()`, morate kreirati bazni slučaj koji zaustavlja izvršavanje funkcije kada je dužina `array` manja ili jednaka `1`.

Ovaj bazni slučaj će zaustaviti rekurzivni poziv. Bez njega, operacija spajanja (merge sort) bi nastavila da radi čak i kada je lista sortirana ili nema elemenata u njoj.

Odmah nakon deklaracije funkcije, kreirajte izjavu `if` sa ovim uslovom: `len(array) <= 1`. Unutar tela izjave `if`, dodajte izjavu `return` za zaustavljanje izvršavanja funkcije.

# --hints--

You should create an `if` statement to check if `len(array) <= 1`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

You should have a `return` statement in the body of the `if` statement.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_ifs()[0].find_body().has_stmt("return")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
--fcc-editable-region--
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
