---
id: 657f59751d5a5c9b069ae0f3
title: Korak 21
challengeType: 20
dashedName: step-21
---

# --description--

Kao što ste naučili u prethodnoj lekciji, možete koristiti varijablu `__name__` da utvrdite da li se Python skripta pokreće kao glavni program ili je uvezena kao modul (kod napisan u drugoj Python datoteci).

Ako je vrednost `__name__` postavljena na `'__main__'`, podrazumeva se da je trenutna skripta glavni program, a ne modul.

U ovom projektu, koristićete trenutnu skriptu kao glavni program.

Kreirajte izjavu `if` koja proverava da li je vrednost `__name__` jednaka `'__main__'`.
Unutar izjave `if`, kreirajte listu nazvanu `numbers` i dodelite joj ove vrednosti: `[4, 10, 6, 14, 2, 1, 8, 5]`.

# --hints--

Your `if` statement should check if `__name__` has a value of `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

You should have a `numbers` list with these values: `[4, 10, 6, 14, 2, 1, 8, 5]`

```js
({
  test: () => assert(runPython(`_Node(_code).find_ifs()[0].find_body().find_variable("numbers").is_equivalent("numbers = [4, 10, 6, 14, 2, 1, 8, 5]")`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
