---
id: 564944c91be2204b269d51e3
title: Promeni tekst unutar elementa koristeći jQuery
challengeType: 6
forumTopicId: 16773
dashedName: change-text-inside-an-element-using-jquery
---

# --description--

Koristeći jQuery, možete promeniti tekst između početnih i krajnjih tagova elementa. Možete čak promeniti atribute za pisanje HTML.

jQuery ima funkciju pod nazivom `.html()` koja omogućava dodavanje HTML oznaka i teksta unutar elementa. Sadržaj koji se ranije nalazio u tom elementu biće potpuno zamenjen sadržajem koji prosledite ovoj funkciji.

Evo kako možete da ponovo napišete i naglasite tekst našeg naslova sekcije:

```js
$("h3").html("<em>jQuery Playground</em>");
```

jQuery takođe ima sličnu funkciju nazvanu `.text()` koja menja samo tekst bez dodavanja labela. Drugim rečima, ova funkcija neće evaluirati nijednu oznaku od HTML koje joj se šalju, već će ga tretirati kao tekst koji želite da promenite postojeći sadržaj.

Promeni dugme sa identifikatorom `target4` da bi naglasio njegov tekst.

<a href="https://www.freecodecamp.org/news/html-elements-explained-what-are-html-tags/#em-element" target="_blank" rel="noopener noreferrer nofollow">Pogledajte naš članak o vestima o &lt;em&gt;</a> razliku između `<i>` i `<em>` i njihove upotrebe.

Zapamtite da, iako je oznaka `<i>` obično korišćena za naglašavanje teksta, sada se koristi kao oznaka za markere slika. Oznaka `<em>` je sada veoma prihvaćena kao oznaka za naglašavanje. Moja od njih će raditi za ovaj izazov.

# --hints--

Treba da naglasiš tekst na svom dugmetu ``target4`` dodavanjem oznaka `HTML`.

```js
assert.isTrue(
  /<em>|<i>\s*#target4\s*<\/em>|<\/i>/gi.test($('#target4').html())
);
```

Ostali tekstovi treba da ostanu nepromenjeni.

```js
assert($('#target4') && $('#target4').text().trim() === '#target4');
```

Ne smijete mijenjati nijedan drugi tekst.

```js
assert.isFalse(/<em>|<i>/gi.test($('h3').html()));
```

Trebalo bi da koristiš `.html()`, a ne `.text()`.

```js
assert(code.match(/\.html\(/g));
```

Trebalo bi da izaberete `button id="target4"` koristeći jQuery.

```js
assert(code.match(/\$\(\s*?(\"|\')#target4(\"|\')\s*?\)\.html\(/));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target4").html('<em>#target4</em>');
  });
</script>

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
