---
id: bad87fee1348bd9aedc08826
title: Ciljaj elemente za klasu elemenata koristeći jQuery
challengeType: 6
forumTopicId: 18316
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-class-using-jquery
---

# --description--

Da li vidiš kako smo sve stvari za `button` napravili da skoče? Izabrali smo ih pomoću `$("button")`, zatim smo dodali klase CSS koristeći `.addClass("animated bounce");`.

Koristeći funkciju `.addClass()` iz jQuery-ja, koja vam omogućava dodavanje klasa elementima.

Prvo, selektujte vaše elemente `div` koji imaju klasu od `well` koristeći selektor za `$(".well")`.

Zapamtite da, kao i kod izjave za CSS, pišete `.` pre imena klase.

Zatim koristi funkciju `.addClass()` iz jQuery-ja da doda klase `animated` i `shake`.

Na primer, možete učiniti da svi elementi koji imaju klasu `text-primary` trese dodavivši sledeće na vaš `document ready function`:

```js
$(".text-primary").addClass("animated shake");
```

# --hints--

Trebalo bi da koristite funkciju `addClass()` iz jQuery-ja da dodate klase `animated` i `shake` svim vašim elementima koji imaju klasu `well`.

```js
assert($('.well').hasClass('animated') && $('.well').hasClass('shake'));
```

Moraš koristiti jQuery samo da dodaš ove klase elementu.

```js
assert(!code.match(/class\.\*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
