---
id: bad87fee1348bd9bedc08826
title: Manipulacija HTML elementima korišćenjem selektora pomoću jQuery-ja
challengeType: 6
forumTopicId: 18319
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-html-elements-with-selectors-using-jquery
---

# --description--

Sada imamo funkcionalnost od `document ready`.

Sada napišite naš prvi jQuery deklarativni izraz. Sve jQuery funkcije počinju sa `$`, često se nazivaju znak dolar, ili bling.

jQuery često bira element sa HTML za <dfn>selektori</dfn>, zatim radi nešto s tim elementom.

Na primer, neka svi tvoji elementi za `button` preskoče/skaču. Dodaj ovaj kod unutar svoje funkcije za 'document ready':

```js
$("button").addClass("animated bounce");
```

Zapamtite da smo već uključili biblioteku jQuery i biblioteku Animate.css u pozadini kako biste mogli da ih koristite u editoru. Stoga koristite jQuery da dodelite klasu `bounce` iz Animate.css za vaše elemente `button`.

# --hints--

Trebalo bi da koristite funkciju ``addClass()`` iz jQuery-ja da dodelite klase ``animated`` i ``bounce`` vašim elementima ``button``.

```js
assert($('button').hasClass('animated') && $('button').hasClass('bounce'));
```

Trebalo bi da koristite jQuery samo za dodavanje ovih klasa elementu.

```js
assert(!code.match(/class.*animated/g));
```

Vaš jQuery kod treba da bude unutar funkcije `$(document).ready();`.

```js
assert(
  code.replace(/\s/g, '').match(/\$\(document\)\.ready\(function\(\)\{\$/g)
);
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
