---
id: bad87fee1348bd9aed208826
title: Selektuj potomke elementa koristeći jQuery
challengeType: 6
forumTopicId: 18320
dashedName: target-the-children-of-an-element-using-jquery
---

# --description--

Kada elementi HTML budu na istom nivou ispod drugog, nazivaju se <dfn>deca</dfn> tog elementa. Na primer, elementi dugmeta u ovom izazovu sa tekstovima `#target1`, `#target2`, i `#target3` su sva deca elementa `<div class="well" id="left-well">`.

jQuery ima funkciju nazvanu `children()` koja vam omogućava pristup potomcima bilo kog odabranog elementa.

Evo primera kako da koristiš funkciju ``children()`` ili dodaš boju ``left-well`` za potomke tvog elementa ``blue``:

```js
$("#left-well").children().css("color", "blue")
```

# --instructions--

Postavi narandžastu boju za sve elemente tvog dela `right-well`.

# --hints--

Svi deca za `#right-well` treba da imaju narandžastu boju.

```js
assert($('#right-well').children().css('color') === 'rgb(255, 165, 0)');
```

Trebalo bi da koristite funkciju `children()` za izmenu ovih elemenata.

```js
assert(code.match(/\.children\(\)\.css/g));
```

Trebalo bi da koristite jQuery samo za dodavanje ovih klasa elementu.

```js
assert(code.match(/<div class="well" id="right-well">/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
    $("#target1").parent().css("background-color", "red");
    $("#right-well").children().css("color", "orange");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
