---
id: bad87fee1348bd9acdd08826
title: Naučite kako rade skripte u Document Ready događaju
challengeType: 6
forumTopicId: 18224
dashedName: learn-how-script-tags-and-document-ready-work
---

# --description--

Sada smo spremni da naučimo jQuery, najpoznatiji alat za JavaScript zauvek.

Pre nego što smo počeli da koristiti jQuery, moramo dodati neke stvari u naš HTML.

Prvo, dodajte element `script` na sam vrh vaše stranice. Uverite se da ga završavate/dodajete na sledećoj liniji.

Tvoj pregledač će pokrenuti JavaScript bilo šta unutar elementa `script`, uključujući jQuery.

Unutar vašeg elementa `script`, dodajte ovaj kod: `$(document).ready(function() {` u svoj `script`. Zatim ga zatvorite na sledećem redu (i dalje unutar vašeg elementa `script`) sa: `});`

Naučićemo više o `functions` kasnije. Važno je znati da kod postavljen unutar `function` će biti izvršen čim vaš pregledač učita vašu stranicu.

Ovo je važno jer bez tvog `document ready function` koda, tvoj kod se može izvršiti pre nego što će biti prikazan HTML, što će uzrokovati grešku.

# --hints--

Treba da kreiraš element `script` osiguravajući da je validan i ima oznaku za zatvaranje.

```js
assert(
  code.match(/<\/script\s*>/g) &&
    code.match(
      /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
    ) &&
    code.match(/<\/script\s*>/g).length ===
      code.match(
        /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
      ).length
);
```

Treba da dodaš `$(document).ready(function() {` na početak svog elementa `script`.

```js
assert(
  code.match(
    /\$\s*?\(\s*?document\s*?\)\.ready\s*?\(\s*?function\s*?\(\s*?\)\s*?\{/g
  )
);
```

Treba da zatvoriš svoju funkciju `$(document).ready(function() {` za `});`

```js
assert(code.match(/\n*?\s*?\}\s*?\);/g));
```

# --seed--

## --seed-contents--

```html
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
  });
</script>
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
