---
id: bad87fee1348bd9aed908626
title: Ciljaj na jedan element koristeći više jQuery selektora
challengeType: 6
forumTopicId: 18322
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-the-same-element-with-multiple-jquery-selectors
---

# --description--

Sada znate tri načina za ciljanje elemenata: po tipu: `$("button")`, po klasi objekata: `$(".btn")`, i po identifikatoru `$("#target1")`.

Iako je moguće dodati mnogo klasama jednim pozivom od `.addClass()`, dodajmo ga sa jednim elementom na *tri različita načina*.

Ako se oslanjate na `.addClass()`, dodajte samo jednu klasu u jednom trenutku za jednu funkciju, na tri različita načina:

Dodaj klasu za `animated` za sve elemente tipa `button`.

Dodaj klasu la `shake` za sve dugmad sa klasom la `.btn`.

Dodajte klasu za `btn-primary` na dugme sa identifikatorom `#target1`.

**Napomena:** Treba da se fokusirate samo na jedan element i dodajete samo jednu klasu odjednom. Zajedno, vaši tri selektora će dodati tri klase `shake`, `animated`, i `btn-primary` za `#target1`.

# --hints--

Kod treba da koristi selektor `$("button")`.

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?button\s*?(?:'|")/gi));
```

Vaš kod treba da koristi selector za `$(".btn")`.

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?\.btn\s*?(?:'|")/gi));
```

Vaš kod mora da koristi selektor `$("#target1")`.

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?#target1\s*?(?:'|")/gi));
```

Morate dodati samo jedan kurs/predmet za svaki vaš treći izbor.

```js
assert(
  code.match(/addClass/g) &&
    code.match(/addClass\s*?\(\s*?('|")\s*?[\w-]+\s*?\1\s*?\)/g).length > 2
);
```

Tvoj element za `#target1` mora imati klase `animated`, `shake` i `btn-primary`.

```js
assert(
  $('#target1').hasClass('animated') &&
    $('#target1').hasClass('shake') &&
    $('#target1').hasClass('btn-primary')
);
```

Trebalo bi da koristite jQuery samo za dodavanje ovih klasa elementu.

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated");
    $(".btn").addClass("shake");
    $("#target1").addClass("btn-primary");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
