---
id: bad87fee1348bd9aed508826
title: Kopiraj element koristeći jQuery
challengeType: 6
forumTopicId: 16780
dashedName: clone-an-element-using-jquery
---

# --description--

Pored pomeranja elemenata, možete takođe da ih kopirate sa jedne lokacije na drugu.

jQuery ima funkciju nazvanu `clone()` koja pravi kopiju elementa.

Na primer, ako bismo hteli da kopiramo `target2` iz našeg `left-well` do `right-well`, koristimo:

```js
$("#target2").clone().appendTo("#right-well");
```

Da li ste videli da ovo uključuje povezivanje dve jQuery funkcije? Ovo se zove <dfn>lančano spajanje funkcija</dfn> i to je jednostavan način za rad sa jQuery-jem.

Kopiraj svoj element `target5` i zalepi ga u svoj `left-well`.

# --hints--

Tvoja komponenta za `target5` bi trebalo da bude unutar tvog `right-well`.

```js
assert($('#right-well').children('#target5').length > 0);
```

Kopija vašeg elementa sa `target5` takođe bi trebalo da bude unutar vašeg `left-well`.

```js
assert($('#left-well').children('#target5').length > 0);
```

Trebalo bi koristiti samo jQuery za pomeranje ovih elemenata.

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
    $("#target2").appendTo("#right-well");
    $("#target5").clone().appendTo("#left-well");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
