---
id: 67172b43f84bcd2dec238a3d
title: Napravi bibliotekar (ili katalogizator) za knjige
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

Trebalo bi da imaš niz podataka stavki nazvanih `books` gde svaka stavka u tom nizu treba da ima sekvencu slova `title`, drugu sekvencu slova `authorName` i broj `releaseYear`.

2. Vaš red od `books` mora imati najmanje tri stavke.

3. Mora da ima funkciju koja se može pozvati, nazvanu `sortByYear`, koja prima dva parametra za sortiranje redova podataka.

4. Funkcija za `sortByYear` treba da vrati `-1` ako je `releaseYear` iz prve knjige manje od onoga iz druge knjige.

5. Funkcija za `sortByYear` treba da vrati `1` ako je `releaseYear` iz prve knjige veće od onog iz druge knjige.

6. Funkcija za `sortByYear` treba da vrati `0` ako su sve vrednosti `releaseYear` jednake.

7. Treba da filtrirate knjige napisane nakon određene godine, kao što je 1950, iz kolone `books` i sačuvate filtrirani red u novu kolonu pod nazivom `filteredBooks`.

8. Treba da sortirate knjige u redu `filteredBooks` prema `releaseYear` po uzlaznom poretku. Naučili ste u prethodnom kursu da će metoda `sort()` sortirati taj red na svom mestu. To znači da će red `filteredBooks` biti izmenjen.

# --hints--

Trebalo bi da imaš funkciju `sortByYear` u svom kodu.

```js
assert.isFunction(sortByYear);
```

Tvoja funkcija za `sortByYear` mora da prihvati dva parametra.

```js
assert.lengthOf(sortByYear, 2);
```

Vaša funkcija za `sortByYear` treba da vrati `-1` ako je `releaseYear` od stavke `book1` manje od onoga za stavku `book2`, `1` ako je `releaseYear` od stavke `book1` veće od onoga za stavku `book2`, i `0` u svim drugim slučajevima.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

Trebalo bi da imaš red podataka `books` u svom kodu.

```js
assert.isArray(books)
```

Vaš red za `books` mora imati najmanje tri stavke knjiga.

```js
assert.isAtLeast(books.length, 3)
```

Vaš red za `books` treba da ima elemente sa sekvencom slova `title`, drugi niz slova `authorName`, i broj `releaseYear`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

Trebalo bi da imaš red podataka `filteredBooks` u svom kodu.

```js
assert.isArray(filteredBooks)
```

Safu ya `filteredBooks` bi trebalo da sadrži neke knjige filtrirane iz safu ya `books` i ne sme biti prazna.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

Serija `filteredBooks` mora da sadrži samo knjige iz prethodne serije `books`.

```js
filteredBooks.forEach(book => {
  assert.isTrue(
    books.some(original => 
      original.title === book.title &&
      original.authorName === book.authorName &&
      original.releaseYear === book.releaseYear
    )
  );
});
```

Linija od `filteredBooks` treba da uključuje knjige objavljene te godine ili ranije.

```js
const maxYear = Math.max(...filteredBooks.map(book => book.releaseYear));
const removedYears = books.filter(book => !filteredBooks.includes(book)).map(book => book.releaseYear);
assert.isTrue(removedYears.every(year => year > maxYear));
```

Trebalo bi da pozovete funkciju višegrađenog nivoa `sort` prosleđivanjem callback funkcije `sortByYear` u niz `filteredBooks`.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
