---
id: 641cdf57c3f7ee276e1d9b32
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Vrati se na svoj listener događaja, moraš da ispraviš tekst elementa `result`. Možeš koristiti ternarni operator da završiš ovu vežbu.

Ovo je primer za postavljanje vrednosti rezultata ternarnog operatera u tekstualni sadržaj elementa:

```js
el.textContent = condition ? "Use this text if the condition is true" : "Use this text if the condition is false";
```

Nakon deklaracije `if`, ažuriraj tekst polja `result` u skladu sa istinom poziva `isSpam()` uz `messageInput.value` kao argument. Ako je tačno, postavi atribut/svojstvo `textContent` na polje `result` na `"Oh no! This looks like a spam message."`. Inače, postavi na `"This message does not seem to contain any spam."`.

Zatim postavi atribut `messageInput` elementa `value` na prazan niz karaktera.

# --hints--

Kada kliknete na `checkMessageButton` i `isSpam(messageInput.value)`, ako je tačno, trebalo bi da postavite svojstvo `textContent` od `result` na `"Oh no! This looks like a spam message."`.

```js
const temp = isSpam;
isSpam = (msg) => true;
messageInput.value = "spam";
try {
  checkMessageButton.dispatchEvent(new Event("click"));
  assert.equal(document.querySelector("#result-message").textContent, "Oh no! This looks like a spam message.");
} finally {
  isSpam = temp;
}
```

Kada kliknete `checkMessageButton` na `isSpam(messageInput.value)`, ako nije tačno, morate postaviti svojstvo `textContent` od `result` na `"This message does not seem to contain any spam."`.

```js
messageInput.value = "spam";
checkMessageButton.dispatchEvent(new Event("click"));
assert.equal(document.querySelector("#result-message").textContent, "This message does not seem to contain any spam.");
```

Nakon postavljanja `textContent` od `result`, treba da ispraviš svojstvo `value` u elementu `messageInput` da bude prazan niz karaktera.

```js
messageInput.value = "spam";
checkMessageButton.dispatchEvent(new Event("click"));
assert.isEmpty(messageInput.value);

messageInput.value = "spam";
const temp = isSpam;
isSpam = (msg) => true;
try {
  checkMessageButton.dispatchEvent(new Event("click"));
  assert.isEmpty(messageInput.value);
} finally {
  isSpam = temp;
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const isSpam = (msg) => false;

--fcc-editable-region--
checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }


});
--fcc-editable-region--
```
