---
id: 6421f98f4999d1179ce37cb4
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Nizovi podataka imaju metodu `some()`. Kao metoda `filter()`, `some()` prima funkciju povratnog poziva koja bi trebalo da uzme element iz niza podataka kao argument. Metoda `some()` vraća `true` ako funkcija povratnog poziva vrati `true` za barem jedan element u nizu podataka.

Ovde je primer pozivanja metode za `.some()` da bi se proverilo da li je bilo koji element u nizu podataka veliko slovo.

```js
const arr = ["A", "b", "C"];
arr.some(letter => letter === letter.toUpperCase());
```

Validirati funkciju `isSpam()` korišćenjem metode `some()`, koja proverava da li je potrebno testiranje `msg` naspram jednog od uobičajenih izraza u `denyList` vraćajući `true`.

# --hints--

Vaša funkcija za `isSpam` treba da koristi metodu `.some()`.

```js
let flag = false;
const temp = Array.prototype.some;
Array.prototype.some = (arg) => {flag = true};
try {
  isSpam("");
  assert.isTrue(flag);
} finally {
  Array.prototype.some = temp;
}
```

Tvoja funkcija za `isSpam` treba da pozove `denyList.some()`.

```js
let flag = false;
const temp = denyList.some;
denyList.some = (arg) => {flag = true};
try {
  isSpam("");
  assert.isTrue(flag);
} finally {
  denyList.some = temp;
} 
```

Tvoja funkcija ``isSpam`` treba da vrati rezultat od ``denyList.some()``, gde njegova funkcija za povratni poziv (callback) treba da proceni ``.test(msg)`` za svaki element od ``denyList``.

```js
const help = "plEase hElp";
const assist = "AssIst ME";
const notSpam = "not spam";
const mockSpam = (msg) => denyList.some((regex) => regex.test(msg));
denyList = [helpRegex, /test1/, /test2/]
assert.strictEqual(isSpam(help), mockSpam(help));
assert.strictEqual(isSpam(assist), mockSpam(assist));
assert.strictEqual(isSpam(notSpam), mockSpam(notSpam));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;

const denyList = [helpRegex];

--fcc-editable-region--
const isSpam = (msg) => helpRegex.test(msg);
--fcc-editable-region--

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
