---
id: 68c3482278a14e3cf225e897
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Unutar vašeg elementa/komponente `tr`, dodajte komponentu `th` sa tekstom `9:00 AM`. Zatim ispod komponente `th`, dodajte tri elementa/komponente `td` sa tekstovima `Keynote: Tech Future`, `Intro to Web Dev` i `UX for All`.

# --hints--

Trebalo bi da imaš element `th` unutar svog elementa `tr`.

```js
assert.exists(document.querySelector("tbody tr th"));
```

Polje za `th` mora imati tekstualni sadržaj od `9:00 AM`.

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

Trebalo bi da imaš ukupno tri elementa za `td` ispod tvog elementa `th`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

Tvoj prvi element za `td` treba da ima tekstualni sadržaj od `Keynote: Tech Future`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

Vaš drugi element za `td` trebao bi da ima tekstualni sadržaj od `Intro to Web Dev`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

Tvoj treći element od `td` treba da ima tekstualni sadržaj `UX for All`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
