---
id: 68c34ace730ca43e5fa7b8d6
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Sada je vreme za dodavanje drugog reda u tabelu.

Počnite dodavanjem drugog elementa za `tr`. Unutarja u tog elementa `tr`, dodajte element `th` sa svojstvom elementa `scope` postavljeno na `"row"` i tekstom od `10:00 AM`.

Zatim, dodajte tri elementa od `td` sa sledećim tekstom:

- `Accessibility Deep Dive`
- `CSS for Beginners`
- `Inclusive Design Principles`

# --hints--

Trebalo bi da imaš element `tr` unutar elementa `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 2);
```

Trebalo bi da imaš element `th` unutar drugog elementa `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[1].querySelector("th"));
```

Vaše polje za `th` treba da ima tekst sa `10:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.textContent.trim(), "10:00 AM");
```

Vaše polje `th` trebalo bi imati atribut polja `scope` postavljen na `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.getAttribute("scope"), "row");
```

Vaš drugi element vezan za `tr` bi trebalo da ima tri komponente od `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[1].querySelectorAll("td"), 3);
```

Vaš element `td` prvi treba da sadrži tekst `Accessibility Deep Dive`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[0]?.textContent.trim(), "Accessibility Deep Dive");
```

Vaše drugo polje za `td` mora da sadrži tekst od `CSS for Beginners`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[1]?.textContent.trim(), "CSS for Beginners");
```   

Treći element vašeg `td` treba da ima tekst `Inclusive Design Principles`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[2]?.textContent.trim(), "Inclusive Design Principles");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
