---
id: 68c346d03709343bff132707
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

U narednim koracima, kreiraćeš telo tabele.

Počnite dodavanjem elementa `tbody` ispod elementa `thead`. Zatim unutar `tbody`, dodajte element `tr`.

# --hints--

Trebalo bi da imaš element `tbody` ispod elementa `thead`.

```js
assert.exists(document.querySelector("tbody"));
```

Trebalo bi da imaš element `tr` unutar elementa `tbody`.

```js
assert.exists(document.querySelector("tbody tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>
    
    --fcc-editable-region--
    
    --fcc-editable-region--
  </table>
</body>
</html>
```
