---
id: 68c34329d430583a6cd5f7f1
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Unutar vašeg elementa sa `tr`, dodajte četiri elementa od `th`. Prvi element od `th` treba da ima tekst `Time`, drugi element treba da ima tekst `Track A`, treći element treba da ima tekst `Track B`, i četvrti element treba da ima tekst `Track C`.

# --hints--

Vaša funkcija za `tr` mora imati sledeće komponente od `th`.

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

Vaš prvi element `th`, treba da ima tekst od `Time`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

Vaš drugi element `th` trebalo bi da ima tekst od `Track A`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

Tvoj treći element za `th` treba da sadrži tekst od `Track B`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

Četvrti vaš element za `th` treba da ima tekst od `Track C`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
