---
id: 68c3420cb0176439a95af1b8
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sljedeći korak je dodavanje naslova tabele.

Počnite dodavanjem elementa `thead` ispod elementa `caption`. Unutar `thead`, dodajte element `tr`.

# --hints--

Trebalo bi da imaš element `thead`.

```js
assert.exists(document.querySelector("table thead"));
```

Trebalo bi da imaš element `tr` unutar `thead`.

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
