---
id: 68c34c608cf9e53f5d36b1d5
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Sledeće, dodaj treći red u tabelu. Počni sa dodavanjem drugog elementa za `tr`. Unutar tog elementa od `tr`, dodaj element `th` sa atributom `scope` postavljenim na `"row"` i tekstom `11:00 AM`. Zatim ispod tog elementa od `th`, dodaj element `td` sa tekstom `Break`.

# --hints--

Treba da imaš treći element `tr` unutar elementa `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

Trebalo bi da imaš element `th` unutar trećeg elementa `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

Vaše polje `th` trebalo bi da sadrži tekst iz/od `11:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

Vaše polje `th` mora imati svojstvo `scope` postavljeno na `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

Trebalo bi da imaš element `td` unutar trećeg elementa `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

Vaš element za `td` treba da sadrži tekst od `Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
