---
id: 68c33e502cd97f2fd52cb63b
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Na ovoj radionici, naučićete kako da kreirate elemente tabele dostupne svima za izradu programa tehnološke konferencije.

Počnite dodavanjem elementa `h1` koji sadrži tekst `Tech Conference 2025 Schedule`.

# --hints--

Trebalo bi da imaš element `h1`.

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

Tvoj element za `h1` mora imati tekst od `Tech Conference 2025 Schedule`.

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
