---
id: 68c34e201b347c40acda8111
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Sada je vreme da dodamo još jedan red u tabelu.

Počnite dodavanjem drugog elementa `tr`. Unutar tog elementa `tr`, dodajte element `th` sa svojstvom elementa `scope` postavljeno na `"row"` i tekstualnim sadržajem `11:30 AM`.

Zatim, dodaj tri komponente ``td`` koje sadrže sljedeći tekstualni sadržaj:

- `AR/VR in Education`
- `JavaScript Fundamentals`
- `Design Systems at Scale`

# --hints--

Trebalo bi da ima četvrti element `tr` unutar elemenata `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 4);
```

Trebalo bi da imaš element `th` unutar četvrtog elementa `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[3].querySelector("th"));
```

Vaše polje `th` treba da ima tekstualni sadržaj od `11:30 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.textContent.trim(), "11:30 AM");
```

Vaše polje od `th` mora imati atribut polja `scope`, postavljeno na `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.getAttribute("scope"), "row");
```

Vaš četvrti element za `tr` treba da ima tri elementa od `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[3].querySelectorAll("td"), 3);
```

Tvoj prvi element od `td` treba da ima tekstualni sadržaj `AR/VR in Education`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[0]?.textContent.trim(), "AR/VR in Education");
```

Vaš drugi element ``td`` trebao bi da ima tekstualni sadržaj ``JavaScript Fundamentals``.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[1]?.textContent.trim(), "JavaScript Fundamentals");
```   

Treći vaš element za `td` treba da ima tekstualni sadržaj iz `Design Systems at Scale`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[2]?.textContent.trim(), "Design Systems at Scale");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
