---
id: 68c34d57ed874e4004f7b5d4
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Trenutno, element sa tekstualnim sadržajem `td` koji ima tekstualni sadržaj `Break` uzima samo jednu liniju. Međutim, bilo bi bolje da uzme sve tri linije.

Kao što se sećate sa prethodnog radionica i obuke, možete koristiti svojstvo elementa ``colspan`` da ćelija tabele zauzme više redova.

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

Dodaj atribut elementa `colspan` na element `td` i postavi mu vrednost na `3`.

# --hints--

Tvoj element `td` sa tekstualnim sadržajem `Break` treba da ima svojstvo elementa `colspan` postavljeno na `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        --fcc-editable-region--
        <td>Break</td>
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
