---
id: 68c34f4aeb2b1a4170506751
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Sledeće, dodajte peti red u tabelu. Počnite sa dodavanjem još jednog elementa `tr`. Unutar tog elementa `tr`, dodajte element `th` sa atributom elementa `scope` postavljenim na `"row"` i tekstom `12:30 PM`. Zatim ispod tog elementa `th`, dodajte element `td` sa tekstom `Lunch Break`. Vaš element `td` takođe bi trebalo da ima atribut elementa `colspan` postavljen na `3` kako bi se mogao proširiti na sva tri puta.

# --hints--

Trebalo bi da imaš element `tr` unutar elementa `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

Trebalo bi da imaš element `th` unutar petog elementa `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

Vaše polje za `th` trebalo bi da ima tekst od/sa `12:30 PM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

Polje vašeg `th` treba da ima atribut polja `scope` postavljeno na `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

Trebalo bi da imaš element `td` unutar petog elementa `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

Tvoj element za `td` treba da ima tekst `Lunch Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
``` 

Vaš element za `td` mora imati svojstvo elementa `colspan` postavljeno na `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
