---
id: 66a37f37ef5823a313de8c26
title: Izgradi ciklus meseca
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

U ovoj laboratoriji, kreiraćeš jednostavnu animaciju kretanja Meseca oko Zemlje koristeći HTML i CSS. Zemlja će biti u centru sistema, a Mesec će kružiti oko nje.

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Treba da poravnate sve elemente u centru stranice postavljanjem visine na `100%` od vidljivog područja i podešavanjem rasporeda (display) na flexbox za element `body`.
   
2. Treba da imate element `div` sa klasom elemenata `space`.

3. Unutar elementa `.space`, bi trebalo da ima dva elementa `div` sa klasama elemenata `earth` i `orbit`, u tom redosledu.

4. Unutar elementa `.orbit`, treba da postoji drugi `div` sa klasom objekata `moon`.
   
5. Element `div` sa atributom `class` od `space` treba da bude centriran na stranici i da ima širinu i visinu od `200px`.

6. Funkcija za `div` sa svojstvom `class` od `space` treba da koristi raspored za `relative`.

7. Element `.earth` treba da koristi raspored (layout) `absolute`. Postavite u centar elementa `.earth` na pola pozicije njegovog roditelja u vertikalnom smeru (gore) i horizontalnom (levo). Nakon toga, pomerite element `.earth` unazad za pola njegove širine i visine, tako da bude centriran unutar svog roditelja, elementa `.space`.

8. Element `.earth` treba da ima širinu i visinu `100px`.

9. Element `.orbit` mora imati širinu i visinu od `200px`.

10. Element `.orbit` treba biti postavljen koristeći raspored od `absolute`. Njegov donji desni ugao treba da bude centriran unutar elementa `.space` koristeći svojstvo `transform` koje gura ka `-50%` u vertikalnom i horizontalnom smeru.

11. Orbita Meseca oko Zemlje bi trebalo da bude kružnica (ili: treba da bude kružnog oblika).

12. Element `.moon` treba da se postavi koristeći raspored `absolute` i da ima širinu i visinu od `30px`. Element `.moon` treba da se pozicionira iznad elementa `.orbit` i da bude savršeno centriran.

13. Treba da dodatno ispraviš poravnanje elementa `.moon` pomeranjem elementa ulevo za pola njegove širine.

14. Izabirnik za `.earth` mora imati pozadinsku boju i `border-radius` od `50%`.

15. Selektor `.moon` treba da ima boju pozadine i `border-radius` od `50%`.

16. Treba definisati animaciju `@keyframes orbit` koja rotira element `.orbit` za 360 stepeni oko svog centra. Ova animacija treba da se koristi za element `.orbit` na period od `5` sekundi, sa linearnom funkcijom vremena i beskonačnim petljama.

**Napomena:** Obavezno povežite vaš list modela sa HTML i koristite svoj CSS.

# --hints--

Sadržaj u elementu `body` treba da bude centriran na stranici postavljanjem visine na `100%` visinu vidljivog područja, `display` na `flex`, `justify-content`, i `align-items` na `center`.

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

Trebalo bi da imaš glavni `div` sa klasom objekata `space`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

Element za `.space` mora imati širinu i visinu od `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

Element za `.space` treba da ima svojstvo `position` od `relative`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

Treba da ima `div` sa klasom objekata `earth` unutar `div` sa svojstvom `class` od `space`.

```js
assert.exists(document.querySelector('div.space div.earth'));
```

Element `.earth` mora biti prvo dete elementa `.space`.

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

Element `.earth` treba da ima širinu i visinu od `100px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

Funkcija `.earth` treba da koristi raspored od `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

Element za `.earth` mora imati atribute od `top` i `left` za lokaciju `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

Element [`.earth`] treba biti premešten nazad i više u odnosu na [`50%`] uz postavljanje atributa [`transform`] na [`translate(-50%, -50%)`] .

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

Trebalo bi da imaš `div` sa klasom objekata `orbit` unutar elementa `.space`.

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

Element sa `.orbit` treba da dođe nakon elementa sa `.earth`.

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

Element `.orbit` bi trebalo da ima širinu i visinu od `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

Komponenta za `.orbit` treba da koristi formatiranje od `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

Element `.orbit` treba da se postavi u sredinu elementa `.space` koristeći atribut `transform` koji gura za `-50%` u vertikalnom i horizontalnom smeru.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

Trebalo bi da imaš `div` sa klasom objekata `moon` unutar elementa `.orbit`.

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

Element `.moon` mora imati širinu i visinu od `30px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

Element `.moon` treba da se postavi koristeći formatiranje `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

Gornji ugao elementa `.moon` treba da se poravna sa gornjim uglom njegovog bloka.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

Morate da premestite levi ugao elementa `.moon` do centra bloka (`50%` sa leve strane).

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

Element `.moon` treba da bude izjednačeno prilagođen kada se transformiše u `-50%` na x-osi.

```js
const transformValue = new __helpers.CSSHelp(document).getStyle('.moon')?.transform;
const validTransforms = ['translateX(-50%)', 'translate(-50%)'];
assert.oneOf(transformValue, validTransforms);
```

Selektor za `.earth` treba da ima pozadinsku boju.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

Selektor za `.earth` treba da ima `border-radius` od `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

Selektor za `.moon` bi trebalo da ima boju pozadine.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

Selektor za `.moon` mora imati `border-radius` od `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

Trebalo bi da imaš pravilo `@keyframes`.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

Tvoj novi zakon `@keyframes` bi trebalo da se zove `orbit`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

Tvoja pravilo za `@keyframes orbit` treba da ima izabirnik `0%`.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

Tvoj zakon za `@keyframes orbit` bi trebalo da ima selector za `100%`.

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

Selektor za `0%` mora imati svojstvo `transform` postavljeno na `rotate(0deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

Selektor za `100%` treba da sledi nakon selektora za `0%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

Selektor za `100%` mora imati svojstvo `transform` postavljeno na `rotate(360deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

Selektor za `.orbit` mora imati svojstvo `animation` postavljeno na `orbit 5s linear infinite`.

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```
