---
id: 587d7fb6367417b2b2512c06
title: Instaliraj i Konfiguriši Mongoose
challengeType: 2
forumTopicId: 301540
dashedName: install-and-set-up-mongoose
---

# --description--

Rad na ovim izazovima će uključivati pisanje vašeg koda korišćenjem jednog od sledećih načina:

- Kopiraj <a href="https://github.com/freeCodeCamp/boilerplate-mongomongoose/" target="_blank" rel="noopener noreferrer nofollow">ovaj repozitorijum od GitHub</a> i završi ove izazove interno.
- Koristi omiljeni alat za izradu sajtova da završiš projekat. Proveri da uključuješ sve fajlove iz našeg repozitorijuma GitHub.

U ovom izazovu, kreiraćete bazu podataka MongoDB Atlas i dostaviti pakete potrebne za povezivanje s njom.

Pratite <a href='https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/' target="_blank" rel="noopener noreferrer nofollow"> ove lekcije </a> da inicijalizujete hostovanu bazu podataka na MongoDB Atlas.

# --instructions--

`mongoose@^5.11.15` Dodato je u fajl vašeg projekta `package.json`. Prvo, mora biti mongoose kao `mongoose` u `myApp.js`. Zatim, kreirajte fajl `.env` i dodajte parametar `MONGO_URI` unutar njega. Njegova vrednost bi trebalo da bude URI vaše baze podataka MongoDB Atlas. Uverite se da obuhvatate URI sa jednostrukim ili dvostrukim navodnicima, i zapamtite da ne možete koristiti razmake blizu `=` u varijablama okruženja. Na primer, `MONGO_URI='VALUE'`.

Kada naiđete na ovaj zadatak, povežite se sa bazom podataka pozivanjem puta `connect` unutar vaše datoteke `myApp.js` koristeći sledeću sintaksu:

```js
mongoose.connect(<Your URI>, { useNewUrlParser: true, useUnifiedTopology: true });
```

# --hints--

"mongoose version ^5.11.15" zavisnost mora biti u package.json

```js
  const response = await fetch(code + '/_api/file/package.json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  var packJson = JSON.parse(data);
  assert.property(packJson.dependencies, 'mongoose');
  assert.match(
    packJson.dependencies.mongoose,
    /^\^5\.11\.15/,
    'Wrong version of "mongoose". It should be ^5.11.15'
  );
```

"mongoose" treba da se poveže sa bazom podataka

```js
  const response = await fetch(code + '/_api/is-mongoose-ok');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.isMongooseOk, 'mongoose is not connected');
```
