---
id: 587d7fb7367417b2b2512c0c
title: Koristi `model.findOne()` da vrati jedan dokument koji odgovara iz vaše baze podataka
challengeType: 2
forumTopicId: 301545
dashedName: use-model-findone-to-return-a-single-matching-document-from-your-database
---

# --description--

`Model.findOne()` radi kao `Model.find()`, ali vraća samo jedan dokument (nije red podataka), čak i kada postoji mnogo stavki. To je posebno važno prilikom pretraživanja po kriterijumima koje ste deklarisali kao jedinstvene.

# --instructions--

Promeni funkciju `findOneByFood` ili pronađi samo jednu osobu koja voli određenu hranu, koristeći `Model.findOne() -> Person`. Koristi argument funkcije `food` kao ključ pretrage.

# --hints--

Dobijanje jednog elementa bi trebalo da uspe

```js
  const response = await fetch(code + '/_api/find-one-by-food', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Gary',
      age: 46,
      favoriteFoods: ['chicken salad']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Gary', 'item.name is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['chicken salad'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```
