---
id: 587d7fb7367417b2b2512c0a
title: Kreiranje mnogo zapisa koristeći model.create()
challengeType: 2
forumTopicId: 301537
dashedName: create-many-records-with-model-create
---

# --description--

Ponekad je potrebno kreirati mnogo instanci vaših modela, na primer prilikom inicijalizacije baze podataka sa početnim podacima. `Model.create()` prima niz objekata kao `[{name: 'John', ...}, {...}, ...]` kao prvi argument i sačuva sve u bazu podataka.

# --instructions--

Promeni funkciju `createManyPeople` ili kreiraj mnogo korisnika koristeći `Model.create()` i argumente `arrayOfPeople`.

**Napomena:** Možete ponovo koristiti model koji ste kreirali u prethodnoj vežbi.

# --hints--

Istovremeno kreiranje više objekata baze podataka trebalo bi da uspe.

```js
  const response = await fetch(code + '/_api/create-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'John', age: 24, favoriteFoods: ['pizza', 'salad'] },
      { name: 'Mary', age: 21, favoriteFoods: ['onions', 'chicken'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an array');
  assert.equal(
    data.length,
    2,
    'the response does not contain the expected number of items'
  );
  assert.equal(data[0].name, 'John', 'The first item is not correct');
  assert.equal(
    data[0].__v,
    0,
    'The first item should be not previously edited'
  );
  assert.equal(data[1].name, 'Mary', 'The second item is not correct');
  assert.equal(
    data[1].__v,
    0,
    'The second item should be not previously edited'
  );
```
