---
id: 587d7fb7367417b2b2512c0d
title: Koristi model.findById() za pretraživanje vaše baze podataka po _id
challengeType: 2
forumTopicId: 301544
dashedName: use-model-findbyid-to-search-your-database-by-id
---

# --description--

Kada se čuva dokument, MongoDB automatski dodaje polje `_id` i čini ga jedinstvenim ključem od slova i brojeva. Pretraživanje za `_id` je vrlo često ponavljana operacija, pa Mongoose nudi posebnu metodu za to.

# --instructions--

Izmenite ``findPersonById`` da traži samo osobu koja ima određeni ``_id``, koristeći ``Model.findById() -> Person``. Koristite argument funkcije ``personId`` kao ključ pretrage.

# --hints--

Pretraga elementa po ID-u bi trebalo da bude uspešna

```js
  const response = await fetch(code + '/_api/find-by-id');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'test', 'item.name is not what expected');
  assert.equal(data.age, 0, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['none'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```
