---
id: 587d7fb6367417b2b2512c07
title: Na primer
challengeType: 2
forumTopicId: 301535
dashedName: create-a-model
---

# --description--

**C**RUD Dio 1 - UNDA

Prvo svega, nam treba dijagram šeme. Svaka šema odgovara kolekciji MongoDB. Ona određuje strukturu dokumenata unutar te kolekcije. Šema je osnovni deo Modela (Models). Može se organizovati unutar da bi stvorila složene strukture, ali u ovom slučaju, zadržaćemo stvari jednostavnim. Na primer, omogućava vam da kreirate instance vaših objekata, koji se nazivaju dokumenti.

Na serveru, interakcija sa bazom podataka se dešava u *handler* funkcijama. Ove funkcije se izvršavaju kada dođe određeni događaj (primer, kada neko klikne na link na vašem API). Pratićemo isti metod u ovim vežbama. Funkcija `done()` je povratni poziv koji nam govori da možemo nastaviti nakon završetja asinhronih operacija kao što su unos, pretraga, ispravljanje ili brisanje. Prati ugovor Node-a i treba ga pozvati kao `done(null, data)` kada je uspešno, ili `done(err)` kada postoji greška.

Upozorenje - Kada koristite udaljene servise, mogući su greške!

```js
/* Example */

const someFunc = function(done) {
  //... do something (risky) ...
  if (error) return done(error);
  done(null, result);
};
```

# --instructions--

Kreirajte šemu entiteta pod nazivom `personSchema` sa sledećom strukturom:

* Deo `name` koji je potreban tipa `String`
* Deo `age` koji je potreban tipa `Number`
* Deo `favoriteFoods` koji je potreban tipa `[String]`

Koristite osnovne tipove šema za Mongoose. Ako želite, možete dodati i više polja, koristite jednostovne validatore kao što su required ili unique, i postaviti podrazumevane vrednosti. Pogledajte <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">makala yetu ya Mongoose</a>.

Sada, kreiraj primer iz `personSchema` i postavi njegovu vrednost u postojeći parametar `Person`.

# --hints--

Kreiranje događaja iz `mongoose` šeme bi trebalo da uspe

```js
  const response = await fetch(code + '/_api/mongoose-model', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Mike',
      age: 28,
      favoriteFoods: ['pizza', 'cheese']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Mike', '"model.name" is not what expected');
  assert.equal(data.age, '28', '"model.age" is not what expected');
  assert.isArray(
    data.favoriteFoods,
    '"model.favoriteFoods" is not an Array'
  );
  assert.include(
    data.favoriteFoods,
    'pizza',
    '"model.favoriteFoods" does not include the expected items'
  );
  assert.include(
    data.favoriteFoods,
    'cheese',
    '"model.favoriteFoods" does not include the expected items'
  );
```
