---
id: 587d7fb8367417b2b2512c0e
title: Izvršite standardno ažuriranje pokretanjem Traži, Izmeni, zatim Sačuvaj'
challengeType: 2
forumTopicId: 301541
dashedName: perform-classic-updates-by-running-find-edit-then-save
---

# --description--

Ranije, ovo je ono što ste trebali da uradite ako želite da izmenite dokument i da ga koristite na neki način (primer, ponovno slanje u server odgovor). Mongoose ima poseban način za ažuriranje: `Model.update()`. On je povezan sa niske razine Mongo drajverom. Može da izmeni više dokumenata koji zadovoljavaju određene kriterijume, ali više ne koristi ažurirani dokument, već samo poruku 'stanje'. Štaviše, čini validaciju modela komplikovanom jer direktno poziva Mongo drajver.

# --instructions--

Promeni funkciju ``findEditThenSave`` ili pretražiti osobu po ``_id`` (koristi bilo koju od gore navedenih metoda) koristeći parametar ``personId`` kao ključ pretrage. Dodaj ``"hamburger"`` na listu ``favoriteFoods`` te osobe (možeš koristiti ``Array.push()``). Zatim - unutar poziva za pretragu - ažurirani ``save()` `Person``.

**Napomena:** Ovo može biti teško, ako u vašem Šemi (Schema), deklarisali ste `favoriteFoods` kao polje podataka, bez definisanja tipa (tj. `[String]`). U tom slučaju, `favoriteFoods` se koristi kao tip Mixed, i morate ručno označiti koristeći `document.markModified('edited-field')`. Pogledajte naš članak o Mongoose-u na <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">...</a>.

# --hints--

Traži-uredi-ažuriraj element, treba da uspe

```js
  const response = await fetch(code + '/_api/find-edit-save', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Poldo',
      age: 40,
      favoriteFoods: ['spaghetti']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Poldo', 'item.name is not what is expected');
  assert.equal(data.age, 40, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['spaghetti', 'hamburger'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 1, 'The item should be previously edited');
```
