---
id: 587d7fb7367417b2b2512c0b
title: Koristi model.find() za pretraživanje u svojoj bazi podataka
challengeType: 2
forumTopicId: 301543
dashedName: use-model-find-to-search-your-database
---

# --description--

U svojoj najjednostavnijoj upotrebi, `Model.find()` prima objekat pretrage (stvar od JSON) kao prvi argument, a zatim poziv za dohvaćanje. Vraća niz podataka rezultata koji odgovaraju. Podržava mnogo vrsta opcija pretrage. Pročitajte više u dokumentaciji.

# --instructions--

Promenite funkciju `findPeopleByName` ili dobijete sve osobe sa određenim imenom, koristeći <code>Model.find() -\> [Person]</code>

Koristi akcionu karticu `personName` kao ključ pretrage.

# --hints--

Pronalaženje svih stavki koje odgovaraju kriterijumima trebalo bi da uspe.

```js
  const response = await fetch(code + '/_api/find-all-by-name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'r@nd0mN4m3',
      age: 24,
      favoriteFoods: ['pizza']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data[0].name,
    'r@nd0mN4m3',
    'item.name is not what expected'
  );
  assert.equal(data[0].__v, 0, 'The item should be not previously edited');
```
