---
id: 587d7fb8367417b2b2512c11
title: Obrišti više dokumenata koristeći model.remove()
challengeType: 2
forumTopicId: 301538
dashedName: delete-many-documents-with-model-remove
---

# --description--

`Model.remove()` Važno je obrisati sve dokumente koji zadovoljavaju navedenim kriterijumima.

# --instructions--

Promeni funkciju `removeManyPeople` da obriše sve entitete čiji je naziv unutar varijable `nameToRemove`, koristeći `Model.remove()`. Prosledi u izraz za pretragu koji ima polje `name` postavljeno, zajedno sa povratnim pozivom.

**Napomena:** `Model.remove()` ne vraća izbrisan dokument, već objekat JSON sa rezultatima operacije i brojem dodirnutih stavki. Nemoj zaboraviti da ga proslediš sa pozivom za vraćanje `done()`, jer ga koristimo u testovima.

# --hints--

Brisanje više stavki mora da uspe

```js
  const response = await fetch(code + '/_api/remove-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Mary', age: 16, favoriteFoods: ['lollipop'] },
      { name: 'Mary', age: 21, favoriteFoods: ['steak'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(!!data.ok, 'The mongo stats are not what expected');
  assert.equal(
    data.n,
    2,
    'The number of items affected is not what expected'
  );
  assert.equal(data.count, 0, 'the db items count is not what expected');
```
