---
id: 587d7fb8367417b2b2512c10
title: Obriši jedan zapis koristeći model.findByIdAndRemove
challengeType: 2
forumTopicId: 301539
dashedName: delete-one-document-using-model-findbyidandremove
---

# --description--

`findByIdAndRemove` na `findOneAndRemove` su kao prethodne metode ažuriranja. Prolazak je kroz izbrisani zapis u bazi podataka. Kao uobičajeno, koristite argument funkcije `personId` kao ključ pretrage.

# --instructions--

Promenite funkciju za `removeById` brisanje jednog korisnika za `_id` tog korisnika. Trebalo bi da koristite jednu od metoda `findByIdAndRemove()` ili `findOneAndRemove()`.

# --hints--

Brisanje elementa bi trebalo da uspe

```js
  const response = await fetch(code + '/_api/remove-one-person', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Jason Bourne',
      age: 36,
      favoriteFoods: ['apples']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Jason Bourne', 'item.name is not what expected');
  assert.equal(data.age, 36, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['apples'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0);
  assert.equal(data.count, 0, 'the db items count is not what expected');
```
