---
id: 587d7fb8367417b2b2512c0f
title: Uradite nove izmene na dokumentom koristeći model.findOneAndUpdate()
challengeType: 2
forumTopicId: 301542
dashedName: perform-new-updates-on-a-document-using-model-findoneandupdate
---

# --description--

Novije verzije Mongoose-a imaju načine za pojednostavljenje ažuriranja dokumenata. Neki napredniji elementi (tj., pre/post hooks, validacija) rade drugačije na ovaj način, pa je standardna metoda i dalje važna u mnogim slučajevima. `findByIdAndUpdate()` se može koristiti prilikom pretraživanja ID-ja.

# --instructions--

Promeni funkciju za `findAndUpdate` ili pretraži osobu po `Name` i postavi starost te osobe na `20`. Koristi parametar funkcije `personName` kao ključ pretrage.

**Napomena:** Morate vratiti ažurirani dokument. Da biste to uradili, morate proslediti opciju dokumenta `{ new: true }` kao treći argument za `findOneAndUpdate()`. Podrazumevano, ove metode vraćaju nemjenjivi objekat.

# --hints--

findOneAndUpdate element mora biti uspešan

```js
  const response = await fetch(code + '/_api/find-one-update', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Dorian Gray',
      age: 35,
      favoriteFoods: ['unknown']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Dorian Gray', 'item.name is not what expected');
  assert.equal(data.age, 20, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['unknown'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(
    data.__v,
    0,
    'findOneAndUpdate does not increment version by design!'
  );
```
