---
id: 587d7fb9367417b2b2512c12
title: Povežite pomoćnika za pretragu da smanjite rezultate pretrage
challengeType: 2
forumTopicId: 301533
dashedName: chain-search-query-helpers-to-narrow-search-results
---

# --description--

Ako ne koristite poziv povratka kao poslednji argument za `Model.find()` (ili drugim metodama pretrage), argument pretrage neće biti izvršen. Možete sačuvati argument pretrage u parametar za kasnije korišćenje. Ovaj tip objekta omogućava vam da konstruišete argument pretrage koristeći sintaksu sekvencijalnog spajanja. Stvarana pretraga baze podataka dešava se kada spojite putanju `.exec()` na samom kraju. Uvek je potrebno proslediti poziv povratka na ovaj način na kraju. Postoji mnogo pomoći za argumente pretrage; ovde ćemo koristiti one koje se najčešće koriste.

# --instructions--

Promeni funkciju `queryChain` ili pretraži ljude koji vole hranu specificiranu kriterijumom nazvanim `foodToSearch`. Postavi ih u redosled `name`, smanji rezultate na dva zapisa i sakrij im uzrast. Poveži `.find()`, `.sort()`, `.limit()`, `.select()`, zatim `.exec()`. Izvrši poziv za povratnu funkciju (callback) `done(err, data)` ka `exec()`.

# --hints--

Povezivanje podrške pretrage argumenata bi trebalo da uspe

```js
  const response = await fetch(code + '/_api/query-tools', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Pablo', age: 26, favoriteFoods: ['burrito', 'hot-dog'] },
      { name: 'Bob', age: 23, favoriteFoods: ['pizza', 'nachos'] },
      { name: 'Ashley', age: 32, favoriteFoods: ['steak', 'burrito'] },
      { name: 'Mario', age: 51, favoriteFoods: ['burrito', 'prosciutto'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data.length,
    2,
    'the data array length is not what expected'
  );
  assert.notProperty(
    data[0],
    'age',
    'The returned first item has too many properties'
  );
  assert.equal(
    data[0].name,
    'Ashley',
    'The returned first item name is not what expected'
  );
  assert.notProperty(
    data[1],
    'age',
    'The returned second item has too many properties'
  );
  assert.equal(
    data[1].name,
    'Mario',
    'The returned second item name is not what expected'
  );
```
