---
id: 587d7fb6367417b2b2512c09
title: Kreiraj i sačuvaj rekord primjera
challengeType: 2
forumTopicId: 301536
dashedName: create-and-save-a-record-of-a-model
---

# --description--

U ovom izazovu ćete morati da kreirate i sačuvate primer zapis (ili: zapis primera).

# --instructions--

Unutar funkcije `createAndSavePerson`, kreirajte instancu dokumenta koristeći model komponentu `Person` koju ste prethodno kreirali. Prođite tu komponentu sa delovima/sekcijama `name`, `age` i `favoriteFoods`. Njihovi tipovi moraju odgovarati onima koji postoje u `personSchema`. Zatim, pozovite metodu `document.save()` na vraćenu instancu dokumenta. Prođite tako poziv za preuzimanje koristeći Node ugovor (contract). Ovo je uobičajen obrazac; sve naredne CRUD metode uzimaju funkciju odgovora kao poslednji argument.

```js
/* Example */

// ...
person.save(function(err, data) {
  //   ...do your stuff here...
});
```

# --hints--

Kreiranje i čuvanje zapisa baze podataka treba da uspe

```js
  const response = await fetch(code + '/_api/create-and-save-person');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isString(data.name, '"item.name" should be a String');
  assert.isNumber(data.age, '28', '"item.age" should be a Number');
  assert.isArray(
    data.favoriteFoods,
    '"item.favoriteFoods" should be an Array'
  );
  assert.equal(data.__v, 0, 'The db item should be not previously edited');
```
