---
id: 63507c19151201368ee3e16c
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Vrednost polja za unos je uvek niz znakova, čak i ako je tip unosa `number`. Morate konvertovati ovaj niz znakova u numerički niz. Da biste to učinili, možete koristiti metodu `.map()`.

Kreiraj razliku od `numbers` i postavi njenu vrednost na `array.map()`. Zapamti da `.map()` kreira novi red, umesto modifikovanja prvobitnog reda.

# --hints--

Vaša funkcija za `calculate` treba da ima razliku od `numbers`.

```js
assert.match(calculate.toString(), /(let|const|var)\s+numbers/);
```

Trebalo bi da koristiš putanju `.map()` na svojoj verziji `array`.

```js
assert.match(calculate.toString(), /array\.map\(\s*\)/);
```

Treba da umetnete rezultate vašeg poziva za `.map()` u vašu varijantu od `numbers`.

```js
assert.match(calculate.toString(), /numbers\s*=\s*array\.map()/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
}
--fcc-editable-region--
```
