---
id: 661899113600ef307d8a7487
title: Korak 23
challengeType: 0
dashedName: step-23
---

# --description--

Da biste dobili medijanu skupa podataka sa neparnim brojem elemenata, moraćete da potražite i vratite medijanu.

Ovo je kako dobijete medijan (srednju vrednost) niza podataka koji ima neparan broj elemenata:

```js
arr[Math.floor(arr.length / 2)];
```

Evo još dužeg primera za dobijanje srednje vrednosti niza podataka sa elementima `5`:

```js
const numbers = [1, 2, 3, 4, 5];
const middleNumber = numbers[Math.floor(numbers.length / 2)];
console.log(middleNumber); // 3
```

Razlog za korišćenje `Math.floor` je zato što će rezultat deljenja nepravilnog broja sa `2` biti decimalni broj. `Math.floor` će zaokružiti nadole do najbližeg celog broja.

Objavi kriterijum za `oddListMedian` i unesi njegovu vrednost rezultata dobijanja srednjeg broja od `testArr1`. Zatim napiši kriterijum za `oddListMedian` u konzoli.

Otvorite konzolu da biste videli rezultate.

# --hints--

Trebalo bi da imate kriterijum za `oddListMedian`.

```js
assert.isDefined(oddListMedian);
```

Treba da unesete rezultat dobijanja srednjeg broja od `testArr1` u parametar `oddListMedian`. Zapamtite da se osvrnete na prvobitni primer: `arr[Math.floor(arr.length / 2)];`

```js
assert.strictEqual(oddListMedian, 3);
```

Trebalo bi da ima konzolnu poruku koja piše parametar `oddListMedian`.

```js
assert.match(code, /console\.log\(\s*oddListMedian\s*\);?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
    <form onsubmit="calculate(); return false;">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

--fcc-editable-region--
const testArr1 = [1, 2, 3, 4, 5];
const testArr2 = [1, 2, 3, 4, 5, 6];
const isEven = testArr2.length % 2 === 0;
console.log(isEven);


--fcc-editable-region--
const getMedian = (array) => {
  const sorted = array.sort((a, b) => a - b);
}

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));
  
  const mean = getMean(numbers);

  document.querySelector("#mean").textContent = mean;
}
```
