---
id: 56105e7b514f539506016a5e
title: Broj Unazad Pomocu for Petlje
challengeType: 1
forumTopicId: 16808
dashedName: count-backwards-with-a-for-loop
---

# --description--

for petlja moze i da broji unazad, dokle god pravilno postavimo uslove.

Da bismo smanjivali za dva u svakoj iteraciji, potrebno je da promenimo inicijalizaciju, uslov i finalni izraz.

Pocecemo sa `i = 10` i nastavicemo petlju dok je `i > 0`. Smanjivacemo `i` za 2 u svakoj iteraciji koristeci `i -= 2`.

```js
const ourArray = [];

for (let i = 10; i > 0; i -= 2) {
  ourArray.push(i);
}
```

`ourArray` ce sada biti `[10, 8, 6, 4, 2]`. Hajde da promenimo inicijalizaciju i finalni izraz kako bismo brojali unazad po neparnim brojevima.

# --instructions--

Ubacite neparne brojeve od 9 do 1 u `myArray` koristeci `for` petlju.

# --hints--

Za ovo treba da koristite `for` petlju.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

Treba da koristite niz metodu `push`.

```js
assert(__helpers.removeJSComments(code).match(/myArray.push/));
```

`myArray` treba da bude jednak `[9, 7, 5, 3, 1]`.

```js
assert.deepEqual(myArray, [9, 7, 5, 3, 1]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 9; i > 0; i -= 2) {
  myArray.push(i);
}
```
