---
id: 56104e9e514f539506016a5c
title: Iteriraj Neparne Brojeve Pomocu for Petlje
challengeType: 1
forumTopicId: 18212
dashedName: iterate-odd-numbers-with-a-for-loop
---

# --description--

for petlje ne moraju iterirati jedan po jedan korak. Menjanjem `final-expression` dela mozemo brojati po parnim brojevima.

Pocinjemo sa `i = 0` i petlja se izvrsava dok je `i < 10`. Povecavacemo `i` za 2 u svakoj iteraciji koristeci `i += 2`.

```js
const ourArray = [];

for (let i = 0; i < 10; i += 2) {
  ourArray.push(i);
}
```

`ourArray` ce sada biti `[0, 2, 4, 6, 8]`. Hajde da promenimo `initialization` kako bismo brojali neparne brojeve.

# --instructions--

Ubacite neparne brojeve od 1 do 9 u `myArray` koristeci `for` petlju.

# --hints--

Za ovo treba da koristite `for` petlju.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

`myArray` treba da bude jednak `[1, 3, 5, 7, 9]`.

```js
assert.deepEqual(myArray, [1, 3, 5, 7, 9]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 1; i < 10; i += 2) {
  myArray.push(i);
}
```
