---
id: 56533eb9ac21ba0edf2244cd
title: Pristup Ugnjezdenim Nizovima
challengeType: 1
forumTopicId: 16160
dashedName: accessing-nested-arrays
---

# --description--

Kao sto smo videli u prethodnim primerima, objekti mogu imati ugnjezdene objekte i ugnjezdene nizove. Kao i kod pristupa ugnjezdenim objektima, notacija niza sa uglastim zagradama moze se kombinovati za pristup ugnjezdenim nizovima.

Evo primera kako se pristupa ugnjezdenom nizu:

```js
const ourPets = [
  {
    animalType: "cat",
    names: [
      "Meowzer",
      "Fluffy",
      "Kit-Cat"
    ]
  },
  {
    animalType: "dog",
    names: [
      "Spot",
      "Bowser",
      "Frankie"
    ]
  }
];

ourPets[0].names[1];
ourPets[1].names[0];
```

`ourPets[0].names[1]` ce biti string `Fluffy`, a `ourPets[1].names[0]` ce biti string `Spot`.

# --instructions--

Koristeci tacka i uglastu notaciju, postavite `secondTree` na drugi element niza `list` iz drugog objekta u nizu `myPlants`.

# --hints--

`secondTree` treba da bude jednak stringu `pine`.

```js
assert(secondTree === 'pine');
```

Vas kod treba da koristi i tacka i uglastu notaciju za pristup `myPlants`.

```js
assert(/=\s*myPlants\[1\].list\[1\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = "";
```

# --solutions--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = myPlants[1].list[1];
```
