---
id: 56533eb9ac21ba0edf2244b6
title: Escape Sekvence u Niskama
challengeType: 1
forumTopicId: 17567
dashedName: escape-sequences-in-strings
---

# --description--

Navodnici nisu jedini znakovi koje mozete <dfn>escape-ovati</dfn> u stringu. Escape sekvence vam omogucavaju da koristite znakove koje inace ne biste mogli direktno da upisete.

<table><thead><tr><th>Kod</th><th>Rezultat</th></tr></thead><tbody><tr><td><code>\'</code></td><td>jednostruki navodnik</td></tr><tr><td><code>\"</code></td><td>dvostruki navodnik</td></tr><tr><td><code>\\</code></td><td>obrnuta kosa crta</td></tr><tr><td><code>\n</code></td><td>novi red</td></tr><tr><td><code>\t</code></td><td>tabulator</td></tr><tr><td><code>\r</code></td><td>povratak na početak reda</td></tr><tr><td><code>\b</code></td><td>brisanje unazad (backspace)</td></tr><tr><td><code>\f</code></td><td>prelazak na novu stranicu (form feed)</td></tr></tbody></table>

*Napomena: obrnuta kosa crta mora biti escape-ovana da bi se prikazala kao obrnuta kosa crta.*

# --instructions--

Postavite sledece tri linije teksta u jednu promenljivu `myStr` koristeci escape sekvence.

<pre>
FirstLine
    \SecondLine
ThirdLine
</pre>

Moracete da koristite escape sekvence da tacno unesete specijalne znakove. Takodje morate pratiti razmake kao gore, bez dodatnih razmaka izmedju sekvenci i reci.

**Napomena:** Uvlacenje ispred `SecondLine` postize se tab karakterom, ne obicnim razmakom.

# --hints--

`myStr` ne treba da sadrzi nijedan obican razmak

```js
assert(!/ /.test(myStr));
```

`myStr` treba da sadrzi stringove `FirstLine`, `SecondLine` i `ThirdLine` (vodite racuna o velikim i malim slovima)

```js
assert(
  /FirstLine/.test(myStr) && /SecondLine/.test(myStr) && /ThirdLine/.test(myStr)
);
```

`FirstLine` treba da bude pracen novim redom `\n`

```js
assert(/FirstLine\n/.test(myStr));
```

`myStr` treba da ima tab karakter `\t` odmah nakon novog reda

```js
assert(/\n\t/.test(myStr));
```

`SecondLine` treba da bude ispredjena obrnutom kosom crtom `\`

```js
assert(/\\SecondLine/.test(myStr));
```

Treba da postoji novi red izmedju `SecondLine` i `ThirdLine`

```js
assert(/SecondLine\nThirdLine/.test(myStr));
```

`myStr` treba da sadrzi samo karaktere prikazane u uputstvu

```js
assert(myStr === 'FirstLine\n\t\\SecondLine\nThirdLine');
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "FirstLine\n\t\\SecondLine\nThirdLine";
```
