---
id: 56533eb9ac21ba0edf2244d1
title: Poredjenje sa Operatorom Stroge Jednakosti
challengeType: 1
forumTopicId: 16790
dashedName: comparison-with-the-strict-equality-operator
---

# --description--

Stroga jednakost (`===`) je slicna operatoru jednakosti (`==`). Medjutim, za razliku od operatora jednakosti koji pokusava da konvertuje obe vrednosti u isti tip, operator stroge jednakosti ne radi konverziju tipa.

Ako poredjene vrednosti imaju razlicite tipove, smatraju se nejednakim i operator stroge jednakosti vraca false.

**Mifano**

```js
3 ===  3  // true
3 === '3' // false
```

U drugom primeru, `3` je tipa `Number`, a `'3'` tipa `String`.

# --instructions--

Koristite operator stroge jednakosti u `if` iskazu tako da funkcija vraca string `Equal` kada je `val` striktno jednako `7`.

# --hints--

`testStrict(10)` treba da vrati string `Not Equal`

```js
assert(testStrict(10) === 'Not Equal');
```

`testStrict(7)` treba da vrati string `Equal`

```js
assert(testStrict(7) === 'Equal');
```

`testStrict("7")` treba da vrati string `Not Equal`

```js
assert(testStrict('7') === 'Not Equal');
```

Treba da koristite operator `===`

```js
assert(__helpers.removeJSComments(code).match(/(val\s*===\s*\d+)|(\d+\s*===\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrict(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testStrict(10);
```

# --solutions--

```js
function testStrict(val) {
  if (val === 7) {
    return "Equal";
  }
  return "Not Equal";
}
```
