---
id: 56533eb9ac21ba0edf2244b1
title: Slozena Dodela sa Uvecanim Mnozenjem
challengeType: 1
forumTopicId: 16662
dashedName: compound-assignment-with-augmented-multiplication
---

# --description--

Operator `*=` mnozi promenljivu brojem.

```js
myVar = myVar * 5;
```

ce pomnoziti `myVar` sa `5`. Ovo se moze prepisati kao:

```js
myVar *= 5;
```

# --instructions--

Promenite dodele vrednosti za `a`, `b` i `c` tako da koriste operator `*=`.

# --hints--

`a` treba da bude jednak `25`.

```js
assert(a === 25);
```

`b` treba da bude jednak `36`.

```js
assert(b === 36);
```

`c` treba da bude jednak `46`.

```js
assert(c === 46);
```

Treba da koristite operator `*=` za svaku promenljivu.

```js
assert(__helpers.removeJSComments(code).match(/\*=/g).length === 3);
```

Ne treba da menjate kod iznad navedenog komentara.

```js
assert(
  /let a = 5;/.test(__helpers.removeJSComments(code)) &&
    /let b = 12;/.test(__helpers.removeJSComments(code)) &&
    /let c = 4\.6;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 12;
let c = 4.6;

// Only change code below this line
a = a * 5;
b = 3 * b;
c = c * 10;
```

# --solutions--

```js
let a = 5;
let b = 12;
let c = 4.6;

a *= 5;
b *= 3;
c *= 10;
```
