---
id: 56533eb9ac21ba0edf2244b2
title: Slozena Dodela sa Uvecanim Deljenjem
challengeType: 1
forumTopicId: 16659
dashedName: compound-assignment-with-augmented-division
---

# --description--

Operator `/=` deli promenljivu drugim brojem.

```js
myVar = myVar / 5;
```

Podelice `myVar` sa `5`. Ovo se moze prepisati kao:

```js
myVar /= 5;
```

# --instructions--

Promenite dodele vrednosti za `a`, `b` i `c` tako da koriste operator `/=`.

# --hints--

`a` treba da bude jednak `4`.

```js
assert(a === 4);
```

`b` treba da bude jednak `27`.

```js
assert(b === 27);
```

`c` treba da bude jednak `3`.

```js
assert(c === 3);
```

Treba da koristite operator `/=` za svaku promenljivu.

```js
assert(__helpers.removeJSComments(code).match(/\/=/g).length === 3);
```

Ne treba da menjate kod iznad navedenog komentara.

```js
assert(
  /let a = 48;/.test(__helpers.removeJSComments(code)) &&
    /let b = 108;/.test(__helpers.removeJSComments(code)) &&
    /let c = 33;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 48;
let b = 108;
let c = 33;

// Only change code below this line
a = a / 12;
b = b / 4;
c = c / 11;
```

# --solutions--

```js
let a = 48;
let b = 108;
let c = 33;

a /= 12;
b /= 4;
c /= 11;
```
