---
id: 56533eb9ac21ba0edf2244d6
title: Poredjenje sa Operatorom Manje Od
challengeType: 1
forumTopicId: 16789
dashedName: comparison-with-the-less-than-operator
---

# --description--

Operator manje od (`<`) poredi vrednosti dva broja. Ako je broj sa leve strane manji od broja sa desne strane, vraca `true`. U suprotnom vraca `false`. Kao i operator jednakosti, operator manje od konvertuje tipove podataka pri poredjenju.

**Primeri**

```js
2   < 5 // true
'3' < 7 // true
5   < 5 // false
3   < 2 // false
'8' < 4 // false
```

# --instructions--

Dodajte operator manje od na oznacene linije kako bi return iskazi imali smisla.

# --hints--

`testLessThan(0)` treba da vrati string `Under 25`

```js
assert(testLessThan(0) === 'Under 25');
```

`testLessThan(24)` treba da vrati string `Under 25`

```js
assert(testLessThan(24) === 'Under 25');
```

`testLessThan(25)` treba da vrati string `Under 55`

```js
assert(testLessThan(25) === 'Under 55');
```

`testLessThan(54)` treba da vrati string `Under 55`

```js
assert(testLessThan(54) === 'Under 55');
```

`testLessThan(55)` treba da vrati string `55 or Over`

```js
assert(testLessThan(55) === '55 or Over');
```

`testLessThan(99)` treba da vrati string `55 or Over`

```js
assert(testLessThan(99) === '55 or Over');
```

Treba da koristite operator `<` najmanje dva puta

```js
assert(__helpers.removeJSComments(code).match(/val\s*<\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testLessThan(val) {
  if (val) {  // Change this line
    return "Under 25";
  }

  if (val) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}

testLessThan(10);
```

# --solutions--

```js
function testLessThan(val) {
  if (val < 25) {  // Change this line
    return "Under 25";
  }

  if (val < 55) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}
```
