---
id: 56533eb9ac21ba0edf2244d0
title: Poredjenje sa Operatorom Jednakosti
challengeType: 1
forumTopicId: 16784
dashedName: comparison-with-the-equality-operator
---

# --description--

U JavaScript-u postoji vise <dfn>operatora poredjenja</dfn>. Svi oni vracaju booleansku vrednost `true` ili `false`.

Najosnovniji je operator jednakosti `==`. On poredi dve vrednosti i vraca `true` ako su jednake, odnosno `false` ako nisu. Zapamtite da je jednakost razlicita od dodele (`=`), koja postavlja vrednost sa desne strane operatora u promenljivu sa leve strane.

```js
function equalityTest(myVal) {
  if (myVal == 10) {
    return "Equal";
  }
  return "Not Equal";
}
```

Ako je `myVal` jednako `10`, operator jednakosti vraca `true`, pa se izvrsava kod unutar viticastih zagrada i funkcija vraca `Equal`. U suprotnom, funkcija vraca `Not Equal`. Da bi JavaScript uporedio dva razlicita <dfn>tipa podataka</dfn> (na primer, `numbers` i `strings`), mora da konvertuje jedan tip u drugi. To se zove prisilna konverzija tipa (type coercion). Nakon toga moze poredjenje kao u nastavku:

```js
1   ==  1  // true
1   ==  2  // false
1   == '1' // true
"3" ==  3  // true
```

# --instructions--

Dodajte operator jednakosti na prikazanoj liniji tako da funkcija vraca string `Equal` kada je `val` jednako `12`.

# --hints--

`testEqual(10)` treba da vrati string `Not Equal`

```js
assert(testEqual(10) === 'Not Equal');
```

`testEqual(12)` treba da vrati string `Equal`

```js
assert(testEqual(12) === 'Equal');
```

`testEqual("12")` treba da vrati string `Equal`

```js
assert(testEqual('12') === 'Equal');
```

Treba da koristite operator `==`

```js
assert(__helpers.removeJSComments(code).match(/==/g) && !__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function testEqual(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testEqual(10);
```

# --solutions--

```js
function testEqual(val) {
  if (val == 12) {
    return "Equal";
  }
  return "Not Equal";
}
```
