---
id: 56533eb9ac21ba0edf2244c8
title: Pristup Svojstvima Objekta Pomocu Uglaste Notacije
challengeType: 1
forumTopicId: 16163
dashedName: accessing-object-properties-with-bracket-notation
---

# --description--

Drugi nacin za pristup svojstvima objekta je uglasta notacija (`[]`). Ako svojstvo kojem pristupate ima razmak u imenu, morate koristiti uglastu notaciju.

Ipak, uglastu notaciju mozete koristiti i za svojstva bez razmaka.

Evo primera koriscenja uglaste notacije za citanje svojstva objekta:

```js
const myObj = {
  "Space Name": "Kirk",
  "More Space": "Spock",
  "NoSpace": "USS Enterprise"
};

myObj["Space Name"];
myObj['More Space'];
myObj["NoSpace"];
```

`myObj["Space Name"]` ce biti string `Kirk`, `myObj['More Space']` ce biti string `Spock`, a `myObj["NoSpace"]` ce biti string `USS Enterprise`.

Zapamtite da nazivi svojstava sa razmacima moraju biti pod navodnicima (jednostrukim ili dvostrukim).

# --instructions--

Procitajte vrednosti svojstava `an entree` i `the drink` iz `testObj` koristeci uglastu notaciju i smestite ih u `entreeValue` i `drinkValue`.

# --hints--

`entreeValue` treba da bude string

```js
assert(typeof entreeValue === 'string');
```

Vrednost `entreeValue` treba da bude string `hamburger`

```js
assert(entreeValue === 'hamburger');
```

`drinkValue` treba da bude string

```js
assert(typeof drinkValue === 'string');
```

Vrednost `drinkValue` treba da bude string `water`

```js
assert(drinkValue === 'water');
```

Treba da koristite uglastu notaciju dva puta

```js
assert(__helpers.removeJSComments(code).match(/testObj\s*?\[('|")[^'"]+\1\]/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};

// Only change code below this line
const entreeValue = testObj;   // Change this line
const drinkValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};
const entreeValue = testObj["an entree"];
const drinkValue = testObj['the drink'];
```
