---
id: 56533eb9ac21ba0edf2244da
title: Uvod u else Iskaz
challengeType: 1
forumTopicId: 18207
dashedName: introducing-else-statements
---

# --description--

Kada je uslov u `if` iskazu istinit, sledeci blok koda se izvrsava. Sta se desava kada je uslov netacan? Obicno se ne desi nista. Uz `else` iskaz, moze se izvrsiti alternativni blok koda.

```js
if (num > 10) {
  return "Bigger than 10";
} else {
  return "10 or Less";
}
```

# --instructions--

Spojite `if` iskaze u jedan `if/else` iskaz.

# --hints--

Treba da imate samo jedan `if` iskaz u editoru

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

Treba da koristite `else` iskaz

```js
assert(/else/g.test(__helpers.removeJSComments(code)));
```

`testElse(4)` treba da vrati string `5 or Smaller`

```js
assert(testElse(4) === '5 or Smaller');
```

`testElse(5)` treba da vrati string `5 or Smaller`

```js
assert(testElse(5) === '5 or Smaller');
```

`testElse(6)` treba da vrati string `Bigger than 5`

```js
assert(testElse(6) === 'Bigger than 5');
```

`testElse(10)` treba da vrati string `Bigger than 5`

```js
assert(testElse(10) === 'Bigger than 5');
```

Ne treba da menjate kod iznad ili ispod navedenih komentara.

```js
assert(/let result = "";/.test(__helpers.removeJSComments(code)) && /return result;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function testElse(val) {
  let result = "";
  // Only change code below this line

  if (val > 5) {
    result = "Bigger than 5";
  }

  if (val <= 5) {
    result = "5 or Smaller";
  }

  // Only change code above this line
  return result;
}

testElse(4);
```

# --solutions--

```js
function testElse(val) {
  let result = "";
  if(val > 5) {
    result = "Bigger than 5";
  } else {
    result = "5 or Smaller";
  }
  return result;
}
```
