---
id: 56533eb9ac21ba0edf2244d2
title: Poredjenje sa Operatorom Nejednakosti
challengeType: 1
forumTopicId: 16787
dashedName: comparison-with-the-inequality-operator
---

# --description--

Operator nejednakosti (`!=`) je suprotnost operatoru jednakosti. Nejednakost znaci nije jednako. Operator nejednakosti vraca `false` kada bi operator jednakosti vratio `true`, i obrnuto. Kao i operator jednakosti, operator nejednakosti konvertuje tipove podataka tokom poredjenja.

**Primeri**

```js
1 !=  2    // true
1 != "1"   // false
1 != '1'   // false
1 != true  // false
0 != false // false
```

# --instructions--

Dodajte operator nejednakosti `!=` u `if` izraz tako da funkcija vraca string `Not Equal` kada `val` nije jednako `99`.

# --hints--

`testNotEqual(99)` treba da vrati string `Equal`

```js
assert(testNotEqual(99) === 'Equal');
```

`testNotEqual("99")` treba da vrati string `Equal`

```js
assert(testNotEqual('99') === 'Equal');
```

`testNotEqual(12)` treba da vrati string `Not Equal`

```js
assert(testNotEqual(12) === 'Not Equal');
```

`testNotEqual("12")` treba da vrati string `Not Equal`

```js
assert(testNotEqual('12') === 'Not Equal');
```

`testNotEqual("bob")` treba da vrati string `Not Equal`

```js
assert(testNotEqual('bob') === 'Not Equal');
```

Treba da koristite operator `!=`

```js
assert(__helpers.removeJSComments(code).match(/(?!!==)!=/));
```

# --seed--

## --seed-contents--

```js
// Setup
function testNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testNotEqual(10);
```

# --solutions--

```js
function testNotEqual(val) {
  if (val != 99) {
    return "Not Equal";
  }
  return "Equal";
}
```
