---
id: 56533eb9ac21ba0edf2244cc
title: Pristup Ugnjezdenim Objektima
challengeType: 1
forumTopicId: 16161
dashedName: accessing-nested-objects
---

# --description--

Pod-svojstvima objekata moze se pristupiti kombinovanjem tacka i uglaste notacije.

Evo ugnjezdenog objekta:

```js
const ourStorage = {
  "desk": {
    "drawer": "stapler"
  },
  "cabinet": {
    "top drawer": { 
      "folder1": "a file",
      "folder2": "secrets"
    },
    "bottom drawer": "soda"
  }
};

ourStorage.cabinet["top drawer"].folder2;
ourStorage.desk.drawer;
```

`ourStorage.cabinet["top drawer"].folder2` ce biti string `secrets`, a `ourStorage.desk.drawer` ce biti string `stapler`.

# --instructions--

Pristupite objektu `myStorage` i postavite vrednost sadrzaja svojstva `glove box` u promenljivu `gloveBoxContents`. Koristite tacka notaciju gde god je moguce, a inace uglastu notaciju.

# --hints--

`gloveBoxContents` treba da bude jednak stringu `maps`.

```js
assert(gloveBoxContents === 'maps');
```

Vas kod treba da koristi tacka notaciju, gde je moguce, za pristup `myStorage`.

```js
assert.match(code, /myStorage\.car\.inside/);
```

`gloveBoxContents` i dalje treba da bude deklarisan sa `const`.

```js
assert.match(code, /const\s+gloveBoxContents\s*=/);
```

Ne treba da menjate objekat `myStorage`.

```js
const expectedMyStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
assert.deepStrictEqual(myStorage, expectedMyStorage);
```

# --seed--

## --seed-contents--

```js
const myStorage = {
  "car": {
    "inside": {
      "glove box": "maps",
      "passenger seat": "crumbs"
     },
    "outside": {
      "trunk": "jack"
    }
  }
};

const gloveBoxContents = undefined;
```

# --solutions--

```js
const myStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
const gloveBoxContents = myStorage.car.inside["glove box"];
```
