---
id: 56533eb9ac21ba0edf2244b5
title: Escape Literalnih Navodnika u Niskama
challengeType: 1
forumTopicId: 17568
dashedName: escaping-literal-quotes-in-strings
---

# --description--

Kada definisete string, morate ga zapoceti i zavrsiti jednostrukim ili dvostrukim navodnicima. Sta ako zelite literalni navodnik: `"` ili `'` unutar stringa?

U JavaScript-u mozete da <dfn>escape-ujete</dfn> navodnik tako da se ne smatra zavrsetkom stringa tako sto stavite <dfn>obrnutu kosu crtu</dfn> (`\`) ispred navodnika.

```js
const sampleStr = "Alan said, \"Peter is learning JavaScript\".";
```

Ovo govori JavaScript-u da sledeci navodnik nije kraj stringa, vec treba da se prikaze kao deo stringa. Ako biste to ispisali u konzoli, dobili biste:

```js
Alan said, "Peter is learning JavaScript".
```

# --instructions--

Koristite <dfn>obrnute kose crte</dfn> da postavite string vrednost promenljivoj `myStr` tako da bi ispis u konzoli bio:

```js
I am a "double quoted" string inside "double quotes".
```

# --hints--

Treba da koristite dva ne-escape-ovana dvostruka navodnika (`"`) i cetiri escape-ovana (`\"`).

```js
assert(__helpers.removeJSComments(code).match(/\\"/g).length === 4 && __helpers.removeJSComments(code).match(/[^\\]"/g).length === 2);
```

Promenljiva `myStr` treba da sadrzi string: `I am a "double quoted" string inside "double quotes".`

```js
assert(/I am a "double quoted" string inside "double quotes(\."|"\.)$/.test(myStr));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "I am a \"double quoted\" string inside \"double quotes\".";
```
