---
id: 56533eb9ac21ba0edf2244c0
title: Globalni Naspram Lokalnog Opsega u Funkcijama
challengeType: 1
forumTopicId: 18194
dashedName: global-vs--local-scope-in-functions
---

# --description--

Moguce je imati <dfn>lokalne</dfn> i <dfn>globalne</dfn> promenljive sa istim imenom. Kada se to desi, lokalna promenljiva ima prednost nad globalnom.

U ovom primeru:

```js
const someVar = "Hat";

function myFun() {
  const someVar = "Head";
  return someVar;
}
```

Funkcija `myFun` vraca string `Head` jer postoji lokalna verzija promenljive.

# --instructions--

Dodajte lokalnu promenljivu u funkciji `myOutfit` da prebrisete vrednost `outerWear` stringom `sweater`.

# --hints--

Ne treba da menjate globalnu vrednost `outerWear`.

```js
assert(outerWear === 'T-Shirt');
```

`myOutfit` treba da vrati string `sweater`.

```js
assert(myOutfit() === 'sweater');
```

Ne treba da menjate `return` naredbu.

```js
assert(/return outerWear/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const outerWear = "T-Shirt";

function myOutfit() {
  // Only change code below this line

  // Only change code above this line
  return outerWear;
}

myOutfit();
```

# --solutions--

```js
const outerWear = "T-Shirt";
function myOutfit() {
  const outerWear = "sweater";
  return outerWear;
}
```
