---
id: 56533eb9ac21ba0edf2244e1
title: Ugnjezdene for Petlje
challengeType: 1
forumTopicId: 18248
dashedName: nesting-for-loops
---

# --description--

Ako imate visedimenzionalni niz, mozete primeniti istu logiku kao ranije da prolazite kroz niz i sve njegove podnizove. Evo primera:

```js
const arr = [
  [1, 2], [3, 4], [5, 6]
];

for (let i = 0; i < arr.length; i++) {
  for (let j = 0; j < arr[i].length; j++) {
    console.log(arr[i][j]);
  }
}
```

Ovo ispisuje svaki pod-element iz `arr` redom. Obratite paznju da u unutrasnjoj petlji proveravamo `.length` od `arr[i]`, jer je `arr[i]` sam niz.

# --instructions--

Izmenite funkciju `multiplyAll` tako da vraca proizvod svih brojeva u podnizovima `arr`.

# --hints--

`multiplyAll([[1], [2], [3]])` treba da vrati `6`

```js
assert(multiplyAll([[1], [2], [3]]) === 6);
```

`multiplyAll([[1, 2], [3, 4], [5, 6, 7]])` treba da vrati `5040`

```js
assert(
  multiplyAll([
    [1, 2],
    [3, 4],
    [5, 6, 7]
  ]) === 5040
);
```

`multiplyAll([[5, 1], [0.2, 4, 0.5], [3, 9]])` treba da vrati `54`

```js
assert(
  multiplyAll([
    [5, 1],
    [0.2, 4, 0.5],
    [3, 9]
  ]) === 54
);
```

# --seed--

## --seed-contents--

```js
function multiplyAll(arr) {
  let product = 1;
  // Only change code below this line

  // Only change code above this line
  return product;
}

multiplyAll([[1, 2], [3, 4], [5, 6, 7]]);
```

# --solutions--

```js
function multiplyAll(arr) {
  let product = 1;
  for (let i = 0; i < arr.length; i++) {
    for (let j = 0; j < arr[i].length; j++) {
      product *= arr[i][j];
    }
  }
  return product;
}
```
