---
id: 56533eb9ac21ba0edf2244c9
title: Pristup Svojstvima Objekta Pomocu Promenljivih
challengeType: 1
forumTopicId: 16165
dashedName: accessing-object-properties-with-variables
---

# --description--

Jos jedna upotreba uglaste notacije na objektima je pristup svojstvu koje je sacuvano kao vrednost promenljive. Ovo je posebno korisno pri iteriranju kroz svojstva objekta ili pri dohvatanju vrednosti iz lookup tabele.

Evo primera koriscenja promenljive za pristup svojstvu:

```js
const dogs = {
  Fido: "Mutt",
  Hunter: "Doberman",
  Snoopie: "Beagle"
};

const myDog = "Hunter";
const myBreed = dogs[myDog];
console.log(myBreed);
```

String `Doberman` ce biti ispisan u konzoli.

Obratite paznju da ne koristimo navodnike oko imena promenljive, jer koristimo *vrednost* promenljive, a ne njeno *ime*.

# --instructions--

Postavite vrednost `playerNumber` na `16`. Zatim upotrebite tu promenljivu da dohvatite ime igraca i smestite ga u `player`.

# --hints--

`playerNumber` treba da bude broj

```js
assert(typeof playerNumber === 'number');
```

Promenljiva `player` treba da bude string

```js
assert(typeof player === 'string');
```

Vrednost `player` treba da bude string `Montana`

```js
assert(player === 'Montana');
```

Treba da koristite uglastu notaciju za pristup `testObj`

```js
assert(/testObj\s*?\[.*?\]/.test(__helpers.removeJSComments(code)));
```

Ne treba direktno da dodelite vrednost `Montana` promenljivoj `player`

```js
assert(!__helpers.removeJSComments(code).match(/player\s*=\s*"|\'\s*Montana\s*"|\'\s*;/gi));
```

Treba da koristite promenljivu `playerNumber` unutar uglaste notacije

```js
assert(/testObj\s*?\[\s*playerNumber\s*\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};

// Only change code below this line
const playerNumber = 42;  // Change this line
const player = testObj;   // Change this line
```

# --solutions--

```js
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};
const playerNumber = 16;
const player = testObj[playerNumber];
```
