---
id: 56533eb9ac21ba0edf2244dd
title: Izbor Izmedju Vise Opcija Pomocu switch Iskaza
challengeType: 1
forumTopicId: 18277
dashedName: selecting-from-many-options-with-switch-statements
---

# --description--

Ako treba da uporedite jednu vrednost sa vise opcija, mozete koristiti <dfn>switch</dfn> iskaz. `switch` iskaz poredi vrednost sa <dfn>case</dfn> granama koje predstavljaju moguce vrednosti. Bilo koji validan JavaScript kod moze se izvrsiti unutar <dfn>case</dfn> grane i izvrsavanje traje od prve poklopljene `case` vrednosti do `break` naredbe.

Evo primera `switch` iskaza:

```js
switch (fruit) {
  case "apple":
    console.log("The fruit is an apple");
    break;
  case "orange":
    console.log("The fruit is an orange");
    break;
}
```

`case` vrednosti se proveravaju strogom jednakošcu (`===`). `break` govori JavaScript-u da prekine izvrsavanje. Ako `break` nedostaje, sledeca naredba ce se izvrsiti.

# --instructions--

Napisi `switch` iskaz koji testira `val` i postavlja `answer` prema sledecim uslovima:
`1` - `alpha`
`2` - `beta`
`3` - `gamma`
`4` - `delta`

# --hints--

`caseInSwitch(1)` treba da vrati string `alpha`

```js
assert(caseInSwitch(1) === 'alpha');
```

`caseInSwitch(2)` treba da vrati string `beta`

```js
assert(caseInSwitch(2) === 'beta');
```

`caseInSwitch(3)` treba da vrati string `gamma`

```js
assert(caseInSwitch(3) === 'gamma');
```

`caseInSwitch(4)` treba da vrati string `delta`

```js
assert(caseInSwitch(4) === 'delta');
```

Ne treba da koristite `if` ili `else` iskaze

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Treba da imate najmanje 3 `break` naredbe

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function caseInSwitch(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

caseInSwitch(1);
```

# --solutions--

```js
function caseInSwitch(val) {
  let answer = "";

  switch (val) {
    case 1:
      answer = "alpha";
      break;
    case 2:
      answer = "beta";
      break;
    case 3:
      answer = "gamma";
      break;
    case 4:
      answer = "delta";
  }
  return answer;
}
```
