---
id: 56533eb9ac21ba0edf2244db
title: Uvod u else if Iskaze
challengeType: 1
forumTopicId: 18206
dashedName: introducing-else-if-statements
---

# --description--

Ako imate vise uslova koje treba obraditi, mozete povezati `if` iskaze sa `else if` iskazima.

```js
if (num > 15) {
  return "Bigger than 15";
} else if (num < 5) {
  return "Smaller than 5";
} else {
  return "Between 5 and 15";
}
```

# --instructions--

Izmenite logiku tako da koristi `else if` iskaze.

# --hints--

Treba da imate bar dva `else` iskaza

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 1);
```

Treba da imate bar dva `if` iskaza

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 1);
```

Treba da imate odgovarajuce vitičaste zagrade za svaki `if else` blok.

```js
assert(
  __helpers.removeJSComments(code).match(
    /if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s+if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s*\{[\s\S]+\s*\}/
  )
);
```

`testElseIf(0)` treba da vrati string `Smaller than 5`

```js
assert(testElseIf(0) === 'Smaller than 5');
```

`testElseIf(5)` treba da vrati string `Between 5 and 10`

```js
assert(testElseIf(5) === 'Between 5 and 10');
```

`testElseIf(7)` treba da vrati string `Between 5 and 10`

```js
assert(testElseIf(7) === 'Between 5 and 10');
```

`testElseIf(10)` treba da vrati string `Between 5 and 10`

```js
assert(testElseIf(10) === 'Between 5 and 10');
```

`testElseIf(12)` treba da vrati string `Greater than 10`

```js
assert(testElseIf(12) === 'Greater than 10');
```

# --seed--

## --seed-contents--

```js
function testElseIf(val) {
  if (val > 10) {
    return "Greater than 10";
  }

  if (val < 5) {
    return "Smaller than 5";
  }

  return "Between 5 and 10";
}

testElseIf(7);
```

# --solutions--

```js
function testElseIf(val) {
  if(val > 10) {
    return "Greater than 10";
  } else if(val < 5) {
    return "Smaller than 5";
  } else {
    return "Between 5 and 10";
  }
}
```
