---
id: 56533eb9ac21ba0edf2244dc
title: Povezivanje If Else Iskaza
challengeType: 1
forumTopicId: 16772
dashedName: chaining-if-else-statements
---

# --description--

`if/else` iskazi se mogu povezati za slozeniju logiku. Evo <dfn>pseudokoda</dfn> za vise povezanih `if` / `else if` iskaza:

```js
if (condition1) {
  statement1
} else if (condition2) {
  statement2
} else if (condition3) {
  statement3
. . .
} else {
  statementN
}
```

# --instructions--

Napisi povezane `if`/`else if` iskaze koji ispunjavaju sledece uslove:

`num < 5` - vrati `Tiny`
`num < 10` - vrati `Small`
`num < 15` - vrati `Medium`
`num < 20` - vrati `Large`
`num >= 20` - vrati `Huge`

# --hints--

Treba da imate najmanje cetiri `else` iskaza

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 3);
```

Treba da imate najmanje cetiri `if` iskaza

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 3);
```

Treba da imate najmanje jedan `return` iskaz

```js
assert(__helpers.removeJSComments(code).match(/return/g).length >= 1);
```

`testSize(0)` treba da vrati string `Tiny`

```js
assert(testSize(0) === 'Tiny');
```

`testSize(4)` treba da vrati string `Tiny`

```js
assert(testSize(4) === 'Tiny');
```

`testSize(5)` treba da vrati string `Small`

```js
assert(testSize(5) === 'Small');
```

`testSize(8)` treba da vrati string `Small`

```js
assert(testSize(8) === 'Small');
```

`testSize(10)` treba da vrati string `Medium`

```js
assert(testSize(10) === 'Medium');
```

`testSize(14)` treba da vrati string `Medium`

```js
assert(testSize(14) === 'Medium');
```

`testSize(15)` treba da vrati string `Large`

```js
assert(testSize(15) === 'Large');
```

`testSize(17)` treba da vrati string `Large`

```js
assert(testSize(17) === 'Large');
```

`testSize(20)` treba da vrati string `Huge`

```js
assert(testSize(20) === 'Huge');
```

`testSize(25)` treba da vrati string `Huge`

```js
assert(testSize(25) === 'Huge');
```

# --seed--

## --seed-contents--

```js
function testSize(num) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

testSize(7);
```

# --solutions--

```js
function testSize(num) {
  if (num < 5) {
    return "Tiny";
  } else if (num < 10) {
    return "Small";
  } else if (num < 15) {
    return "Medium";
  } else if (num < 20) {
    return "Large";
  } else {
    return "Huge";
  }
}
```
