---
id: 56533eb9ac21ba0edf2244df
title: Vise Istih Opcija u switch Iskazima
challengeType: 1
forumTopicId: 18242
dashedName: multiple-identical-options-in-switch-statements
---

# --description--

Ako `break` nedostaje u `switch` `case` grani, sledece `case` grane ce se izvrsavati sve dok se ne naidje na `break`. Ako vise ulaza ima isti rezultat, mozete ih predstaviti u `switch` iskazu ovako:

```js
let result = "";
switch (val) {
  case 1:
  case 2:
  case 3:
    result = "1, 2, or 3";
    break;
  case 4:
    result = "4 alone";
}
```

Slucajevi 1, 2 i 3 daju isti rezultat.

# --instructions--

Napisite `switch` iskaz koji postavlja `answer` za sledece opsege:
`1-3` - `Low`
`4-6` - `Mid`
`7-9` - `High`

**Napomena:** Potrebna je po jedna `case` grana za svaki broj u opsegu.

# --hints--

`sequentialSizes(1)` treba da vrati string `Low`

```js
assert(sequentialSizes(1) === 'Low');
```

`sequentialSizes(2)` treba da vrati string `Low`

```js
assert(sequentialSizes(2) === 'Low');
```

`sequentialSizes(3)` treba da vrati string `Low`

```js
assert(sequentialSizes(3) === 'Low');
```

`sequentialSizes(4)` treba da vrati string `Mid`

```js
assert(sequentialSizes(4) === 'Mid');
```

`sequentialSizes(5)` treba da vrati string `Mid`

```js
assert(sequentialSizes(5) === 'Mid');
```

`sequentialSizes(6)` treba da vrati string `Mid`

```js
assert(sequentialSizes(6) === 'Mid');
```

`sequentialSizes(7)` treba da vrati string `High`

```js
assert(sequentialSizes(7) === 'High');
```

`sequentialSizes(8)` treba da vrati string `High`

```js
assert(sequentialSizes(8) === 'High');
```

`sequentialSizes(9)` treba da vrati string `High`

```js
assert(sequentialSizes(9) === 'High');
```

Ne treba da koristite `if` ili `else` iskaze

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Treba da imate devet `case` grana

```js
assert(__helpers.removeJSComments(code).match(/case/g).length === 9);
```

# --seed--

## --seed-contents--

```js
function sequentialSizes(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

sequentialSizes(1);
```

# --solutions--

```js
function sequentialSizes(val) {
  let answer = "";

  switch (val) {
    case 1:
    case 2:
    case 3:
      answer = "Low";
      break;
    case 4:
    case 5:
    case 6:
      answer = "Mid";
      break;
    case 7:
    case 8:
    case 9:
      answer = "High";
  }

  return answer;
}
```
