---
id: 56533eb9ac21ba0edf2244ab
title: Razumevanje Osetljivosti na Mala i Velika Slova u Promenljivim
challengeType: 1
forumTopicId: 18334
dashedName: understanding-case-sensitivity-in-variables
---

# --description--

U JavaScript-u su sve promenljive i imena funkcija osetljivi na mala i velika slova. To znaci da je bitno kako pisete slova.

`MYVAR` nije isto sto i `MyVar` niti `myvar`. Moguce je imati vise razlicitih promenljivih sa slicnim nazivima koji se razlikuju samo po slovima. Radi preglednosti, preporucuje se da to izbegavate.

**Najbolja praksa**

Pisite nazive promenljivih u JavaScript-u koristeci <dfn>camelCase</dfn>. U <dfn>camelCase</dfn> zapisu, nazivi sa vise reci pocinju malim slovom, a svaka sledeca rec pocinje velikim slovom.

**Primeri:**

```js
var someVariable;
var anotherVariableName;
var thisVariableNameIsSoLong;
```

# --instructions--

Promenite postojecu deklaraciju i dodelu vrednosti tako da nazivi koriste <dfn>camelCase</dfn>.

Ne kreirajte nove promenljive.

# --hints--

`studlyCapVar` treba da bude deklarisan i da ima vrednost `10`.

```js
assert(typeof studlyCapVar !== 'undefined' && studlyCapVar === 10);
```

`properCamelCase` treba da bude deklarisan i da ima vrednost stringa `A String`.

```js
assert(
  typeof properCamelCase !== 'undefined' && properCamelCase === 'A String'
);
```

`titleCaseOver` treba da bude deklarisan i da ima vrednost `9000`.

```js
assert(typeof titleCaseOver !== 'undefined' && titleCaseOver === 9000);
```

`studlyCapVar` treba da koristi camelCase i u deklaraciji i u dodeli vrednosti.

```js
assert(__helpers.removeJSComments(code).match(/studlyCapVar/g).length === 2);
```

`properCamelCase` treba da koristi camelCase i u deklaraciji i u dodeli vrednosti.

```js
assert(__helpers.removeJSComments(code).match(/properCamelCase/g).length === 2);
```

`titleCaseOver` treba da koristi camelCase i u deklaraciji i u dodeli vrednosti.

```js
assert(__helpers.removeJSComments(code).match(/titleCaseOver/g).length === 2);
```

# --seed--

## --seed-contents--

```js
// Variable declarations
var StUdLyCapVaR;
var properCamelCase;
var TitleCaseOver;

// Variable assignments
STUDLYCAPVAR = 10;
PRoperCAmelCAse = "A String";
tITLEcASEoVER = 9000;
```

# --solutions--

```js
var studlyCapVar;
var properCamelCase;
var titleCaseOver;

studlyCapVar = 10;
properCamelCase = "A String";
titleCaseOver = 9000;
```
