---
id: 56533eb9ac21ba0edf2244c4
title: Obrazac Ranog Povratka iz Funkcije
challengeType: 1
forumTopicId: 18272
dashedName: return-early-pattern-for-functions
---

# --description--

Kada se naidje na `return` naredbu, izvrsavanje trenutne funkcije staje i kontrola se vraca mestu poziva.

**Mfano**

```js
function myFun() {
  console.log("Hello");
  return "World";
  console.log("byebye")
}
myFun();
```

Ovo ce ispisati string `Hello` u konzoli i vratiti string `World`. String `byebye` nikada nece biti ispisan, jer funkcija izlazi na `return` naredbi.

# --instructions--

Izmenite funkciju `abTest` tako da, ako je `a` ili `b` manje od `0`, funkcija odmah vrati `undefined`.

**Podsetnik**
Zapamtite da je <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-javascript/understanding-uninitialized-variables" target="_blank" rel="noopener noreferrer nofollow"><code>undefined</code> kljucna rec</a>, a ne string.

# --hints--

`abTest(2, 2)` treba da vrati broj

```js
assert(typeof abTest(2, 2) === 'number');
```

`abTest(2, 2)` treba da vrati `8`

```js
assert(abTest(2, 2) === 8);
```

`abTest(-2, 2)` treba da vrati `undefined`

```js
assert(abTest(-2, 2) === undefined);
```

`abTest(2, -2)` treba da vrati `undefined`

```js
assert(abTest(2, -2) === undefined);
```

`abTest(2, 8)` treba da vrati `18`

```js
assert(abTest(2, 8) === 18);
```

`abTest(3, 3)` treba da vrati `12`

```js
assert(abTest(3, 3) === 12);
```

`abTest(0, 0)` treba da vrati `0`

```js
assert(abTest(0, 0) === 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function abTest(a, b) {
  // Only change code below this line



  // Only change code above this line

  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}

abTest(2,2);
```

# --solutions--

```js
function abTest(a, b) {
  if(a < 0 || b < 0) {
    return undefined;
  }
  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}
```
