---
id: 56533eb9ac21ba0edf2244bb
title: Popunjavanje Reci
challengeType: 1
forumTopicId: 18377
dashedName: word-blanks
---

# --description--

Dobicete recenice sa nekoliko nedostajucih reci, kao sto su imenice, glagoli, pridevi i prilozi. Zatim popunjavate prazna mesta recima po izboru tako da zavrsena recenica ima smisla.

Razmotri sledeću rečenicu:

```md
It was really ____, and we ____ ourselves ____.
```

Ova recenica ima tri prazna mesta - pridev, glagol i prilog, i mozemo dodati reci po izboru da je dovrsimo. Zatim gotovu recenicu mozemo smestiti u promenljivu ovako:

```js
const sentence = "It was really " + "hot" + ", and we " + "laughed" + " ourselves " + "silly" + ".";
```

# --instructions--

U ovom izazovu dobijate imenicu, glagol, pridev i prilog. Potrebno je da napravite celu recenicu koristeci reci po izboru, zajedno sa recima koje su vam date.

Potrebno je da koristite `+` operator za spajanje stringova da napravite novi string koristeci date promenljive: `myNoun`, `myAdjective`, `myVerb` i `myAdverb`. Zatim taj string dodelite promenljivoj `wordBlanks`. Ne treba da menjate reci dodeljene tim promenljivim.

Takodje obratite paznju na razmake u stringu, tako da konacna recenica ima razmak izmedju svih reci. Rezultat treba da bude kompletna recenica.

# --before-each--

```js
const removeAssignments = str => str
  .replace(/myNoun\s*=\s*["']dog["']/g, '')
  .replace(/myAdjective\s*=\s*["']big["']/g, '')
  .replace(/myVerb\s*=\s*["']ran["']/g, '')
  .replace(/myAdverb\s*=\s*["']quickly["']/g, '');
```

# --hints--

`wordBlanks` treba da bude string.

```js
assert(typeof wordBlanks === 'string');
```

Ne treba da menjate vrednosti dodeljene promenljivim `myNoun`, `myVerb`, `myAdjective` ili `myAdverb`.

```js
assert(
  myNoun === 'dog' &&
    myVerb === 'ran' &&
    myAdjective === 'big' &&
    myAdverb === 'quickly'
);
```

Ne treba direktno da koristite vrednosti `dog`, `ran`, `big` ili `quickly` za pravljenje `wordBlanks`.

```js
const newCode = removeAssignments(__helpers.removeJSComments(code));
assert(
  !/dog/.test(newCode) &&
    !/ran/.test(newCode) &&
    !/big/.test(newCode) &&
    !/quickly/.test(newCode)
);
```

`wordBlanks` treba da sadrzi sve reci dodeljene promenljivim `myNoun`, `myVerb`, `myAdjective` i `myAdverb`, odvojene ne-rec karakterima (uz bilo koje dodatne reci po izboru).

```js
assert(
  /\bdog\b/.test(wordBlanks) &&
    /\bbig\b/.test(wordBlanks) &&
    /\bran\b/.test(wordBlanks) &&
    /\bquickly\b/.test(wordBlanks)
);
```

# --seed--

## --seed-contents--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

// Only change code below this line
const wordBlanks = ""; // Change this line
// Only change code above this line
```

# --solutions--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

let wordBlanks = "Once there was a " + myNoun + " which was very " + myAdjective + ". ";
wordBlanks += "It " + myVerb + " " + myAdverb + " around the yard.";
```
