---
id: 56533eb9ac21ba0edf2244d8
title: Poredjenja sa Logickim AND Operatorom
challengeType: 1
forumTopicId: 16799
dashedName: comparisons-with-the-logical-and-operator
---

# --description--

Ponekad je potrebno proveriti vise uslova odjednom. <dfn>Logicki AND</dfn> operator (`&&`) vraca `true` samo ako su <dfn>operandi</dfn> i levo i desno od njega tacni.

Isto ponasanje se moze dobiti ugnjezdenjem jednog `if` iskaza unutar drugog.

```js
if (num > 5) {
  if (num < 10) {
    return "Yes";
  }
}
return "No";
```

Ovaj kod vraca `Yes` ako je `num` veci od `5` i manji od `10`. Ista logika moze se napisati koristeci <dfn>logicki AND</dfn> operator.

```js
if (num > 5 && num < 10) {
  return "Yes";
}
return "No";
```

# --instructions--

Zamenite dva `if` iskaza jednim iskazom koji koristi `&&` i vraca string `Yes` ako je `val` manje ili jednako `50` i vece ili jednako `25`. U suprotnom, treba da vrati string `No`.

# --hints--

Treba da koristite operator `&&` jednom

```js
assert(__helpers.removeJSComments(code).match(/&&/g).length === 1);
```

Treba da imate samo jedan `if` iskaz

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalAnd(0)` treba da vrati string `No`

```js
assert(testLogicalAnd(0) === 'No');
```

`testLogicalAnd(24)` treba da vrati string `No`

```js
assert(testLogicalAnd(24) === 'No');
```

`testLogicalAnd(25)` treba da vrati string `Yes`

```js
assert(testLogicalAnd(25) === 'Yes');
```

`testLogicalAnd(30)` treba da vrati string `Yes`

```js
assert(testLogicalAnd(30) === 'Yes');
```

`testLogicalAnd(50)` treba da vrati string `Yes`

```js
assert(testLogicalAnd(50) === 'Yes');
```

`testLogicalAnd(51)` treba da vrati string `No`

```js
assert(testLogicalAnd(51) === 'No');
```

`testLogicalAnd(75)` treba da vrati string `No`

```js
assert(testLogicalAnd(75) === 'No');
```

`testLogicalAnd(80)` treba da vrati string `No`

```js
assert(testLogicalAnd(80) === 'No');
```

# --seed--

## --seed-contents--

```js
function testLogicalAnd(val) {
  // Only change code below this line

  if (val) {
    if (val) {
      return "Yes";
    }
  }

  // Only change code above this line
  return "No";
}

testLogicalAnd(10);
```

# --solutions--

```js
function testLogicalAnd(val) {
  if (val >= 25 && val <= 50) {
    return "Yes";
  }
  return "No";
}
```
