---
id: 56533eb9ac21ba0edf2244ed
title: Dodavanje Promenljivih u String
challengeType: 1
forumTopicId: 16656
dashedName: appending-variables-to-strings
---

# --description--

Kao sto mozemo da gradimo string kroz vise linija koristeci <dfn>string literale</dfn>, mozemo i da dodajemo promenljive u string pomocu plus-equals operatora (`+=`).

Mfano:

```js
const anAdjective = "awesome!";
let ourStr = "freeCodeCamp is ";
ourStr += anAdjective;
```

`ourStr` bi imao vrednost `freeCodeCamp is awesome!`.

# --instructions--

Postavite `someAdjective` na string od najmanje 3 karaktera i dodajte ga u `myStr` koristeci operator `+=`.

# --hints--

`someAdjective` treba da bude postavljen na string od najmanje 3 karaktera.

```js
assert(typeof someAdjective !== 'undefined' && someAdjective.length > 2);
```

Treba da dodate `someAdjective` u `myStr` koristeci operator `+=`.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*someAdjective\s*/).length > 0);
```

# --seed--

## --seed-contents--

```js
// Change code below this line
const someAdjective = "";
let myStr = "Learning to code is ";
```

# --solutions--

```js
const someAdjective = "neat";
let myStr = "Learning to code is ";
myStr += someAdjective;
```
