---
id: 56533eb9ac21ba0edf2244cf
title: Kolekcija Zapisa
challengeType: 1
forumTopicId: 18261
dashedName: record-collection
---

# --description--

Pravite funkciju koja pomaze u odrzavanju kolekcije muzickih albuma. Kolekcija je organizovana kao objekat koji sadrzi vise albuma, pri cemu je svaki album takodje objekat. Svaki album u kolekciji predstavlja jedinstven `id` kao naziv svojstva. Unutar svakog albuma postoje razlicita svojstva sa informacijama o tom albumu. Nemaju svi albumi potpune informacije.

Funkcija `updateRecords` prima 4 argumenta predstavljena sledecim parametrima:

-   `records` - objekat koji sadrzi vise pojedinacnih albuma
-   `id` - broj koji predstavlja odredjeni album u objektu `records`
-   `prop` - string koji predstavlja naziv svojstva albuma koje se azurira
-   `value` - string koji sadrzi informacije kojima se azurira svojstvo albuma

Dovrsite funkciju koristeci pravila ispod da izmenite objekat prosledjen funkciji.

-   Funkcija uvek mora da vrati ceo objekat `records`.
-   Ako je `value` prazan string, obrisite dato svojstvo `prop` iz albuma.
-   Ako `prop` nije `tracks` i `value` nije prazan string, postavite `value` kao vrednost svojstva `prop` tog albuma.
-   Ako je `prop` jednako `tracks`, `value` nije prazan string, a album nema svojstvo `tracks`, napravite prazan niz i dodajte `value` u njega.
-   Ako je `prop` jednako `tracks` i `value` nije prazan string, dodajte `value` na kraj postojeceg niza `tracks` tog albuma.

**Napomena:** Kopija objekta `recordCollection` koristi se za testove. Ne treba direktno da menjate objekat `recordCollection`.

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

Nakon `updateRecords(recordCollection, 5439, "artist", "ABBA")`, `artist` treba da bude string `ABBA`

```js
assert(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'] ===
    'ABBA'
);
```

Nakon `updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")`, `tracks` treba da sadrzi string `Take a Chance on Me` kao jedini i poslednji element.

```js
assert(
  updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me') &&
  _recordCollection[5439]['tracks'].length === 1 &&
  _recordCollection[5439]['tracks'].pop() === 'Take a Chance on Me'
);
```

Nakon `updateRecords(recordCollection, 2548, "artist", "")`, `artist` ne treba da postoji

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert(!_recordCollection[2548].hasOwnProperty('artist'));
```

Nakon `updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")`, `tracks` treba da sadrzi string `Addicted to Love` kao poslednji element.

```js
assert(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop() === 'Addicted to Love'
);
```

Nakon `updateRecords(recordCollection, 2468, "tracks", "Free")`, `tracks` treba da sadrzi string `1999` kao prvi element.

```js
assert(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0] === '1999'
);
```

Nakon `updateRecords(recordCollection, 2548, "tracks", "")`, `tracks` ne treba da postoji

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert(!_recordCollection[2548].hasOwnProperty('tracks'));
```

Nakon `updateRecords(recordCollection, 1245, "albumTitle", "Riptide")`, `albumTitle` treba da bude string `Riptide`

```js
assert(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ] === 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
// Setup
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

// Only change code below this line
function updateRecords(records, id, prop, value) {
  return records;
}

updateRecords(recordCollection, 5439, 'artist', 'ABBA');
```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

// Only change code below this line
function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
