---
id: 56533eb9ac21ba0edf2244d9
title: Poredjenja sa Logickim OR Operatorom
challengeType: 1
forumTopicId: 16800
dashedName: comparisons-with-the-logical-or-operator
---

# --description--

<dfn>Logicki OR</dfn> operator (`||`) vraca `true` ako je makar jedan od <dfn>operanada</dfn> `true`. U suprotnom vraca `false`.

<dfn>Logicki OR</dfn> operator ima dve uspravne crte: (`||`). Obicno se nalazi izmedju tastera Backspace i Enter.

Primer ispod bi trebalo da vam bude poznat iz prethodnih lekcija.

```js
if (num > 10) {
  return "No";
}
if (num < 5) {
  return "No";
}
return "Yes";
```

Ovaj kod ce vratiti `Yes` ako je `num` izmedju `5` i `10` (ukljucujuci `5` i `10`). Ista logika moze se zapisati koristeci <dfn>logicki OR</dfn> operator.

```js
if (num > 10 || num < 5) {
  return "No";
}
return "Yes";
```

# --instructions--

Spojite dva `if` iskaza u jedan koji vraca string `Outside` ako `val` nije izmedju `10` i `20`, ukljucujuci granice. U suprotnom, vratite string `Inside`.

# --hints--

Treba da koristite operator `||` jednom

```js
assert(__helpers.removeJSComments(code).match(/\|\|/g).length === 1);
```

Treba da imate samo jedan `if` iskaz

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalOr(0)` treba da vrati string `Outside`

```js
assert(testLogicalOr(0) === 'Outside');
```

`testLogicalOr(9)` treba da vrati string `Outside`

```js
assert(testLogicalOr(9) === 'Outside');
```

`testLogicalOr(10)` treba da vrati string `Inside`

```js
assert(testLogicalOr(10) === 'Inside');
```

`testLogicalOr(15)` treba da vrati string `Inside`

```js
assert(testLogicalOr(15) === 'Inside');
```

`testLogicalOr(19)` treba da vrati string `Inside`

```js
assert(testLogicalOr(19) === 'Inside');
```

`testLogicalOr(20)` treba da vrati string `Inside`

```js
assert(testLogicalOr(20) === 'Inside');
```

`testLogicalOr(21)` treba da vrati string `Outside`

```js
assert(testLogicalOr(21) === 'Outside');
```

`testLogicalOr(25)` treba da vrati string `Outside`

```js
assert(testLogicalOr(25) === 'Outside');
```

# --seed--

## --seed-contents--

```js
function testLogicalOr(val) {
  // Only change code below this line

  if (val) {
    return "Outside";
  }

  if (val) {
    return "Outside";
  }

  // Only change code above this line
  return "Inside";
}

testLogicalOr(15);
```

# --solutions--

```js
function testLogicalOr(val) {
  if (val < 10 || val > 20) {
    return "Outside";
  }
  return "Inside";
}
```
