---
id: 56533eb9ac21ba0edf2244c3
title: Dodela Pomocu Vrucene Return Vrednosti
challengeType: 1
forumTopicId: 16658
dashedName: assignment-with-a-returned-value
---

# --description--

Kao što pamtiš iz lekcije <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/storing-values-with-the-assignment-operator" target="_blank" rel="noopener noreferrer nofollow">Čuvanje vrednosti pomoću operatora dodele</a>, sve sa desne strane znaka jednakosti se izračunava pre dodele. To znači da možemo uzeti povratnu vrednost funkcije i dodeliti je promenljivoj.

Deklarisimo funkciju `sum` koja sabira dva broja.

```js
ourSum = sum(5, 12);
```

Poziv funkcije `sum` sa argumentima `5` i `12` daje return vrednost `17`. Ta return vrednost se dodeljuje promenljivoj `ourSum`.

# --instructions--

Pozovite funkciju `processArg` sa argumentom `7` i dodelite njenu return vrednost promenljivoj `processed`.

# --hints--

`processed` treba da ima vrednost `2`

```js
assert(processed === 2);
```

Treba da dodelite rezultat `processArg` promenljivoj `processed`

```js
assert(/processed\s*=\s*processArg\(\s*7\s*\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

// Only change code below this line

```

# --solutions--

```js
var processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

processed = processArg(7);
```
