---
id: 56533eb9ac21ba0edf2244b0
title: Slozena Dodela sa Uvecanim Oduzimanjem
challengeType: 1
forumTopicId: 16660
dashedName: compound-assignment-with-augmented-subtraction
---

# --description--

Kao i operator `+=`, operator `-=` oduzima broj od promenljive.

```js
myVar = myVar - 5;
```

ce oduzeti `5` od `myVar`. Ovo se moze prepisati kao:

```js
myVar -= 5;
```

# --instructions--

Promenite dodele vrednosti za `a`, `b` i `c` tako da koriste operator `-=`.

# --hints--

`a` treba da bude jednak `5`.

```js
assert(a === 5);
```

`b` treba da bude jednak `-6`.

```js
assert(b === -6);
```

`c` treba da bude jednak `2`.

```js
assert(c === 2);
```

Treba da koristite operator `-=` za svaku promenljivu.

```js
assert(__helpers.removeJSComments(code).match(/-=/g).length === 3);
```

Ne treba da menjate kod iznad navedenog komentara.

```js
assert(
  /let a = 11;/.test(__helpers.removeJSComments(code)) && /let b = 9;/.test(__helpers.removeJSComments(code)) && /let c = 3;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 11;
let b = 9;
let c = 3;

// Only change code below this line
a = a - 6;
b = b - 15;
c = c - 1;
```

# --solutions--

```js
let a = 11;
let b = 9;
let c = 3;

a -= 6;
b -= 15;
c -= 1;
```
