---
id: 56533eb9ac21ba0edf2244aa
title: Razumevanje Neinicijalizovanih Promenljivih
challengeType: 1
forumTopicId: 18335
dashedName: understanding-uninitialized-variables
---

# --description--

Kada se JavaScript promenljive deklarisu, imaju pocetnu vrednost `undefined`. Ako uradite matematicku operaciju nad promenljivom koja je `undefined`, rezultat ce biti `NaN`, sto znaci <dfn>"Not a Number"</dfn>. Ako spojite string sa `undefined` promenljivom, dobicete string `undefined`.

# --instructions--

Inicijalizujte tri promenljive `a`, `b` i `c` sa `5`, `10` i `"I am a"` redom, kako ne bi bile `undefined`.

# --hints--

`a` treba da bude deklarisana i da ima konacnu vrednost `6`.

```js
assert(typeof a === 'number' && a === 6);
```

`b` treba da bude deklarisana i da ima konacnu vrednost `15`.

```js
assert(typeof b === 'number' && b === 15);
```

`c` ne treba da sadrzi `undefined` i treba da ima konacnu vrednost stringa `I am a String!`

```js
assert(!/undefined/.test(c) && c === 'I am a String!');
```

Ne treba da menjate kod ispod navedenog komentara.

```js
assert(
  /a = a \+ 1;/.test(__helpers.removeJSComments(code)) &&
    /b = b \+ 5;/.test(__helpers.removeJSComments(code)) &&
    /c = c \+ " String!";/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
var a;
var b;
var c;
// Only change code above this line

a = a + 1;
b = b + 5;
c = c + " String!";
```

# --solutions--

```js
var a = 5;
var b = 10;
var c = "I am a";
a = a + 1;
b = b + 5;
c = c + " String!";
```
