---
id: 56533eb9ac21ba0edf2244c7
title: Pristup Svojstvima Objekta Pomocu Tacka Notacije
challengeType: 1
forumTopicId: 16164
dashedName: accessing-object-properties-with-dot-notation
---

# --description--

Postoje dva nacina za pristup svojstvima objekta: tacka notacija (`.`) i uglasta notacija (`[]`), kao kod nizova.

Tacka notacija se koristi kada unapred znate ime svojstva kojem pristupate.

Evo primera koriscenja tacka notacije (`.`) za citanje svojstva objekta:

```js
const myObj = {
  prop1: "val1",
  prop2: "val2"
};

const prop1val = myObj.prop1;
const prop2val = myObj.prop2;
```

`prop1val` ce imati string vrednost `val1`, a `prop2val` string vrednost `val2`.

# --instructions--

Procitajte vrednosti svojstava objekta `testObj` koristeci tacka notaciju. Postavite promenljivu `hatValue` na svojstvo `hat`, a promenljivu `shirtValue` na svojstvo `shirt`.

# --hints--

`hatValue` treba da bude string

```js
assert(typeof hatValue === 'string');
```

Vrednost `hatValue` treba da bude string `ballcap`

```js
assert(hatValue === 'ballcap');
```

`shirtValue` treba da bude string

```js
assert(typeof shirtValue === 'string');
```

Vrednost `shirtValue` treba da bude string `jersey`

```js
assert(shirtValue === 'jersey');
```

Treba da koristite tacka notaciju dva puta

```js
assert(__helpers.removeJSComments(code).match(/testObj\.\w+/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

// Only change code below this line
const hatValue = testObj;      // Change this line
const shirtValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

const hatValue = testObj.hat;
const shirtValue = testObj.shirt;
```
