---
id: 56533eb9ac21ba0edf2244b8
title: Spajanje Stringova Operatorom Plus-Jednako
challengeType: 1
forumTopicId: 16803
dashedName: concatenating-strings-with-the-plus-equals-operator
---

# --description--

Mozemo koristiti i operator `+=` da <dfn>spojimo</dfn> string na kraj postojeceg stringa. To je korisno kada dugacak string delite u vise linija.

**Napomena:** Obratite paznju na razmake. Spajanje ne ubacuje razmake automatski, pa ih morate sami dodati.

Mfano:

```js
let ourStr = "I come first. ";
ourStr += "I come second.";
```

`ourStr` sada ima vrednost stringa `I come first. I come second.`.

# --instructions--

Sastavite `myStr` u vise linija spajajuci ova dva stringa: `This is the first sentence.` i `This is the second sentence.` koristeci operator `+=`. Koristite `+=` kao u primeru i obavezno ostavite razmak izmedju dva stringa. Prvo dodelite prvi string promenljivoj `myStr`, zatim dodajte drugi string.

# --hints--

`myStr` treba da ima jedan razmak izmedju dva stringa.

```js
assert(/sentence\. This/.test(myStr));
```

`myStr` treba da ima vrednost stringa `This is the first sentence. This is the second sentence.`

```js
assert(myStr === 'This is the first sentence. This is the second sentence.');
```

Treba da koristite operator `+=` za pravljenje `myStr`.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*(["']).*\1/g));
```

# --seed--

## --seed-contents--

```js
let myStr;
```

# --solutions--

```js
let myStr = "This is the first sentence. ";
myStr += "This is the second sentence.";
```
