---
id: 56533eb9ac21ba0edf2244bf
title: Lokalni Opseg i Funkcije
challengeType: 1
forumTopicId: 18227
dashedName: local-scope-and-functions
---

# --description--

Promenljive deklarisane unutar funkcije, ukljucujuci parametre funkcije, imaju <dfn>lokalni</dfn> opseg. To znaci da su vidljive samo unutar te funkcije.

Evo funkcije `myTest` sa lokalnom promenljivom `loc`.

```js
function myTest() {
  const loc = "foo";
  console.log(loc);
}

myTest();
console.log(loc);
```

Poziv funkcije `myTest()` ce ispisati string `foo` u konzoli. Linija `console.log(loc)` (izvan funkcije `myTest`) ce baciti gresku, jer `loc` nije deklarisana van funkcije.

# --instructions--

Editor ima dva `console.log` poziva da vam pomogne da vidite sta se desava. Pratite konzolu dok pokrecete kod. Deklarisite lokalnu promenljivu `myVar` unutar `myLocalScope` i pokrenite testove.

**Napomena:** Konzola ce i dalje prikazati `ReferenceError: myVar is not defined`, ali to nece oboriti testove.

# --hints--

Kod ne treba da ima globalnu promenljivu `myVar`.

```js
function declared() {
  myVar;
}

assert.throws(declared, ReferenceError);
```

Treba da dodate lokalnu promenljivu `myVar`.

```js
assert(
  /functionmyLocalScope\(\)\{.*(var|let|const)myVar[\s\S]*}/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

# --seed--

## --seed-contents--

```js
function myLocalScope() {
  // Only change code below this line

  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```

# --solutions--

```js
function myLocalScope() {
  // Only change code below this line
  let myVar;
  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```
