---
id: 56533eb9ac21ba0edf2244ae
title: Pronalazenje Ostatka pri Deljenju u JavaScript-u
challengeType: 1
forumTopicId: 18184
dashedName: finding-a-remainder-in-javascript
---

# --description--

<dfn>Operator ostatka</dfn> `%` daje ostatak pri deljenju dva broja.

**Mfano**

<pre>
5 % 2 = 1
5 / 2 = 2 mabaki 1
2 * 2 = 4
5 - 4 = 1
</pre>

**Upotreba**
U matematici se moze proveriti da li je broj paran ili neparan proverom ostatka pri deljenju sa `2`. Parni brojevi imaju ostatak `0`, a neparni ostatak `1`.

<pre>
17 % 2 = 1
48 % 2 = 0
</pre>

**Napomena:** <dfn>Operator ostatka</dfn> se ponekad pogresno naziva modulo operatorom. Vrlo je slican modulu, ali ne radi isto za negativne brojeve.

# --instructions--

Postavite `remainder` na ostatak pri deljenju `11` sa `3` koristeci <dfn>operator ostatka</dfn> (`%`).

# --hints--

Promenljiva `remainder` treba da bude inicijalizovana.

```js
assert(/(const|let|var)\s+?remainder/.test(__helpers.removeJSComments(code)));
```

Vrednost `remainder` treba da bude `2`.

```js
assert(remainder === 2);
```

Treba da koristite `%` operator.

```js
assert(/\s+?remainder\s*?=\s*?.*%.*;?/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const remainder = 0;
```

# --solutions--

```js
const remainder = 11 % 3;
```
