---
id: 56533eb9ac21ba0edf2244af
title: Slozena Dodela sa Uvecanim Sabiranjem
challengeType: 1
forumTopicId: 16661
dashedName: compound-assignment-with-augmented-addition
---

# --description--

U programiranju je uobicajeno koristiti dodelu da bi se promenio sadrzaj promenljive. Zapamtite da se sve sa desne strane znaka jednakosti prvo izracuna, pa mozemo napisati:

```js
myVar = myVar + 5;
```

da dodamo `5` na `myVar`. Posto je ovo vrlo cest obrazac, postoje operatori koji matematicku operaciju i dodelu obavljaju u jednom koraku.

Jedan od tih operatora je `+=`.

```js
let myVar = 1;
myVar += 5;
console.log(myVar);
```

`6` ce biti ispisano u konzoli.

# --instructions--

Promenite dodele vrednosti za `a`, `b` i `c` tako da koriste `+=` operator.

# --hints--

`a` treba da bude jednak `15`.

```js
assert(a === 15);
```

`b` treba da bude jednak `26`.

```js
assert(b === 26);
```

`c` treba da bude jednak `19`.

```js
assert(c === 19);
```

Treba da koristite `+=` operator za svaku promenljivu.

```js
assert(__helpers.removeJSComments(code).match(/\+=/g).length === 3);
```

Ne treba da menjate postojeci kod iznad navedenog komentara.

```js
assert(
  /let a = 3;/.test(__helpers.removeJSComments(code)) &&
    /let b = 17;/.test(__helpers.removeJSComments(code)) &&
    /let c = 12;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 3;
let b = 17;
let c = 12;

// Only change code below this line
a = a + 12;
b = 9 + b;
c = c + 7;
```

# --solutions--

```js
let a = 3;
let b = 17;
let c = 12;

a += 12;
b += 9;
c += 7;
```
