---
id: 56533eb9ac21ba0edf2244b7
title: Spajanje Stringova Plus Operatorom
challengeType: 1
forumTopicId: 16802
dashedName: concatenating-strings-with-plus-operator
---

# --description--

U JavaScript-u, kada se operator `+` koristi nad `String` vrednostima, to se zove <dfn>konkatenacija</dfn>. Mozete napraviti novi string spajanjem drugih stringova.

**Mfano**

```js
'My name is Alan,' + ' I concatenate.'
```

**Napomena:** Obratite paznju na razmake. Konkatenacija ne dodaje razmake izmedju stringova, pa ih morate dodati sami.

Mfano:

```js
const ourStr = "I come first. " + "I come second.";
```

String `I come first. I come second.` bice ispisan u konzoli.
# --instructions--

Napravite `myStr` od stringova `This is the start.` i `This is the end.` koristeci operator `+`. Obavezno ubacite razmak izmedju ta dva stringa.

# --hints--

`myStr` treba da ima jedan razmak izmedju dva stringa.

```js
assert(/start\. This/.test(myStr));
```

`myStr` treba da ima vrednost stringa `This is the start. This is the end.`

```js
assert(myStr === 'This is the start. This is the end.');
```

Treba da koristite operator `+` za pravljenje `myStr`.

```js
assert(__helpers.removeJSComments(code).match(/(["']).*\1\s*\+\s*(["']).*\2/g));
```

`myStr` treba da bude napravljen koristeci kljucnu rec `const`.

```js
assert(/const\s+myStr/.test(__helpers.removeJSComments(code)));
```

Treba da dodelite tu vrednost promenljivoj `myStr`.

```js
assert(/myStr\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "This is the start. " + "This is the end.";
```
