---
id: 56533eb9ac21ba0edf2244ca
title: Koriscenje Objekata za Lookup
challengeType: 1
forumTopicId: 18373
dashedName: using-objects-for-lookups
---

# --description--

Objekti se mogu posmatrati kao skladiste kljuc/vrednost, slicno recniku. Ako imate tabelarne podatke, mozete koristiti objekat za lookup umesto `switch` naredbe ili lanca `if/else`. Ovo je posebno korisno kada znate da je ulaz ogranicen na odredjeni skup vrednosti.

Evo primera objekta za clanak:

```js
const article = {
  "title": "How to create objects in JavaScript",
  "link": "https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/",
  "author": "Kaashan Hussain",
  "language": "JavaScript",
  "tags": "TECHNOLOGY",
  "createdAt": "NOVEMBER 28, 2018"
};

const articleAuthor = article["author"];
const articleLink = article["link"];

const value = "title";
const valueLookup = article[value];
```

`articleAuthor` je string `Kaashan Hussain`, `articleLink` je string `https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/`, a `valueLookup` je string `How to create objects in JavaScript`.

# --instructions--

Zamenite `switch` naredbu objektom pod nazivom `lookup`. Iskoristite ga za lookup vrednosti `val` i smestite odgovarajuci string u promenljivu `result`.

# --hints--

`phoneticLookup("alpha")` treba da bude jednak stringu `Adams`

```js
assert(phoneticLookup('alpha') === 'Adams');
```

`phoneticLookup("bravo")` treba da bude jednak stringu `Boston`

```js
assert(phoneticLookup('bravo') === 'Boston');
```

`phoneticLookup("charlie")` treba da bude jednak stringu `Chicago`

```js
assert(phoneticLookup('charlie') === 'Chicago');
```

`phoneticLookup("delta")` treba da bude jednak stringu `Denver`

```js
assert(phoneticLookup('delta') === 'Denver');
```

`phoneticLookup("echo")` treba da bude jednak stringu `Easy`

```js
assert(phoneticLookup('echo') === 'Easy');
```

`phoneticLookup("foxtrot")` treba da bude jednak stringu `Frank`

```js
assert(phoneticLookup('foxtrot') === 'Frank');
```

`phoneticLookup("")` treba da bude jednak `undefined`

```js
assert(typeof phoneticLookup('') === 'undefined');
```

Ne treba da menjate `return` naredbu

```js
assert(__helpers.removeJSComments(code).match(/return\sresult;/));
```

Ne treba da koristite `case`, `switch` ili `if` naredbe

```js
assert(
  !/case|switch|if/g.test(__helpers.removeJSComments(code).replace(/([/]{2}.*)|([/][*][^/*]*[*][/])/g, ''))
);
```

# --seed--

## --seed-contents--

```js
// Setup
function phoneticLookup(val) {
  let result = "";

  // Only change code below this line
  switch(val) {
    case "alpha":
      result = "Adams";
      break;
    case "bravo":
      result = "Boston";
      break;
    case "charlie":
      result = "Chicago";
      break;
    case "delta":
      result = "Denver";
      break;
    case "echo":
      result = "Easy";
      break;
    case "foxtrot":
      result = "Frank";
  }

  // Only change code above this line
  return result;
}

phoneticLookup("charlie");
```

# --solutions--

```js
function phoneticLookup(val) {
  let result = "";

  const lookup = {
    alpha: "Adams",
    bravo: "Boston",
    charlie: "Chicago",
    delta: "Denver",
    echo: "Easy",
    foxtrot: "Frank"
  };

  result = lookup[val];

  return result;
}
```
