---
id: 56533eb9ac21ba0edf2244ba
title: Razumi Neizmenjivost Stringova
challengeType: 1
forumTopicId: 18331
dashedName: understand-string-immutability
---

# --description--

U JavaScript-u, `String` vrednosti su <dfn>neizmenjive</dfn>, sto znaci da se ne mogu menjati nakon sto se naprave.

Na primer, sledeci kod ce prijaviti gresku jer se slovo `B` u stringu `Bob` ne moze promeniti u `J`:

```js
let myStr = "Bob";
myStr[0] = "J";
```

Ovo ne znaci da `myStr` ne moze dobiti novu vrednost. Jedini nacin da promenite `myStr` je da joj dodelite potpuno novu vrednost, kao ovde:

```js
let myStr = "Bob";
myStr = "Job";
```

# --instructions--

Ispravite dodelu za `myStr` tako da dobije string vrednost `Hello World`, koristeci pristup prikazan iznad.

# --hints--

`myStr` treba da ima string vrednost `Hello World`.

```js
assert(myStr === 'Hello World');
```

Ne treba da menjate kod iznad navedenog komentara.

```js
assert(/myStr = "Jello World"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let myStr = "Jello World";

// Only change code below this line
myStr[0] = "H"; // Change this line
// Only change code above this line
```

# --solutions--

```js
let myStr = "Jello World";
myStr = "Hello World";
```
