---
id: 56533eb9ac21ba0edf2244d4
title: Poredjenje sa Operatorom Vece Od
challengeType: 1
forumTopicId: 16786
dashedName: comparison-with-the-greater-than-operator
---

# --description--

Operator vece od (`>`) poredi vrednosti dva broja. Ako je broj sa leve strane veci od broja sa desne strane, vraca `true`. U suprotnom vraca `false`.

Kao i operator jednakosti, operator vece od konvertuje tipove podataka tokom poredjenja.

**Primeri**

```js
5   >  3  // true
7   > '3' // true
2   >  3  // false
'1' >  9  // false
```

# --instructions--

Dodajte operator vece od na oznacene linije kako bi return iskazi imali smisla.

# --hints--

`testGreaterThan(0)` treba da vrati string `10 or Under`

```js
assert(testGreaterThan(0) === '10 or Under');
```

`testGreaterThan(10)` treba da vrati string `10 or Under`

```js
assert(testGreaterThan(10) === '10 or Under');
```

`testGreaterThan(11)` treba da vrati string `Over 10`

```js
assert(testGreaterThan(11) === 'Over 10');
```

`testGreaterThan(99)` treba da vrati string `Over 10`

```js
assert(testGreaterThan(99) === 'Over 10');
```

`testGreaterThan(100)` treba da vrati string `Over 10`

```js
assert(testGreaterThan(100) === 'Over 10');
```

`testGreaterThan(101)` treba da vrati string `Over 100`

```js
assert(testGreaterThan(101) === 'Over 100');
```

`testGreaterThan(150)` treba da vrati string `Over 100`

```js
assert(testGreaterThan(150) === 'Over 100');
```

Treba da koristite operator `>` najmanje dva puta

```js
assert(__helpers.removeJSComments(code).match(/val\s*>\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterThan(val) {
  if (val) {  // Change this line
    return "Over 100";
  }

  if (val) {  // Change this line
    return "Over 10";
  }

  return "10 or Under";
}

testGreaterThan(10);
```

# --solutions--

```js
function testGreaterThan(val) {
  if (val > 100) {  // Change this line
    return "Over 100";
  }
  if (val > 10) {  // Change this line
    return "Over 10";
  }
  return "10 or Under";
}
```
