---
id: 56533eb9ac21ba0edf2244b9
title: Sastavljanje Stringova sa Promenljivim
challengeType: 1
forumTopicId: 16805
dashedName: constructing-strings-with-variables
---

# --description--

Ponekad cete morati da sastavite string. Koristeci operator spajanja (`+`), mozete ubaciti jednu ili vise promenljivih u string koji gradite.

Mfano:

```js
const ourName = "freeCodeCamp";
const ourStr = "Hello, our name is " + ourName + ", how are you?";
```

`ourStr` bi imao vrednost stringa `Hello, our name is freeCodeCamp, how are you?`.

# --instructions--

Postavite `myName` na string koji odgovara vasem imenu i napravite `myStr` tako da `myName` bude izmedju stringova `My name is ` i ` and I am well!`

# --hints--

`myName` treba da bude string od najmanje 3 karaktera.

```js
assert(typeof myName !== 'undefined' && myName.length > 2);
```

Treba da koristite dva `+` operatora da napravite `myStr` sa `myName` u sredini.

```js
assert(__helpers.removeJSComments(code).match(/["']\s*\+\s*myName\s*\+\s*["']/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const myName = "";
const myStr = "";
```

# --solutions--

```js
const myName = "Bob";
const myStr = "My name is " + myName + " and I am well!";
```
