---
id: 56533eb9ac21ba0edf2244d3
title: Poredjenje sa Strogim Operatorom Nejednakosti
challengeType: 1
forumTopicId: 16791
dashedName: comparison-with-the-strict-inequality-operator
---

# --description--

Operator stroge nejednakosti (`!==`) je logicka suprotnost operatoru stroge jednakosti. Znaci "nije tacno jednako" i vraca `false` tamo gde bi operator stroge jednakosti vratio `true`, i obrnuto. Operator stroge nejednakosti ne konvertuje tipove podataka.

**Primeri**

```js
3 !==  3  // false
3 !== '3' // true
4 !==  3  // true
```

# --instructions--

Dodajte operator stroge nejednakosti u `if` izraz tako da funkcija vraca string `Not Equal` kada `val` nije striktno jednako `17`

# --hints--

`testStrictNotEqual(17)` treba da vrati string `Equal`

```js
assert(testStrictNotEqual(17) === 'Equal');
```

`testStrictNotEqual("17")` treba da vrati string `Not Equal`

```js
assert(testStrictNotEqual('17') === 'Not Equal');
```

`testStrictNotEqual(12)` treba da vrati string `Not Equal`

```js
assert(testStrictNotEqual(12) === 'Not Equal');
```

`testStrictNotEqual("bob")` treba da vrati string `Not Equal`

```js
assert(testStrictNotEqual('bob') === 'Not Equal');
```

Treba da koristite operator `!==`

```js
assert(__helpers.removeJSComments(code).match(/(val\s*!==\s*\d+)|(\d+\s*!==\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrictNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testStrictNotEqual(10);
```

# --solutions--

```js
function testStrictNotEqual(val) {
  if (val !== 17) {
    return "Not Equal";
  }
  return "Equal";
}
```
