---
id: 56533eb9ac21ba0edf2244c2
title: Vrati Vrednost iz Funkcije Pomocu Return
challengeType: 1
forumTopicId: 18271
dashedName: return-a-value-from-a-function-with-return
---

# --description--

Mozemo proslediti vrednosti funkciji koristeci <dfn>argumente</dfn>. Mozete koristiti `return` naredbu da posaljete vrednost nazad iz funkcije.

**Mfano**

```js
function plusThree(num) {
  return num + 3;
}

const answer = plusThree(5);
```

`answer` ima vrednost `8`.

`plusThree` prima <dfn>argument</dfn> `num` i vraca vrednost `num + 3`.

# --instructions--

Napravite funkciju `timesFive` koja prima jedan argument, mnozi ga sa `5` i vraca novu vrednost.

# --hints--

`timesFive` treba da bude funkcija

```js
assert(typeof timesFive === 'function');
```

`timesFive(5)` treba da vrati `25`

```js
assert(timesFive(5) === 25);
```

`timesFive(2)` treba da vrati `10`

```js
assert(timesFive(2) === 10);
```

`timesFive(0)` treba da vrati `0`

```js
assert(timesFive(0) === 0);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function timesFive(num) {
  return num * 5;
}
timesFive(10);
```
