---
id: 56533eb9ac21ba0edf2244bd
title: Prosledjivanje Vrednosti Funkcijama Pomocu Argumenata
challengeType: 1
forumTopicId: 18254
dashedName: passing-values-to-functions-with-arguments
---

# --description--

<dfn>Parametri</dfn> su promenljive koje sluze kao mesta za vrednosti koje funkcija treba da primi kada se pozove. Kada se funkcija definise, obicno se definise sa jednim ili vise parametara. Stvarne vrednosti koje se daju (ili <dfn>"prosledjuju"</dfn>) funkciji pri pozivu zovu se <dfn>argumenti</dfn>.

Evo funkcije sa dva parametra, `param1` i `param2`:

```js
function testFun(param1, param2) {
  console.log(param1, param2);
}
```

Zatim mozemo pozvati `testFun` ovako: `testFun("Hello", "World");`. Prosledili smo dva string argumenta, `Hello` i `World`. Unutar funkcije, `param1` ce biti `Hello`, a `param2` ce biti `World`. Mozete ponovo pozvati `testFun` sa drugim argumentima i parametri ce dobiti nove vrednosti.

# --instructions--

<ol><li>Napravite funkciju pod nazivom <code>functionWithArgs</code> koja prima dva argumenta i ispisuje njihov zbir u developerskoj konzoli.</li><li>Pozovite funkciju sa dva broja kao argumentima.</li></ol>

# --before-each--

```js
var logOutput;
```

# --hints--

`functionWithArgs` treba da bude funkcija.

```js
assert(typeof functionWithArgs === 'function');
```

`functionWithArgs(1,2)` treba da ispise `3`.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(1, 2);
assert(logOutput == 3);
```

`functionWithArgs(7,9)` treba da ispise `16`.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(7, 9);
assert(logOutput == 16);
```

Treba da pozovete `functionWithArgs` sa dva broja nakon definicije.

```js
assert(
  /functionWithArgs\([-+]?\d*\.?\d*,[-+]?\d*\.?\d*\)/.test(
    __helpers.removeJSComments(code).replace(/\s/g, '')
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function functionWithArgs(a, b) {
  console.log(a + b);
}
functionWithArgs(10, 5);
```
