---
id: 56533eb9ac21ba0edf2244a8
title: Cuvanje Vrednosti Pomocu Operatera Dodele
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

U JavaScript-u mozete cuvati vrednost u promenljivoj koristeci <dfn>operater dodele</dfn> (`=`).

```js
myVariable = 5;
```

Ovo dodeljuje broj `5` promenljivoj `myVariable`.

Ako postoji bilo kakva racunica sa desne strane operatera `=`, ona se izvrsi pre nego sto se vrednost dodeli promenljivoj sa leve strane.

```js
var myVar;
myVar = 5;
```

Prvo, ovaj kod kreira promenljivu `myVar`. Zatim kod dodeljuje `5` promenljivoj `myVar`. Ako se `myVar` ponovo pojavi u kodu, program ce je tretirati kao da je `5`.

# --instructions--

Dodelite vrednost `7` promenljivoj `a`.

# --hints--

Ne treba da menjate kod iznad navedenog komentara.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

`a` treba da ima vrednost 7.

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
