---
id: 56533eb9ac21ba0edf2244be
title: Globalni Opseg i Funkcije
challengeType: 1
forumTopicId: 18193
dashedName: global-scope-and-functions
---

# --description--

U JavaScript-u, <dfn>opseg (scope)</dfn> se odnosi na vidljivost promenljivih. Promenljive deklarisane van funkcije imaju <dfn>globalni</dfn> opseg. To znaci da su vidljive bilo gde u JavaScript kodu.

Promenljive deklarisane bez kljucnih reci `let` ili `const` automatski nastaju u `global` opsegu. To moze izazvati neocekivano ponasanje u drugim delovima koda ili pri ponovnom pokretanju funkcije. Uvek deklarisite promenljive koristeci `let` ili `const`.

# --instructions--

Koristeci `let` ili `const`, deklarisite globalnu promenljivu `myGlobal` van bilo koje funkcije. Inicijalizujte je vrednoscu `10`.

Unutar funkcije `fun1`, dodelite vrednost `5` promenljivoj `oopsGlobal` ***bez*** upotrebe `var`, `let` ili `const`.

# --before-each--

```js
var oopsGlobal;
```

# --hints--

`myGlobal` treba da bude deklarisana

```js
assert(typeof myGlobal != 'undefined');
```

`myGlobal` treba da ima vrednost `10`

```js
assert(myGlobal === 10);
```

`myGlobal` treba da bude deklarisana koristeci `let` ili `const`

```js
assert(/(let|const)\s+myGlobal/.test(__helpers.removeJSComments(code)));
```

`oopsGlobal` treba da bude globalna promenljiva sa vrednoscu `5`

```js
fun1();
assert(typeof oopsGlobal != 'undefined');
```

# --seed--

## --seed-contents--

```js
// Declare the myGlobal variable below this line


function fun1() {
  // Assign 5 to oopsGlobal here

}

// Only change code above this line

function fun2() {
  let output = "";
  if (typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if (typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```

# --solutions--

```js
const myGlobal = 10;

function fun1() {
  oopsGlobal = 5;
}

function fun2() {
  let output = "";
  if(typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if(typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```
