---
id: 56533eb9ac21ba0edf2244c6
title: Stani u Red
challengeType: 1
forumTopicId: 18307
dashedName: stand-in-line
---

# --description--

U racunarskim naukama, <dfn>queue</dfn> je apstraktna <dfn>struktura podataka</dfn> u kojoj se elementi cuvaju po redosledu. Novi elementi se dodaju na kraj reda, a stariji uklanjaju sa pocetka reda.

# --instructions--

Napisite funkciju `nextInLine` koja prima niz (`arr`) i broj (`item`) kao argumente.

Dodajte broj na kraj niza, zatim uklonite prvi element niza.

Funkcija `nextInLine` treba da vrati uklonjeni element.

# --hints--

`nextInLine([], 5)` treba da vrati broj.

```js
assert.isNumber(nextInLine([], 5));
```

`nextInLine([], 1)` treba da vrati `1`

```js
assert(nextInLine([], 1) === 1);
```

`nextInLine([2], 1)` treba da vrati `2`

```js
assert(nextInLine([2], 1) === 2);
```

`nextInLine([5,6,7,8,9], 1)` treba da vrati `5`

```js
assert(nextInLine([5, 6, 7, 8, 9], 1) === 5);
```

# --seed--

## --seed-contents--

```js
function nextInLine(arr, item) {
  // Only change code below this line
  
  return item;
  // Only change code above this line
}

// Setup
let testArr = [1, 2, 3, 4, 5];

// Display code
console.log("Before: " + JSON.stringify(testArr));
console.log(nextInLine(testArr, 6));
console.log("After: " + JSON.stringify(testArr));
```

# --solutions--

```js
let testArr = [1, 2, 3, 4, 5];

function nextInLine(arr, item) {
    arr.push(item);
    return arr.shift();
}
```
