---
id: 56533eb9ac21ba0edf2244ad
title: Umanji Broj u JavaScript-u
challengeType: 1
forumTopicId: 17558
dashedName: decrement-a-number-with-javascript
---

# --description--

Mozete lako <dfn>umanjiti</dfn> vrednost promenljive za jedan koristeci `--` operator.

```js
i--;
```

ni sawa na

```js
i = i - 1;
```

**Napomena:** Cela linija je `i--;`, tako da nije potreban znak jednakosti.

# --instructions--

Promenite kod tako da koristi `--` operator na `myVar`.

# --hints--

`myVar` treba da bude jednak `10`.

```js
assert(myVar === 10);
```

`myVar = myVar - 1;` treba da bude zamenjeno.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*myVar\s*[-]\s*1.*?;?/));
```

Ne treba direktno da postavite `myVar` na `10`.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*10.*?;?/));
```

Treba da koristite `--` operator na `myVar`.

```js
assert(/[-]{2}\s*myVar|myVar\s*[-]{2}/.test(__helpers.removeJSComments(code)));
```

Ne treba da menjate kod iznad navedenog komentara.

```js
assert(/let myVar = 11;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 11;

// Only change code below this line
myVar = myVar - 1;
```

# --solutions--

```js
let myVar = 11;
myVar--;
```
