---
id: 56533eb9ac21ba0edf2244a9
title: Inicijalizacija Promenljivih Operaterom Dodele
challengeType: 1
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

Uobicajeno je da promenljivu inicijalizujete pocetnom vrednoscu u istoj liniji u kojoj je deklarisete.

```js
var myVar = 0;
```

Ovo kreira novu promenljivu `myVar` i dodeljuje joj pocetnu vrednost `0`.

# --instructions--

Deklarisite promenljivu `a` pomocu `var` i inicijalizujte je vrednoscu `9`.

# --hints--

Treba da inicijalizujete `a` vrednoscu `9`.

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
