---
id: 56533eb9ac21ba0edf2244e0
title: Pretvaranje If Else Lanaca u switch
challengeType: 1
forumTopicId: 18266
dashedName: replacing-if-else-chains-with-switch
---

# --description--

Ako imate mnogo opcija za izbor, `switch` izraz je cesto laksi za pisanje od dugih lanaca `if`/`else if`. Sledece:

```js
if (val === 1) {
  answer = "a";
} else if (val === 2) {
  answer = "b";
} else {
  answer = "c";
}
```

moze se zameniti sa:

```js
switch (val) {
  case 1:
    answer = "a";
    break;
  case 2:
    answer = "b";
    break;
  default:
    answer = "c";
}
```

# --instructions--

Pretvorite `if`/`else if` lanac u `switch` izraz.

# --hints--

Ne treba da koristite `else` izraze nigde u editoru

```js
assert(!/else/g.test(__helpers.removeJSComments(code)));
```

Ne treba da koristite `if` izraze nigde u editoru

```js
assert(!/if/g.test(__helpers.removeJSComments(code)));
```

Treba da imate najmanje cetiri `break` naredbe

```js
assert(__helpers.removeJSComments(code).match(/break/g).length >= 4);
```

`chainToSwitch("bob")` treba da vrati string `Marley`

```js
assert(chainToSwitch('bob') === 'Marley');
```

`chainToSwitch(42)` treba da vrati string `The Answer`

```js
assert(chainToSwitch(42) === 'The Answer');
```

`chainToSwitch(1)` treba da vrati string `There is no #1`

```js
assert(chainToSwitch(1) === 'There is no #1');
```

`chainToSwitch(99)` treba da vrati string `Missed me by this much!`

```js
assert(chainToSwitch(99) === 'Missed me by this much!');
```

`chainToSwitch(7)` treba da vrati string `Ate Nine`

```js
assert(chainToSwitch(7) === 'Ate Nine');
```

`chainToSwitch("John")` treba da vrati `""` (prazan string)

```js
assert(chainToSwitch('John') === '');
```

`chainToSwitch(156)` treba da vrati `""` (prazan string)

```js
assert(chainToSwitch(156) === '');
```

# --seed--

## --seed-contents--

```js
function chainToSwitch(val) {
  let answer = "";
  // Only change code below this line

  if (val === "bob") {
    answer = "Marley";
  } else if (val === 42) {
    answer = "The Answer";
  } else if (val === 1) {
    answer = "There is no #1";
  } else if (val === 99) {
    answer = "Missed me by this much!";
  } else if (val === 7) {
    answer = "Ate Nine";
  }

  // Only change code above this line
  return answer;
}

chainToSwitch(7);
```

# --solutions--

```js
function chainToSwitch(val) {
  let answer = "";

  switch (val) {
    case "bob":
      answer = "Marley";
      break;
    case 42:
      answer = "The Answer";
      break;
    case 1:
      answer = "There is no #1";
      break;
    case 99:
      answer = "Missed me by this much!";
      break;
    case 7:
      answer = "Ate Nine";
  }
  return answer;
}
```
