---
id: 56533eb9ac21ba0edf2244d5
title: Poredjenje sa Operatorom Vece Ili Jednako
challengeType: 1
forumTopicId: 16785
dashedName: comparison-with-the-greater-than-or-equal-to-operator
---

# --description--

Operator vece ili jednako (`>=`) poredi vrednosti dva broja. Ako je broj sa leve strane veci ili jednak broju sa desne strane, vraca `true`. U suprotnom vraca `false`.

Kao i operator jednakosti, operator vece ili jednako konvertuje tipove podataka pri poredjenju.

**Primeri**

```js
6   >=  6  // true
7   >= '3' // true
2   >=  3  // false
'7' >=  9  // false
```

# --instructions--

Dodajte operator vece ili jednako na oznacene linije da bi return iskazi imali smisla.

# --hints--

`testGreaterOrEqual(0)` treba da vrati string `Less than 10`

```js
assert(testGreaterOrEqual(0) === 'Less than 10');
```

`testGreaterOrEqual(9)` treba da vrati string `Less than 10`

```js
assert(testGreaterOrEqual(9) === 'Less than 10');
```

`testGreaterOrEqual(10)` treba da vrati string `10 or Over`

```js
assert(testGreaterOrEqual(10) === '10 or Over');
```

`testGreaterOrEqual(11)` treba da vrati string `10 or Over`

```js
assert(testGreaterOrEqual(11) === '10 or Over');
```

`testGreaterOrEqual(19)` treba da vrati string `10 or Over`

```js
assert(testGreaterOrEqual(19) === '10 or Over');
```

`testGreaterOrEqual(100)` treba da vrati string `20 or Over`

```js
assert(testGreaterOrEqual(100) === '20 or Over');
```

`testGreaterOrEqual(21)` treba da vrati string `20 or Over`

```js
assert(testGreaterOrEqual(21) === '20 or Over');
```

Treba da koristite operator `>=` najmanje dva puta

```js
assert(__helpers.removeJSComments(code).match(/val\s*>=\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterOrEqual(val) {
  if (val) {  // Change this line
    return "20 or Over";
  }

  if (val) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}

testGreaterOrEqual(10);
```

# --solutions--

```js
function testGreaterOrEqual(val) {
  if (val >= 20) {  // Change this line
    return "20 or Over";
  }

  if (val >= 10) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}
```
