---
id: 56533eb9ac21ba0edf2244bc
title: Lista Za Kupovinu
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

Napravite listu za kupovinu u promenljivoj `myList`. Lista treba da bude visedimenzionalni niz koji sadrzi vise podnizova.

Prvi element svakog podniza treba da bude string koji predstavlja naziv proizvoda. Drugi element treba da bude broj koji predstavlja kolicinu, na primer

```js
["Chocolate Bar", 15]
```

Lista treba da sadrzi najmanje 5 podnizova.

# --hints--

`myList` treba da bude niz.

```js
assert(Array.isArray(myList));
```

Prvi elementi u svakom podnizu treba da budu stringovi.

```js
assert(myList.every(elem => elem && typeof elem[0] === 'string'));
```

Drugi elementi u svakom podnizu treba da budu brojevi.

```js
assert(myList.every(elem => elem && typeof elem[1] === 'number'));
```

Treba da imate najmanje 5 stavki na listi.

```js
assert(myList.length > 4);
```

# --seed--

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
