---
id: 56533eb9ac21ba0edf2244cb
title: Manipulisanje Slozenim Objektima
challengeType: 1
forumTopicId: 18208
dashedName: manipulating-complex-objects
---

# --description--

Ponekad je potrebno cuvati podatke u fleksibilnoj <dfn>strukturi podataka</dfn>. JavaScript objekat je jedan nacin da se to uradi. Omogucava kombinaciju <dfn>stringova</dfn>, <dfn>brojeva</dfn>, <dfn>boolean</dfn> vrednosti, <dfn>nizova</dfn>, <dfn>funkcija</dfn> i <dfn>objekata</dfn>.

Evo primera slozene strukture podataka:

```js
const ourMusic = [
  {
    "artist": "Daft Punk",
    "title": "Homework",
    "release_year": 1997,
    "formats": [ 
      "CD", 
      "Cassette", 
      "LP"
    ],
    "gold": true
  }
];
```

Ovo je niz koji sadrzi jedan objekat. Taj objekat ima razlicite delove <dfn>metapodataka</dfn> o albumu. Sadrzi i ugnjezdeni niz `formats`. Ako zelite da dodate jos zapisa albuma, uradite to dodavanjem zapisa u niz najviseg nivoa. Objekti cuvaju podatke u svojstvima po obrascu kljuc-vrednost. U primeru iznad, `"artist": "Daft Punk"` je svojstvo sa kljucem `artist` i vrednoscu `Daft Punk`.

**Napomena:** Potrebno je da stavite zarez posle svakog objekta u nizu, osim posle poslednjeg.

# --instructions--

Dodajte novi album u niz `myMusic`. Dodajte string svojstva `artist` i `title`, broj `release_year` i niz stringova `formats`.

# --hints--

`myMusic` treba da bude niz

```js
assert(Array.isArray(myMusic));
```

`myMusic` treba da ima najmanje dva elementa

```js
assert(myMusic.length > 1);
```

Elementi niza `myMusic` treba da budu objekti

```js
myMusic.forEach(object => {assert.typeOf(object, 'object')})
```

Vas objekat u `myMusic` treba da ima najmanje cetiri svojstva

```js
myMusic.forEach(object => {assert(Object.keys(object).length > 3); });
```

Vas objekat u `myMusic` treba da ima svojstvo `artist` koje je string

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['artist']);
  assert.typeOf(object.artist, 'string')
})
```

Vas objekat u `myMusic` treba da ima svojstvo `title` koje je string

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['title']);
  assert.typeOf(object.title, 'string')
})
```

Vas objekat u `myMusic` treba da ima svojstvo `release_year` koje je broj

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['release_year']);
  assert.typeOf(object.release_year, 'number')
})
```

Vas objekat u `myMusic` treba da ima svojstvo `formats` koje je niz

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['formats']);
  assert.typeOf(object.formats, 'array')
})
```

`formats` treba da bude niz stringova sa najmanje dva elementa

```js
myMusic.forEach(object => {
  object.formats.forEach(format => {
    assert.typeOf(format, 'string')
  });
  assert.isAtLeast(object.formats.length, 2)
})
```

# --seed--

## --seed-contents--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CD",
      "8T",
      "LP"
    ],
    "gold": true
  }
];
```

# --solutions--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP" ],
    "gold": true
  },
  {
    "artist": "ABBA",
    "title": "Ring Ring",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP",
    "CD",
  ]
  }
];
```
