---
id: 56533eb9ac21ba0edf2244de
title: Dodavanje Podrazumevane Opcije u switch Iskazima
challengeType: 1
forumTopicId: 16653
dashedName: adding-a-default-option-in-switch-statements
---

# --description--

U `switch` iskazu mozda necete navesti sve moguce vrednosti kao `case` grane. Umesto toga, mozete dodati `default` granu koja ce se izvrsiti ako nijedna `case` grana ne odgovara. Posmatrajte je kao poslednji `else` u `if/else` lancu.

`default` treba da bude poslednja grana.

```js
switch (num) {
  case value1:
    statement1;
    break;
  case value2:
    statement2;
    break;
...
  default:
    defaultStatement;
    break;
}
```

# --instructions--

Napisite `switch` iskaz koji postavlja `answer` po sledecim pravilima:
`a` - `apple`
`b` - `bird`
`c` - `cat`
`default` - `stuff`

# --hints--

`switchOfStuff("a")` treba da vrati string `apple`

```js
assert(switchOfStuff('a') === 'apple');
```

`switchOfStuff("b")` treba da vrati string `bird`

```js
assert(switchOfStuff('b') === 'bird');
```

`switchOfStuff("c")` treba da vrati string `cat`

```js
assert(switchOfStuff('c') === 'cat');
```

`switchOfStuff("d")` treba da vrati string `stuff`

```js
assert(switchOfStuff('d') === 'stuff');
```

`switchOfStuff(4)` treba da vrati string `stuff`

```js
assert(switchOfStuff(4) === 'stuff');
```

Ne treba da koristite `if` ili `else` iskaze

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Treba da koristite `default` granu

```js
assert(switchOfStuff('string-to-trigger-default-case') === 'stuff');
```

Treba da imate najmanje 3 `break` naredbe

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function switchOfStuff(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

switchOfStuff(1);
```

# --solutions--

```js
function switchOfStuff(val) {
  let answer = "";

  switch(val) {
    case "a":
      answer = "apple";
      break;
    case "b":
      answer = "bird";
      break;
    case "c":
      answer = "cat";
      break;
    default:
      answer = "stuff";
  }
  return answer;
}
```
