---
id: 56533eb9ac21ba0edf2244ac
title: Uvecaj Broj u JavaScript-u
challengeType: 1
forumTopicId: 18201
dashedName: increment-a-number-with-javascript
---

# --description--

Mozete lako <dfn>uvecati</dfn> promenljivu za jedan koristeci `++` operator.

```js
i++;
```

ni sawa na

```js
i = i + 1;
```

**Napomena:** Cela linija je `i++;`, tako da nije potreban znak jednakosti.

# --instructions--

Promenite kod tako da koristi `++` operator na `myVar`.

# --hints--

`myVar` treba da bude jednak `88`.

```js
assert(myVar === 88);
```

Ne treba da koristite operater dodele.

```js
assert(
  /let\s+myVar\s*=\s*87;\s*\/*.*\s*([+]{2}\s*myVar|myVar\s*[+]{2})/.test(__helpers.removeJSComments(code))
);
```

Treba da koristite `++` operator.

```js
assert(/[+]{2}\s*myVar|myVar\s*[+]{2}/.test(__helpers.removeJSComments(code)));
```

Ne treba da menjate kod iznad navedenog komentara.

```js
assert(/let myVar = 87;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 87;

// Only change code below this line
myVar = myVar + 1;
```

# --solutions--

```js
let myVar = 87;
myVar++;
```
