---
id: 56592a60ddddeae28f7aa8e1
title: Pristup Visedimenzionalnim Nizovima Pomocu Indeksa
challengeType: 1
forumTopicId: 16159
dashedName: access-multi-dimensional-arrays-with-indexes
---

# --description--

Jedan nacin da razmisljate o <dfn>visedimenzionalnom</dfn> nizu je kao o *nizu nizova*. Kada koristite uglaste zagrade za pristup nizu, prvi par zagrada se odnosi na elemente spoljnog nivoa, a svaki sledeci par zagrada na naredni unutrasnji nivo.

**Primer**

```js
const arr = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14]
];

const subarray = arr[3];
const nestedSubarray = arr[3][0];
const element = arr[3][0][1];
```

U ovom primeru, `subarray` ima vrednost `[[10, 11, 12], 13, 14]`,
`nestedSubarray` ima vrednost `[10, 11, 12]`, a `element` ima vrednost `11`.

**Napomena:** Ne treba da bude razmaka izmedju imena niza i uglastih zagrada; oblici poput `array [0][0]` ili `array [0] [0]` nisu prihvatljivi. Iako JavaScript to moze obraditi, to moze zbuniti druge programere koji citaju kod.

# --instructions--

Koristeci notaciju sa uglastim zagradama izaberite element iz `myArray` tako da `myData` bude jednako `8`.

# --hints--

`myData` treba da bude jednako `8`.

```js
assert(myData === 8);
```

Treba da koristite notaciju sa uglastim zagradama da procitate tacnu vrednost iz `myArray`.

```js
assert(/myData=myArray\[2\]\[1\]/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code))));
```

# --seed--

## --seed-contents--

```js
const myArray = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14],
];

const myData = myArray[0][0];
```

# --solutions--

```js
const myArray = [[1, 2, 3], [4, 5, 6], [7, 8, 9], [[10, 11, 12], 13, 14]];
const myData = myArray[2][1];
```
