---
id: 565bbe00e9cc8ac0725390f4
title: Brojanje Karata
challengeType: 1
forumTopicId: 16809
dashedName: counting-cards
---

# --description--

U kazino igri Blackjack, igrac moze proceniti da li ima prednost u sledecoj ruci protiv kuce tako sto prati odnos visokih i niskih karata preostalih u spilu. Ovo se zove brojanje karata.

Ako je u spilu ostalo vise visokih karata, to ide u korist igracu. Svakom kartonu je dodeljena vrednost prema tabeli ispod. Kada je brojac pozitivan, igrac treba da ulozi vise. Kada je brojac nula ili negativan, igrac treba da ulozi manje.

<table><thead><tr><th>Mabadiliko ya hesabu</th><th>Kadi</th></tr></thead><tbody><tr><td>+1</td><td>2, 3, 4, 5, 6</td></tr><tr><td>0</td><td>7, 8, 9</td></tr><tr><td>-1</td><td>10, 'J', 'Q', 'K', 'A'</td></tr></tbody></table>

Napisacete funkciju za brojanje karata. Prima parametar `card` koji moze biti broj ili string i uvecava ili umanjuje promenljivu `count` u zavisnosti od vrednosti karte (vidi tabelu). Funkcija treba da vrati string koji prikazuje trenutno stanje brojaca i rec `Bet` ako je brojac pozitivan, ili `Hold` ako je nula ili negativan. Trenutni brojac i odluka igraca (`Bet` ili `Hold`) moraju biti odvojeni jednim razmakom.

**Primer izlaza:** `-3 Hold` ili `5 Bet`

**Savet**
Nemojte resetovati `count` na 0 kada je karta 7, 8 ili 9.
Ne vracajte niz.
Ne ukljucujte navodnike (jednostruke ili dvostruke) u izlaz.

# --hints--

Vasa funkcija treba da vrati vrednost brojaca i tekst (`Bet` ili `Hold`) sa jednim razmakom izmedju.

```js
assert(//
  (function () {
    count = 0;
    let out = cc(10);
    const hasSpace = /-?\d+ (Bet|Hold)/.test('' + out);
    return hasSpace;
  })()
);
```

Niz karata 2, 3, 4, 5, 6 treba da vrati string `5 Bet`

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc(3);
    cc(4);
    cc(5);
    var out = cc(6);
    if (out === '5 Bet') {
      return true;
    }
    return false;
  })()
);
```

Niz karata 7, 8, 9 treba da vrati string `0 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(7);
    cc(8);
    var out = cc(9);
    if (out === '0 Hold') {
      return true;
    }
    return false;
  })()
);
```

Niz karata 10, J, Q, K, A treba da vrati string `-5 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(10);
    cc('J');
    cc('Q');
    cc('K');
    var out = cc('A');
    if (out === '-5 Hold') {
      return true;
    }
    return false;
  })()
);
```

Niz karata 3, 7, Q, 8, A treba da vrati string `-1 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(3);
    cc(7);
    cc('Q');
    cc(8);
    var out = cc('A');
    if (out === '-1 Hold') {
      return true;
    }
    return false;
  })()
);
```

Niz karata 2, J, 9, 2, 7 treba da vrati string `1 Bet`

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc('J');
    cc(9);
    cc(2);
    var out = cc(7);
    if (out === '1 Bet') {
      return true;
    }
    return false;
  })()
);
```

Niz karata 2, 2, 10 treba da vrati string `1 Bet`

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc(2);
    var out = cc(10);
    if (out === '1 Bet') {
      return true;
    }
    return false;
  })()
);
```

Niz karata 3, 2, A, 10, K treba da vrati string `-1 Hold`

```js
assert(
  (function () {
    count = 0;
    cc(3);
    cc(2);
    cc('A');
    cc(10);
    var out = cc('K');
    if (out === '-1 Hold') {
      return true;
    }
    return false;
  })()
);
```

# --seed--

## --seed-contents--

```js
let count = 0;

function cc(card) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

cc(2); cc(3); cc(7); cc('K'); cc('A');
```

# --solutions--

```js
let count = 0;
function cc(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
