---
id: 5664820f61c48e80c9fa476c
title: Golf Kod
challengeType: 1
forumTopicId: 18195
dashedName: golf-code
---

# --description--

U golfu, svaka rupa ima `par`, sto predstavlja prosecan broj `strokes` (udaraca) koji se ocekuje da igrac napravi da ubaci lopticu u rupu. U zavisnosti od toga koliko su vasi `strokes` iznad ili ispod `par` vrednosti, koristi se razlicit naziv.

Vasa funkcija prima argumente `par` i `strokes`. Vratite odgovarajuci string prema sledecoj tabeli koja navodi ishode po prioritetu, od vrha (najvisi prioritet) ka dnu (najnizi prioritet):

<table><thead><tr><th>Michoro</th><th>Rudisha</th></tr></thead><tbody><tr><td>1</td><td>"Hole-in-one!"</td></tr><tr><td>&#x3C;= par - 2</td><td>"Eagle"</td></tr><tr><td>par - 1</td><td>"Birdie"</td></tr><tr><td>par</td><td>"Par"</td></tr><tr><td>par + 1</td><td>"Bogey"</td></tr><tr><td>par + 2</td><td>"Double Bogey"</td></tr><tr><td>>= par + 3</td><td>"Go Home!"</td></tr></tbody></table>

`par` i `strokes` ce uvek biti pozitivni brojevi. Dodat je niz sa svim nazivima radi lakseg rada.

# --hints--

`golfScore(4, 1)` treba da vrati string `Hole-in-one!`

```js
assert(golfScore(4, 1) === 'Hole-in-one!');
```

`golfScore(4, 2)` treba da vrati string `Eagle`

```js
assert(golfScore(4, 2) === 'Eagle');
```

`golfScore(5, 2)` treba da vrati string `Eagle`

```js
assert(golfScore(5, 2) === 'Eagle');
```

`golfScore(4, 3)` treba da vrati string `Birdie`

```js
assert(golfScore(4, 3) === 'Birdie');
```

`golfScore(4, 4)` treba da vrati string `Par`

```js
assert(golfScore(4, 4) === 'Par');
```

`golfScore(1, 1)` treba da vrati string `Hole-in-one!`

```js
assert(golfScore(1, 1) === 'Hole-in-one!');
```

`golfScore(5, 5)` treba da vrati string `Par`

```js
assert(golfScore(5, 5) === 'Par');
```

`golfScore(4, 5)` treba da vrati string `Bogey`

```js
assert(golfScore(4, 5) === 'Bogey');
```

`golfScore(4, 6)` treba da vrati string `Double Bogey`

```js
assert(golfScore(4, 6) === 'Double Bogey');
```

`golfScore(4, 7)` treba da vrati string `Go Home!`

```js
assert(golfScore(4, 7) === 'Go Home!');
```

`golfScore(5, 9)` treba da vrati string `Go Home!`

```js
assert(golfScore(5, 9) === 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

function golfScore(par, strokes) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

golfScore(5, 4);
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
