---
id: 5675e877dbd60be8ad28edc6
title: Iteriraj Kroz Niz Pomocu for Petlje
challengeType: 1
forumTopicId: 18216
dashedName: iterate-through-an-array-with-a-for-loop
---

# --description--

Uobicajena vezba u JavaScript-u je prolazak kroz elemente niza. Jedan nacin za to je koriscenje `for` petlje. Ovaj kod ispisuje svaki element niza `arr` u konzoli:

```js
const arr = [10, 9, 8, 7, 6];

for (let i = 0; i < arr.length; i++) {
   console.log(arr[i]);
}
```

Zapamtite da nizovi koriste indeksiranje od nule, sto znaci da je poslednji indeks `length - 1`. Uslov petlje je `i < arr.length`, sto zaustavlja petlju kada `i` postane jednako `length`. U ovom slucaju poslednji prolaz je `i === 4`, odnosno kada je `i` jednako `arr.length - 1`, pa se u konzoli prikazuje `6`. Zatim se `i` povecava na `5`, i petlja se zavrsava jer je `i < arr.length` netacno.

# --instructions--

Deklarisite i inicijalizujte promenljivu `total` sa `0`. Koristite `for` petlju da dodate vrednost svakog elementa niza `myArr` u `total`.

# --hints--

`total` treba da bude deklarisan i inicijalizovan na 0.

```js
assert(__helpers.removeJSComments(code).match(/(var|let|const)\s*?total\s*=\s*0.*?;?/));
```

`total` treba da bude jednako 20.

```js
assert(total === 20);
```

Treba da koristite `for` petlju za iteraciju kroz `myArr`.

```js
assert(/for\s*\(/g.test(__helpers.removeJSComments(code)) && /myArr\s*\[/g.test(__helpers.removeJSComments(code)));
```

Ne treba direktno da dodelite vrednost 20 promenljivoj `total`.

```js
assert(!__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/total[=+-]0*[1-9]+/gm));
```

# --seed--

## --seed-contents--

```js
// Setup
const myArr = [2, 3, 4, 5, 6];

// Only change code below this line

```

# --solutions--

```js
const myArr = [2, 3, 4, 5, 6];
let total = 0;

for (let i = 0; i < myArr.length; i++) {
  total += myArr[i];
}
```
