---
id: 5679ceb97cbaa8c51670a16b
title: Vracanje Boolean Vrednosti iz Funkcija
challengeType: 1
forumTopicId: 18273
dashedName: returning-boolean-values-from-functions
---

# --description--

Mozda se secate iz lekcije <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/comparison-with-the-equality-operator" target="_blank" rel="noopener noreferrer nofollow">Poredjenje sa operatorom jednakosti</a> da svi operatori poredjenja vracaju boolean vrednosti `true` ili `false`.

Ponekad ljudi koriste `if/else` iskaz da obave poredjenje, ovako:

```js
function isEqual(a, b) {
  if (a === b) {
    return true;
  } else {
    return false;
  }
}
```

Ali postoji bolji nacin. Posto `===` vec vraca `true` ili `false`, mozemo direktno vratiti rezultat poredjenja:

```js
function isEqual(a, b) {
  return a === b;
}
```

# --instructions--

Izmenite funkciju `isLess` tako da uklonite `if/else` iskaz.

# --hints--

`isLess(10, 15)` treba da vrati `true`

```js
assert(isLess(10, 15) === true);
```

`isLess(15, 10)` treba da vrati `false`

```js
assert(isLess(15, 10) === false);
```

Ne treba da koristite nijedan `if` niti `else` iskaz

```js
assert(!/if|else/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function isLess(a, b) {
  // Only change code below this line
  if (a < b) {
    return true;
  } else {
    return false;
  }
  // Only change code above this line
}

isLess(10, 15);
```

# --solutions--

```js
function isLess(a, b) {
  return a < b;
}
```
