---
id: 567af2437cbaa8c51670a16c
title: Provera Svojstava Objekta
challengeType: 1
forumTopicId: 18324
dashedName: testing-objects-for-properties
---

# --description--

Da proverite da li neko svojstvo postoji u objektu, mozete koristiti metodu `.hasOwnProperty()`. `someObject.hasOwnProperty(someProperty)` vraca `true` ili `false` u zavisnosti od toga da li je svojstvo definisano u objektu.

**Primer**

```js
function checkForProperty(object, property) {
  return object.hasOwnProperty(property);
}

checkForProperty({ top: 'hat', bottom: 'pants' }, 'top'); // true
checkForProperty({ top: 'hat', bottom: 'pants' }, 'middle'); // false
```

Prvi poziv funkcije `checkForProperty` vraca `true`, dok drugi vraca `false`.

# --instructions--

Izmenite funkciju `checkObj` tako da proverava da li objekat prosledjen kroz parametar `obj` sadrzi svojstvo prosledjeno kroz parametar `checkProp`. Ako svojstvo postoji u `obj`, vratite njegovu vrednost. Ako ne postoji, vratite `Not Found`.

# --hints--

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "gift")` treba da vrati string `pony`.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'gift') === 'pony'
);
```

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "pet")` treba da vrati string `kitten`.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'pet') === 'kitten'
);
```

`checkObj({gift: "pony", pet: "kitten", bed: "sleigh"}, "house")` treba da vrati string `Not Found`.

```js
assert(
  checkObj({ gift: 'pony', pet: 'kitten', bed: 'sleigh' }, 'house') ===
    'Not Found'
);
```

`checkObj({city: "Seattle"}, "city")` treba da vrati string `Seattle`.

```js
assert(checkObj({ city: 'Seattle' }, 'city') === 'Seattle');
```

`checkObj({city: "Seattle"}, "district")` treba da vrati string `Not Found`.

```js
assert(checkObj({ city: 'Seattle' }, 'district') === 'Not Found');
```

`checkObj({pet: "kitten", bed: "sleigh"}, "gift")` treba da vrati string `Not Found`.

```js
assert(checkObj({ pet: 'kitten', bed: 'sleigh' }, 'gift') === 'Not Found');
```

# --seed--

## --seed-contents--

```js
function checkObj(obj, checkProp) {
  // Only change code below this line
  return "Change Me!";
  // Only change code above this line
}
```

# --solutions--

```js
function checkObj(obj, checkProp) {
  if(obj.hasOwnProperty(checkProp)) {
    return obj[checkProp];
  } else {
    return "Not Found";
  }
}
```
