---
id: 5690307fddb111c6084545d7
title: Logicki Redosled u if else Iskazima
challengeType: 1
forumTopicId: 18228
dashedName: logical-order-in-if-else-statements
---

# --description--

Redosled je bitan u `if`, `else if` iskazima.

Funkcija se izvrsava odozgo nadole, pa morate paziti koji uslov dolazi prvi.

Pogledajte ova dva primera funkcija.

Ovo je prvi:

```js
function foo(x) {
  if (x < 1) {
    return "Less than one";
  } else if (x < 2) {
    return "Less than two";
  } else {
    return "Greater than or equal to two";
  }
}
```

A drugi samo menja redosled uslova:

```js
function bar(x) {
  if (x < 2) {
    return "Less than two";
  } else if (x < 1) {
    return "Less than one";
  } else {
    return "Greater than or equal to two";
  }
}
```

Iako ove dve funkcije izgledaju gotovo isto, prosledjivanje brojeva daje razlicite rezultate.

```js
foo(0)
bar(0)
```

`foo(0)` ce vratiti string `Less than one`, a `bar(0)` ce vratiti string `Less than two`.

# --instructions--

Izmenite logicki redosled u funkciji tako da vraca tacan rezultat u svakom slucaju.

# --hints--

`orderMyLogic(4)` treba da vrati string `Less than 5`

```js
assert(orderMyLogic(4) === 'Less than 5');
```

`orderMyLogic(6)` treba da vrati string `Less than 10`

```js
assert(orderMyLogic(6) === 'Less than 10');
```

`orderMyLogic(11)` treba da vrati string `Greater than or equal to 10`

```js
assert(orderMyLogic(11) === 'Greater than or equal to 10');
```

# --seed--

## --seed-contents--

```js
function orderMyLogic(val) {
  if (val < 10) {
    return "Less than 10";
  } else if (val < 5) {
    return "Less than 5";
  } else {
    return "Greater than or equal to 10";
  }
}

orderMyLogic(7);
```

# --solutions--

```js
function orderMyLogic(val) {
  if(val < 5) {
    return "Less than 5";
  } else if (val < 10) {
    return "Less than 10";
  } else {
    return "Greater than or equal to 10";
  }
}
```
