---
id: 56bbb991ad1ed5201cd392cc
title: Menjanje Nizova Pomocu pop
challengeType: 1
forumTopicId: 18236
dashedName: manipulate-arrays-with-pop
---

# --description--

Jos jedan nacin za izmenu podataka u nizu je upotreba funkcije `.pop()`.

`.pop()` uklanja vrednost sa kraja niza. Tu uklonjenu vrednost mozemo sacuvati dodeljivanjem novoj promenljivoj. Drugim recima, `.pop()` uklanja poslednji element niza i vraca taj element.

Bilo koji tip elementa moze se ukloniti iz niza - broj, string, pa cak i ugnjezdeni nizovi.

```js
const threeArr = [1, 4, 6];
const oneDown = threeArr.pop();
console.log(oneDown);
console.log(threeArr);
```

Prvi `console.log` ce prikazati vrednost `6`, a drugi vrednost `[1, 4]`.

# --instructions--

Koristite funkciju `.pop()` da uklonite poslednji element iz `myArray` i smestite vracenu vrednost u novu promenljivu `removedFromMyArray`.

# --hints--

`myArray` treba da sadrzi samo `[["John", 23]]`.

```js
assert(
  (function (d) {
    if (d[0][0] == 'John' && d[0][1] === 23 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

Treba da koristite `pop()` na `myArray`.

```js
assert(/removedFromMyArray\s*=\s*myArray\s*.\s*pop\s*(\s*)/.test(__helpers.removeJSComments(code)));
```

`removedFromMyArray` treba da sadrzi samo `["cat", 2]`.

```js
assert(
  (function (d) {
    if (d[0] == 'cat' && d[1] === 2 && d[2] == undefined) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
const removedFromMyArray = myArray.pop();
```
