---
id: 56bbb991ad1ed5201cd392cb
title: Menjanje Nizova Pomocu push
challengeType: 1
forumTopicId: 18237
dashedName: manipulate-arrays-with-push
---

# --description--

Jednostavan nacin za dodavanje podataka na kraj niza je metoda `push()`.

Metoda `push()` prima jedan ili vise argumenata i dodaje ih na kraj niza redosledom kojim su navedeni. Vraca novu duzinu niza.

Primeri:

```js
const arr1 = [1, 2, 3];
arr1.push(4, 5);

const arr2 = ["Stimpson", "J", "cat"];
arr2.push(["happy", "joy"]);
```

`arr1` sada ima vrednost `[1, 2, 3, 4, 5]`, a `arr2` ima vrednost `["Stimpson", "J", "cat", ["happy", "joy"]]`.

# --instructions--

Dodajte `["dog", 3]` na kraj promenljive `myArray`.

# --hints--

`myArray` sada treba da bude `[["John", 23], ["cat", 2], ["dog", 3]]`.

```js
assert(
  (function (d) {
    if (
      d[2] != undefined &&
      d[0][0] == 'John' &&
      d[0][1] === 23 &&
      d[2][0] == 'dog' &&
      d[2][1] === 3 &&
      d[2].length == 2
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
myArray.push(["dog",3]);
```
