---
id: 56bbb991ad1ed5201cd392d2
title: Dodajte Nova Svojstva JavaScript Objektu
challengeType: 1
forumTopicId: 301169
dashedName: add-new-properties-to-a-javascript-object
---

# --description--

Nova svojstva postojecim JavaScript objektima mozete dodati na isti nacin kao sto ih menjate.

Ovako dodajemo svojstvo `bark` objektu `ourDog`:

```js
ourDog.bark = "bow-wow";
```

au

```js
ourDog["bark"] = "bow-wow";
```

Kada sada proverimo `ourDog.bark`, dobicemo zvuk lajanja, `bow-wow`.

Primer:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};

ourDog.bark = "bow-wow";
```

# --instructions--

Dodajte svojstvo `bark` objektu `myDog` i postavite ga na zvuk psa, kao "woof". Mozete koristiti notaciju tacke ili uglastih zagrada.

# --hints--

Treba da dodate svojstvo `bark` objektu `myDog`.

```js
assert(myDog.bark !== undefined);
```

Ne treba da dodate `bark` u samoj inicijalizaciji `myDog`.

```js
assert(!/bark[^\n]:/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};


```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.bark = "Woof Woof";
```
