---
id: 56bbb991ad1ed5201cd392ce
title: Menjanje Nizova Pomocu unshift
challengeType: 1
forumTopicId: 18239
dashedName: manipulate-arrays-with-unshift
---

# --description--

Ne samo da mozete da uklonite elemente sa pocetka niza pomocu `shift`, vec mozete i da ih dodate na pocetak pomocu `unshift`.

`.unshift()` radi slicno kao `.push()`, ali umesto dodavanja elementa na kraj niza, `unshift()` dodaje element na pocetak niza.

Primer:

```js
const ourArray = ["Stimpson", "J", "cat"];
ourArray.shift();
ourArray.unshift("Happy");
```

Posle `shift`, `ourArray` ce imati vrednost `["J", "cat"]`. Posle `unshift`, `ourArray` ce imati vrednost `["Happy", "J", "cat"]`.

# --instructions--

Dodajte `["Paul", 35]` na pocetak promenljive `myArray` koristeci `unshift()`.

# --hints--

`myArray` sada treba da bude `[["Paul", 35], ["dog", 3]]`.

```js
assert(
  (function (d) {
    if (
      typeof d[0] === 'object' &&
      d[0][0] == 'Paul' &&
      d[0][1] === 35 &&
      d[1][0] != undefined &&
      d[1][0] == 'dog' &&
      d[1][1] != undefined &&
      d[1][1] == 3
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();
myArray.unshift(["Paul", 35]);
```
