---
id: 56bbb991ad1ed5201cd392ca
title: Pristup Podacima Niza Pomocu Indeksa
challengeType: 1
forumTopicId: 16158
dashedName: access-array-data-with-indexes
---

# --description--

Mozemo pristupiti podacima unutar nizova koristeci <dfn>indekse</dfn>.

Indeksi niza zapisuju se uglastim zagradama, slicno kao kod stringova, ali umesto karaktera oznacavaju element niza. Kao i stringovi, nizovi koriste indeksiranje <dfn>od nule</dfn>, pa prvi element ima indeks `0`.

<br>

**Primer**

```js
const array = [50, 60, 70];
console.log(array[0]);
const data = array[1];
```

`console.log(array[0])` ispisuje `50`, a `data` ima vrednost `60`.

# --instructions--

Napravite promenljivu pod nazivom `myData` i postavite je na prvu vrednost iz `myArray` koristeci uglaste zagrade.

# --hints--

Promenljiva `myData` treba da bude jednaka prvoj vrednosti iz `myArray`.

```js
assert(
  (function () {
    if (
      typeof myArray !== 'undefined' &&
      typeof myData !== 'undefined' &&
      myArray[0] === myData
    ) {
      return true;
    } else {
      return false;
    }
  })()
);
```

Podacima u promenljivoj `myArray` treba pristupiti pomocu uglaste notacije.

```js
assert(
  (function () {
    if (__helpers.removeJSComments(code).match(/\s*=\s*myArray\[0\]/g)) {
      return true;
    } else {
      return false;
    }
  })()
);
```

# --seed--

## --seed-contents--

```js
const myArray = [50, 60, 70];


```

# --solutions--

```js
const myArray = [50, 60, 70];
const myData = myArray[0];
```
