---
id: 56bbb991ad1ed5201cd392cd
title: Menjanje Nizova Pomocu shift
challengeType: 1
forumTopicId: 18238
dashedName: manipulate-arrays-with-shift
---

# --description--

`pop()` uvek uklanja poslednji element niza. Sta ako zelite da uklonite prvi?

Tu na scenu stupa `.shift()`. Radi slicno kao `.pop()`, samo uklanja prvi element umesto poslednjeg.

Primer:

```js
const ourArray = ["Stimpson", "J", ["cat"]];
const removedFromOurArray = ourArray.shift();
```

`removedFromOurArray` bi imao vrednost stringa `Stimpson`, a `ourArray` bi bio `["J", ["cat"]]`.

# --instructions--

Koristite funkciju `.shift()` da uklonite prvi element iz `myArray` i smestite uklonjenu vrednost u novu promenljivu `removedFromMyArray`.

# --hints--

`myArray` sada treba da bude `[["dog", 3]]`.

```js
assert(
  (function (d) {
    if (d[0][0] == 'dog' && d[0][1] === 3 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

`removedFromMyArray` treba da sadrzi `["John", 23]`.

```js
assert(
  (function (d) {
    if (
      d[0] == 'John' &&
      d[1] === 23 &&
      typeof removedFromMyArray === 'object'
    ) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line
const removedFromMyArray = myArray.shift();
```
