---
id: 56bbb991ad1ed5201cd392d1
title: Azuriranje Svojstava Objekta
challengeType: 1
forumTopicId: 18336
dashedName: updating-object-properties
---

# --description--

Nakon kreiranja JavaScript objekta, njegova svojstva mozete azurirati u bilo kom trenutku, kao i kod drugih tipova podataka. Mozete koristiti notaciju tacke ili uglastih zagrada.

Na primer, uzmimo `ourDog`:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```

Pošto je veoma srećan pas, promenimo njegovo ime u string `Happy Camper`. Evo kako ažuriramo svojstvo `name`: `ourDog.name = "Happy Camper";` ili `ourDog["name"] = "Happy Camper";`. Kada sada proverimo `ourDog.name`, umesto `Camper` dobićemo njegovo novo ime, `Happy Camper`.

# --instructions--

Azurirajte svojstvo `name` objekta `myDog`. Promenite ime iz `Coder` u `Happy Coder`. Mozete koristiti notaciju tacke ili uglastih zagrada.

# --hints--

Treba da azurirate svojstvo `name` objekta `myDog` tako da bude string `Happy Coder`.

```js
assert(/happy coder/gi.test(myDog.name));
```

Ne smete menjati definiciju `myDog`.

```js
assert(/"name": "Coder"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.name = "Happy Coder";
```
