---
id: 56bbb991ad1ed5201cd392d0
title: Kreirajte JavaScript Objekte
challengeType: 1
forumTopicId: 16769
dashedName: build-javascript-objects
---

# --description--

Verovatno ste vec culi za rec `object`.

Objekti su slicni `arrays`, osim sto se umesto indeksa za pristup i izmenu podataka koriste ono sto se zove `properties`.

Objekti su korisni za cuvanje podataka na organizovan nacin i mogu predstavljati stvarne entitete, kao sto je macka.

Evo primera objekta macke:

```js
const cat = {
  "name": "Whiskers",
  "legs": 4,
  "tails": 1,
  "enemies": ["Water", "Dogs"]
};
```

U ovom primeru, sva svojstva su sacuvana kao stringovi, kao `name`, `legs` i `tails`. Medjutim, mozete koristiti i brojeve kao svojstva. Cak mozete izostaviti navodnike za svojstva od jedne reci, kao ovde:

```js
const anotherObject = {
  make: "Ford",
  5: "five",
  "model": "focus"
};
```

Ipak, ako objekat ima svojstva koja nisu string, JavaScript ce ih automatski pretvoriti u string.

# --instructions--

Kreirajte objekat koji predstavlja psa pod nazivom `myDog` sa svojstvima `name` (string), `legs`, `tails` i `friends`.

Ovim svojstvima mozete dodeliti bilo koje vrednosti, pod uslovom da je `name` string, `legs` i `tails` brojevi, a `friends` niz.

# --hints--

`myDog` treba da ima svojstvo `name` i ono treba da bude `string`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('name') &&
      z.name !== undefined &&
      typeof z.name === 'string'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` treba da ima svojstvo `legs` i ono treba da bude `number`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('legs') &&
      z.legs !== undefined &&
      typeof z.legs === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` treba da ima svojstvo `tails` i ono treba da bude `number`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('tails') &&
      z.tails !== undefined &&
      typeof z.tails === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` treba da ima svojstvo `friends` i ono treba da bude `array`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('friends') &&
      z.friends !== undefined &&
      Array.isArray(z.friends)
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` treba da ima samo navedena svojstva.

```js
assert(
  (function (z) {
    return Object.keys(z).length === 4;
  })(myDog)
);
```

# --seed--

## --seed-contents--

```js
const myDog = {
  // Only change code below this line


  // Only change code above this line
};
```

# --solutions--

```js
const myDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```
