---
id: 56bbb991ad1ed5201cd392cf
title: Napisite Ponovo Iskoristiv JavaScript Pomocu Funkcije
challengeType: 1
forumTopicId: 18378
dashedName: write-reusable-javascript-with-functions
---

# --description--

U JavaScript-u, mozemo podeliti kod na ponovo iskoristive delove koji se zovu <dfn>funkcije</dfn>.

Evo primera funkcije:

```js
function functionName() {
  console.log("Hello World");
}
```

Mozete pozvati ovu funkciju tako sto koristite njeno ime praceno obicnim zagradama, ovako: `functionName();` Svaki put kada se funkcija pozove, ispisace poruku `Hello World` u developerskoj konzoli. Sav kod unutar viticastih zagrada se izvrsi svaki put kada se funkcija pozove.

# --instructions--

<ol>
  <li>
    Napravite funkciju pod nazivom <code>reusableFunction</code> koja ispisuje string <code>Hi World</code> u developerskoj konzoli.
  </li>
  <li>
    Pozovite tu funkciju.
  </li>
</ol>

# --before-each--

```js
function testConsole() {
  var originalConsole = console;
  var nativeLog = console.log;
  var hiWorldWasLogged = false;
  console.log = function (message) {
    if (message === 'Hi World') {
      console.warn(message);
      hiWorldWasLogged = true;
    }
    if (nativeLog.apply) {
      nativeLog.apply(originalConsole, arguments);
    } else {
      var nativeMsg = Array.prototype.slice.apply(arguments).join(' ');
      nativeLog(nativeMsg);
    }
  };
  reusableFunction();
  console.log = nativeLog;
  return hiWorldWasLogged;
}
```

# --hints--

`reusableFunction` treba da bude funkcija.

```js
assert(typeof reusableFunction === 'function');
```

Kada se `reusableFunction` pozove, treba da ispisuje string `Hi World` u konzoli.

```js
assert(testConsole());
```

Treba da pozovete `reusableFunction` nakon sto je definisana.

```js
const functionStr = reusableFunction && __helpers.removeWhiteSpace(reusableFunction.toString());
const codeWithoutFunction = __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).replace(/reusableFunction\(\)\{/g, '');
assert(/reusableFunction\(\)/.test(codeWithoutFunction));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reusableFunction() {
  console.log("Hi World");
}
reusableFunction();
```
