---
id: 56bbb991ad1ed5201cd392d3
title: Obrisite Svojstva iz JavaScript Objekta
challengeType: 1
forumTopicId: 17560
dashedName: delete-properties-from-a-javascript-object
---

# --description--

Svojstva iz objekata mozemo i obrisati ovako:

```js
delete ourDog.bark;
```

Primer:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"],
  "bark": "bow-wow"
};

delete ourDog.bark;
```

Nakon poslednje prikazane linije iznad, `ourDog` izgleda ovako:

```js
{
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
}
```

# --instructions--

Obrisite svojstvo `tails` iz `myDog`. Mozete koristiti notaciju tacke ili uglastih zagrada.

# --hints--

Treba da obrisete svojstvo `tails` iz `myDog`.

```js
assert(typeof myDog === 'object' && myDog.tails === undefined);
```

Ne treba da menjate postavku `myDog`.

```js
assert(__helpers.removeJSComments(code).match(/"tails": 1/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};
delete myDog.tails;
```
