---
id: 587d7b7e367417b2b2512b22
title: Koristite parseInt Funkciju sa Osnovom Brojnog Sistema
challengeType: 1
forumTopicId: 301182
dashedName: use-the-parseint-function-with-a-radix
---

# --description--

Funkcija `parseInt()` parsira string i vraca ceo broj. Prima drugi argument za osnovu brojnog sistema (radix), koji opisuje osnovu broja u stringu. Osnova moze biti ceo broj izmedju 2 i 36.

Poziv funkcije izgleda ovako:

```js
parseInt(string, radix);
```

A evo primera:

```js
const a = parseInt("11", 2);
```

Argument osnove pokazuje da je `11` u binarnom sistemu, odnosno bazi 2. Ovaj primer konvertuje string `11` u ceo broj `3`.

# --instructions--

Koristite `parseInt()` u funkciji `convertToInteger` da konvertujete binarni broj u ceo broj i vratite rezultat.

# --hints--

`convertToInteger` treba da koristi funkciju `parseInt()`

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("10011")` treba da vrati broj

```js
assert(typeof convertToInteger('10011') === 'number');
```

`convertToInteger("10011")` treba da vrati 19

```js
assert(convertToInteger('10011') === 19);
```

`convertToInteger("111001")` treba da vrati 57

```js
assert(convertToInteger('111001') === 57);
```

`convertToInteger("JamesBond")` treba da vrati `NaN`

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("10011");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str, 2);
}
```
