---
id: 587d7b7e367417b2b2512b21
title: Koristite Vise Ternarnih Operatora
challengeType: 1
forumTopicId: 301179
dashedName: use-multiple-conditional-ternary-operators
---

# --description--

U prethodnom izazovu koristili ste jedan uslovni operator. Takodje mozete povezati vise takvih operatora kako biste proverili vise uslova.

Sledeca funkcija koristi `if`, `else if` i `else` da proveri vise uslova:

```js
function findGreaterOrEqual(a, b) {
  if (a === b) {
    return "a and b are equal";
  }
  else if (a > b) {
    return "a is greater";
  }
  else {
    return "b is greater";
  }
}
```

Funkcija iznad moze se prepisati koristeci vise ternarnih operatora:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" 
    : (a > b) ? "a is greater" 
    : "b is greater";
}
```

Smatra se dobrom praksom da vise ternarnih operatora formatirate tako da svaki uslov bude u svom redu, kao iznad. Koriscenje vise ternarnih operatora bez dobrog formatiranja moze uciniti kod tezim za citanje. Na primer:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" : (a > b) ? "a is greater" : "b is greater";
}
```

# --instructions--

U funkciji `checkSign`, koristite vise ternarnih operatora - prateci preporuceni raspored iz `findGreaterOrEqual` - da proverite da li je broj pozitivan, negativan ili nula. Funkcija treba da vrati `positive`, `negative` ili `zero`.

# --hints--

`checkSign` treba da koristi vise ternarnih operatora

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?\s*?\?\s*?.+?\s*?:\s*?.+?/gi.test(__helpers.removeJSComments(code)));
```

`checkSign(10)` treba da vrati string `positive`. Velicina slova je bitna

```js
assert(checkSign(10) === 'positive');
```

`checkSign(-12)` treba da vrati string `negative`. Velicina slova je bitna

```js
assert(checkSign(-12) === 'negative');
```

`checkSign(0)` treba da vrati string `zero`. Velicina slova je bitna

```js
assert(checkSign(0) === 'zero');
```

# --seed--

## --seed-contents--

```js
function checkSign(num) {

}

checkSign(10);
```

# --solutions--

```js
function checkSign(num) {
  return (num > 0) ? 'positive' : (num < 0) ? 'negative' : 'zero';
}
```
