---
id: 587d7b7e367417b2b2512b24
title: Tumia opereta ya kikokotoo cha masharti cha pande tatu
challengeType: 1
forumTopicId: 301181
dashedName: use-the-conditional-ternary-operator
---

# --description--

<Uslovni operator</dfn>, poznat i kao <dfn>ternarni operator</dfn>, može se koristiti kao jednolinijski `if-else` izraz.

Sintaksa je `a ? b : c`, gde je `a` uslov, `b` je kod koji se izvršava kada uslov vrati `true`, a `c` je kod koji se izvršava kada uslov vrati `false`.

Sledeća funkcija koristi naredbu `if/else` da proveri uslov:

```js
function findGreater(a, b) {
  if(a > b) {
    return "a is greater";
  }
  else {
    return "b is greater or equal";
  }
}
```

Ovo se može prepisati korišćenjem uslovnog operatora:

```js
function findGreater(a, b) {
  return a > b ? "a is greater" : "b is greater or equal";
}
```

# --instructions--

Koristi uslovni operator u funkciji `checkEqual` da proveri da li su dva broja jednaka. Funkcija treba da vrati string `Equal` ili string `Not Equal`.

# --hints--

`checkEqual` inapaswa kutumia opereta ya masharti

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?/.test(__helpers.removeJSComments(code)));
```

`checkEqual(1, 2)` treba da vrati string `Not Equal`

```js
assert(checkEqual(1, 2) === 'Not Equal');
```

`checkEqual(1, 1)` treba da vrati string `Equal`

```js
assert(checkEqual(1, 1) === 'Equal');
```

`checkEqual(1, -1)` treba da vrati string `Not Equal`

```js
assert(checkEqual(1, -1) === 'Not Equal');
```

# --seed--

## --seed-contents--

```js
function checkEqual(a, b) {

}

checkEqual(1, 2);
```

# --solutions--

```js
function checkEqual(a, b) {
  return a === b ? "Equal" : "Not Equal";
}
```
