---
id: 587d7b7e367417b2b2512b23
title: Koristite parseInt Funkciju
challengeType: 1
forumTopicId: 301183
dashedName: use-the-parseint-function
---

# --description--

Funkcija `parseInt()` parsira string i vraca ceo broj. Evo primera:

```js
const a = parseInt("007");
```

Gornja funkcija pretvara string `007` u ceo broj `7`. Ako prvi znak u stringu ne moze da se pretvori u broj, funkcija vraca `NaN`.

# --instructions--

Koristite `parseInt()` u funkciji `convertToInteger` da konvertujete ulazni string `str` u ceo broj i vratite ga.

# --hints--

`convertToInteger` treba da koristi funkciju `parseInt()`

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("56")` treba da vrati broj

```js
assert(typeof convertToInteger('56') === 'number');
```

`convertToInteger("56")` treba da vrati 56

```js
assert(convertToInteger('56') === 56);
```

`convertToInteger("77")` treba da vrati 77

```js
assert(convertToInteger('77') === 77);
```

`convertToInteger("JamesBond")` treba da vrati `NaN`

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("56");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str);
}
```
