---
id: 587d7b87367417b2b2512b3f
title: Istrazite Razlike Izmedju var i let Kljucnih Reci
challengeType: 1
forumTopicId: 301202
dashedName: explore-differences-between-the-var-and-let-keywords
---

# --description--

Jedan od vecih problema pri deklarisanju promenljivih sa `var` je to sto lako mozete ponovo deklarisati promenljivu:

```js
var camper = "James";
var camper = "David";
console.log(camper);
```

U kodu iznad, promenljiva `camper` je prvo deklarisana kao `James`, a zatim ponovo kao `David`. Konzola prikazuje string `David`.

U malim programima mozda necete naici na ovaj problem. Ali kako kod raste, mozete slucajno ponovo deklarisati promenljivu koju niste nameravali. Posto ovo ponasanje ne baca gresku, pronalazenje i ispravljanje bagova postaje teze.

Kljucna rec `let` uvedena je u ES6, velikom JavaScript azuriranju, kako bi se resio ovaj potencijalni problem sa `var`. O drugim ES6 funkcionalnostima ucicete u narednim izazovima.

Ako u gornjem kodu zamenite `var` sa `let`, dobicete gresku:

```js
let camper = "James";
let camper = "David";
```

Greska ce biti vidljiva u konzoli vaseg pregledača.

Dakle, za razliku od `var`, kada koristite `let`, promenljiva sa istim imenom moze biti deklarisana samo jednom.

# --instructions--

Ispravite kod da koristi samo kljucnu rec `let`.

# --hints--

`var` ne treba da postoji u kodu.

```js
assert.notMatch(code, /var/g);
```

`catName` treba da bude string `Oliver`.

```js
assert.equal(catName, 'Oliver');
```

`catSound` treba da bude string `Meow!`

```js
assert.equal(catSound, 'Meow!');
```

# --seed--

## --seed-contents--

```js
var catName = "Oliver";
var catSound = "Meow!";
```

# --solutions--

```js
let catName = "Oliver";
let catSound = "Meow!";
```
