---
id: 598e8944f009e646fc236146
title: Razumevanje undefined Vrednosti Koju Funkcija Vraca
challengeType: 1
forumTopicId: 301177
dashedName: understanding-undefined-value-returned-from-a-function
---

# --description--

Funkcija moze da sadrzi `return` iskaz, ali ne mora. Kada funkcija nema `return`, pri pozivu izvrsava unutrasnji kod, ali vracena vrednost je `undefined`.

**Primer**

```js
let sum = 0;

function addSum(num) {
  sum = sum + num;
}

addSum(3);
```

`addSum` je funkcija bez `return` iskaza. Funkcija menja globalnu promenljivu `sum`, ali njena povratna vrednost je `undefined`.

# --instructions--

Napravite funkciju `addFive` bez argumenata. Ova funkcija treba da doda 5 promenljivoj `sum`, ali povratna vrednost treba da bude `undefined`.

# --hints--

`addFive` treba da bude funkcija.

```js
assert(typeof addFive === 'function');
```

Nakon izvrsavanja svih funkcija, `sum` treba da bude jednako `8`.

```js
assert(sum === 8);
```

Povratna vrednost funkcije `addFive` treba da bude `undefined`.

```js
assert(addFive() === undefined);
```

Unutar funkcije `addFive` treba da dodate `5` promenljivoj `sum`.

```js
assert(
  __helpers.removeWhiteSpace(addFive.toString()).match(/sum=sum\+5|sum\+=5/)
);
```

# --seed--

## --seed-contents--

```js
// Setup
let sum = 0;

function addThree() {
  sum = sum + 3;
}

// Only change code below this line


// Only change code above this line

addThree();
addFive();
```

# --solutions--

```js
let sum = 0;

function addThree() {
  sum = sum + 3;
}

function addFive() {
  sum = sum + 5;
}

addThree();
addFive();
```
