---
id: 599a789b454f2bbd91a3ff4d
title: Vezbanje Poredjenja Razlicitih Vrednosti
challengeType: 1
forumTopicId: 301174
dashedName: practice-comparing-different-values
---

# --description--

U prethodna dva izazova ucili smo o operatoru jednakosti (`==`) i operatoru stroge jednakosti (`===`). Hajde da uradimo kratak pregled i dodatno uvezbamo ove operatore.

Ako vrednosti koje se porede nisu istog tipa, operator jednakosti ce uraditi konverziju tipa i zatim uporediti vrednosti. Operator stroge jednakosti poredi i tip podatka i vrednost bez konverzije.

**Primeri**

`3 == '3'` vraca `true` zato sto JavaScript radi konverziju tipa iz stringa u broj. `3 === '3'` vraca `false` zato sto su tipovi razliciti i ne radi se konverzija.

**Napomena:** U JavaScript-u tip promenljive ili vrednosti mozete proveriti pomocu operatora `typeof`, na primer:

```js
typeof 3
typeof '3'
```

`typeof 3` vraca string `number`, a `typeof '3'` vraca string `string`.

# --instructions--

Funkcija `compareEquality` u editoru poredi dve vrednosti koristeci operator jednakosti. Izmenite funkciju da vraca string `Equal` samo kada su vrednosti strogo jednake.

# --hints--

`compareEquality(10, "10")` treba da vrati string `Not Equal`

```js
assert(compareEquality(10, '10') === 'Not Equal');
```

`compareEquality("20", 20)` treba da vrati string `Not Equal`

```js
assert(compareEquality('20', 20) === 'Not Equal');
```

Treba da koristite operator `===`

```js
assert(__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function compareEquality(a, b) {
  if (a == b) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

compareEquality(10, "10");
```

# --solutions--

```js
function compareEquality(a,b) {
  if (a === b) {
    return "Equal";
  }
  return "Not Equal";
}
```
