---
id: 5cc0bd7a49b71cb96132e54c
title: Koristite Rekurziju za Kreiranje Opsega Brojeva
challengeType: 1
forumTopicId: 301180
dashedName: use-recursion-to-create-a-range-of-numbers
---

# --description--

Nastavljajuci prethodni izazov, ovo je jos jedna prilika da upotrebite rekurzivnu funkciju za resavanje problema.

# --instructions--

Napravili smo funkciju `rangeOfNumbers` sa dva parametra. Funkcija treba da vrati niz celih brojeva koji pocinje brojem iz `startNum`, a zavrsava brojem iz `endNum`. Pocetni broj ce uvek biti manji ili jednak zavrsnom broju. Funkcija mora koristiti rekurziju (pozivati samu sebe) i ne sme koristiti nijedan tip petlje. Takodje treba da radi i kada su `startNum` i `endNum` jednaki.

# --hints--

Vasa funkcija treba da vraca niz.

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

Vas kod ne sme koristiti sintaksu petlji (`for` ili `while`) niti funkcije viseg nivoa kao sto su `forEach`, `map`, `filter` ili `reduce`.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

`rangeOfNumbers` treba da koristi rekurziju (poziv same sebe) za resavanje ovog izazova.

```js
assert(
  rangeOfNumbers.toString().match(/rangeOfNumbers\s*\(.+\)/)
);
```

`rangeOfNumbers(1, 5)` treba da vrati `[1, 2, 3, 4, 5]`.

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)` treba da vrati `[6, 7, 8, 9]`.

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)` treba da vrati `[4]`.

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

Globalne promenljive ne treba koristiti za cuvanje ovog niza.

```js
rangeOfNumbers(1, 3)
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js
function rangeOfNumbers(startNum, endNum) {
  return [];
};
```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
