---
id: 5cfa3679138e7d9595b9d9d4
title: Zamenite Petlje Rekurzijom
challengeType: 1
videoUrl: >-
  https://www.freecodecamp.org/news/how-recursion-works-explained-with-flowcharts-and-a-video-de61f40cb7f9/
forumTopicId: 301175
dashedName: replace-loops-using-recursion
---

# --description--

Rekurzija je ideja da se funkcija moze definisati pomocu same sebe. Da biste ovo bolje razumeli, razmotrite sledecu vezbu: pomnozite prvih `n` elemenata niza da biste dobili njihov proizvod. Koristeci `for` petlju, to moze izgledati ovako:

```js
  function multiply(arr, n) {
    let product = 1;
    for (let i = 0; i < n; i++) {
      product *= arr[i];
    }
    return product;
  }
```

Medjutim, obratite paznju da je `multiply(arr, n) == multiply(arr, n - 1) * arr[n - 1]`. To znaci da mozete prepisati `multiply` koristeci samu sebe i bez upotrebe petlje.

```js
  function multiply(arr, n) {
    if (n <= 0) {
      return 1;
    } else {
      return multiply(arr, n - 1) * arr[n - 1];
    }
  }
```

Rekurzivna verzija funkcije `multiply` radi ovako. U <dfn>osnovnom slucaju</dfn>, kada je `n <= 0`, vraca 1. Za vece vrednosti `n`, funkcija poziva samu sebe sa `n - 1`. Taj poziv se evaluira na isti nacin, pozivajuci `multiply` ponovo sve dok `n <= 0`. U tom trenutku svi pozivi mogu da se vrate, a pocetni `multiply` vraca odgovor.

**Napomena:** Rekurzivne funkcije moraju imati osnovni slucaj u kom vracaju vrednost bez novog poziva funkcije (u ovom primeru kada `n <= 0`), inace se nikada nece zavrsiti.

# --instructions--

Napisite rekurzivnu funkciju `sum(arr, n)` koja vraca zbir prvih `n` elemenata niza `arr`.

# --hints--

`sum([1], 0)` treba da bude jednako 0.

```js
assert.equal(sum([1], 0), 0);
```

`sum([2, 3, 4], 1)` treba da bude jednako 2.

```js
assert.equal(sum([2, 3, 4], 1), 2);
```

`sum([2, 3, 4, 5], 3)` treba da bude jednako 9.

```js
assert.equal(sum([2, 3, 4, 5], 3), 9);
```

Vas kod ne sme zavisiti od bilo kakvih petlji (`for` ili `while`) niti od funkcija viseg nivoa kao sto su `forEach`, `map`, `filter` ili `reduce`.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

Treba da koristite rekurziju za resavanje ovog izazova.

```js
assert(
  sum.toString().match(/sum\(.*\)/g).length > 1
);
```

# --seed--

## --seed-contents--

```js
function sum(arr, n) {
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function sum(arr, n) {
  // Only change code below this line
  if(n <= 0) {
    return 0;
  } else {
    return sum(arr, n - 1) + arr[n - 1];
  }
  // Only change code above this line
}
```
