---
id: 5ee127a03c3b35dd45426493
title: Dodeljivanje Vrednosti Jedne Promenljive Drugoj
challengeType: 1
forumTopicId: 418265
dashedName: assigning-the-value-of-one-variable-to-another
---

# --description--

Nakon sto dodelite vrednost promenljivoj koristeci <dfn>operator dodele</dfn>, mozete dodeliti vrednost te promenljive drugoj promenljivoj istim <dfn>operatorom dodele</dfn>.

```js
var myVar;
myVar = 5;
var myNum;
myNum = myVar;
```

Gore se deklarise promenljiva `myVar` bez vrednosti, zatim joj se dodeljuje `5`. Sledece, promenljiva `myNum` se deklarise bez vrednosti. Potom se sadrzaj promenljive `myVar` (sto je `5`) dodeljuje promenljivoj `myNum`. Sada i `myNum` ima vrednost `5`.

# --instructions--

Dodelite sadrzaj promenljive `a` promenljivoj `b`.

# --hints--

Ne treba da menjate kod iznad navedene napomene.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)) && /a = 7;/.test(__helpers.removeJSComments(code)) && /var b;/.test(__helpers.removeJSComments(code)));
```

`b` treba da ima vrednost `7`.

```js
assert(typeof b === 'number' && b === 7);
```

`a` treba da bude dodeljeno promenljivoj `b` koristeci `=`.

```js
assert(/b\s*=\s*a\s*/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
var a;
a = 7;
var b;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
var b;
b = a;
```
