---
id: bd7993c9c69feddfaeb7bdef
title: Pomnozite Dva Decimalna Broja u JavaScript-u
challengeType: 1
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

U JavaScript-u mozete racunati i sa decimalnim brojevima, kao i sa celim brojevima.

Hajde da pomnozimo dva decimalna broja da dobijemo njihov proizvod.

# --instructions--

Promenite `0.0` tako da proizvod bude jednak `5.0`.

# --hints--

Promenljiva `product` treba da bude jednaka `5.0`.

```js
assert(product === 5.0);
```

Treba da koristite operator `*`

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 2.0 * 0.0;
```

# --solutions--

```js
const product = 2.0 * 2.5;
```
