---
id: cf1111c1c11feddfaeb3bdef
title: Saberite Dva Broja Pomocu JavaScript-a
challengeType: 1
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` je tip podataka u JavaScript-u koji predstavlja numericke podatke.

Sada pokusajte da saberete dva broja koristeci JavaScript.

JavaScript koristi simbol `+` kao operator sabiranja kada se postavi izmedju dva broja.

**Primer:**

```js
const myVar = 5 + 10;
```

`myVar` sada ima vrednost `15`.

# --instructions--

Promenite `0` tako da zbir bude `20`.

# --hints--

`sum` treba da bude jednako `20`.

```js
assert(sum === 20);
```

Treba da koristite operator `+`.

```js
assert(/\+/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
