---
id: cf1231c1c11feddfaeb5bdef
title: Mnozenje Dva Broja u JavaScript-u
challengeType: 1
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

Mozemo i mnoziti jedan broj drugim.

JavaScript koristi simbol `*` za mnozenje dva broja.

**Primer**

```js
const myVar = 13 * 13;
```

`myVar` ce imati vrednost `169`.

# --instructions--

Promenite `0` tako da proizvod bude `80`.

# --hints--

Promenljiva `product` treba da bude jednaka 80.

```js
assert(product === 80);
```

Treba da koristite operator `*`.

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 8 * 0;
```

# --solutions--

```js
const product = 8 * 10;
```
