---
id: cf1391c1c11feddfaeb4bdef
title: Kreiranje Decimalnih Brojeva u JavaScript-u
challengeType: 1
forumTopicId: 16826
dashedName: create-decimal-numbers-with-javascript
---

# --description--

Mozemo cuvati decimalne brojeve i u promenljivima. Decimalni brojevi ponekad se nazivaju <dfn>floating point</dfn> brojevi ili <dfn>float</dfn>-ovi.

**Napomena:** kada radite racunanja sa brojevima, izvrsavaju se sa ogranicenom tacnoscu. Operacije sa decimalnim brojevima mogu dati neocekivane rezultate. Ako dodjete do takvog slucaja, otvorite temu na <a href="https://forum.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">freeCodeCamp forumu</a>.

# --instructions--

Napravite promenljivu `myDecimal` i dodelite joj decimalnu vrednost sa decimalnim delom (na primer `5.7`).

# --hints--

`myDecimal` treba da bude broj.

```js
assert(typeof myDecimal === 'number');
```

`myDecimal` treba da ima decimalnu tacku

```js
assert(myDecimal % 1 != 0);
```

# --seed--

## --seed-contents--

```js
const ourDecimal = 5.7;

// Only change code below this line

```

# --solutions--

```js
const myDecimal = 9.9;
```
