---
id: 66fe4f33a2cc9b33f4d5cd9b
title: Izgradnja generatora permutacija
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

U ovom laboratoriju, kreiraćeš generator permutacija koji uzima sekvencu slova i vraća sve moguće permutacije slova u toj sekvenci. Na primer, moguće permutacije sekvence slova `cat` su `cat`, `cta`, `act`, `atc`, `tac`, i `tca`.

Rekurzivna metoda za kreiranje permutacija sekvence slova radi tako što čuva početni deo sekvence (prefiks) i kreira permutacije preostalog dela.

Na primer, uzmimo reč `machine`. Prvi ciklus kreiranja permutacija biće postavljanje `m` kao prefiksa sekvence, a zatim kreiranje permutacija preostalog dela sekvence, `achine`.

Za preostali deo sekvence, permutacije nastavljaju na isti način. Dodaje se slovo na prefiks, možda `c`, pa prefiks postaje `mc`. Zatim se svaka permutacija `ahine` spaja sa prefiksom.

Ovo se nastavlja sve dok prefiks ne sadrži sva slova, a preostali deo sekvence bude prazan, što znači da je formirana jedna permutacija.

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Treba da kreiraš funkciju nazvanu `permuteString`.
2. Funkcija `permuteString` treba da primi jedan parametar, niz karaktera, a zatim dva parametra sa podrazumevanim vrednostima: vrednost prefiksa i prazan niz podataka za skladištenje i vraćanje rezultata. Vrednost prefiksa bi trebalo da bude niz karaktera koji se koristi za prikupljanje karaktera radi formiranja permutacije. Funkcija će biti pozvana samo sa jednim argumentom, kao što je `permuteString("cat")`.
3. Unutar funkcije, treba da proveriš da li je dužina poslanog niza `0`. Ako jeste, dodaj trenutni prefiks na rezultate i vrati rezultate.
4. Iteriraj kroz svaki karakter u ulaznom nizu i za svaku iteraciju, ukloni trenutni karakter iz niza i pozovi funkciju `permuteString` rekursivno sa ažuriranim argumentima kako bi se izgradile preostale permutacije.
5. Treba da vratiš niz podataka sa konačnim rezultatima.
6. Treba da osiguraš da su permutacije jedinstvene uklanjanjem duplikata.

# --hints--

Trebalo bi da imaš funkciju `permuteString`.

```js
assert.isFunction(permuteString);
```

Trebalo bi da koristiš rekurziju u svojoj funkciji `permuteString`.

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` treba da vrati `[ "far", "fra", "afr", "arf", "rfa", "raf" ]`.

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` trebalo bi da vrati `[ "fcc", "cfc", "ccf" ]`.

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` treba da vrati `[ "p" ]`.

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` treba da vrati `[""]`.

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);

```

`permuteString("walk")` trebalo bi da vrati `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]`.

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString` Trebalo bi da vrati tačne rezultate.

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```
