---
id: 587d824a367417b2b2512c45
title: Telo nepoznate poruke
challengeType: 4
forumTopicId: 301568
dashedName: anonymous-message-board
---

# --description--

Napravi kompletan program za JavaScript koji funkcionalno sliči ovom: <a href="https://anonymous-message-board.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://anonymous-message-board.freecodecamp.rocks/</a>.

Rad na ovom projektu će zahtevati pisanje vašeg koda koristeći jednu od sledećih metoda:

-   Kopiraj <a href="https://github.com/freeCodeCamp/boilerplate-project-messageboard/" target="_blank" rel="noopener noreferrer nofollow">hii repo od GitHub</a> i završi svoj projekat lokalno.
-   Koristi alat za izradu veb-stranice koji ti se sviđa da završiš projekat. Uveri se da uključuješ sve fajlove iz našeg repozitorijuma GitHub.

# --instructions--

1. Postavite `NODE_ENV` za testiranje bez navodnika kada budete spremni da pišete testove sa bazom podataka za sekvencu vašeg konektora baze podataka (u `.env`)
2. Savetuje se kreiranje kontrolera/handler-a i rukovanje rutama za `routes/api.js`
3. Dodati ćete bilo koje sigurnosne funkcije na `server.js`

Napišite sledeće testove u `tests/2_functional-tests.js`:

-   Kreiranje nove teme: POST zahtev za `/api/threads/{board}`
-   Pregled 10 najnovijih tema i po 3 odgovora: GET zahtev za `/api/threads/{board}`
-   Brisanje teme sa neispravnom lozinkom: DELETE zahtev za `/api/threads/{board}` i `delete_password` (nevažeće)
-   Brisanje teme sa ispravnom lozinkom: DELETE zahtev za `/api/threads/{board}` i `delete_password` (važeće)
-   Prijavljivanje teme: PUT zahtev za `/api/threads/{board}`
-   Kreiranje novog odgovora: POST zahtev za `/api/replies/{board}`
-   Pregled jedne teme i svih odgovora: GET zahtev za `/api/replies/{board}`
-   Brisanje odgovora sa neispravnom lozinkom: DELETE zahtev za `/api/replies/{board}` i `delete_password` (nevažeće)
-   Brisanje odgovora sa ispravnom lozinkom: DELETE zahtev za `/api/replies/{board}` i `delete_password` (važeće)
-   Prijavljivanje odgovora: PUT zahtev za `/api/replies/{board}`

# --hints--

Možeš da koristi svoj projekat, ne mora biti primer URL.

```js
  assert(
    !/.*\/anonymous-message-board\.freecodecamp\.rocks/.test(
      code
    )
  );
```

Dopustite da je vaš sajt ugrađen samo u iFrame na vašim sopstvenim stranicama.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-frame-options']?.includes('SAMEORIGIN'));
```

Omogućite DNS prefetčing.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-dns-prefetch-control']?.includes('off'));
```

Dozvolite da vaš sajt šalje referer samo na vaše sopstvene stranice.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['referrer-policy']?.includes('same-origin'));
```

Možete poslati POST zahtev na `/api/threads/{board}` sa podacima forme koji uključuju `text` i `delete_password`. Rekord sačuvan u bazi podataka će imati najmanje polja za `_id`, `text`, `created_on` (datum i vreme), `bumped_on` (datum i vreme, počinje isto kao `created_on`), `reported` (vrednost tačno ili netačno), `delete_password`, i `replies` (red podataka).

```js
  const date = new Date();
  const text = `fcc_test_${date}`;
  const deletePassword = 'delete_me';
  const data = { text, delete_password: deletePassword };
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(url + '/api/threads/fcc_test');
    const parsed = await checkData.json();
    try {
      assert.equal(parsed[0].text, text);
      assert.isNotNull(parsed[0]._id);
      assert.equal(new Date(parsed[0].created_on).toDateString(), date.toDateString());
      assert.equal(parsed[0].bumped_on, parsed[0].created_on);
      assert.isArray(parsed[0].replies);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Možete poslati POST zahtev na `/api/replies/{board}` sa podacima forme koji uključuju `text`, `delete_password` i `thread_id`. Ovo će ažurirati datum `bumped_on` u datum komentara. U nizu niti `replies`, element će biti sačuvan koji ima barem svojstva `_id`, `text`, `created_on`, `delete_password` i `reported`.

```js
  const url = code;
  const body = await fetch(url + '/api/threads/fcc_test');
  const thread = await body.json();

  const date = new Date();
  const text = `fcc_test_reply_${date}`;
  const delete_password = 'delete_me';
  const thread_id = thread[0]._id;
  const replyCount = thread[0].replies.length;

  const data = { text, delete_password, thread_id };
  const res = await fetch(url + '/api/replies/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
    const parsed = await checkData.json();
    try {
      assert.equal(parsed.replies.length, replyCount + 1);
      assert.equal(parsed.replies[0].text, text);
      assert.equal(parsed._id, thread_id);
      assert.equal(parsed.bumped_on, parsed.replies[0].created_on);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Možete poslati GET zahtev na `/api/threads/{board}`. To će vratiti listu od 10 najnovijih niti objavljenih na tabli i 3 najnovija odgovora za svaku od njih. Sekcije `reported` i `delete_password` neće biti poslati klijentu.

```js
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test');

  if (res.ok) {
    const threads = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isAtMost(threads.length, 10);
      for (let i = 0; i < threads.length; i++) {
        assert.containsAllKeys(threads[i], ["_id", "text", "created_on", "bumped_on", "replies"]);
        assert.isAtMost(threads[i].replies.length, 3);
        assert.notExists(threads[i].delete_password);
        assert.notExists(threads[i].reported);
        for (let j = 0; j < threads[i].replies.length; j++) {
          assert.notExists(threads[i].replies[j].delete_password);
          assert.notExists(threads[i].replies[j].reported);
        }
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Možete poslati GET zahtev na `/api/replies/{board}?thread_id={thread_id}`. To će vratiti celu niti i sve odgovore, takođe uklanjajući iste delove sa klijenta kao deo prethodnog pokušaja.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);

  if (res.ok) {
    const thread = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isObject(thread);
      assert.containsAllKeys(thread, ["_id", "text", "created_on", "bumped_on", "replies"]);
      assert.isArray(thread.replies);
      assert.notExists(thread.delete_password);
      assert.notExists(thread.reported);
      for (let i = 0; i < thread.replies.length; i++) {
        assert.notExists(thread.replies[i].delete_password);
        assert.notExists(thread.replies[i].reported);
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Možete poslati DELETE zahtev na `/api/threads/{board}` i proslediti `thread_id` i `delete_password` da obrišete niti. Vraćaće sekvencu karaktera `incorrect password` ili `success`.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  let data = { thread_id, delete_password: "wrong_password" };
  const res_invalid = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  data = { thread_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    const not_deleted = await res_invalid.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      assert.equal(not_deleted, "incorrect password");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Možete poslati DELETE zahtev na `/api/replies/{board}` i proslediti `thread_id`, `reply_id`, i `delete_password`. Vraćaće niz karaktera `incorrect password` ili `success`. Uspešno, tekst `reply_id` će biti promenjen u `[deleted]`.

```js
  const url = code;

  const thread_data = {
    text: "fcc_test_thread",
    delete_password: "delete_me",
  };
  await fetch(`${url}/api/threads/fcc_test`, {
    method: "POST",
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(thread_data)
  });
  let res = await fetch(`${url}/api/threads/fcc_test`);
  let threads = await res.json();
  const thread_id = threads[0]._id;
  
  const reply_data = { thread_id, text: "fcc_test_reply", delete_password: "delete_me" };
  await fetch(`${url}/api/replies/fcc_test`, {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(reply_data)
  });
  res = await fetch(`${url}/api/threads/fcc_test`);
  threads = await res.json();
  const reply_id = threads[0].replies[0]._id;

  const data = { thread_id, reply_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/replies/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
      const thread = await res.json();
      assert.equal(thread._id, thread_id);
      assert.equal(thread.replies[0]._id, reply_id);
      assert.equal(thread.replies[0].text, "[deleted]");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Možete poslati PUT zahtev na `/api/threads/{board}` i proslediti `thread_id`. Vraća niz znakova `reported`. Vrednost `reported` od `thread_id` će biti promenjena u `true`.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const data = { thread_id };
  
  res = await fetch(`${url}/api/threads/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Možete poslati PUT zahtev za `/api/replies/{board}` i proslediti `thread_id` i `reply_id`. Vraćaće niz znakova `reported`. Vrednost `reported` od `reply_id` će biti promenjena u `true`.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const reply_id = threads[0].replies[0]._id;
  const data = { thread_id, reply_id };
  
  res = await fetch(`${url}/api/replies/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Sva deset testova performansi je završena i prošla.

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 10);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
    assert.isAtLeast(test.assertions.length, 1);
  });
```
