---
id: 685239f0cad43c7a909cbd98
title: Ispravite web stranicu za promociju ljubimaca
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Sally, vlasnica prodavnice za ljubimce, kreirala je svoju prvu web stranicu, ali postoji više problema.

Tvoj zadatak je da ispraviš sve greške kako bi Sali nastavila sa izgradnjom svoje stranice.

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Sally želi da koristi slike mačaka, ali ne izgledaju dobro. Moraš ispraviti sledeće u elementu `img`:
   - Promeni atribut elementa `href` u odgovarajući atribut izvora slike.
   - Promeni atribut elementa `att` u odgovarajući atribut koji predstavlja kratki opis teksta za sliku.
   - Ukloni zatvarajući tag `</img>` jer su elementi `img` elementi bez zatvarajućeg taga.
2. Sally želi da koristi linkove kako bi usmerila korisnike na stranice o psima i mačkama. Ali linkovi ne rade dobro. Moraš ispraviti sledeće u elementima `a`:
   - Promeni oba atributa elementa `src` u odgovarajuće atribute koji se koriste za definisanje URL.

# --hints--

Element vašeg `img` treba da ima svojstvo elemenata `src` umesto svojstva elemenata `href`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

Tvoj element za `img` treba da ima svojstvo elementa `alt` umesto postojećeg svojstva elementa `att`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

Vaš element za `img` ne sme imati zatvorni atribut/tag `</img>`.

```js
assert.notMatch(code, /<\/img>/);
```

Element vašeg `a` sa tekstom `Visit cats page` mora imati atribut elementa `href` umesto atributa elementa `src`.

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

Vaš element za `a` sa tekstom `Visit dogs page` mora imati atribut elementa `href` umesto atributa elementa `src`.

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
