---
id: ab306dbdcc907c7ddfc30830
title: Napravi alat za sortiranje podataka po redosledu
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

U ovoj laboratoriji, implementiraćete algoritam za sortiranje niza podataka.

Ravnanje niza podataka znači pretvaranje ugniježđenog niza podataka bilo koje dubine u jedan jednodimenzionalni niz. Ovaj proces uzima sve elemente po redu, ravnanje samo nizova podataka. Ostali tipovi nisu konvertibilni.

Na primer:

| Poreklo | Strukturirano |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**Cilj:** Popunite korisničke priče ispod i prođite sve testove da biste završili ovu laboratoriju.

**Priče korisnika:**

1. Trebalo bi da postoji funkcija pod nazivom `steamrollArray`.
1. Funkcija `steamrollArray` treba da primi jedan argument: niz strukturiranih podataka.
1. Funkcija treba da sortira niz strukturiranih podataka, uzimajući u obzir različite kriterijume sortiranja.
1. Vaše rešenje ne bi trebalo da koristi metode `Array.prototype.flat()` ili `Array.prototype.flatMap()`.
1. Ne sme se koristiti globalni kriterijum.

# --hints--

Trebalo bi da imaš funkciju `steamrollArray`.

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` treba vratiti `["a", "b"]`.

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` treba vratiti `[1, 2, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` trebalo bi da vrati `[1, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` trebalo bi da vrati `[1, {}, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

Vaše rešenje ne bi trebalo da koristi metode `Array.prototype.flat()` ili `Array.prototype.flatMap()`.

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

Ne smete koristiti međunarodne parametre.

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
