---
id: 587d7da9367417b2b2512b66
title: Povežite dve linije podataka pomoću metode concat
challengeType: 1
forumTopicId: 301229
dashedName: combine-two-arrays-using-the-concat-method
---

# --description--

<dfn>Concatenation</dfn> to znači spajanje stvari kraj na kraj. JavaScript pruža način da se `concat` za nizove karaktera i nizove podataka koji rade na isti način. Za nizove podataka, ova metoda se poziva na jedan, a zatim drugi niz se prosleđuje kao argument za `concat`, koji se dodaje na kraj prvog niza. Vraća novi niz i ne izaziva nikakve promene u prvobitnim nizovima. Evo primera:

```js
[1, 2, 3].concat([4, 5, 6]);
```

Vraćeni red će biti `[1, 2, 3, 4, 5, 6]`.

# --instructions--

Koristi putanju `concat` u funkciji `nonMutatingConcat` da spoji `attach` na kraj `original`. Funkcija treba vratiti spojeni red.

# --hints--

Tvoj kod bi trebalo da koristi metodu `concat`.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

Safu od `first` ne sme da se menja.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

Reda `second` ne bi trebalo da se menja.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingConcat([1, 2, 3], [4, 5])` treba da vrati `[1, 2, 3, 4, 5]`.

```js
assert(
  JSON.stringify(nonMutatingConcat([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingConcat(original, attach) {
  // Only change code below this line


  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingConcat(first, second);
```

# --solutions--

```js
function nonMutatingConcat(original, attach) {
  return original.concat(attach);
}
const first = [1, 2, 3];
const second = [4, 5];
```
