---
id: 587d7b8f367417b2b2512b62
title: Implementiranje mape na prototipskom modelu
challengeType: 1
forumTopicId: 301230
dashedName: implement-map-on-a-prototype
---

# --description--

Kao što ste videli korišćenjem `Array.prototype.map()`, ili lako ranije sa `map()`, metoda `map` vraća red podataka iste dužine kao pozvani. Takođe ne menja originalni red podataka, pod uslovom da njegova funkcija povratne pozive to ne promeni.

Drugim rečima, `map` je čista funkcija, a njen rezultat zavisi isključivo od njenog ulaza. Štaviše, uzima drugu funkciju kao svoj argument.

Možete naučiti mnogo o načinu `map` implementirajući sopstvenu verziju te metode. Preporučuje se korišćenje ciklusa `for` ili `Array.prototype.forEach()`.

# --instructions--

Napišite sopstveni `Array.prototype.myMap()`, koji bi trebalo da radi potpuno isto kao i `Array.prototype.map()`. Ne smete koristiti ugrađeni način `map`. Primer za `Array` može se pristupiti putem `myMap` koristeći `this`.

# --hints--

`[23, 65, 98, 5, 13].myMap(item => item * 2)` treba da bude jednak `[46, 130, 196, 10, 26]`.

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item * 2;
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`["naomi", "quincy", "camperbot"].myMap(element => element.toUpperCase())` treba da vrati `["NAOMI", "QUINCY", "CAMPERBOT"]`.

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element.toUpperCase();
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

`[1, 1, 2, 5, 2].myMap((element, index, array) => array[index + 1] || array[0])` treba da vrati `[1, 2, 5, 2, 1]`.

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array[index + 1] || array[0];
assert(JSON.stringify(_test_s.map(_callback)) === JSON.stringify(_test_s.myMap(_callback)));
```

Tvoj kod ne bi trebalo da koristi putanju `map`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?map/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myMap = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    newArray.push(callback(this[i], i, this));
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const doubled_s = s.myMap(item => item * 2);
```
