---
id: 587d7da9367417b2b2512b6a
title: Vratite obrađeni podudni red bez menjanja izvornog reda podataka
challengeType: 1
forumTopicId: 301237
dashedName: return-a-sorted-array-without-changing-the-original-array
---

# --description--

Jedan efekat metode `sort` je što menja redos elemenata u originalnom nizu podataka. Drugim rečima, menja niz podataka na mestu. Jedan način da se to izbegne je prvo konkatere prazan niz podataka sa onim koji se sortira (zapamtite da `slice` i `concat` vraćaju novi niz podataka), a zatim koristite metodu `sort`.

# --instructions--

Koristi putanju `sort` u funkciji `nonMutatingSort` za sortiranje elemenata niza podataka po rastućem redosledu. Funkcija treba da vrati novi niz podataka, a ne da promeni vrednost `globalArray`.

# --hints--

Vaš kod treba da koristi metodu `sort`.

```js
assert(nonMutatingSort.toString().match(/\.sort/g));
```

Vrednost `globalArray` ne bi trebalo da se menja.

```js
assert(JSON.stringify(globalArray) === JSON.stringify([5, 6, 3, 2, 9]));
```

`nonMutatingSort(globalArray)` treba da vrati `[2, 3, 5, 6, 9]`.

```js
assert(
  JSON.stringify(nonMutatingSort(globalArray)) ===
    JSON.stringify([2, 3, 5, 6, 9])
);
```

`nonMutatingSort(globalArray)` Ne sme biti ugrađeni kod.

```js
assert(!nonMutatingSort.toString().match(/\[.*?[23569].*?\]/gs));
```

Funkcija treba da vrati novi red podataka, a ne prosleđeni red podataka.

```js
assert(nonMutatingSort(globalArray) !== globalArray);
```

`nonMutatingSort([1, 30, 4, 21, 100000])` trebalo bi da vrati `[1, 4, 21, 30, 100000]`.

```js
assert(JSON.stringify(nonMutatingSort([1, 30, 4, 21, 100000])) ===
    JSON.stringify([1, 4, 21, 30, 100000]))
```

`nonMutatingSort([140000, 104, 99])` treba da vrati `[99, 104, 140000]`.

```js
assert(JSON.stringify(nonMutatingSort([140000, 104, 99])) ===
    JSON.stringify([99, 104, 140000]))
```

# --seed--

## --seed-contents--

```js
const globalArray = [5, 6, 3, 2, 9];

function nonMutatingSort(arr) {
  // Only change code below this line


  // Only change code above this line
}

nonMutatingSort(globalArray);
```

# --solutions--

```js
const globalArray = [5, 6, 3, 2, 9];
function nonMutatingSort(arr) {
  return [].concat(arr).sort((a,b) => a-b);
}
```
