---
id: 587d7b8e367417b2b2512b5d
title: Razumite rizike korišćenja komandnog koda
challengeType: 1
forumTopicId: 301241
dashedName: understand-the-hazards-of-using-imperative-code
---

# --description--

Funkcionalnost je dobra karakteristika. Čini vaš kod lakšim za održavanje i štiti vas od sitnih grešaka. Ali pre nego što dođemo tamo, hajde da pregledamo stil komandi u programu kako bismo pokazali gde možete naići na probleme.

Engleski (i na mnogim drugim jezicima), imperativ se koristi za davanje naredbi. Slično tome, imperativni način u programiranju je onaj koji računaru daje set instrukcija za izvršavanje zadatka.

Često ovi izvještaji mijenjaju stanje programa, kao što je ispravljanje globalnih parametara. Uobijedan primjer je pisanje petlje `for` koja daje potpune uputstva za iteriranje preko indeksa niza podataka.

Naspravno toga, proceduralni program je vrsta deklarativnog programa. Kažete računaru šta želite da uradi pozivanjem funkcije ili procedure.

JavaScript pruža mnogo već deklarisanih metoda koje obrađuju uobičajene zadatke, tako da ne morate da pišete kako računar treba da ih izvrši. Na primer, umesto korišćenja petlje ``for`` opisane gore, možete pozvati metodu ``map`` koja obrađuje ponavljajuće detalje preko niza podataka. Ovo pomaže u izbegavanju logičkih grešaka, kao što je "Off By One Greška", diskutovana u sekciji za rešavanje problema.

Zamislite situaciju: imate više tabova u vašem pretraživaču i želite da pratite otvorene kartice. Pokušaćemo da kreiramo ovaj primer koristeći jednostavan kod objekta/klase.

Prozor se sastoji od kartica, i obično imate više otvorenih prozora. Naslovi svake stranice otvorene u svakom prozoru čuvaju se u nizu podataka. Nakon rada u pretraživaču (otvaranje nove kartice, povezivanje prozora i zatvaranje kartica), želite da prikažete otvorene kartice. Zatvorene kartice se uklanjaju iz niza podataka, a nove kartice (lako) dodaju se na kraj tog niza.

Kodni editor prikazuje izvršavanje ove funkcije koristeći funkcije `tabOpen()`, `tabClose()`, i `join()`. Red podataka `tabs` je deo elementa prozora koji čuva imena otvorenih stranica.

# --instructions--

Proverite kod u editoru. Koristite metod koji ima nuspojave u aplikaciji, što uzrokuje neispravno ponašanje. Poslednji spisak otvorenih kartica, sačuvan u `finalTabs.tabs`, trebalo bi da bude `['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']`, ali je spisak koji pruža kod malo drugačiji.

Promenite `Window.prototype.tabClose` ili uklonite pogrešni tab.

# --hints--

`finalTabs.tabs` treba biti `['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium', 'new tab', 'Netflix', 'YouTube', 'Vine', 'GMail', 'Work mail', 'Docs', 'freeCodeCamp', 'new tab']`

```js
assert.deepEqual(finalTabs.tabs, [
  'FB',
  'Gitter',
  'Reddit',
  'Twitter',
  'Medium',
  'new tab',
  'Netflix',
  'YouTube',
  'Vine',
  'GMail',
  'Work mail',
  'Docs',
  'freeCodeCamp',
  'new tab'
]);
```

# --seed--

## --seed-contents--

```js
// tabs is an array of titles of each site open within the window
const Window = function(tabs) {
  this.tabs = tabs; // We keep a record of the array inside the object
};

// When you join two windows into one window
Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

// When you open a new tab at the end
Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab'); // Let's open a new tab for now
  return this;
};

// When you close a tab
Window.prototype.tabClose = function(index) {

  // Only change code below this line

  const tabsBeforeIndex = this.tabs.splice(0, index); // Get the tabs before the tab
  const tabsAfterIndex = this.tabs.splice(index + 1); // Get the tabs after the tab

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex); // Join them together

  // Only change code above this line

  return this;
 };

// Let's create three browser windows
const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']); // Your mailbox, drive, and other work sites
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']); // Social sites
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']); // Entertainment sites

// Now perform the tab opening, closing, and other operations
const finalTabs = socialWindow
  .tabOpen() // Open a new tab for cat memes
  .join(videoWindow.tabClose(2)) // Close third tab in video window, and join
  .join(workWindow.tabClose(1).tabOpen());
console.log(finalTabs.tabs);
```

# --solutions--

```js
const Window = function(tabs) {
  this.tabs = tabs;
};

Window.prototype.join = function(otherWindow) {
  this.tabs = this.tabs.concat(otherWindow.tabs);
  return this;
};

Window.prototype.tabOpen = function(tab) {
  this.tabs.push('new tab');
  return this;
};

Window.prototype.tabClose = function(index) {
  const tabsBeforeIndex = this.tabs.slice(0, index);
  const tabsAfterIndex = this.tabs.slice(index + 1);

  this.tabs = tabsBeforeIndex.concat(tabsAfterIndex);
  return this;
 };

const workWindow = new Window(['GMail', 'Inbox', 'Work mail', 'Docs', 'freeCodeCamp']);
const socialWindow = new Window(['FB', 'Gitter', 'Reddit', 'Twitter', 'Medium']);
const videoWindow = new Window(['Netflix', 'YouTube', 'Vimeo', 'Vine']);

const finalTabs = socialWindow
  .tabOpen()
  .join(videoWindow.tabClose(2))
  .join(workWindow.tabClose(1).tabOpen());
```
