---
id: 587d7b8f367417b2b2512b64
title: Implementiraj filtracioni metod na primeru
challengeType: 1
forumTopicId: 301231
dashedName: implement-the-filter-method-on-a-prototype
---

# --description--

Možete naučiti mnogo o metodi `filter` implementiranjem sopstvene verzije te metode. Preporučuje se da koristite ciklus `for` ili `Array.prototype.forEach()`.

# --instructions--

Napišite sopstveni `Array.prototype.myFilter()`, koji bi trebalo da radi potpuno isto kao `Array.prototype.filter()`. Ne smete koristiti ugrađeni način `filter`. Primer za `Array` može se pristupiti na putanji `myFilter` koristeći `this`.

# --hints--

`[23, 65, 98, 5, 13].myFilter(item => item % 2)` mora biti jednako `[23, 65, 5, 13]`.

```js
const _test_s = [23, 65, 98, 5, 13];
const _callback = item => item % 2;
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

`["naomi", "quincy", "camperbot"].myFilter(element => element === "naomi")` treba da vrati `["naomi"]`.

```js
const _test_s = ["naomi", "quincy", "camperbot"];
const _callback = element => element === "naomi";
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

`[1, 1, 2, 5, 2].myFilter((element, index, array) => array.indexOf(element) === index)` treba da vrati `[1, 2, 5]`.

```js
const _test_s = [1, 1, 2, 5, 2];
const _callback = (element, index, array) => array.indexOf(element) === index;
assert(JSON.stringify(_test_s.filter(_callback)) === JSON.stringify(_test_s.myFilter(_callback)));
```

Vaš kod ne bi trebalo da koristi metodu `filter`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?filter/g));
```

# --seed--

## --seed-contents--

```js
Array.prototype.myFilter = function(callback) {
  const newArray = [];
  // Only change code below this line

  // Only change code above this line
  return newArray;
};
```

# --solutions--

```js
Array.prototype.myFilter = function(callback) {
  const newArray = [];
  for (let i = 0; i < this.length; i++) {
    if (callback(this[i], i, this)) newArray.push(this[i]);
  }
  return newArray;
};

// Test case
const s = [23, 65, 98, 5];
const odd_s = s.myFilter(item => item % 2 === 1);
```
