---
id: 587d7b8f367417b2b2512b60
title: Promenite globalnu parametru iz funkcija
challengeType: 1
forumTopicId: 301235
dashedName: refactor-global-variables-out-of-functions
---

# --description--

Sada smo videli dva različita principa performativnog programiranja:

1) Ne menjajte parametar ili vrednost – kreirajte nove parametre i vrednosti i referencirajte ih ako je potrebno iz funkcije. Napomena: korišćenjem nečega kao `const newArr = arrVar`, gde je `arrVar` red podataka samo za kreiranje reference na postojeći parametar, a ne kopija. Stoga, promena vrednosti u `newArr` će promeniti vrednost u `arrVar`.

2) Deklariši parametre funkcije – bilo kakva izračunavanja unutar funkcije zavise samo od argumenata prosleđenih funkciji, a ne od globalnih varijabli ili parametara.

Dodavanje jedinice po broju nije nešto posebno uzbudljivo, ali možemo koristiti ove principe kada radimo sa nizovima podataka ili nečim složenijim.

# --instructions--

Prepišite kod tako da globalni niz podataka `bookList` ne bude modifikovan unutar bilo koje funkcije. Funkcija `add` treba da doda/dodaje `bookName` koji je dodan na kraj prosluđenog niza podataka i da vrati novi niz (listu). Funkcija `remove` treba da ukloni `bookName` koji je izvučen iz prosluđenog niza podataka.

**Napomena:** Sve funkcije moraju se odnositi na red podataka, a novi parametri moraju biti dodati pre parametra `bookName`.

# --hints--

`bookList` ne bi trebalo da se menja, a da istovremeno bude jednako `["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"]`.

```js
add(bookList, "Test");
remove(bookList, "The Hound of the Baskervilles");
assert(
  JSON.stringify(bookList) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'On The Electrodynamics of Moving Bodies',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae'
    ])
);
```

`add(bookList, "A Brief History of Time")` treba da se odnosi na `["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae", "A Brief History of Time"]`.

```js
assert(
  JSON.stringify(add(bookList, "A Brief History of Time")) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'On The Electrodynamics of Moving Bodies',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae',
      'A Brief History of Time'
    ])
);
```

`remove(bookList, "On The Electrodynamics of Moving Bodies")` treba da se odnosi na `["The Hound of the Baskervilles", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"]`.

```js
assert(
  JSON.stringify(remove(bookList, 'On The Electrodynamics of Moving Bodies')) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae'
    ])
);
```

`remove(add(bookList, "A Brief History of Time"), "On The Electrodynamics of Moving Bodies");` trebalo bi da bude jednako `["The Hound of the Baskervilles", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae", "A Brief History of Time"]`.

```js
assert(
  JSON.stringify(remove(add(bookList, 'A Brief History of Time'), 'On The Electrodynamics of Moving Bodies')) ===
    JSON.stringify([
      'The Hound of the Baskervilles',
      'Philosophiæ Naturalis Principia Mathematica',
      'Disquisitiones Arithmeticae',
      'A Brief History of Time'
    ])
);
```

# --seed--

## --seed-contents--

```js
// The global variable
const bookList = ["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"];

// Change code below this line
function add(bookName) {

  bookList.push(bookName);
  return bookList;
  
  // Change code above this line
}

// Change code below this line
function remove(bookName) {
  const book_index = bookList.indexOf(bookName);
  if (book_index >= 0) {

    bookList.splice(book_index, 1);
    return bookList;

    // Change code above this line
    }
}
```

# --solutions--

```js
// The global variable
const bookList = ["The Hound of the Baskervilles", "On The Electrodynamics of Moving Bodies", "Philosophiæ Naturalis Principia Mathematica", "Disquisitiones Arithmeticae"];

function add(bookList, bookName) {
  return [...bookList, bookName];
}

function remove(bookList, bookName) {
  const bookListCopy = [...bookList];
  const bookNameIndex = bookList.indexOf(bookName);
  if (bookNameIndex >= 0) {
    bookListCopy.splice(bookNameIndex, 1);
  }
  return bookListCopy;
}
```
