---
id: 587d7dab367417b2b2512b6f
title: Koristi putanju `some` da proveriš da li svi elementi u nizu podataka zadovoljavaju kriterijume.
challengeType: 1
forumTopicId: 301314
dashedName: use-the-some-method-to-check-that-any-elements-in-an-array-meet-a-criteria
---

# --description--

Funkcija `some` radi sa nizovima podataka da proveri da li bilo koji element prolazi kroz određeni test. Vraća vrednost tačno ili netačno - `true` ako bilo koja vrednost ispunjava kriterijume, `false` inače.

Na primer, sledeći kod proverava da li je bilo koji element u redovima podataka ``numbers`` manji od 10:

```js
const numbers = [10, 50, 8, 220, 110, 11];

numbers.some(function(currentValue) {
  return currentValue < 10;
});
```

Putanja za `some` koristiće `true`.

# --instructions--

Koristi putanju `some` unutar funkcije `checkPositive` da proveri da li je bilo koji element u `arr` pozitivan. Funkcija bi trebalo da vrati vrednost tačno ili netačno.

# --hints--

Vaš kod bi trebalo da koristi metod `some`.

```js
assert(__helpers.removeJSComments(code).match(/\.some/g));
```

`checkPositive([1, 2, 3, -4, 5])` bi trebalo da vrati `true`.

```js
assert(checkPositive([1, 2, 3, -4, 5]));
```

`checkPositive([1, 2, 3, 4, 5])` treba da vrati `true`.

```js
assert(checkPositive([1, 2, 3, 4, 5]));
```

`checkPositive([-1, -2, -3, -4, -5])` bi trebalo da vrati `false`.

```js
assert(!checkPositive([-1, -2, -3, -4, -5]));
```

# --seed--

## --seed-contents--

```js
function checkPositive(arr) {
  // Only change code below this line


  // Only change code above this line
}

checkPositive([1, 2, 3, -4, 5]);
```

# --solutions--

```js
function checkPositive(arr) {
  return arr.some(elem => elem > 0);
}
```
