---
id: 587d7dab367417b2b2512b70
title: Uvod u Kuringa i Korak-po-korak upotrebu
challengeType: 1
forumTopicId: 301232
dashedName: introduction-to-currying-and-partial-application
---

# --description--

<dfn>Arity</dfn> funkcije je broj potrebnih argumenata. <dfn>Kuriranje</dfn> funkcije znači pretvaranje funkcije sa aritetom N u N funkcija sa aritetom 1.

Drugim rečima, reorganizuje funkciju da uzme jedan argument, pa vraća drugu funkciju koja uzima sledeći argument, i tako dalje.

Evo primera:

```js
function unCurried(x, y) {
  return x + y;
}

function curried(x) {
  return function(y) {
    return x + y;
  }
}

const curried = x => y => x + y

curried(1)(2)
```

`curried(1)(2)` koristi `3`.

Ovo je važno u vašem programu ako ne možete pružiti sve argumente funkciji odjednom. Možete sačuvati svaki poziv funkcije u varijablu koja će držati referencu na vraćenu funkciju koja uzima sledeći argument kada bude dostupan. Evo primera korišćenjem funkcije umotane u primer iznad:

```js
const funcForY = curried(1);
console.log(funcForY(2)); // 3
```

Slično, <dfn>upotreba delova</dfn> može se opisati kao korišćenje malog broja argumenata za funkciju istovremeno i vraćanje druge funkcije koja se koristi sa više argumenata. Evo primera:

```js
function impartial(x, y, z) {
  return x + y + z;
}

const partialFn = impartial.bind(this, 1, 2);
partialFn(10); // 13
```

# --instructions--

Popunite telo funkcije za `add` ili ga koristite da uključite parametre `x`, `y`, i `z`.

# --hints--

`add(10)(20)(30)` trebao bi vratiti `60`.

```js
assert(add(10)(20)(30) === 60);
```

`add(1)(2)(3)` trebalo bi da vrati `6`.

```js
assert(add(1)(2)(3) === 6);
```

`add(11)(22)(33)` bi trebalo da vrati `66`.

```js
assert(add(11)(22)(33) === 66);
```

Tvoj kod treba da sadrži završnu deklaraciju koja vraća `x + y + z`.

```js
assert(__helpers.removeJSComments(code).match(/[xyz]\s*?\+\s*?[xyz]\s*?\+\s*?[xyz]/g));
```

# --seed--

## --seed-contents--

```js
function add(x) {
  // Only change code below this line


  // Only change code above this line
}

add(10)(20)(30);
```

# --solutions--

```js
const add = x => y => z => x + y + z
```
