---
id: 587d7b8d367417b2b2512b5b
title: Naučite o funkcionalnom programiranju
challengeType: 1
forumTopicId: 301233
dashedName: learn-about-functional-programming
---

# --description--

Funkcionalno programiranje je stil programiranja gde su rešenja jednostavne, slobodne funkcije, bez sporednih efekata van opsega funkcije: `ULAZ -> PROCES -> IZLAZ`

Funkcionalno programiranje je o:

1) Slobodnim funkcijama - nema zavisnosti od stanja programa, što uključuje promenljive globalne vrednosti koje se mogu menjati

2) Čistim funkcijama - isti ulaz uvek daje isti izlaz

3) Funkcijama sa malo sporednih efekata - bilo kakve promene, ili modifikacije, na stanju programa van funkcije se pažljivo kontrolišu

# --instructions--

Članovi freeCodeCamp vole čaj.

U editoru koda, funkcije `prepareTea` i `getTea` su već definisane za vas. Pozovite funkciju `getTea` da dobijete 40 šoljica čaja za tim, i sačuvajte ih u promenljivu `tea4TeamFCC`.

# --hints--

Promenljiva `tea4TeamFCC` treba da sačuva 40 šoljica čaja za tim.

```js
assert(tea4TeamFCC.length === 40);
```

Promenljiva `tea4TeamFCC` treba da sačuva šoljice zelenog čaja.

```js
assert(tea4TeamFCC[0] === 'greenTea');
```

# --seed--

## --seed-contents--

```js
// Function that returns a string representing a cup of green tea
const prepareTea = () => 'greenTea';

/*
Given a function (representing the tea type) and number of cups needed, the
following function returns an array of strings (each representing a cup of
a specific type of tea).
*/
const getTea = (numOfCups) => {
  const teaCups = [];

  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }
  return teaCups;
};

// Only change code below this line
const tea4TeamFCC = null;
// Only change code above this line
```

# --solutions--

```js
const prepareTea = () => 'greenTea';

const getTea = (numOfCups) => {
  const teaCups = [];
  
  for(let cups = 1; cups <= numOfCups; cups += 1) {
    const teaCup = prepareTea();
    teaCups.push(teaCup);
  }

  return teaCups;
};

const tea4TeamFCC = getTea(40);