---
id: 587d7da9367417b2b2512b67
title: Dodaj elemente na kraj niza podataka koristeći `concat` umesto `push`
challengeType: 1
forumTopicId: 301226
dashedName: add-elements-to-the-end-of-an-array-using-concat-instead-of-push
---

# --description--

Funkcionalno programiranje je o kreiranju i korišćenju funkcija koje ne menjaju podatke.

Najnovija funkcionalnost je predstavila način `concat` kao način spajanja redova podataka u novi red bez menjanja originalnih redova podataka. Uporedite `concat` sa metodom `push`. `push` dodaje stavke na kraju istog redast koji menja taj red. Evo primera:

```js
const arr = [1, 2, 3];
arr.push(4, 5, 6);
```

`arr` ima modifikovanu vrednost od `[1, 2, 3, 4, 5, 6]`, što nije metoda izvršavanja programa.

`concat` Omogućava način za dodavanje novih elemenata na kraj niza podataka bez bočnih efekata koji menjaju (strukturu).

# --instructions--

Promenite funkciju `nonMutatingPush` ili koristite `concat` da spojite `newItem` na kraj `original` bez menjanja redova podataka za `original` ili `newItem`. Funkcija bi trebalo da vrati red podataka.

# --hints--

Tvoj kod bi trebalo da koristi metodu `concat`.

```js
assert(__helpers.removeJSComments(code).match(/\.concat/g));
```

Vaš kod ne sme koristiti putanju `push`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?push/g));
```

Podatni red za `first` ne bi trebalo da se menja.

```js
assert(JSON.stringify(first) === JSON.stringify([1, 2, 3]));
```

Red podataka za `second` ne sme se menjati.

```js
assert(JSON.stringify(second) === JSON.stringify([4, 5]));
```

`nonMutatingPush([1, 2, 3], [4, 5])` treba da vrati `[1, 2, 3, 4, 5]`.

```js
assert(
  JSON.stringify(nonMutatingPush([1, 2, 3], [4, 5])) ===
    JSON.stringify([1, 2, 3, 4, 5])
);
```

# --seed--

## --seed-contents--

```js
function nonMutatingPush(original, newItem) {
  // Only change code below this line
  return original.push(newItem);

  // Only change code above this line
}

const first = [1, 2, 3];
const second = [4, 5];
nonMutatingPush(first, second);
```

# --solutions--

```js
function nonMutatingPush(original, newItem) {
  return original.concat(newItem);
}
const first = [1, 2, 3];
const second = [4, 5];
```
