---
id: 587d7b88367417b2b2512b45
title: 'Koristi funkciju višeg reda kao što su map, filter ili reduce za rešavanje složenih izazova'
challengeType: 1
forumTopicId: 301311
dashedName: use-higher-order-functions-map-filter-or-reduce-to-solve-a-complex-problem
---

# --description--

Sada, nakon što ste završili nekoliko izazova koristeći napredne funkcije kao `map()`, `filter()`, i `reduce()`, sada ih možete koristiti za rešavanje još težih izazova.

# --instructions--

Potpuni kod funkcije `squareList` koristeći bilo koju kombinaciju od `map()`, `filter()`, i `reduce()`. Funkcija mora vratiti novi red podataka koji sadrži kvadrat *samo pozitivnih celih brojeva* (decimalni brojevi nisu ceo broj) kada primi red stvarnih brojeva. Primer reda stvarnih brojeva je `[-3, 4.8, 5, 3, -3.2]`.

**Napomena:** Vaša funkcija ne sme da koristi nijedan tip od `for` ili `while` ili funkciju za `forEach()`.

# --hints--

`squareList` treba da bude `function`.

```js
assert.typeOf(squareList, 'function'),
  '<code>squareList</code> should be a <code>function</code>';
```

`for`, `while`, na `forEach` ne bi trebalo koristiti.

```js
assert(!__helpers.removeJSComments(code).match(/for|while|forEach/g));
```

`map`, `filter`, ili `reduce` treba da se koristi.

```js
assert(
  __helpers
    .removeWhiteSpace(__helpers.removeJSComments(code))
    .match(/\.(map|filter|reduce)\(/g)
);
```

Funkcija mora da vrati `array`.

```js
assert(Array.isArray(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])));
```

`squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2])` trebalo bi da vrati `[16, 1764, 36]`.

```js
assert.deepStrictEqual(squareList([4, 5.6, -9.8, 3.14, 42, 6, 8.34, -2]), [
  16,
  1764,
  36
]);
```

`squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3])` trebalo bi da vrati `[9, 100, 49]`.

```js
assert.deepStrictEqual(squareList([-3.7, -5, 3, 10, 12.5, 7, -4.5, -17, 0.3]), [
  9,
  100,
  49
]);
```

# --seed--

## --seed-contents--

```js
const squareList = arr => {
  // Only change code below this line
  return arr;
  // Only change code above this line
};

const squaredIntegers = squareList([-3, 4.8, 5, 3, -3.2]);
console.log(squaredIntegers);
```

# --solutions--

```js
const squareList = arr => {
  const positiveIntegers = arr.filter(num => {
    return num >= 0 && Number.isInteger(num);
  });
  const squaredIntegers = positiveIntegers.map(num => {
    return num ** 2;
  });
  return squaredIntegers;
};
```
