---
id: 587d7da9367417b2b2512b69
title: Sortiraj red podataka po abecednom redosledu koristeći metodu sortiranja
challengeType: 1
forumTopicId: 18303
dashedName: sort-an-array-alphabetically-using-the-sort-method
---

# --description--

Metoda `sort` organizuje elemente redova podataka prema funkciji povratnog poziva.

Na primer:

```js
function ascendingOrder(arr) {
  return arr.sort(function(a, b) {
    return a - b;
  });
}

ascendingOrder([1, 5, 2, 3, 4]);
```

Ovo će vratiti vrednost od `[1, 2, 3, 4, 5]`.

```js
function reverseAlpha(arr) {
  return arr.sort(function(a, b) {
    return a === b ? 0 : a < b ? 1 : -1;
  });
}

reverseAlpha(['l', 'h', 'z', 'b', 's']);
```

Ovo će vratiti vrednost `['z', 's', 'l', 'h', 'b']`.

Standardna metoda sortiranja u JavaScript je na osnovu Unicode tačkaste vrednosti sekvence karaktera, što može vratiti nepouzdane rezultate. Stoga se preporučuje pružanje funkcije za poređenje kako bi se specificisao način sortiranja elemenata niza podataka. Kada se ta funkcija za poređenje, obično nazvana `compareFunction`, postavi, elementi niza podataka se sortiraju prema vrednosti return od `compareFunction`: Ako `compareFunction(a,b)` vrati vrednost manju od 0 za dva elementa `a` i `b`, onda će `a` doći pre `b`. Ako `compareFunction(a,b)` vrati vrednost veću od 0 za dva elementa `a` i `b`, onda će `b` doći pre `a`. Ako `compareFunction(a,b)` vrati vrednost jednaku 0 za dva elementa `a` i `b`, onda će `a` i `b` ostati kako su.

# --instructions--

Koristi putanju `sort` u funkciji `alphabeticalOrder` da sortira elemente `arr` po abecednom redosledu. Funkcija bi trebalo da vrati sortirovani niz podataka.

# --hints--

Vaš kod treba da koristi `sort`.

```js
assert(__helpers.removeJSComments(code).match(/\.sort/g));
```

`alphabeticalOrder(["a", "d", "c", "a", "z", "g"])` treba da vrati `["a", "a", "c", "d", "g", "z"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'd', 'c', 'a', 'z', 'g'])) ===
    JSON.stringify(['a', 'a', 'c', 'd', 'g', 'z'])
);
```

`alphabeticalOrder(["x", "h", "a", "m", "n", "m"])` trebao bi da vrati `["a", "h", "m", "m", "n", "x"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['x', 'h', 'a', 'm', 'n', 'm'])) ===
    JSON.stringify(['a', 'h', 'm', 'm', 'n', 'x'])
);
```

`alphabeticalOrder(["a", "a", "a", "a", "x", "t"])` treba da vrati `["a", "a", "a", "a", "t", "x"]`.

```js
assert(
  JSON.stringify(alphabeticalOrder(['a', 'a', 'a', 'a', 'x', 't'])) ===
    JSON.stringify(['a', 'a', 'a', 'a', 't', 'x'])
);
```

# --seed--

## --seed-contents--

```js
function alphabeticalOrder(arr) {
  // Only change code below this line

  return arr
  // Only change code above this line
}

alphabeticalOrder(["a", "d", "c", "a", "z", "g"]);
```

# --solutions--

```js
function alphabeticalOrder(arr) {
  return arr.sort();
}
```
