---
id: 587d7b90367417b2b2512b65
title: Vrati deo niza podataka koristeći `slice` metod
challengeType: 1
forumTopicId: 301239
dashedName: return-part-of-an-array-using-the-slice-method
---

# --description--

Funkcija ``slice`` vraća kopiju određenih elemenata niza podataka. Može prihvatiti dva argumenta: prvi navodi indeks za početak sečenja, a drugi indeks za kraj sečenja (ne uključujući). Ako nisu navedeni argumenti, podrazumevano je da počinje od početka niza podataka do kraja, što je jednostavan način za kreiranje kopije celog niza. Funkcija ``slice`` ne modifikuje originalni niz podataka, već vraća novi.

Evo primera:

```js
const arr = ["Cat", "Dog", "Tiger", "Zebra"];
const newArray = arr.slice(1, 3);
```

`newArray` ima vrednost `["Dog", "Tiger"]`.

# --instructions--

Koristi putanju `slice` u funkciji `sliceArray` da vrati deo redova podataka iz `anim` uzimajući u obzir indikatore `beginSlice` i `endSlice` koji su dati. Funkcija treba da vrati redove podataka.

# --hints--

Tvoj kod bi trebalo da koristi metod `slice`.

```js
assert(__helpers.removeJSComments(code).match(/\.slice/g));
```

Parametar za `inputAnim` se ne sme menjati.

```js
assert(
  JSON.stringify(inputAnim) ===
    JSON.stringify(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 3)` trebalo bi da vrati `["Dog", "Tiger"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 3)) ===
    JSON.stringify(['Dog', 'Tiger'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 0, 1)` treba da vrati `["Cat"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 0, 1)) ===
    JSON.stringify(['Cat'])
);
```

`sliceArray(["Cat", "Dog", "Tiger", "Zebra", "Ant"], 1, 4)` treba da vrati `["Dog", "Tiger", "Zebra"]`.

```js
assert(
  JSON.stringify(sliceArray(['Cat', 'Dog', 'Tiger', 'Zebra', 'Ant'], 1, 4)) ===
    JSON.stringify(['Dog', 'Tiger', 'Zebra'])
);
```

# --seed--

## --seed-contents--

```js
function sliceArray(anim, beginSlice, endSlice) {
  // Only change code below this line


  // Only change code above this line
}

const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
sliceArray(inputAnim, 1, 3);
```

# --solutions--

```js
function sliceArray(anim, beginSlice, endSlice) {
  return anim.slice(beginSlice, endSlice);
}
const inputAnim = ["Cat", "Dog", "Tiger", "Zebra", "Ant"];
```
