---
id: 587d7daa367417b2b2512b6c
title: Povežite red podatke u niz znakova koristeći JOIN metodu
challengeType: 1
forumTopicId: 18221
dashedName: combine-an-array-into-a-string-using-the-join-method
---

# --description--

Njia `join` se koristi za spajanje elemenata niza podataka kako bi se stvorila sekvenca karaktera. Uzima argument funkcije separatora koji se koristi za razdvajanje elemenata niza podataka unutar sekvence karaktera.

Evo primera:

```js
const arr = ["Hello", "World"];
const str = arr.join(" ");
```

`str` ima vrednost niza karaktera `Hello World`.# --instructions--

Koristi metodu `join` (među ostalnim) unutar funkcije `sentensify` da napraviš rečenicu od reči koje se nalaze u nizu slova `str`. Funkcija treba da vrati niz slova. Na primer, `I-like-Star-Wars` će biti konvertovan u `I like Star Wars`. Za ovaj izazov, ne koristi metodu `replace`.

# --hints--

Vaš kod bi trebalo da koristi metodu `join`.

```js
assert(__helpers.removeJSComments(code).match(/\.join/g));
```

Vaš kod ne bi trebalo da koristi način `replace`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?replace/g));
```

`sentensify("May-the-force-be-with-you")` Trebalo bi da vrati niz karaktera.

```js
assert(typeof sentensify('May-the-force-be-with-you') === 'string');
```

`sentensify("May-the-force-be-with-you")` mora da vrati sekvencu karaktera `May the force be with you`.

```js
assert(sentensify('May-the-force-be-with-you') === 'May the force be with you');
```

`sentensify("The.force.is.strong.with.this.one")` treba da vrati sekvencu karaktera `The force is strong with this one`.

```js
assert(
  sentensify('The.force.is.strong.with.this.one') ===
    'The force is strong with this one'
);
```

`sentensify("There,has,been,an,awakening")` treba da vrati niz karaktera `There has been an awakening`.

```js
assert(
  sentensify('There,has,been,an,awakening') === 'There has been an awakening'
);
```

# --seed--

## --seed-contents--

```js
function sentensify(str) {
  // Only change code below this line


  // Only change code above this line
}

sentensify("May-the-force-be-with-you");
```

# --solutions--

```js
function sentensify(str) {
  return str.split(/\W/).join(' ');
}
```
