---
id: 587d7daa367417b2b2512b6b
title: Razdvoj niz karaktera u niz podataka koristeći metodu `split`
challengeType: 1
forumTopicId: 18305
dashedName: split-a-string-into-an-array-using-the-split-method
---

# --description--

Metoda `split` deliše niz znakova u niz podataka niza znakova. Prihvata argument delimiatora, koji može biti znak korišćen za deljenje niza znakova ili regularni izraz. Na primer, ako je delimiator razmak, dobijate niz reči, a ako je delimiator prazan niz znakova, dobijate niz svakog znaka u nizu znakova.

Ovde su prikazani dva primera koje dele jednu sekvencu pomoću razmaka, a drugi deo pomoću brojeva, koristeći standardnu terminologiju:

```js
const str = "Hello World";
const bySpace = str.split(" ");

const otherString = "How9are7you2today";
const byDigits = otherString.split(/\d/);
```

`bySpace` će imati vrednost `["Hello", "World"]` i `byDigits` će imati vrednost `["How", "are", "you", "today"]`.

Pošto sekvenca karaktera ne može da se promeni, metoda `split` olakšava rad sa njima.

# --instructions--

Koristi putanju `split` unutar funkcije `splitify` da razdvoji `str` u niz reči. Funkcija treba da vrati taj niz. Zapamti da se reči ne uvek razdvajaju razmakom, i da niz ne bi trebalo da sadrži interpunkcijske znakove.

# --hints--

Tvoj kod bi trebalo da koristi način `split`.

```js
assert(__helpers.removeJSComments(code).match(/\.split/g));
```

`splitify("Hello World,I-am code")` treba da vrati `["Hello", "World", "I", "am", "code"]`.

```js
assert(
  JSON.stringify(splitify('Hello World,I-am code')) ===
    JSON.stringify(['Hello', 'World', 'I', 'am', 'code'])
);
```

`splitify("Earth-is-our home")` treba da vrati `["Earth", "is", "our", "home"]`.

```js
assert(
  JSON.stringify(splitify('Earth-is-our home')) ===
    JSON.stringify(['Earth', 'is', 'our', 'home'])
);
```

`splitify("This.is.a-sentence")` bi trebalo da vrati `["This", "is", "a", "sentence"]`.

```js
assert(
  JSON.stringify(splitify('This.is.a-sentence')) ===
    JSON.stringify(['This', 'is', 'a', 'sentence'])
);
```

# --seed--

## --seed-contents--

```js
function splitify(str) {
  // Only change code below this line


  // Only change code above this line
}

splitify("Hello World,I-am code");
```

# --solutions--

```js
function splitify(str) {
  return str.split(/\W/);
}
```
