---
id: 641da5462576784453146ec2
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Sada kada imate podatke koje želite, možete ih koristiti za popunjavanje korisničkog interfejsa (UI).

Međutim, podaci koje dohvatate će imati ukupno 26 autora. Umesto da prikažete svih 26 autora odjednom, trebalo bi da prikazujete po 8 autora.

Počnite korišćenjem ključne reči `let` za kreiranje dve varijable nazvane `startingIndex` i `endingIndex`, i dodelite im brojeve `0` i `8`, respektivno. Biće vam potrebno da ovde koristite `let` jer ćete kasnije ponovo dodeljivati ove vrednosti.

Zatim, kreirajte varijablu `authorDataArr` i dodelite joj prazan niz (array).

Deklarišite i inicijalizujte svaku varijablu na novoj liniji.

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

Trebalo bi da koristite `let` za deklarisanje varijable pod nazivom `startingIndex`.

```js
assert.match(code, /let\s+startingIndex/);
```

Vaša varijabla `startingIndex` bi trebalo da bude broj.

```js
assert.isNumber(startingIndex);
```

Vaša promenljiva `startingIndex` bi trebalo da bude set da bi bila `0`.

```js
assert.strictEqual(startingIndex, 0);
```

Trebalo bi da koristite `let` za deklarisanje varijable pod nazivom `endingIndex`.

```js
assert.match(code, /let\s+endingIndex/);
```

Vaša varijabla `endingIndex` treba da bude broj.```js
assert.isNumber(endingIndex);
```

You should set your `endingIndex` variable to `8`.

```js
assert.strictEqual(endingIndex, 8)
```

Trebalo bi da imate varijablu pod nazivom `authorDataArr`.

```js
assert.isDefined(authorDataArr);
```

Vaša varijabla `authorDataArr` bi trebalo da bude prazan niz.

```js
assert.isArray(authorDataArr);
assert.isEmpty(authorDataArr);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
