---
id: 587d7db9367417b2b2512ba6
title: Navedite samo nekoliko primjera
challengeType: 1
forumTopicId: 301366
dashedName: specify-only-the-lower-number-of-matches
---

# --description--

Možete navesti opseg brojeva sistema koristeći indikatore opsega sa uglastim zagradama. Ponekad želite samo da navedete niski broj sistema bez gornje granice.

Da navedete samo nekoliko sistemskih brojeva, stavite prvi broj praćen zarezom.

Na primer, za uporedbu samo sekvence znakova `hah` sa znakom `a` ako se pojavljuje najmanje puta `3`, vaš regex će biti `/ha{3,}h/`.

```js
let A4 = "haaaah";
let A2 = "haah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleA = /ha{3,}h/;
multipleA.test(A4);
multipleA.test(A2);
multipleA.test(A100);
```

Po podešavanju, tri poziva za `test` će vratiti `true`, `false`, na `true`.

# --instructions--

Promeni regex `haRegex` ili pronađi reč `Hazzah` samo kada sadrži četiri ili više slova `z`.

# --hints--

Tvoj regex treba da koristi kvadesne zagrade.

```js
assert(haRegex.source.match(/{.*?}/).length > 0);
```

Regex ne bi trebalo da odgovara nizu karaktera `Hazzah`

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzah'));
```

Tvoj regex ne bi trebalo da upoređuje niz slova `Hazzzah`

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzzah'));
```

Regex mora da odgovara sekvenci karaktera `Hazzzzah`

```js
assert('Hazzzzah'.match(haRegex)[0].length === 8);
```

Tvoj regex treba da odgovara nizu slova `Hazzzzzah`

```js
assert('Hazzzzzah'.match(haRegex)[0].length === 9);
```

Vaš regex treba da odgovara nizu znakova `Hazzzzzzah`

```js
assert('Hazzzzzzah'.match(haRegex)[0].length === 10);
```

Regex vaš bi trebalo da poklapa sekvencu karaktera `Hazzah` koja sadrži karaktere `z` 30 u sebi.

```js
assert('Hazzzzzzzzzzzzzzzzzzzzzzzzzzzzzzah'.match(haRegex)[0].length === 34);
```

# --seed--

## --seed-contents--

```js
let haStr = "Hazzzzah";
let haRegex = /change/; // Change this line
let result = haRegex.test(haStr);
```

# --solutions--

```js
let haStr = "Hazzzzah";
let haRegex = /Haz{4,}ah/; // Change this line
let result = haRegex.test(haStr);
```
