---
id: 587d7db3367417b2b2512b8f
title: Uporedite sekvencu stvarnih slova
challengeType: 1
forumTopicId: 301355
dashedName: match-literal-strings
---

# --description--

U prethodnom izazovu, tražili ste reč `Hello` koristeći regularni izraz `/Hello/`. Taj regularni izraz je tražio dosločno poklapanje sekvence karaktera `Hello`. Evo još jednog primera pretraživanja dosločnog poklapanja sekvence karaktera `Kevin`:

```js
let testStr = "Hello, my name is Kevin.";
let testRegex = /Kevin/;
testRegex.test(testStr);
```

Ovaj poziv za `test` će vratiti `true`.

Drugi bilo koji deo `Kevin` neće odgovarati. Na primer, uobičajeni izraz `/Kevin/` ne odgovara `kevin` niti `KEVIN`.

```js
let wrongRegex = /kevin/;
wrongRegex.test(testStr);
```

Ovaj poziv sa `test` će vratiti `false`.

Predstojeći izazovi će pokazati kako da uporedite i te druge vrste.

# --instructions--

Popunite uobičajenu frazu `waldoRegex` da biste dobili `"Waldo"` u nizu slova `waldoIsHiding` za tačno poklapanje.

# --hints--

Vaše uobičajeno korisničko ime `waldoRegex` mora da dobije niz slova `Waldo`

```js
waldoRegex.lastIndex = 0;
assert(waldoRegex.test(waldoIsHiding));
```

Vaša uobičajena upotreba `waldoRegex` ne bi trebalo da traži ništa drugo.

```js
waldoRegex.lastIndex = 0;
assert(!waldoRegex.test('Somewhere is hiding in this text.'));
```

Treba da izvršiš doslovno poređenje sekvence slova za svoj uobičajeni izgovor.

```js
assert(!/\/.*\/i/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /search/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```

# --solutions--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /Waldo/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```
