---
id: 587d7db5367417b2b2512b97
title: Povežite brojeve i slova abecede
challengeType: 1
forumTopicId: 301356
dashedName: match-numbers-and-letters-of-the-alphabet
---

# --description--

Korišćenje simbola križa (`-`) za podudarnost redova slova nije ograničeno samo na slova. Takođe radi i za podudarnost redova brojeva.

Na primer, `/[0-5]/` pokriva bilo koji broj između `0` i `5`, uključujući `0` i `5`.

Takođe, moguće je kombinovati niz slova i brojeva u jedan skup karaktera.

```js
let jennyStr = "Jenny8675309";
let myRegex = /[a-z0-9]/ig;
jennyStr.match(myRegex);
```

# --instructions--

Tvorite regularni izraz koji odgovara nizu znakova između `h` i `s`, i nizu brojeva između `2` i `6`. Zapamtite da uključite odgovarajuće parametre u regularni izraz.

# --hints--

Regularni izraz `myRegex` treba da pronađe 17 podudaranja.

```js
assert(result.length == 17);
```

Regularni izraz `myRegex` treba da koristi globalnu zastavicu.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

Vaš standardni način korišćenja `myRegex` treba da koristi parametar koji nije osetljiv na veličinu slova.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /[h-s2-6]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
