---
id: 587d7db7367417b2b2512b9d
title: Uporedi šablone karaktera na početku
challengeType: 1
forumTopicId: 301349
dashedName: match-beginning-string-patterns
---

# --description--

Prethodni izazovi su pokazali da se regularni izrazi mogu koristiti za pronalaženje više podudaranja. Takođe se koriste za pronalaženje obrazaca na određenim pozicijama unutar niza.

U prethodnom izazovu koristio si znak kareta (`^`) unutar skupa znakova da napraviš negirani skup znakova u obliku `[^thingsThatWillNotBeMatched]`. Van skupa znakova, znak kareta služi za pronalaženje obrazaca na početku niza.

```js
let firstString = "Ricky is first and can be found.";
let firstRegex = /^Ricky/;
firstRegex.test(firstString);
let notFirst = "You can't find Ricky now.";
firstRegex.test(notFirst);
```

Prvi poziv za `test` vraća `true`, dok drugi poziv vraća `false`.

# --instructions--

Koristi caret znak u regularnom izrazu da tražiš `Cal` samo na početku sekvence znaka `rickyAndCal`.

# --hints--

Tvoj regularni izraz treba da pronađe niz `Cal` sa velikim početnim slovom.

```js
assert(calRegex.source == '^Cal');
```

Tvoj regularni izraz ne treba da koristi nijednu zastavicu.

```js
assert(calRegex.flags == '');
```

Tvoj regularni izraz treba da pronađe niz `Cal` na početku niza.

```js
calRegex.lastIndex = 0;
assert(calRegex.test('Cal and Ricky both like racing.'));
```

Vaša standardna upotreba ne bi trebalo da uporedi sekvencu slova `Cal` u sredini niza.

```js
calRegex.lastIndex = 0;
assert(!calRegex.test('Ricky and Cal both like racing.'));
```

# --seed--

## --seed-contents--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /change/; // Change this line
let result = calRegex.test(rickyAndCal);
```

# --solutions--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /^Cal/; // Change this line
let result = calRegex.test(rickyAndCal);
```
