---
id: 587d7db8367417b2b2512ba3
title: Popuni prazno mesto
challengeType: 1
forumTopicId: 301359
dashedName: match-whitespace
---

# --description--

Prethodni izazovi su se odnosili na poređenje alfabetskih znakova i brojeva. Takođe možete upoređivati prazan prostor ili razmak između slova.

Možete pronaći prazan prostor koristeći `\s`, što je mala slova `s`. Ovaj primer ne odgovara samo praznom prostoru, već uključuje znakove povratka karriđere (carriage return), tab, form feed i nove linije. Možete ga smatrati kao klasu elemenata `[ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let spaceRegex = /\s/g;
whiteSpace.match(spaceRegex);
```

Ovaj poziv `match` će vratiti `[" ", " "]`.# --instructions--

Promenite uobičajeni izraz `countWhiteSpace` ili potražite više praznih razmaka u nizu karaktera.

# --hints--

Vaša uobičajena izjava treba da koristi globalnu zastavicu.

```js
assert(countWhiteSpace.global);
```

Vaša uobičajena upotreba bi trebalo da koristi skraćenicu slova `\s` ili da uporedi sve razmake.

```js
assert(/\\s/.test(countWhiteSpace.source));
```

Vaša uobičajena upotreba bi trebalo da dobije osam pozicija u nizu slova `Men are from Mars and women are from Venus.`

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countWhiteSpace).length ==
    8
);
```

Vaša uobičajena upotreba trebalo bi da dobije tri mesta u nizu slova `Space: the final frontier.`

```js
assert('Space: the final frontier.'.match(countWhiteSpace).length == 3);
```

Vaš uobičajeni unos ne bi trebalo da pronađe nikakvo mesto u sekvenci karaktera `MindYourPersonalSpace`

```js
assert('MindYourPersonalSpace'.match(countWhiteSpace) == null);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /change/; // Change this line
let result = sample.match(countWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /\s/g;
let result = sample.match(countWhiteSpace);
```
