---
id: 587d7db7367417b2b2512b9f
title: Povežite sve slova i brojeve
challengeType: 1
forumTopicId: 301346
dashedName: match-all-letters-and-numbers
---

# --description--

Korišćenjem klasa znakova mogao si da pronađeš sva slova alfabeta pomoću `[a-z]`. Ova klasa znakova je toliko česta da postoji i kraći zapis, koji uključuje i nekoliko dodatnih znakova.

Klasa znakova koja u JavaScript-u najviše odgovara slovima je `\w`. Ova prečica je isto što i `[A-Za-z0-9_]`. Obuhvata velika i mala slova, brojeve i donju crtu (`_`).

```js
let longHand = /[A-Za-z0-9_]+/;
let shortHand = /\w+/;
let numbers = "42";
let varNames = "important_var";
longHand.test(numbers);
shortHand.test(numbers);
longHand.test(varNames);
shortHand.test(varNames);
```

Svi četiri Wito iz `test` vraćaju `true`.

Ove klase za kratke karakteristike takođe se poznaju kao <dfn>klase za kratke karakteristike</dfn>.

# --instructions--

Upotrebi prečicu `\w` da izbrojiš broj slova i brojeva u citatima i drugim nizovima.

# --hints--

Regularni izraz treba da koristi globalnu zastavicu.

```js
assert(alphabetRegexV2.global);
```

Vaše uobičajeno korišćenje bi trebalo da koristi kratku klasu karaktera `\w` za poređenje svih znakova koji su slova i brojevi.

```js
assert(/\\w/.test(alphabetRegexV2.source));
```

Regularni izraz treba da pronađe 31 slovo ili broj u nizu `The five boxing wizards jump quickly.`

```js
assert(
  'The five boxing wizards jump quickly.'.match(alphabetRegexV2).length === 31
);
```

Regularni izraz treba da pronađe 32 slova ili broja u nizu `Pack my box with five dozen liquor jugs.`

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(alphabetRegexV2).length ===
    32
);
```

Regularni izraz treba da pronađe 30 slova ili brojeva u nizu `How vexingly quick daft zebras jump!`

```js
assert(
  'How vexingly quick daft zebras jump!'.match(alphabetRegexV2).length === 30
);
```

Regularni izraz treba da pronađe 36 slova ili brojeva u nizu `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.`

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(alphabetRegexV2)
    .length === 36
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /change/; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /\w/g; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```
