---
id: 587d7db9367417b2b2512ba5
title: Navedite broj ispod i iznad meča
challengeType: 1
forumTopicId: 301367
dashedName: specify-upper-and-lower-number-of-matches
---

# --description--

Zapamtite da koristite znak plus `+` za pretraživanje jednog ili više slova i zvezdicu `*` za pretraživanje nule ili više slova. Ovo je jednostavno, ali ponekad želite da uporedite određene obrasce.

Možete specificirati minimalan i maksimalan broj sistema koristeći <dfn>kvantitativne indikatore</dfn>. Indikatori opsega se koriste sa zaokruženim zagradama (`{` i `}`). Unosite dva broja između zaokruženih zagrada – za minimalan i maksimalan broj sistema.

Primer, da uporedimo samo karakter `a` koji se ponavlja između `3` i `5` puta u nizu karaktera `ah`, vaš običan izraz bi bio `/a{3,5}h/`.

```js
let A4 = "aaaah";
let A2 = "aah";
let multipleA = /a{3,5}h/;
multipleA.test(A4);
multipleA.test(A2);
```

Prvi poziv za `test` vraća `true`, a drugi vraća `false`.

# --instructions--

Promeni uobičajenu frazu `ohRegex` ili uporedite ceo niz `Oh no` samo kada sadrži karaktere `3` počevši od `6` do `h`.

# --hints--

Vaše uobičajeno korišćenje trebalo bi da koristi vitičaste zagrade.

```js
assert(ohRegex.source.match(/{.*?}/).length > 0);
```

Vaša uobičajena upotreba ne bi trebalo da poredi sekvencu `Ohh no`

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohh no'));
```

Vaš uobičajeni pristup bi trebalo da odgovara nizu `Ohhh no`

```js
assert('Ohhh no'.match(ohRegex)[0].length === 7);
```

Vaše uobičajeno korišćenje treba da uporedi sekvencu `Ohhhh no`

```js
assert('Ohhhh no'.match(ohRegex)[0].length === 8);
```

Vaša uobičajena upotreba bi trebalo da uporedi sekvencu `Ohhhhh no`

```js
assert('Ohhhhh no'.match(ohRegex)[0].length === 9);
```

Uobičajena upotreba trebalo bi da uporedi sekvencu `Ohhhhhh no`

```js
assert('Ohhhhhh no'.match(ohRegex)[0].length === 10);
```

Vaš uobičajeni tekst ne bi trebalo da poredi sekvencu `Ohhhhhhh no`

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohhhhhhh no'));
```

# --seed--

## --seed-contents--

```js
let ohStr = "Ohhh no";
let ohRegex = /change/; // Change this line
let result = ohRegex.test(ohStr);
```

# --solutions--

```js
let ohStr = "Ohhh no";
let ohRegex = /Oh{3,6} no/; // Change this line
let result = ohRegex.test(ohStr);
```
