---
id: 587d7db6367417b2b2512b98
title: Uporedi jedno neodređeno slovo
challengeType: 1
forumTopicId: 301358
dashedName: match-single-characters-not-specified
---

# --description--

Sada, kreirali ste skup znakova koje želite da uporedite, ali takođe možete kreirati skup znakova koje ne želite da upoređujete. Ove vrste skupa znakova nazivaju se <dfn>seti za herufi zilizokataliwa</dfn>.

Da biste kreirali set isključenih znakova, postavite caret znak (`^`) nakon otvarajućih zagrada i pre karaktera koje ne želite da uporedite.

Na primer, ``/[^aeiou]/gi`` upoređuje sve karaktere koji nisu samoglasnici. Zapamtite da se karaktere kao što su ``.``, ``!``, ``[``, ``@``, ``/``, i razmaci upoređuju – skup isključenih samoglasnika razdvaja samo samoglasne karaktere.

# --instructions--

Jedan običan izraz koji poklapa sva slova koja nisu ni brojevi, niti samoglasnici. Zapamtite da uključite odgovarajuće zastavice u regularni izraz.

# --hints--

Vaša uobičajena upotreba `myRegex` trebalo bi da uporedi 9 stavki.

```js
assert(result.length == 9);
```

Vaša uobičajena upotreba `myRegex` treba da koristi global zastavica.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

Standardno korišćenje `myRegex` bi trebalo da koristi zastavicu koja ignoriše velika i mala slova.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "3 blind mice.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "3 blind mice.";
let myRegex = /[^0-9aeiou]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
