---
id: 587d7db8367417b2b2512ba2
title: Smanite moguće korisničke imena
challengeType: 1
forumTopicId: 301363
dashedName: restrict-possible-usernames
---

# --description--

Korisnička imena se koriste svuda online. Na taj način korisnicima daju jedinstveni identitet na sajtovima koje vole.

Morate proveriti sva korisnička imena u bazi podataka. Evo jednostavnih pravila koja korisnici moraju da prate kada kreiraju svoje korisničko ime.

1) Korisnička imena mogu koristiti samo slova i brojeve.

2) Samo brojke u korisničkom imenu moraju biti na kraju. Na kraju mogu biti nule ili više brojeva. Korisničko ime ne sme počinjati sa brojem.

3) Karakteri za korisničko ime mogu biti mala ili velika slova.

4) Koriska imena moraju imati dužinu od najmanje dva karaktera. Korisko ime sa dve slova može koristiti samo alfabetska slova kao svoje karaktere.

# --instructions--

Promeni običan izraz `userCheck` tako da odgovara kriterijumima navedenim gore.

# --hints--

Vaš uobičajeni korisnik treba da odgovara nizu karaktera `JACK`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('JACK'));
```

Vaš uobičajeni tekst ne bi trebalo da odgovara nizu karaktera `J`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J'));
```

Vaš uobičajeni korisnik treba da odgovara nizu slova `Jo`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Jo'));
```

Vaš uobičajeni unos bi trebalo da odgovara nizu slova `Oceans11`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Oceans11'));
```

Vaš uobičajeni korisnički naziv mora odgovarati nizu znakova `RegexGuru`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('RegexGuru'));
```

Vaš uobičajeni izraz ne bi trebalo da odgovara nizu slova `007`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('007'));
```

Vaš uobičajeni unos ne bi trebalo da odgovara nizu slova `9`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('9'));
```

Vaša uobičajena upotreba ne bi trebalo da odgovara nizu znakova `A1`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('A1'));
```

Vaš uobičajeni niz ne bi trebalo da odgovara sekvenci slova `BadUs3rnam3`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('BadUs3rnam3'));
```

Vaša uobičajena upotreba mora da odgovara sekvenci slova `Z97`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Z97'));
```

Tvoj uobičajeni pristup ne bi trebalo da odgovara nizu karaktera `c57bT3`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('c57bT3'));
```

Vaša uobičajena upotreba mora da odgovara nizu slova `AB1`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('AB1'));
```

Tvoj uobičajeni nalog ne sme odgovarati sekvenci slova `J%4`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J%4'))
```

# --seed--

## --seed-contents--

```js
let username = "JackOfAllTrades";
let userCheck = /change/; // Change this line
let result = userCheck.test(username);
```

# --solutions--

```js
let username = "JackOfAllTrades";
const userCheck = /^[a-z]([0-9]{2,}|[a-z]+\d*)$/i;
let result = userCheck.test(username);
```
