---
id: 587d7db6367417b2b2512b99
title: Identifikujte slova koja se ponavljaju jednom ili više puta
challengeType: 1
forumTopicId: 301350
dashedName: match-characters-that-occur-one-or-more-times
---

# --description--

Ponekad je potrebno uporediti slovo (ili niz slova) koje se pojavljuje jednokratno ili više puta uzastopno. Ovo znači da se dešava barem jednom i može biti ponovljeno.

Možete koristiti znak `+` da proverite da li je to stanje. Zapamtite, slovo ili uzorak mora biti kontinuitetno. To jest, slovo mora ponavljati jedno za drugim.

Na primer, `/a+/g` bi pronašao jedno poklapanje u `abc` i vratio `["a"]`. Zbog `+`, takođe bi pronašao jedno poklapanje u `aabc` i vratio `["aa"]`.

Ako bi umesto toga skenirao niz karaktera `abab`, pronašao bi dve podudarnosti i vratio `["a", "a"]` jer karakteri `a` nisu uzastopni - postoji `b` između njih. Na kraju, pošto nema `a` u nizu karaktera `bcd`, nije pronašao podudarnost.

# --instructions--

Želite da pronađete ekvivalent kada slovo `s` se pojavljuje jednom ili više puta u `Mississippi`. Napišite običan izraz koji koristi znak `+`.

# --hints--

Vaše uobičajeno korišćenje `myRegex` treba da koristi znak za `+` prilikom poređenja slova `s` jednom ili više puta.

```js
assert(/\+/.test(myRegex.source));
```

Vaša uobičajena upotreba `myRegex` bi trebalo da uporedi dve stvari.

```js
assert(result.length == 2);
```

Parametar za `result` treba da bude red podataka sa dve jednačine od `ss`

```js
assert(result[0] == 'ss' && result[1] == 'ss');
```

# --seed--

## --seed-contents--

```js
let difficultSpelling = "Mississippi";
let myRegex = /change/; // Change this line
let result = difficultSpelling.match(myRegex);
```

# --solutions--

```js
let difficultSpelling = "Mississippi";
let myRegex = /s+/g; // Change this line
let result = difficultSpelling.match(myRegex);
```
