---
id: 587d7db5367417b2b2512b95
title: Poveži slovo sa mnogo mogućnosti
challengeType: 1
forumTopicId: 301357
dashedName: match-single-character-with-multiple-possibilities
---

# --description--

Naučili ste kako da upoređujete stvarne obrasce (`/literal/`) sa karakterom za wildkard (`/./`). Ovo su dve potpuno različite vrste izraza, gde jedan traži tačno poklapanje, a drugi poredi sve. Postoje opcije koje su slične između ova dva potpuno različita tipa.

Možete pretraživati doslovnu sekvencu karaktera sa određenom fleksibilnošću korišćenjem <dfn>karakterne klase</dfn>. Karakterne klase vam omogućavaju da definišete grupu karaktera koje želite da uparite tako što ih stavite u uglaste zagrade (`[` i `]`).

Na primer, želite da uporedite `bag`, `big`, sa `bug` ali ne `bog`. Možete kreirati regularni izraz `/b[aiu]g/` da to učinite. `[aiu]` je klasa karaktera koja će podudarati samo sa karakterima `a`, `i`, ili `u`.

```js
let bigStr = "big";
let bagStr = "bag";
let bugStr = "bug";
let bogStr = "bog";
let bgRegex = /b[aiu]g/;
bigStr.match(bgRegex);
bagStr.match(bgRegex);
bugStr.match(bgRegex);
bogStr.match(bgRegex);
```

Redosledom, četiri poziva od `match` vraćaće vrednosti `["big"]`, `["bag"]`, `["bug"]`, na `null`.

# --instructions--

Upotrebi klasu znakova sa samoglasnicima (`a`, `e`, `i`, `o`, `u`) u regularnom izrazu `vowelRegex` da pronađeš sve samoglasnike u nizu `quoteSample`.

**Napomena:** Uverite se da upoređujete samoglasnike za velika i mala slova.

# --hints--

Treba ti da dobiješ sva 25 samoglasnika.

```js
assert(result.length == 25);
```

Vaša uobičajena upotreba `vowelRegex` treba koristiti klasu znakova.

```js
assert(/\[.*\]/.test(vowelRegex.source));
```

Vaše uobičajeno korišćenje `vowelRegex` treba da koristi global zastavicu.

```js
assert(vowelRegex.flags.match(/g/).length == 1);
```

Regularni izraz `vowelRegex` treba da koristi zastavicu za neosetljivost na velika i mala slova (`i`).

```js
assert(vowelRegex.flags.match(/i/).length == 1);
```

Regularni izraz ne treba da pronađe nijedan suglasnik.

```js
assert(!/[b-df-hj-np-tv-z]/gi.test(result.join()));
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /change/; // Change this line
let result = vowelRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /[aeiou]/gi; // Change this line
let result = quoteSample.match(vowelRegex); // Change this line
```
