---
id: 587d7db4367417b2b2512b93
title: Pronađi više prvih mečeva
challengeType: 1
forumTopicId: 301342
dashedName: find-more-than-the-first-match
---

# --description--

Sada ste uspeli da izvučete ili potražite uobičajene fraze odjednom.

```js
let testStr = "Repeat, Repeat, Repeat";
let ourRegex = /Repeat/;
testStr.match(ourRegex);
```

Ovde `match` će vratiti `["Repeat"]`.

Da biste pronašli ili isključili frazu koja se pojavljuje više od jednom, možete koristiti globalni flag za pretragu: `g`.

```js
let repeatRegex = /Repeat/g;
testStr.match(repeatRegex);
```

I ovde `match` vraća vrednost `["Repeat", "Repeat", "Repeat"]`

# --instructions--

Kada uzmete standardni izraz `starRegex`, pronađite i navedite obe reči od `Twinkle` iz sekvence slova `twinkleStar`.

**Napomena**  
Možete imati mnogo zastavica u vašem običnom izrazu, kao što je `/search/gi`

# --hints--

Vaše uobičajeno korišćenje `starRegex` bi trebalo da koristi međunarodnu zastavu `g`

```js
assert(starRegex.flags.match(/g/).length == 1);
```

Regularni izraz `starRegex` treba da koristi zastavicu `i` za neosetljivost na velika i mala slova.

```js
assert(starRegex.flags.match(/i/).length == 1);
```

Podudaranje treba da pronađe obe pojave reči `Twinkle`.

```js
assert(
  result.sort().join() ==
    twinkleStar
      .match(/twinkle/gi)
      .sort()
      .join()
);
```

Meš vaš `result` treba da ima dva elementa unutra.

```js
assert(result.length == 2);
```

# --seed--

## --seed-contents--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /change/; // Change this line
let result = twinkleStar; // Change this line
```

# --solutions--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /twinkle/gi;
let result = twinkleStar.match(starRegex);
```
