---
id: 587d7db5367417b2b2512b96
title: Povežite slova abecede
challengeType: 1
forumTopicId: 301354
dashedName: match-letters-of-the-alphabet
---

# --description--

Videli ste kako možete koristiti <dfn>seti za herufi</dfn> za specifikovanje grupe odgovarajućih znakova, ali to je previše pisanja kada vam treba da pokupite veliki opseg znakova (na primer, svaki znak u alfabetu). Srećom, postoji ugrađena funkcija koja ovo čini kratkim i jednostavnim.

Unutar skupa karakterâ, možete definisati opseg odgovarajućih znakova korišćenjem znaka križa: `-`.

Na primer, da biste uporedili mala slova `a` do `e`, koristićete `[a-e]`.

```js
let catStr = "cat";
let batStr = "bat";
let matStr = "mat";
let bgRegex = /[a-e]at/;
catStr.match(bgRegex);
batStr.match(bgRegex);
matStr.match(bgRegex);
```

Po podrazumevanju, tri poziva za `match` vraćaće vrednosti od `["cat"]`, `["bat"]`, na `null`.

# --instructions--

Uporedite sve slova u nizu slova `quoteSample`.

**Napomena**: Uverite se da poklapate sva velika i mala slova.

# --hints--

Vaše uobičajeno korišćenje `alphabetRegex` treba da uporedi stavke 35.

```js
assert(result.length == 35);
```

Regularni izraz `alphabetRegex` treba da koristi globalnu zastavicu.

```js
assert(alphabetRegex.flags.match(/g/).length == 1);
```

Regularni izraz `alphabetRegex` treba da koristi zastavicu za neosetljivost na velika i mala slova.

```js
assert(alphabetRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /change/; // Change this line
let result = alphabetRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /[a-z]/gi; // Change this line
let result = quoteSample.match(alphabetRegex); // Change this line
```
