---
id: 587d7dba367417b2b2512ba8
title: Gledaj sve ili ništa
challengeType: 1
forumTopicId: 301338
dashedName: check-for-all-or-none
---

# --description--

Ponekad sistemi koje želite da pretražite mogu imati delove koji mogu da postoje ili ne postoje. Međutim, može biti važno proveriti njegovo postojanje ipak.

Možete specificirati mogućnost postojanja elementa sa znakiem pitanja, ``?``. Ovo proverava da li je prisutan nula ili jedan od prethodnih elemenata. Možete smatrati ovaj znak kao rečeno da je prethodni element opcionalan.

Na primer, postoji mala razlika između američkog engleskog i britanskog engleskog, i možete koristiti znak pitanja da odgovara oba pravopisa.

```js
let american = "color";
let british = "colour";
let rainbowRegex= /colou?r/;
rainbowRegex.test(american);
rainbowRegex.test(british);
```

Sva korišćenja putem `test` metode vraćaju `true`.

# --instructions--

Promeni uobičajen izraz `favRegex` ili da odgovara američkoj engleskoj verziji (`favorite`) i britanskoj engleskoj verziji (`favourite`).

# --hints--

Vaša standardna upotreba bi trebalo da koristi opcionalni simbol, `?`.

```js
favRegex.lastIndex = 0;
assert(favRegex.source.match(/\?/).length > 0);
```

Vaš uobičajeni korisnik bi trebalo da odgovara nizu slova `favorite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favorite'));
```

Vaša uobičajena upotreba mora odgovarati nizu karaktera `favourite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favourite'));
```

Vaša uobičajena upotreba ne sme da odgovara nizu slova `fav`

```js
favRegex.lastIndex = 0;
assert(!favRegex.test('fav'));
```

# --seed--

## --seed-contents--

```js
let favWord = "favorite";
let favRegex = /change/; // Change this line
let result = favRegex.test(favWord);
```

# --solutions--

```js
let favWord = "favorite";
let favRegex = /favou?r/;
let result = favRegex.test(favWord);
```
