---
id: 587d7db6367417b2b2512b9a
title: Uporedi slova koja se pojavljuju nula ili više puta
challengeType: 1
forumTopicId: 301351
dashedName: match-characters-that-occur-zero-or-more-times
---

# --description--

Posledji izazov je koristio znak plus `+` za pretraživanje karaktera koji se pojavljuju jedan ili više puta. Postoji i opcija koja odgovara karakterima koji se pojavljuju nula ili više puta.

Karakter koji to radi je zvezda ili asterisk: `*`.

```js
let soccerWord = "gooooooooal!";
let gPhrase = "gut feeling";
let oPhrase = "over the moon";
let goRegex = /go*/;
soccerWord.match(goRegex);
gPhrase.match(goRegex);
oPhrase.match(goRegex);
```

Redosledom, tri poziva `match` će vratiti vrednosti `["goooooooo"]`, `["g"]`, na `null`.

# --instructions--

Kreiraj regularni izraz `chewieRegex` koji koristi karakter `*` za podudarnost velikog slova `A` praćenog odmah malim slovom `a` nula ili više puta u `chewieQuote`. Tvoj regularni izraz ne zahteva zastavice niti klase karaktera, i ne bi trebalo da odgovara bilo kom drugom tekstu.

# --hints--

Vaše standardno korišćenje `chewieRegex` trebalo bi da koristi karakter `*` za poređenje karaktera `a` nula ili više puta.

```js
assert(/\*/.test(chewieRegex.source));
```

Vaše uobičajeno korišćenje treba da uporedi sekvencu karaktera `A` unutar `chewieQuote`.

```js
assert(result[0][0] === 'A');
```

Vaša uobičajena vrednost treba da odgovara sekvenci karaktera `Aaaaaaaaaaaaaaaa` unutar `chewieQuote`.

```js
assert(result[0] === 'Aaaaaaaaaaaaaaaa');
```

Vaša uobičajena vrednost `chewieRegex` treba da odgovara karakteru 16 u `chewieQuote`.

```js
assert(result[0].length === 16);
```

Tvoj uobičajeni deo ne bi trebalo da poredi bilo koji karakter u nizu znakova `He made a fair move. Screaming about it can't help you.`

```js
assert(
  !"He made a fair move. Screaming about it can't help you.".match(chewieRegex)
);
```

Tvoj normalni način ne bi trebalo da poredi bilo koji slovo u nizu slova `Let him have it. It's not wise to upset a Wookiee.`

```js
assert(
  !"Let him have it. It's not wise to upset a Wookiee.".match(chewieRegex)
);
```

# --seed--

## --seed-contents--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
// Only change code below this line
let chewieRegex = /change/; // Change this line
// Only change code above this line

let result = chewieQuote.match(chewieRegex);
```

# --solutions--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
let chewieRegex = /Aa*/;
let result = chewieQuote.match(chewieRegex);
```
