---
id: 587d7db3367417b2b2512b8e
title: Koristeći testni način
challengeType: 1
forumTopicId: 301369
dashedName: using-the-test-method
---

# --description--

Redovni izrazi se koriste u programskim jezicima za poređenje delova niza karaktera. Kreiraju obrasce koji pomažu u izvršavanju tog poređenja.

Ako želite pronaći reč `the` u sekvenci slova `The dog chased the cat`, možete koristiti sledeći regularni izraz: `/the/`. Proverite da znaci navodnika nisu potrebni unutar regularnog izraza.

JavaScript ima mnogo načina za korišćenje regularnih izraza. Jedan način testiranja regularnog izraza je korišćenje metode `.test()`. Metoda `.test()` uzima regularni izraz, primenjuje ga na niz znakova (koji je unutar običnih zagrada) i vraća `true` ili `false` ako vaš sistem pronađe nešto ili ne.

```js
let testStr = "freeCodeCamp";
let testRegex = /Code/;
testRegex.test(testStr);
```

Puta za `test` ovde vraća `true`.

# --instructions--

Koristi uobičajeni izraz `myRegex` u nizu karaktera `myString` koristeći metod `.test()`.

# --hints--

Trebalo bi da koristiš `.test()` za testiranje standardnog govora.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

Vaši rezultati treba da vrate `true`.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex; // Change this line
```

# --solutions--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex.test(myString); // Change this line
```
