---
id: 587d7db8367417b2b2512ba0
title: Pronađi sve osim slova i brojeva
challengeType: 1
forumTopicId: 301353
dashedName: match-everything-but-letters-and-numbers
---

# --description--

Naučio si da pomoću `\w` pronađeš slova, brojeve i donju crtu (`[A-Za-z0-9_]`). Često ćeš želeti da pronađeš upravo suprotno.

Suprotnost od `\w` je `\W`. Obrati pažnju da se za negirani obrazac koristi veliko slovo. Ova prečica je isto što i `[^A-Za-z0-9_]`.

```js
let shortHand = /\W/;
let numbers = "42%";
let sentence = "Coding!";
numbers.match(shortHand);
sentence.match(shortHand);
```

Prvi poziv metode `match` treba da vrati `["%"]`, a drugi `["!"]`.

# --instructions--

Upotrebi prečicu `\W` da izbrojiš broj znakova koji nisu slova ni brojevi u različitim nizovima.

# --hints--

Regularni izraz treba da koristi globalnu zastavicu.

```js
assert(nonAlphabetRegex.global);
```

Regularni izraz treba da pronađe 6 znakova koji nisu slova ni brojevi u nizu `The five boxing wizards jump quickly.`

```js
assert(
  'The five boxing wizards jump quickly.'.match(nonAlphabetRegex).length == 6
);
```

Vaša standardna upotreba bi trebalo da koristi klasu kratkog puta za poređenje znakova koji nisu slova i brojeve.

```js
assert(/\\W/.test(nonAlphabetRegex.source));
```

Vaše standardno korisničko ime mora imati 8 znakova koji nisu slova i brojeve u nizu karaktera `Pack my box with five dozen liquor jugs.`

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(nonAlphabetRegex).length == 8
);
```

Vaše standardno korisničko ime treba da sadrži 6 znakova koji nisu slova ni brojeve u nizu znaka `How vexingly quick daft zebras jump!`

```js
assert(
  'How vexingly quick daft zebras jump!'.match(nonAlphabetRegex).length == 6
);
```

Vaš standardni unos mora da sadrži 12 alfanumeričkih znakova u nizu `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.`

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(nonAlphabetRegex)
    .length == 12
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let nonAlphabetRegex = /change/; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards_jump quickly.";
let nonAlphabetRegex = /\W/g; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```
