---
id: 587d7db7367417b2b2512b9c
title: TraŇĺi jednog ili viŇ°e kriminalaca u prańáenju
challengeType: 1
forumTopicId: 301343
dashedName: find-one-or-more-criminals-in-a-hunt
---

# --description--

Vreme je da malo staniŇ° i probaŇ° svoju novu veŇ°tinu pisanja uobińćajenih fraza. Grupa kriminalaca pobegla je iz zatvora i trńćala, ali ne znaŇ° koliko ih je bilo. MeńĎutim, znaŇ° da ostaju blizu zajedno kada su u blizini drugih ljudi. Ti si odgovoran za pronalaŇĺenje svih kriminalaca odjednom.

Evo primera kako to moŇĺeŇ° da proveriŇ°:

Standardni uzorak ``/z+/`` odgovara slovu ``z`` kada se pojavljuje jednom ili viŇ°e puta zaredom. Pronańáet ńáete podudarnosti u svim sledeńáim sekvencama:

```js
"z"
"zzzzzz"
"ABCzzzz"
"zzzzABC"
"abczzzzzzzzzzzzzzzzzzzzzabc"
```

MeńĎutim, neńáeŇ° dobiti podudaranje u sledeńáoj niski zato Ň°to ne sadrŇĺi slovo `z`.

```js
""
"ABC"
"abcabc"
```

# --instructions--

NapiŇ°ite uobińćajenu izjavu tipa ‚Äěgreedy‚Äú koja traŇĺi jednog ili viŇ°e krivaca unutar grupe drugih ljudi. Krivac je predstavljen velikim slovom `C`.

# --hints--

VaŇ°a uobińćajena upotreba bi trebalo da uporedi jednog kriminalca (`C`) u nizu `C`

```js
assert('C'.match(reCriminals) && 'C'.match(reCriminals)[0] == 'C');
```

Regularni izraz treba da pronańĎe dva kriminalca (`CC`) u nizu `CC`

```js
assert('CC'.match(reCriminals) && 'CC'.match(reCriminals)[0] == 'CC');
```

Regularni izraz treba da pronańĎe tri kriminalca (`CCC`) u nizu `P1P5P4CCCcP2P6P3`.

```js
assert(
  'P1P5P4CCCcP2P6P3'.match(reCriminals) &&
    'P1P5P4CCCcP2P6P3'.match(reCriminals)[0] == 'CCC'
);
```

Regularni izraz treba da pronańĎe pet kriminalaca (`CCCCC`) u nizu `P6P2P7P4P5CCCCCP3P1`

```js
assert(
  'P6P2P7P4P5CCCCCP3P1'.match(reCriminals) &&
    'P6P2P7P4P5CCCCCP3P1'.match(reCriminals)[0] == 'CCCCC'
);
```

Regularni izraz ne treba da pronańĎe nijednog kriminalca u praznom nizu `""`

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test(''));
```

Regularni izraz ne treba da pronańĎe nijednog kriminalca u nizu `P1P2P3`

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test('P1P2P3'));
```

Standardna upotreba treba da uporedi pedeset kriminalaca (`CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC`) u sekvenci `P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3`.

```js
assert(
  'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
    reCriminals
  ) &&
    'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
      reCriminals
    )[0] == 'CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC'
);
```

# --seed--

## --seed-contents--

```js
let reCriminals = /./; // Change this line
```

# --solutions--

```js
let reCriminals = /C+/; // Change this line
```
