---
id: 587d7db4367417b2b2512b90
title: Povežite sekvencu stvarnih i različitih mogućnosti slova
challengeType: 1
forumTopicId: 301345
dashedName: match-a-literal-string-with-different-possibilities
---

# --description--

Korišćenjem regularnog izraza kao što je `/coding/`, možete pronaći obrazac `coding` u drugoj sekvenci znakova.

Ovo je korisno za pronalaženje jednog obrasca, ali je ograničeno na samo jedan obrazac. Više obrazaca možeš da pronađeš pomoću operatora alternacije (`|`), odnosno operatora OR.

Ovaj operater poredi strukture pre ili posle njega. Na primer, ako želite da uporedite sekvencu karaktera `yes` ili `no`, zahtevani standardni izraz je `/yes|no/`.

Takođe možete pretraživati više od samo dve strukture. To možete učiniti dodavanjem više struktura i operatera `OR` između njih, kao što je `/yes|no|maybe/`.

# --instructions--

Popuni uobičajenu frazu `petRegex` ili upari ljubimce `dog`, `cat`, `bird`, ili `fish`.

# --hints--

Vaš uobičajeni unos `petRegex` treba da vrati `true` za sekvencu slova `John has a pet dog.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('John has a pet dog.'));
```

Vaša uobičajena upotreba `petRegex` trebalo bi da vrati `false` za sekvencu slova `Emma has a pet rock.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Emma has a pet rock.'));
```

Vaša standardna upotreba `petRegex` bi trebalo da vrati `true` za sekvencu karaktera `Emma has a pet bird.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Emma has a pet bird.'));
```

Regularni izraz `petRegex` treba da vrati `true` za niz `Liz has a pet cat.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Liz has a pet cat.'));
```

Vaš standardni unos `petRegex` bi trebalo da vrati `false` za niz slova `Kara has a pet dolphin.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Kara has a pet dolphin.'));
```

Tvoj uobičajeni `petRegex` bi trebalo da vrati `true` za niz karaktera `Alice has a pet fish.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Alice has a pet fish.'));
```

Vaš standardni unos `petRegex` bi trebalo da vrati `false` za sekvencu slova `Jimmy has a pet computer.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Jimmy has a pet computer.'));
```

# --seed--

## --seed-contents--

```js
let petString = "James has a pet cat.";
let petRegex = /change/; // Change this line
let result = petRegex.test(petString);
```

# --solutions--

```js
let petString = "James has a pet cat.";
let petRegex = /dog|cat|bird|fish/; // Change this line
let result = petRegex.test(petString);
```
