---
id: 587d7db4367417b2b2512b92
title: Prikaži slične utakmice
challengeType: 1
forumTopicId: 301340
dashedName: extract-matches
---

# --description--

Sada si samo proveravao da li uzorak postoji ili ne postoji unutar niza karaktera. Takođe možete izvući stvarna podudaranja koja ste pronašli koristeći metodu `.match()`.

Da biste koristili metodu `.match()`, koristite tu metodu u nizu karaktera i unesite regularni izraz unutar običnih zagrada.

Evo primera:

```js
"Hello, World!".match(/Hello/);
let ourStr = "Regular expressions";
let ourRegex = /expressions/;
ourStr.match(ourRegex);
```

Ovde `match` prvi će vratiti `["Hello"]` i drugi će vratiti `["expressions"]`.

Imaj na umu da je sintaksa metode `.match()` „obrnuta“ u odnosu na metodu `.test()` koju si do sada koristio.

```js
'string'.match(/regex/);
/regex/.test('string');
```

# --instructions--

Koristi putanju `.match()` da izvuče niz znakova `coding`.

# --hints--

`result` mora imati sekvencu slova `coding`

```js
assert(result.join() === 'coding');
```

Vaš standardni unos `codingRegex` treba da traži sekvencu slova `coding`

```js
assert(codingRegex.source === 'coding');
```

Treba da koristiš metodu `.match()`.

```js
assert(__helpers.removeJSComments(code).match(/\.match\(.*\)/));
```

# --seed--

## --seed-contents--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /change/; // Change this line
let result = extractStr; // Change this line
```

# --solutions--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /coding/; // Change this line
let result = extractStr.match(codingRegex); // Change this line
```
