---
id: 587d7db9367417b2b2512ba7
title: Navedite stvarni broj mečeva
challengeType: 1
forumTopicId: 301365
dashedName: specify-exact-number-of-matches
---

# --description--

Možete definisati minimalni i maksimalni opseg modela koristeći indikatore raspona sa vitičnim zagradama. Ponekad želite samo određeni broj slučajeva.

Da biste naveli određen broj sistema, stavite taj broj između vitičastih zagrada.

Na primer, da bi odgovarao reči `hah` koja sadrži slova `a` `3` puta, vaše standardno korišćenje će biti `/ha{3}h/`.

```js
let A4 = "haaaah";
let A3 = "haaah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleHA = /ha{3}h/;
multipleHA.test(A4);
multipleHA.test(A3);
multipleHA.test(A100);
```

Prema podešavanjima, tri poziva za `test` će preusmeriti `false`, `true`, na `false`.

# --instructions--

Promeni uobičajen izraz `timRegex` da odgovara reči `Timber` samo kada ima četiri slova od `m`.

# --hints--

U vaše uobičajeno korišćenje treba koristiti kvastri kurvete/zagrade.

```js
assert(timRegex.source.match(/{.*?}/).length > 0);
```

Vaš uobičajeni pristup ne bi trebalo da odgovara nizu karaktera `Timber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timber'));
```

Vaš uobičajeni tekst ne bi smega da odgovara nizu karaktera `Timmber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmber'));
```

Vaš uobičajeni korisnik ne bi trebalo da odgovara nizu karaktera `Timmmber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmmber'));
```

Vaše standardno korisničko ime mora da odgovara nizu znakova `Timmmmber`

```js
timRegex.lastIndex = 0;
assert(timRegex.test('Timmmmber'));
```

Vaš standardni tekst ne sme da odgovara sekvenci karaktera `Timber` koja sadrži 30 slova od `m` unutar njega.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Ti' + 'm'.repeat(30) + 'ber'));
```

# --seed--

## --seed-contents--

```js
let timStr = "Timmmmber";
let timRegex = /change/; // Change this line
let result = timRegex.test(timStr);
```

# --solutions--

```js
let timStr = "Timmmmber";
let timRegex = /Tim{4}ber/; // Change this line
let result = timRegex.test(timStr);
```
