---
id: 587d7db4367417b2b2512b91
title: Puuza Mala ili Velika Slova Tokom Poređenja
challengeType: 1
forumTopicId: 301344
dashedName: ignore-case-while-matching
---

# --description--

Do sada smo videli standardan metod za poređenje sekvenci karaktera na nivou pojedinačnih znakova. Ali ponekad, možda želite da uporedite razlike između malih i velikih slova.

Slovo (ili ponekad nazivan tip slova) je razlika između velikih i malih slova. Primeri velikih slova su `A`, `B`, i `C`. Primeri malih slova su `a`, `b`, i `c`.

Možete uporediti oba slova koristeći ono što se naziva zastava (flag). Postoje druge zastave, ali ovde ćete razmatrati zastavu koja ignoriše mala ili velika slova – zastavu `i`. Možete je koristiti dodavanjem na običan izraz. Primer korišćenja ove zastave je `/ignorecase/i`. Ovaj regularni izraz može da uporedi niz slova `ignorecase`, `igNoreCase` i `IgnoreCase`.

# --instructions--

Napiši regularni izraz `fccRegex` koji pronalazi `freeCodeCamp`, bez obzira na velika i mala slova. Regularni izraz ne treba da pronalazi skraćenice niti varijante sa razmacima.

# --hints--

Vaša uobičajena upotreba mora da odgovara sekvenci slova `freeCodeCamp`

```js
fccRegex.lastIndex = 0; 
assert(fccRegex.test('freeCodeCamp'));
```

Vaš uobičajeni unos bi trebalo da odgovara nizu znakova `FreeCodeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodeCamp'));
```

Vaša uobičajena upotreba bi trebalo da odgovara nizu slova `FreecodeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreecodeCamp'));
```

Vaš standardni unos bi trebalo da odgovara nizu karaktera `FreeCodecamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodecamp'));
```

Vaš uobičajeni unos ne sme da odgovara sekvenci znakova `Free Code Camp`

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('Free Code Camp'));
```

Regularni izraz treba da pronađe niz `FreeCOdeCamp`.

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCOdeCamp'));
```

Vaš uobičajeni unos ne sme da odgovara nizu znakova `FCC`

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('FCC'));
```

Vaš standardni unos bi trebalo da odgovara nizu znakova `FrEeCoDeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCoDeCamp'));
```

Vaše uobičajeno korisničko ime treba da odgovara nizu znakova `FrEeCodECamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCodECamp'));
```

Vaš uobičajeni unos bi trebalo da odgovara sekvenci karaktera `FReeCodeCAmp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FReeCodeCAmp'));
```

# --seed--

## --seed-contents--

```js
let myString = "freeCodeCamp";
let fccRegex = /change/; // Change this line
let result = fccRegex.test(myString);
```

# --solutions--

```js
let myString = "freeCodeCamp";
let fccRegex = /freecodecamp/i; // Change this line
let result = fccRegex.test(myString);
```
