---
id: 587d7db8367417b2b2512ba1
title: Uporedite sve nebrojkovne karaktere
challengeType: 1
forumTopicId: 301347
dashedName: match-all-non-numbers
---

# --description--

Poslednji izazov je pokazao kako da pronađeš brojeve pomoću prečice `\d`. Možeš pronaći i znakove koji nisu brojevi pomoću slične prečice sa velikim slovom `D`.

Prečica za pronalaženje znakova koji nisu brojevi je `\D`. To je isto što i klasa znakova `[^0-9]`, koja pronalazi bilo koji znak koji nije cifra od 0 do 9.

# --instructions--

Koristi kratku klasu objekata sa slovima koje nisu brojevi `\D` da izračunaš koliko je slova koja nisu brojevi u nazivima filmova.

# --hints--

Regularni izraz treba da koristi prečicu za pronalaženje znakova koji nisu brojevi.

```js
assert(/\\D/.test(noNumRegex.source));
```

Regularni izraz treba da koristi globalnu zastavicu.

```js
assert(noNumRegex.global);
```

Regularni izraz ne treba da pronađe nijedan znak koji nije broj u nizu `9`.

```js
assert('9'.match(noNumRegex) == null);
```

Regularni izraz treba da pronađe 6 znakova koji nisu brojevi u nizu `Catch 22`.

```js
assert('Catch 22'.match(noNumRegex).length == 6);
```

Regularni izraz treba da pronađe 11 znakova koji nisu brojevi u nizu `101 Dalmatians`.

```js
assert('101 Dalmatians'.match(noNumRegex).length == 11);
```

Regularni izraz treba da pronađe 15 znakova koji nisu brojevi u nizu `One, Two, Three`.

```js
assert('One, Two, Three'.match(noNumRegex).length == 15);
```

Regularni izraz treba da pronađe 12 znakova koji nisu brojevi u nizu `21 Jump Street`.

```js
assert('21 Jump Street'.match(noNumRegex).length == 12);
```

Regularni izraz treba da pronađe 17 znakova koji nisu brojevi u nizu `2001: A Space Odyssey`.

```js
assert('2001: A Space Odyssey'.match(noNumRegex).length == 17);
```

# --seed--

## --seed-contents--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /change/; // Change this line
let result = movieName.match(noNumRegex).length;
```

# --solutions--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /\D/g; // Change this line
let result = movieName.match(noNumRegex).length;
```
