---
id: 587d7db5367417b2b2512b94
title: Uporedi bilo ЕЎta sa wildkardom za period
challengeType: 1
forumTopicId: 301348
dashedName: match-anything-with-wildcard-period
---

# --description--

Ponekad ne znaЕЎ (ili nema potrebe da znaЕЎ) taДЌne karaktere u svojim sistemima. RazmiЕЎljanje o svim poklapanjima, recimo, sa greЕЎkama u kucanju Д‡e potrajati dugo. SreД‡om, moЕѕeЕЎ uЕЎtedeti vreme koriЕЎД‡enjem wildkarta: `.`

Wildcard karakter `.` poklapa se sa bilo kojim pojedinaДЌnim karakterom. Wildcard se takoД‘e naziva `dot` i `period`. MoЕѕete koristiti wildcard karakter kao bilo koji drugi karakter u regularnom izrazu. Na primer, ako Еѕelite da poklopite `hug`, `huh`, `hut` i `hum`, moЕѕete koristiti regularni izraz `/hu./` da poklopite sva ДЌetiri elementa.

```js
let humStr = "I'll hum a song";
let hugStr = "Bear hug";
let huRegex = /hu./;
huRegex.test(humStr);
huRegex.test(hugStr);
```

Oba poziva od `test` Д‡e vratiti `true`.

# --instructions--

DovrЕЎi regularni izraz `unRegex` tako da pronaД‘e nizove `run`, `sun`, `fun`, `pun`, `nun` i `bun`. Regularni izraz treba da koristi dЕѕoker znak.

# --hints--

Treba da koristiЕЎ putanju `.test()`.

```js
assert(__helpers.removeJSComments(code).match(/\.test\(.*\)/));
```

Trebalo bi da koristiЕЎ znak za wildkard u svom obiДЌnom tekstu `unRegex`

```js
assert(/\./.test(unRegex.source));
```

Regularni izraz `unRegex` treba da pronaД‘e `run` u nizu `Let us go on a run.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Let us go on a run.'));
```

VaЕЎa uobiДЌajena vrednost `unRegex` mora da odgovara `sun` u nizu znakova `The sun is out today.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('The sun is out today.'));
```

VaЕЎa uobiДЌajena upotreba `unRegex` trebalo bi da odgovara `fun` u sekvenci karaktera `Coding is a lot of fun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Coding is a lot of fun.'));
```

VaЕЎe uobiДЌajeno koriЕЎД‡enje `unRegex` mora da se poklapa sa `pun` u sekvenci slova `Seven days without a pun makes one weak.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Seven days without a pun makes one weak.'));
```

VaЕЎa uobiДЌajena vrednost `unRegex` mora odgovarati `nun` u nizu slova `One takes a vow to be a nun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('One takes a vow to be a nun.'));
```

VaЕЎ standardni `unRegex` mora da odgovara `bun` u nizu karaktera `She got fired from the hot dog stand for putting her hair in a bun.`

```js
unRegex.lastIndex = 0;
assert(
  unRegex.test(
    'She got fired from the hot dog stand for putting her hair in a bun.'
  )
);
```

Standardno koriЕЎД‡enje `unRegex` ne bi trebalo da uporeД‘uje sekvencu slova `There is a bug in my code.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('There is a bug in my code.'));
```

VaЕЎe uobiДЌajeno koriЕЎД‡enje `unRegex` ne bi trebalo da uporedi niz znakova `Catch me if you can.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('Catch me if you can.'));
```

# --seed--

## --seed-contents--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /change/; // Change this line
let result = unRegex.test(exampleStr);
```

# --solutions--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /.un/; // Change this line
let result = unRegex.test(exampleStr);
```
