---
id: 587d7dbb367417b2b2512bac
title: Uklonite nepotrebne razmake u početku i na kraju
challengeType: 1
forumTopicId: 301362
dashedName: remove-whitespace-from-start-and-end
---

# --description--

Ponekad postoje nepotrebni razmaci u blizini niza karaktera, iako nisu neophodni. Standardna procedura za obradu niza je uklanjanje tih razmaka sa samog početka i kraja niza.

# --instructions--

Napišite regularni izraz i koristite odgovarajuće metode za nizove karaktera kako biste uklonili razmake sa početka i kraja niza.

**Napomena:** Metoda `String.prototype.trim()` bi mogla da radi ovde, ali će vam trebati da završite ovaj izazov koristeći običan jezik.

# --hints--

`result` treba da bude jednak nizu `Hello, World!`

```js
assert(result === 'Hello, World!');
```

Vaše rešenje ne bi trebalo da koristi metodu `String.prototype.trim()`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?trim/));
```

Promenljiva `result` ne treba da bude direktno postavljena na dati niz.

```js
assert(!__helpers.removeJSComments(code).match(/result\s*=\s*["'`].*?["'`]/));
```

Vrednost parametra `hello` ne smije biti izmjenjena.

```js
assert(hello === '   Hello, World!  ');
```

# --seed--

## --seed-contents--

```js
let hello = "   Hello, World!  ";
let wsRegex = /change/; // Change this line
let result = hello; // Change this line
```

# --solutions--

```js
let hello = "   Hello, World!  ";
let wsRegex = /^(\s+)(.+[^\s])(\s+)$/;
let result = hello.replace(wsRegex, '$2');
```
