---
id: 587d7db7367417b2b2512b9e
title: Uporediti sisteme sekvenci slova koje se završavaju
challengeType: 1
forumTopicId: 301352
dashedName: match-ending-string-patterns
---

# --description--

U prethodnom izazovu naučio si da koristiš znak kareta za pronalaženje obrazaca na početku niza. Postoji i način da pronađeš obrasce na kraju niza.

Možete pronaći na kraju sekvence slova koristeći znak dolara `$` na kraju regulare izraza.

```js
let theEnding = "This is a never ending story";
let storyRegex = /story$/;
storyRegex.test(theEnding);
let noEnding = "Sometimes a story will have to end";
storyRegex.test(noEnding);
```

Prvi poziv `test` će vratiti `true`, dok drugi poziv vraća `false`.

# --instructions--

Upotrebi znak sidra (`$`) da pronađeš niz `caboose` na kraju niza `caboose`.

# --hints--

Trebaće vam da pronađete `caboose` koristeći znak dolara `$` kao sidreni tag u vašem običnom tekstu.

```js
assert(lastRegex.source == 'caboose$');
```

Regularni izraz ne treba da koristi nijednu zastavicu.

```js
assert(lastRegex.flags == '');
```

Treba da pronađeš `caboose` na kraju niza `The last car on a train is the caboose`.

```js
lastRegex.lastIndex = 0;
assert(lastRegex.test('The last car on a train is the caboose'));
```

# --seed--

## --seed-contents--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /change/; // Change this line
let result = lastRegex.test(caboose);
```

# --solutions--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /caboose$/; // Change this line
let result = lastRegex.test(caboose);
```
