---
id: 587d7db9367417b2b2512ba4
title: Uporedi slova bez razmaka
challengeType: 1
forumTopicId: 18210
dashedName: match-non-whitespace-characters
---

# --description--

Nauči kako da tražiš razmak koristeći `\s`, za mala slova `s`. Takođe možeš da pretražuješ sve osim razmaka.

Traži znakove koji nisu razmaci koristeći `\S`, koji su velika slova `s`. Ovaj primer neće poklopiti razmake, povratak karriđere, tabulator, prazan formular i znakove novog reda. Možeš ga smatrati klasom objekata `[^ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let nonSpaceRegex = /\S/g;
whiteSpace.match(nonSpaceRegex).length;
```

Vrednost koja se vraća putem `.length` biće `32`.

# --instructions--

Promeni uobičajen izraz `countNonWhiteSpace` ili pronađi više razdvojenih slova u nizu slova.

# --hints--

Vaše standardno korišćenje bi trebalo da koristi svetsku zastavu.

```js
assert(countNonWhiteSpace.global);
```

Vaša uobičajena upotreba bi trebalo da koristi kratke karaktere `\S` za poređenje svih neprostripovnih karaktera.

```js
assert(/\\S/.test(countNonWhiteSpace.source));
```

Vaša uobičajena upotreba trebalo bi da sadrži 35 znakova bez razmaka u nizu karaktera `Men are from Mars and women are from Venus.`

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countNonWhiteSpace)
    .length == 35
);
```

Vaše uobičajeno korišćenje bi trebalo da ima 23 znaka bez razmaka u nizu slova `Space: the final frontier.`

```js
assert('Space: the final frontier.'.match(countNonWhiteSpace).length == 23);
```

Vaše standardno korišćenje bi trebalo da primi 21 znak bez razmaka u nizu karaktera `MindYourPersonalSpace`

```js
assert('MindYourPersonalSpace'.match(countNonWhiteSpace).length == 21);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /change/; // Change this line
let result = sample.match(countNonWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /\S/g; // Change this line
let result = sample.match(countNonWhiteSpace);
```
