---
id: 5c3dda8b4d8df89bea71600f
title: Pogledajte kombinaciju grupa slova
challengeType: 1
forumTopicId: 301339
dashedName: check-for-mixed-grouping-of-characters
---

# --description--

Ponekad želimo da pregledamo grupe karaktera koristeći redovne izraze i za to koristimo standardne zagrade `()`.

Ako želite da pronađete bilo `Penguin` ili `Pumpkin` u sekvenci slova, možete koristiti sledeći regularni izraz: `/P(engu|umpk)in/g`

Zatim proverite da li se grupe sekvenci karaktera koje želite nalaze u test sekvenciji koristeći `test()`.

```js
let testStr = "Pumpkin";
let testRegex = /P(engu|umpk)in/;
testRegex.test(testStr);
```

Njia ya `test` hapa itarudisha `true`.

# --instructions--

Izmenite regularni izraz da odgovara imena `Franklin Roosevelt` ili `Eleanor Roosevelt` na način koji prepoznaje velika i mala slova, te takođe dozvoljava srednja imena.

Zatim izmeni kod tako da se regularni izraz koji si napravio proveri nad `myString` i vrati `true` ili `false`, u zavisnosti od toga da li postoji podudaranje.

# --hints--

Vaš uobičajeni unos `myRegex` mora da vrati `true` za niz slova `Franklin D. Roosevelt`

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Franklin D. Roosevelt'));
```

Tvoj uobičajeni `myRegex` treba da vrati `true` niz znakova `Eleanor Roosevelt`

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Eleanor Roosevelt'));
```

Vaša uobičajena upotreba `myRegex` trebalo bi da vrati `false` za sekvencu slova `Franklin Rosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Franklin Rosevelt'));
```

Vaš uobičajeni unos `myRegex` treba da vrati `false` za sekvencu karaktera `Frank Roosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Frank Roosevelt'));
```

Regularni izraz `myRegex` treba da vrati `false` za niz `FranklinRoosevelt`.

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('FranklinRoosevelt'));
```

Vaš uobičajeni unos `myRegex` treba da vrati `false` za sekvencu slova `EleanorRoosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('EleanorRoosevelt'));
```

Trebalo bi da koristiš `.test()` za testiranje standardnog govora.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

Vaši rezultati treba da vrate `true`.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /False/; // Change this line
let result = false; // Change this line
// After passing the challenge experiment with myString and see how the grouping works
```

# --solutions--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /(Franklin|Eleanor) (([A-Z]\.?|[A-Z][a-z]+) )?Roosevelt/;
let result = myRegex.test(myString);
```
