---
id: 698b42492a7f5f737895892e
title: Jenga Konverter Liste susjednosti za Matricu susjednosti
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

U ovom laboratoriju, kreiraćete funkciju koja konvertuje reprezentaciju liste susjednosti grafa u matricu susjednosti. Lista susjednosti je struktura čiji svaki ključ predstavlja čvor, a odgovarajuća vrijednost je niz podataka čvorova sa kojima je čvor tog ključa povezan. Matrica susjednosti je dvodimenzionalni niz podataka gdje je element na poziciji `[i][j]` jednak `1` ako postoji ivica od čvora `i` ka čvoru `j`, a `0` u suprotnom.

Na primer, ukoliko se dostavi spisak susjedstva:

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

Odgovarajuća matrica susjedstva biće:

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**Cilj:** Popunite korisničke priče ispod i prođite sve testove da biste završili ovu laboratoriju.

**Priče korisnika:**

1. Treba definisati funkciju pod nazivom `adjacencyListToMatrix` koja konvertuje listu susjedstva u matricu susjedstva.
2. Funkcija treba da primi nešto što predstavlja listu susjedstva neponderisanog grafa (bilo nedirekcionisan ili direkcionisan) kao svoj argument.
3. Funkcija treba da:
   - Konvertuje listu susjedstva u matricu susjedstva.
   - Ispisuje svaki red matrice susjedstva.
   - Vraća matricu susjedstva.

Na primer, `adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})` trebalo bi da ispisa:

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

da vrati `[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]`.

# --hints--

Trebalo bi da definišete funkciju nazvanu `adjacencyListToMatrix`.

```js
assert.isFunction(adjacencyListToMatrix);
```

Funkcija `adjacencyListToMatrix` treba da ima jedan parametar.

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

Funkcija mora tačno da utvrdi broj čvorova iz liste susjednosti.

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

Funkcija mora tačno da postavi vrednosti matrice na `1` za postojeće konektore.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Функција би требало да испише сваки ред матрице.

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

Funkcija mora da vrati matricu susednosti.

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

Kada se pruži lista susjeda `{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}`, funkcija bi trebalo da vrati `[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]`.

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Pri dohvaćanju liste susjeda `{0: [1], 1: [0]}`, funkcija bi trebalo da vrati `[[0, 1], [1, 0]]`.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Kada se dohvati lista susjedstava `{0: [], 1: [], 2: []}`, funkcija treba da vrati `[[0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
