---
id: 683ecb010f4c6bafd428802c
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovom radionici, validiraćete skup medicinskih podataka kako biste osigurali da su u skladu sa nizom pravila.

Deklarišite varijablu pod imenom `medical_records` i dodelite joj prazan spisak. U sledećim koracima, koristićete je za skladištenje vaših medicinskih podataka.

# --hints--

You should declare a variable named `medical_records`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

You should assign an empty list to your `medical_records` variable.

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

