---
id: 6842d5b537c64266420faf62
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Dodajte ključ `age` sa vrednošću celog broja `34` i ključ `gender` sa vrednošću niza karaktera `Female` u vaš rečnik. Nemojte zaboraviti zarez između parova ključ-vrednost.

# --hints--

Your `medical_records` list should contain a single dictionary.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

You should have a `age` key with the value of the integer `34` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

You should have a `gender` key with the value of the string `Female` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

You should have a `patient_id` key with the value of the string `P1001` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
