---
id: 6842ba0086201089ea268c92
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Prateći istu strukturu koju ste koristili u prethodnim koracima, lista `medical_records` je popunjena podacima drugih pacijenata za vas. Slobodno pogledajte to.

Sledeći korak je da počnete da pišete funkciju za validaciju skupa podataka. Kreirajte funkciju pod nazivom `validate` sa jednim parametrom `data`.

Želite da osigurate da su vaši podaci ili lista ili tupel. Stoga, unutar funkcije `validate`, deklarišite varijablu pod nazivom `is_sequence` i dodelite joj poziv `isinstance`. Prosledite `data` kao prvi argument, a tupel koji sadrži `list` i `tuple` kao drugi argument.

# --hints--

You should have a function named `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

Your `validate` function should have a single parameter `data`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

You should have a variable named `is_sequence` within your `validate` function.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

You should assign `isinstance(data, (list, tuple))` to your `is_sequence` variable.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--
```
