---
id: 6842d659a4ae7973dabab3b0
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Dopunite rečnik dodavanjem sledećih tri para ključ-vrednost:

- Ključ `diagnosis` sa vrednošću `Hypertension`
- Ključ `medications` sa vrednošću liste `['Lisinopril']`
- Ključ `last_visit_id` sa vrednošću `V2301`.

# --hints--

You should have a `diagnosis` key with the value of the string `Hypertension` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

You should have a `medications` key with the value of the list `['Lisinopril']` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

You should have a `last_visit_id` key with the value of the string `V2301` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

You should have a `patient_id` key with the value of the string `P1001` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

You should have a `age` key with the value of the integer `34` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

You should have a `gender` key with the value of the string `Female` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
