---
id: 6842aee5bcaf4ab8883f8191
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Kao što ste naučili u prethodnim lekcijama, rečnik je struktura podataka koja čuva parove ključ-vrednost, gde su ključevi nemenjivi tipovi podataka (kao što su nizovi karaktera) i jedinstveni unutar rečnika:

```py
person = {
    'name': 'John',
    'age': 33
}
```

Lista `medical_records` će čuvati rečnike, od kojih svaki predstavlja pacijenta. Dodajte rečnik sa ključem `patient_id` i vrednošću stringa `P1001` u listu `medical_records`.

# --hints--

Your `medical_records` list should contain a single dictionary.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

You should have a `patient_id` key with the value of the string `P1001` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
--fcc-editable-region--
    
--fcc-editable-region--
]
```
