---
id: 6849311be42a0431b8e59109
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Proveravate da li su podaci prosleđeni vašoj funkciji lista ili tupel. Još morate osigurati da je svaki element u sekvenci rečnik.

Unutar vaše petlje `for`, ako element u `dictionary` nije instanca od `dict`, štampajte `Invalid format: expected a dictionary at position <index>.` (gde se `<index>` treba zameniti trenutnim indeksom) i postavite `is_invalid` na `True`.

# --hints--

You should have an `if` statement inside your `for` loop.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0]
`) })
```

Your `if` statement should check if `dictionary` is not an instance of `dict`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(dictionary, dict)")
`) })
```

You should print `Invalid format: expected a dictionary at position <index>.` (where `<index>` should be replaced by the current index) inside your `if` statement.

```js
({ test: () => assert(runPython(`
built_in_print = print
out = []

def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
validate([{}, "", {}, [], 0])
print = built_in_print

_first = 'Invalid format: expected a dictionary at position 1.'
_second = 'Invalid format: expected a dictionary at position 3.'
_third = 'Invalid format: expected a dictionary at position 4.'
len(out) >= 3 and _first in out and _second in out and _third in out
`)) })
```

You should set `is_invalid` to `True` inside your `if` statement.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("is_invalid = True")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

    for index, dictionary in enumerate(data):
--fcc-editable-region--
        pass
--fcc-editable-region--
```
