---
id: 5895f70cf9fc0f352b528e66
title: Serijalizacija korisničkog objekta
challengeType: 2
forumTopicId: 301563
dashedName: serialization-of-a-user-object
---

# --description--

Serijalizacija i deserijalizacija su važni koncepti u vezi sa autentifikacijom. Serijalizacija objekta znači pretvaranje njegovog sadržaja u mali *ključ* koji se kasnije može deserijalizovati nazad u pravi objekat. To nam omogućava da znamo ko je komunicirao sa serverom bez slanja podataka za autentifikaciju, kao što su korisničko ime i lozinka, pri svakom zahtevu za novu stranicu.

Da bi ovo ispravno podesio, trebaš da imaš funkciju za serijalizaciju i funkciju za deserijalizaciju. U Passport-u, one se mogu kreirati sa:

```javascript
passport.serializeUser(cb);
passport.deserializeUser(cb);
```

Callback funkcija koja se prosleđuje `serializeUser` poziva se sa dva argumenta: ceo korisnički objekat, i callback koji koristi passport.

Callback očekuje dva argumenta: Grešku, ako postoji, i jedinstveni ključ za identifikaciju korisnika koji treba da bude vraćen u callback-u. Koristićeš `_id` korisnika u tom objektu. Ovo je garantovano jedinstveno, jer ga generiše MongoDB.

Slično tome, `deserializeUser` se poziva sa dva argumenta: jedinstveni ključ, i callback funkcija.

Ova callback funkcija očekuje dva argumenta: Grešku, ako postoji, i ceo korisnički objekat. Da bi dobio ceo korisnički objekat, napravi upit za pretragu po Mongo `_id`, kao što je prikazano ispod:

```javascript
passport.serializeUser((user, done) => {
  done(null, user._id);
});

passport.deserializeUser((id, done) => {
  myDataBase.findOne({ _id: new ObjectID(id) }, (err, doc) => {
    done(null, null);
  });
});
```

Dodaj obe gore navedene funkcije na svoj server. Klasa `ObjectID` dolazi iz `mongodb` paketa. `mongodb@~3.6.0` je već dodata kao zavisnost. Deklariši ovu klasu sa:

```javascript
const { ObjectID } = require('mongodb');
```

`deserializeUser` će bacati grešku dok ne uspostaviš konekciju sa bazom podataka. Za sada, iskomentariši `myDatabase.findOne` poziv, i samo pozovi `done(null, null)` u callback-u `deserializeUser`.

Pošalji svoju stranicu kada vidiš da si uradio dobro. Ako naiđeš na grešku, možeš <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#serialization-of-a-user-object-4" target="_blank" rel="noopener noreferrer nofollow">pogledati projekat urađen do ovog koraka</a>.

# --hints--

Treba da serijalizuješ korisnički objekat ispravno.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.serializeUser/gi,
    'You should have created your passport.serializeUser function'
  );
  assert.match(
    data,
    /null,\s*user._id/gi,
    'There should be a callback in your serializeUser with (null, user._id)'
  );
```

Treba da deserijalizuješ korisnički objekat ispravno.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.deserializeUser/gi,
    'You should have created your passport.deserializeUser function'
  );
  assert.match(
    data,
    /null,\s*null/gi,
    'There should be a callback in your deserializeUser with (null, null) for now'
  );
```

MongoDB treba da bude zavisnost.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'mongodb',
    'Your project should list "mongodb" as a dependency'
  );
```

MongoDB treba da bude zahtevan ispravno uključujući ObjectId.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')mongodb\1/gi,
    'You should have required mongodb'
  );
  assert.match(
    data,
    /new ObjectID.*id/gi,
    'Even though the block is commented out, you should use new ObjectID(id) for when we add the database'
  );