---
id: 5895f70cf9fc0f352b528e67
title: Implementiraj serijalizaciju Passport korisnika
challengeType: 2
forumTopicId: 301556
dashedName: implement-the-serialization-of-a-passport-user
---

# --description--

Nisi učitao pravi korisnički objekat jer baza podataka nije podešena. Poveži se sa bazom podataka odmah, kada pokreneš server, i sačuvaj trajnu konekciju za ceo životni ciklus aplikacije. Da bi ovo uradio, dodaj string konekcije tvoje baze podataka (na primer: `mongodb+srv://<username>:<password>@cluster0-jvwxi.mongodb.net/?retryWrites=true&w=majority`) u promenljivu okruženja `MONGO_URI`. Ovo se koristi u `connection.js` fajlu.

*Ako naiđeš na probleme sa postavljanjem besplatne baze podataka na MongoDB Atlas, pogledaj <a href="https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/" target="_blank" rel="noopener noreferrer nofollow">ovaj tutorijal</a>.*

Sada želiš da se povežeš sa svojom bazom podataka, zatim počni da slušaš zahteve. Svrha ovoga je da se ne dozvole zahtevi pre nego što se baza podataka poveže ili ako postoji greška u bazi podataka. Da bi ovo postigao, stavi svoj kod i rute aplikacije unutar sledećeg koda:

```javascript
myDB(async client => {
  const myDataBase = await client.db('database').collection('users');

  // Be sure to change the title
  app.route('/').get((req, res) => {
    // Change the response to render the Pug template
    res.render('index', {
      title: 'Connected to Database',
      message: 'Please login'
    });
  });

  // Serialization and deserialization here...

  // Be sure to add this...
}).catch(e => {
  app.route('/').get((req, res) => {
    res.render('index', { title: e, message: 'Unable to connect to database' });
  });
});
// app.listen out here...
```

Uveri se da si iskomentarisao `myDataBase` kod u `deserializeUser`, i izmeni svoj `done(null, null)` da uključi `doc`.

Pošalji svoju stranicu kada vidiš da si uradio dobro. Ako naiđeš na greške, možeš <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implement-the-serialization-of-a-passport-user-5" target="_blank" rel="noopener noreferrer nofollow">pogledati projekat urađen do ovog koraka</a>.

# --hints--

Konekcija sa bazom podataka treba da postoji.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Connected to Database/gi,
    'You successfully connected to the database!'
  );
```

Deserijalizacija sada treba da koristi bazu podataka ispravno i `done(null, null)` treba da bude pozvan sa `doc`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /null,\s*doc/gi,
    'The callback in deserializeUser of (null, null) should be altered to (null, doc)'
  );