---
id: 5895f70df9fc0f352b528e6a
title: Kreiraj novi middleware
challengeType: 2
forumTopicId: 301551
dashedName: create-new-middleware
---

# --description--

Kao što je sada, bilo koji korisnik može jednostavno otići na `/profile` bio autentifikovan ili ne kucanjem URL-a. Želiš da sprečiš ovo proveravanjem da li je korisnik prvo autentifikovan pre prikazivanja stranice profila. Ovo je dobar primer kada se kreira middleware.

Izazov ovde je kreirati middleware funkciju `ensureAuthenticated(req, res, next)`, koja će proveriti da li je korisnik autentifikovan pozivanjem Passport-ovog `isAuthenticated` metoda na `request` koji proverava da li je `req.user` definisan. Ako postoji, onda `next()` treba da bude pozvan. U suprotnom, možeš odgovoriti na zahtev preusmeravanjem na početnu stranicu za prijavu.

Implementacija ovog middleware-a je:

```javascript
function ensureAuthenticated(req, res, next) {
  if (req.isAuthenticated()) {
    return next();
  }
  res.redirect('/');
};
```

Kreiraj gore pomenutu middleware funkciju, zatim prosledi `ensureAuthenticated` kao middleware za zahteve stranice profila pre argumenta GET zahteva:

```javascript
app
  .route('/profile')
  .get(ensureAuthenticated, (req,res) => {
    res.render('profile');
  });
```

Pošalji svoju stranicu kada vidiš da si uradio dobro. Ako naiđeš na greške, možeš <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#create-new-middleware-8" target="_blank" rel="noopener noreferrer nofollow">pogledati projekat urađen do sada</a>.

# --hints--

Middleware `ensureAuthenticated` treba da bude implementiran i povezan sa `/profile` rutom.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /ensureAuthenticated[^]*req.isAuthenticated/,
    'Your ensureAuthenticated middleware should be defined and utilize the req.isAuthenticated function'
  );
  assert.match(
    data,
    /profile[^]*get[^]*ensureAuthenticated/,
    'Your ensureAuthenticated middleware should be attached to the /profile route'
  );
```

Neautentifikovani GET zahtev za `/profile` treba da bude ispravno preusmeren na `/`.

```js
  const url = new URL("/profile", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/,
    'An attempt to go to the profile at this point should redirect to the homepage since we are not logged in'
  );