---
id: 5895f70ef9fc0f352b528e6b
title: Kako sastaviti profil
challengeType: 2
forumTopicId: 301554
dashedName: how-to-put-a-profile-together
---

# --description--

Sada kada možeš da osiguraš da je korisnik koji dobija pristup `/profile` autentifikovan, možeš koristiti informacije sadržane u `req.user` na tvojoj stranici.

Prosledi objekat sa svojstvom `username` i vrednošću `req.user.username` kao drugi argument `render` metodu prikaza profila.

Zatim, idi na svoj `profile.pug` prikaz, i dodaj sledeću liniju ispod postojećeg `h1` elementa, sa istim nivoom uvlačenja:

```pug
h2.center#welcome Welcome, #{username}!
```

Ovo kreira `h2` element sa `center` klasom i `welcome` ID-jem sa tekstom `Welcome, ` praćenim korisničkim imenom.

Takođe, u `profile.pug`, dodaj link koji referencira `/logout` rutu, koja će biti domaćin logike za odjavu korisnika:

```pug
a(href='/logout') Logout
```

Pošalji svoju stranicu kada vidiš da si uradio dobro. Ako naiđeš na greške, možeš <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-put-a-profile-together-9" target="_blank" rel="noopener noreferrer nofollow">pogledati projekat urađen do ovog koraka</a>.

# --hints--

Treba da dodaš ispravno Pug varijablu prikaza u `/profile`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /username:( |)req.user.username/,
    'You should be passing the variable username with req.user.username into the render function of the profile page'
  );