---
id: 5895f70cf9fc0f352b528e65
title: Podesi Passport
challengeType: 2
forumTopicId: 301565
dashedName: set-up-passport
---

# --description--

Vreme je da podesiš *Passport* kako bi mogao da počneš da dozvoljavaš korisniku da se registruje ili prijavi na nalog. Pored Passport-a, koristićeš Express-session za upravljanje sesijama. Express-session ima mnoge napredne funkcije koje možeš koristiti, ali za sada ćeš koristiti samo osnovne. Korišćenje ovog middleware-a čuva identifikator sesije kao kolačić kod klijenta, i omogućava nam da dobijemo podatke sesije koristeći taj identifikator na serveru. Na ovaj način, čuvaš privatne informacije naloga van kolačića koji klijent koristi da kaže tvom serveru da su klijenti autentifikovani i čuvaš *ključ* za pristup podacima sačuvanim na serveru.

`passport@~0.4.1` i `express-session@~1.17.1` su već instalirani, i oba su navedena kao zavisnosti u tvom `package.json` fajlu.

Moraćeš da podesiš podešavanja sesije i inicijalizuješ Passport. Prvo, napravi varijable `session` i `passport` zahtevanjem `express-session` i `passport` respektivno.

Zatim, podesi svoju Express aplikaciju da koristi sesiju definisanjem sledećih opcija:

```javascript
app.use(session({
  secret: process.env.SESSION_SECRET,
  resave: true,
  saveUninitialized: true,
  cookie: { secure: false }
}));
```

Uveri se da dodaš `SESSION_SECRET` u tvoj `.env` fajl, i daj mu neočekivanu vrednost. Ovo se koristi za izračunavanje heša koji se koristi za zaštitu tvog kolačića!

Nakon što si sve to uradio, reci svojoj express aplikaciji da **koristi** `passport.initialize()` i `passport.session()`.

Pošalji svoju stranicu kada vidiš da si uradio dobro. Ako naiđeš na greške, možeš <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-passport-3" target="_blank" rel="noopener noreferrer nofollow">pogledati projekat koji je stigao do ovog koraka</a>.

# --hints--

Passport i Express-session treba da budu zavisnosti.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'passport',
    'Your project should list "passport" as a dependency'
  );
  assert.property(
    packJson.dependencies,
    'express-session',
    'Your project should list "express-session" as a dependency'
  );
```

Zavisnosti treba da budu zahtevane ispravno.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')passport("|')/,
    'You should have required passport'
  );
  assert.match(
    data,
    /require.*("|')express-session("|')/,
    'You should have required express-session'
  );
```

Express aplikacija treba da koristi nove zavisnosti.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /passport\.initialize/, 'Your express app should use "passport.initialize()"');
  assert.match(data, /passport\.session/, 'Your express app should use "passport.session()"');
```

Sesija i tajna sesije treba da budu konfigurisani ispravno.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /secret *:\s*process\.env(\.SESSION_SECRET|\[(?<q>"|')SESSION_SECRET\k<q>\])/,
    'Your express app should have express-session set up with your secret as process.env.SESSION_SECRET'
  );