---
id: 58965611f9fc0f352b528e6c
title: Odjava korisnika
challengeType: 2
forumTopicId: 301560
dashedName: logging-a-user-out
---

# --description--

Kreiranje logike za odjavu je jednostavno. Ruta treba samo da ukloni autentifikaciju korisnika, i preusmeri na početnu stranicu umesto da prikazuje bilo koji prikaz.

U passport-u, uklanjanje autentifikacije korisnika je jednostavno kao pozivanje `req.logout()` pre preusmeravanja. Dodaj `/logout` rutu da uradi to:

```js
app.route('/logout')
  .get((req, res) => {
    req.logout();
    res.redirect('/');
  });
```

Možda si primetio da nisi rukovao izgubljenim stranicama (404). Uobičajeni način rukovanja ovim u Node-u je korišćenjem sledećeg middleware-a. Nastavi i dodaj ovo nakon svih tvojih ruta:

```js
app.use((req, res, next) => {
  res.status(404)
    .type('text')
    .send('Not Found');
});
```

Pošalji svoju stranicu kada vidiš da si uradio dobro. Ako naiđeš na greške, možeš <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#logging-a-user-out-10" target="_blank" rel="noopener noreferrer nofollow">pogledati projekat koji je stigao do ovog koraka</a>.

# --hints--

`req.logout()` treba da bude pozvan u tvojoj `/logout` ruti.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /req.logout/gi,
    'You should be calling req.logout() in your /logout route'
  );
```

`/logout` treba da preusmeri na početnu stranicu.

```js
  const url = new URL("/logout", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/gi,
    'When a user logs out they should be redirected to the homepage'
  );