---
id: 58966a17f9fc0f352b528e6d
title: Registracija novih korisnika
challengeType: 2
forumTopicId: 301561
dashedName: registration-of-new-users
---

# --description--

Sada treba da dozvoliš novom korisniku na tvom sajtu da registruje nalog. U `res.render` za početnu stranicu dodaj novu varijablu u objekat koji se prosleđuje - `showRegistration: true`. Kada osvežiš svoju stranicu, trebalo bi da vidiš formu za registraciju koja je već kreirana u tvom `index.pug` fajlu. Ova forma je podešena da šalje **POST** na `/register`, pa napravi tu rutu i dodaj korisnički objekat u bazu podataka prateći sledeću logiku.

Logika rute za registraciju treba da bude sledeća:

1. Registruj novog korisnika
2. Autentifikuj novog korisnika
3. Preusmeri na `/profile`

Logika koraka 1 treba da bude sledeća:

1. Napravi upit za pretragu u bazi podataka sa `findOne`
2. Ako postoji greška, pozovi `next` sa tom greškom
3. Ako je korisnik vraćen, preusmeri nazad na početnu
4. Ako korisnik nije pronađen i nema grešaka, onda uradi `insertOne` u bazu podataka sa korisničkim imenom i lozinkom. Sve dok nema grešaka, pozovi `next` za korak 2, da autentifikuje novog korisnika, što si već napisao logiku u svojoj `POST /login` ruti.

```js
app.route('/register')
  .post((req, res, next) => {
    myDataBase.findOne({ username: req.body.username }, (err, user) => {
      if (err) {
        next(err);
      } else if (user) {
        res.redirect('/');
      } else {
        myDataBase.insertOne({
          username: req.body.username,
          password: req.body.password
        },
          (err, doc) => {
            if (err) {
              res.redirect('/');
            } else {
              // The inserted document is held within
              // the ops property of the doc
              next(null, doc.ops[0]);
            }
          }
        )
      }
    })
  },
    passport.authenticate('local', { failureRedirect: '/' }),
    (req, res, next) => {
      res.redirect('/profile');
    }
  );
```

Pošalji svoju stranicu kada vidiš da si uradio dobro. Ako naiđeš na greške, možeš <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#registration-of-new-users-11" target="_blank" rel="noopener noreferrer nofollow">pogledati projekat koji je stigao do ovog koraka</a>.

**NAPOMENA:** Odavde, problemi mogu nastati u vezi sa korišćenjem *picture-in-picture* pretraživača. Ako koristiš online IDE koji pruža pregled aplikacije unutar editora, preporučuje se da otvoriš taj pregled u novom tabu.

# --hints--

Treba da imaš `/register` rutu i da prikažeš formu za registraciju na početnoj stranici.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /showRegistration:( |)true/gi,
    'You should be passing the variable showRegistration as true to your render function for the homepage'
  );
  assert.match(
    data,
    /register[^]*post[^]*findOne[^]*username:( |)req.body.username/gi,
    'You should have a route that accepts a POST request on /register that queries the db with findOne and the query being username: req.body.username'
  );
```

Registracija treba da radi.

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttp = new XMLHttpRequest();
  xhttp.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      test(this);
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttp.open('POST', url + '/register', true);
  xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
  xhttp.send(`username=${user}&password=${user}`);
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Register should work, and redirect successfully to the profile.'
    );
  }
```

Prijava treba da radi.

```js
  const url = code;
  const user = `freeCodeCampTester${Date.now()}`;
  const xhttpReg = new XMLHttpRequest();
  xhttpReg.onreadystatechange = function () {
    if (this.readyState == 4 && this.status == 200) {
      login();
    } else {
      throw new Error(`${this.status} ${this.statusText}`);
    }
  };
  xhttpReg.open('POST', url + '/register', true);
  xhttpReg.setRequestHeader(
    'Content-type',
    'application/x-www-form-urlencoded'
  );
  xhttpReg.send(`username=${user}&password=${user}`);
  function login() {
    const xhttp = new XMLHttpRequest();
    xhttp.onreadystatechange = function () {
      if (this.readyState == 4 && this.status == 200) {
        test(this);
      } else {
        throw new Error(`${this.status} ${this.statusText}`);
      }
    };
    xhttp.open('POST', url + '/login', true);
    xhttp.setRequestHeader('Content-type', 'application/x-www-form-urlencoded');
    xhttp.send(`username=${user}&password=${user}`);
  }
  function test(xhttpRes) {
    const data = xhttpRes.responseText;
    assert.match(
      data,
      /Profile/gi,
      'Login should work if previous test was done successfully and redirect successfully to the profile.'
    );
    assert.match(
      data,
      new RegExp(user, 'g'),
      'The profile should properly display the welcome to the user logged in'
    );
  }
```

Odjava treba da radi.

```js
  $.ajax({
    url: code + '/logout',
    type: 'GET',
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(data, /Home/gi, 'Logout should redirect to home');
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );
```

Profil ne treba više da radi nakon odjave.

```js
  $.ajax({
    url: code + '/profile',
    type: 'GET',
    crossDomain: true,
    xhrFields: { withCredentials: true }
  }).then(
    (data) => {
      assert.match(
        data,
        /Home/gi,
        'Profile should redirect to home when we are logged out now again'
      );
    },
    (xhr) => {
      throw new Error(xhr.statusText);
    }
  );