---
id: 589fc830f9fc0f352b528e74
title: Podesi okruženje
challengeType: 2
forumTopicId: 301566
dashedName: set-up-the-environment
---

# --description--

Sledeći izazovi će koristiti `chat.pug` fajl. Zato, u tvom `routes.js` fajlu, dodaj GET rutu koja cilja `/chat` koja koristi `ensureAuthenticated`, i prikazuje `chat.pug`, sa `{ user: req.user }` prosleđenim kao argument za odgovor. Sada, izmeni tvoju postojeću rutu `/auth/github/callback` da postavi `req.session.user_id = req.user.id`, i preusmeri na `/chat`.

`socket.io@~2.3.0` je već dodat kao zavisnost, pa ga uvezi/instanciraj u svom serveru kao što sledi za `http` (dolazi unutar Node.js):

```javascript
const http = require('http').createServer(app);
const io = require('socket.io')(http);
```

Sada nakon što si http server povezao sa express app-om, treba da slušaš sa http servera. Izmeni liniju koja ima `app.listen` da bude `http.listen`.

Prva stvar koja treba da se upravlja je slušanje novog konekcije od klijenta. Ključna reč <dfn>on</dfn> radi tačno to - sluša specifičan događaj. Zahteva 2 argumenta: niz karaktera koji opisuje ime događaja koji se emitira, i funkcija kroz koju se prosleđuju podaci. U slučaju našeg konekcionog slušača, koristi `socket` za definisanje podataka u drugom argumentu. Socket je jedan klijent koji je konektovan.

Da bi slušao konekciju na tvom serveru, dodaj sledeće unutar tvoje baze podataka konekcije:

```javascript
io.on('connection', socket => {
  console.log('A user has connected');
});
```

Sada da bi klijent konektovan, trebaš samo da dodaš sledeće u `client.js` koji se učitava na stranici nakon autentifikacije:

```js
/*global io*/
let socket = io();
```

Komentari su tamo za sprečavanje grešaka koje obično vidite jer 'io' nije definisan u fajlu. Već si dodao pouzdan CDN za Socket.IO biblioteku na stranicu u `chat.pug`.

Sada probaj da učitaš svoju app i autentifikuj se i trebalo bi da vidiš u konzoli servera `A user has connected`.

**Napomena:** `io()` radi samo kada si konektovan na socket na istom url/serveru. Za konektovanje na spoljašnji socket na drugom mestu, koristićeš `io.connect('URL');`.

Pošalji svoju stranicu kada vidiš da si uradio dobro. Ako naiđeš na greške, možeš <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#set-up-the-environment-6" target="_blank" rel="noopener noreferrer nofollow">pogledati projekat urađen do sada</a>.

# --hints--

`socket.io` treba da bude zavisnost.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'socket.io',
    'Your project should list "socket.io" as a dependency'
  );
```

Treba da zahtevaš i instanciraš `http` kao `http`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /http.*=.*require.*('|")http\1/s,
    'Your project should list "http" as a dependency'
  );
```

Treba da zahtevaš i instanciraš `socket.io` kao `io`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.*=.*require.*('|")socket.io\1.*http/s,
    'You should correctly require and instantiate socket.io as io.'
  );
```

Socket.IO treba da sluša konekcije.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /io.on.*('|")connection\1.*socket/s,
    'io should listen for "connection" and socket should be the 2nd arguments variable'
  );
```

Tvoj klijent treba da se konektuje na server.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.*=.*io/s,
    'Your client should be connection to server with the connection defined as socket'
  );