---
id: 589fc832f9fc0f352b528e79
title: Pošalji i prikaži poruke chata
challengeType: 2
forumTopicId: 301562
dashedName: send-and-display-chat-messages
---

# --description--

Vreme je da počnete da dozvoljavate klijentima da šalju poruke chata na server da bi se prikazale svim klijentima! U tvom `client.js` fajlu, trebalo bi da već imaš deo koda koji rukuje kada se forma za poruku preda.

```js
$('form').submit(function() {
  /*logic*/
});
```

Unutar koda za predaju forme, trebaš da emituješ događaj nakon što definišeš `messageToSend` ali pre nego što očistiš polje za tekst `#m`. Taj događaj treba da se zove `'chat message'` i podaci trebaju da budu samo `messageToSend`.

```js
socket.emit('chat message', messageToSend);
```

Sada, na tvom serveru, trebaš da slušaš socket za događaj `'chat message'` sa podacima nazvanim `message`. Kada se događaj primi, treba da emituješ događaj `'chat message'` za sve sockete koristeći `io.emit`, šaljući objekat podataka sa `username` i `message`.

U `client.js`, sada trebaš da slušaš događaj `'chat message'` i, kada se primi, dodaj element liste na `#messages` sa korisničkim imenom, tačkom, i porukom!

U ovoj tački, chat treba da radi u potpunosti i šalje poruke svim klijentima!

Pošalji svoju stranicu kada vidiš da si uradio dobro. Ako naiđeš na greške, možeš <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#send-and-display-chat-messages-11" target="_blank" rel="noopener noreferrer nofollow">pogledati projekat urađen do sada</a>.

# --hints--

Server treba da sluša `'chat message'` i onda da ga emituje ispravno.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*io.emit.*('|")chat message('|").*username.*message/s,
    'Your server should listen to the socket for "chat message" then emit to all users "chat message" with name and message in the data object'
  );
```

Klijent treba da rukuje i prikaže nove podatke iz `'chat message'` događaja ispravno.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")chat message('|")[^]*messages.*li/s,
    'You should append a list item to #messages on your client within the "chat message" event listener to display the new message'
  );