---
id: 68dbc6b7e8f794a2ae69bc74
title: Algoritam Tekeleze za Malkija N
challengeType: 27
dashedName: implement-the-n-queens-algorithm
---

# --description--

Izazov sa N kraljica traži od tebe da postaviš N kraljica na šahovsku ploču N×N na način da se dve kraljice ne napadaju međusobno (da nijedna od dve ne deli red, kolonu ili dijagonalu).

Na primer, ako postoji daska od 4x4, jedna validna konfiguracija je:

```md
[1, 3, 0, 2]
```

Ovo znači da je kraljica postavljena u kolonu 1 u redu 0; u redu 1, kraljica je postavljena u kolonu 3; u redu 2, kraljica je postavljena u kolonu 0; i u redu 3, kraljica je postavljena u kolonu 2.

Iz vizuelne perspektive, ovaj raspored izgleda kao:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Gde `Q` predstavlja kraljicu, a `.` predstavlja prazno polje.

U ovoj laboratoriji, ćeš implementirati rešenje za izazov Kraljica N koristeći metodu dubinskog pretraživanja.

**Cilj**: Popunite korisničke priče ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Mora da postoji funkcija nazvana `dfs_n_queens`.
2. Funkcija mora da primi samo jedan argument: ceo broj `n`.
3. Ako je `n` manjiše od `1`, funkcija treba da vrati prazan niz (`[]`).
4. Funkcija treba da vrati listu rešenja; svako rešenje je lista dužine `n`, gde element na indeksu iz primjera `i` predstavlja kolonski indeks (počevši od 0) kraljice u redu `i`.

# --hints--

Trebalo bi da imaš funkciju nazvanu `dfs_n_queens` koja prima jedan argument.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs_n_queens)
  sig = inspect.signature(dfs_n_queens)
  assert len(sig.parameters) == 1
`) })
```

Ako je `n` manje od `1`, funkcija treba da vrati prazan spisak.

```js
({ test: () => runPython(`
  assert dfs_n_queens(0) == []
  assert dfs_n_queens(-1) == []
  assert dfs_n_queens(-5) == []
`) })
```

Funkcija treba da vrati listu rešenja, gde je svako rešenje lista dužine `n`.

```js
({ test: () => runPython(`
  result = dfs_n_queens(4)
  assert isinstance(result, list)
  for solution in result:
    assert isinstance(solution, list)
    assert len(solution) == 4
`) })
```

`dfs_n_queens(1)` trebalo bi da vrati `[[0]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(1) == [[0]]
`) })
```

`dfs_n_queens(2)` treba da vrati `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(2) == []
`) })
```

`dfs_n_queens(3)` treba vratiti `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(3) == []
`) })
```

`dfs_n_queens(4)` trebalo bi vratiti `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(4) == [[1, 3, 0, 2], [2, 0, 3, 1]]
`) })
```

`dfs_n_queens(5)` trebalo bi da vrati `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`len(dfs_n_queens(5))` bi trebalo da bude `10`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(5)) == 10
`) })
```

`len(dfs_n_queens(8))` treba da bude `92`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(8)) == 92
`) })
```

`dfs_n_queens` Trebalo bi da vrati tačne rezultate.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(7)) == 40
  assert dfs_n_queens(6) == [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs_n_queens(n):
    if n < 1:
        return []

    results = []

    def is_safe(queens, row, col):
        for r, c in enumerate(queens):
            if c == col or abs(row - r) == abs(col - c):
                return False
        return True

    def dfs(queens):
        row = len(queens)
        if row == n:
            results.append(queens[:])
            return
        for col in range(n):
            if is_safe(queens, row, col):
                queens.append(col)
                dfs(queens)
                queens.pop()

    dfs([])
    return results
```
