---
id: a103376db3ba46b2d50db289
title: Izmena sekvence slova u spinal case
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

Spinal case je struktura sekvence slova gde su sve reči napisane malim slovima i razdvojene crticama (-). `"this-is-spinal-tap"` je primer sekvence slova u strukturi spinal case.

U ovoj laboratoriji, kreiraćete funkciju koja modifikuje sekvencu slova unesenih u case.

**Cilj**: Završite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratoriju.

Priče o korisniku:

1. Treba da kreirate funkciju sa imenom `spinalCase`.
2. Funkcija `spinalCase` mora da primi jedan argument, niz znakova.
3. Funkcija `spinalCase` mora da vrati niz znakova u strukturu spinal case. Na primer, ako je argument `ProductLanding page`, funkcija bi trebalo da vrati `product-landing-page`.

# --hints--

Trebalo bi da kreirate funkciju pod nazivom `spinalCase`.

```js
assert.isFunction(spinalCase);
```

`spinalCase` Trebalo bi da primi samo jedan argument.

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")` treba da vrati niz karaktera `this-is-spinal-tap`.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` mora vratiti niz karaktera `this-is-spinal-tap`.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` treba vratiti sekvencu slova `the-andy-griffith-show`.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` mora da vrati sekvencu karaktera `teletubbies-say-eh-oh`.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` treba da vrati sekvencu karaktera `all-the-small-things`.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
