---
id: 56533eb9ac21ba0edf2244e2
title: Šifra Cezara
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16003
dashedName: caesars-cipher
---

# --description--

Jedan od <dfn>mismbo</dfn> jednostavan i veoma poznat je <dfn>šifra Cezara</dfn>, takođe poznata kao <dfn>šifra zamene</dfn>. U šifri zamene, značenja slova se menjaju po određenom nivo/stepenu.

Uobičajena moderna upotreba je kod <a href="https://www.freecodecamp.org/news/how-to-code-the-caesar-cipher-an-introduction-to-basic-encryption-3bf77b4e19f7/" target="_blank" rel="noopener noreferrer nofollow">ROT13</a>, gde se vrednosti slova pomeraju za 13 pozicija. Takođe `A ↔ N`, `B ↔ O` itd.

Napišite funkciju koja prima niz znakova kodiran sa <a href="https://www.freecodecamp.org/news/how-to-code-the-caesar-cipher-an-introduction-to-basic-encryption-3bf77b4e19f7/" target="_blank" rel="noopener noreferrer nofollow">ROT13</a> kao ulaz i vraća dekodovani niz znakova.

Sva slova će biti velika. Ne menjajte nealfabetičke znakove (tj. razmaci, interpunkcija), već ih propustite onako kako jesu.

# --hints--

`rot13("SERR PBQR PNZC")` trebalo bi da se prevede u niz znakova `FREE CODE CAMP`

```js
assert(rot13('SERR PBQR PNZC') === 'FREE CODE CAMP');
```

`rot13("SERR CVMMN!")` mora da se prevede u niz slova `FREE PIZZA!`

```js
assert(rot13('SERR CVMMN!') === 'FREE PIZZA!');
```

`rot13("SERR YBIR?")` treba biti prevedeno u niz slova `FREE LOVE?`

```js
assert(rot13('SERR YBIR?') === 'FREE LOVE?');
```

`rot13("GUR DHVPX OEBJA SBK WHZCF BIRE GUR YNML QBT.")` treba biti prevedeno u niz slova `THE QUICK BROWN FOX JUMPS OVER THE LAZY DOG.`

```js
assert(
  rot13('GUR DHVPX OEBJA SBK WHZCF BIRE GUR YNML QBT.') ===
    'THE QUICK BROWN FOX JUMPS OVER THE LAZY DOG.'
);
```

# --seed--

## --seed-contents--

```js
function rot13(str) {
  return str;
}

rot13("SERR PBQR PNZC");
```

# --solutions--

```js
var lookup = {
  'A': 'N','B': 'O','C': 'P','D': 'Q',
  'E': 'R','F': 'S','G': 'T','H': 'U',
  'I': 'V','J': 'W','K': 'X','L': 'Y',
  'M': 'Z','N': 'A','O': 'B','P': 'C',
  'Q': 'D','R': 'E','S': 'F','T': 'G',
  'U': 'H','V': 'I','W': 'J','X': 'K',
  'Y': 'L','Z': 'M'
};

function rot13(encodedStr) {
  var codeArr = encodedStr.split("");  // String to Array
  var decodedArr = []; // Your Result goes here
  // Only change code below this line

  decodedArr = codeArr.map(function(letter) {
    if(lookup.hasOwnProperty(letter)) {
      letter = lookup[letter];
    }
    return letter;
  });

  // Only change code above this line
  return decodedArr.join(""); // Array to String
}
```
