---
id: aaa48de84e1ecc7c742e1124
title: Provera palindrome
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16004
dashedName: palindrome-checker
---

# --description--

Vrati `true` ako je dati niz slova palindrom. Inače, vrati `false`.

<dfn>Palindrome</dfn> je reč ili rečenica koja se čita isto napred i nazad, bez obzira na interpunkciju, velika ili mala slova, i razmake.

**Napomena:** Moraćete da uklonite **sve znake koji nisu slova i brojeve** (interpunkcijski znaci, razmaci i simbole) i konvertujete sve u velika ili mala slova kako biste proverili palindrom.

Susretaćemo se sa nizom karaktera različitih struktura, kao `racecar`, `RaceCar`, i `race CAR` među ostalnim.

Pia će prebaciti niz slova sa posebnim znakovima, kao što su `2A3*3a2`, `2A3 3a2`, na `2_A3*3#A2`.

# --hints--

`palindrome("eye")` Trebalo bi da vrati vrednost tačno ili netačno.

```js
assert(typeof palindrome('eye') === 'boolean');
```

`palindrome("eye")` trebalo bi da vrati `true`.

```js
assert(palindrome('eye') === true);
```

`palindrome("_eye")` trebalo bi da vrati `true`.

```js
assert(palindrome('_eye') === true);
```

`palindrome("race car")` trebalo bi vratiti `true`.

```js
assert(palindrome('race car') === true);
```

`palindrome("not a palindrome")` treba da vrati `false`.

```js
assert(palindrome('not a palindrome') === false);
```

`palindrome("A man, a plan, a canal. Panama")` trebalo bi da vrati `true`.

```js
assert(palindrome('A man, a plan, a canal. Panama') === true);
```

`palindrome("never odd or even")` trebalo bi da vrati `true`.

```js
assert(palindrome('never odd or even') === true);
```

`palindrome("nope")` trebalo bi da vrati `false`.

```js
assert(palindrome('nope') === false);
```

`palindrome("almostomla")` trebalo bi da vrati `false`.

```js
assert(palindrome('almostomla') === false);
```

`palindrome("My age is 0, 0 si ega ym.")` treba da vrati `true`.

```js
assert(palindrome('My age is 0, 0 si ega ym.') === true);
```

`palindrome("1 eye for of 1 eye.")` treba da vrati `false`.

```js
assert(palindrome('1 eye for of 1 eye.') === false);
```

`palindrome("0_0 (: /-\ :) 0-0")` trebalo bi da vrati `true`.

```js
assert(palindrome('0_0 (: /- :) 0-0') === true);
```

`palindrome("five|\_/|four")` trebalo bi da vrati `false`.

```js
assert(palindrome('five|_/|four') === false);
```

# --seed--

## --seed-contents--

```js
function palindrome(str) {
  return true;
}

palindrome("eye");
```

# --solutions--

```js
function palindrome(str) {
  var string = str.toLowerCase().split(/[^A-Za-z0-9]/gi).join('');
  var aux = string.split('');
  if (aux.join('') === aux.reverse().join('')){
    return true;
  }

  return false;
}
```
