---
id: 68e6a99a0fd7386fb66e156f
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, vežbaćete stilizovanje polja za čekiranje kreiranjem panela podešavanja igre.

Sav HTML šablon (`DOCTYPE`, `html`, `head` i `body`) je dostavljen za vas.

Unutar `body`, kreirajte element `div` koji sadrži/koji je ugnježdjen sa elementom `h1`. Elementu `h1` dodelite tekst `Game Settings`.

# --hints--

Trebalo bi da imate element `div`.

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

Trebalo bi da imate element `h1`.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Vaš element ``h1`` treba da bude unutar vašeg elementa ``div``.

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

Vaš `h1` element treba da ima tekst `Game Settings`.

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
