---
id: 68e9a6bf1a74721a4e43e217
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Sada ćete ukloniti podrazumevani čekboks koji je primenjen od strane pregledača. Unutar vašeg selektora `input[type='checkbox']`, dodajte `appearance` sa vrednošću `none`.

Postavljanje svojstva `appearance` na `none` će obrisati izgled koji pregledač primenjuje na čekboksove, omogućavajući vam da prikažete željeni stil.

Nakon toga, pošto čekboks više neće biti vidljiv, postavite `border` sa debljinom `2px`, stilom `solid` i heksadecimalnim kodom `#f1be32`.

# --hints--

Trebalo bi da dodate svojstvo `appearance` vašem selektoru `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

Trebalo bi da imate svojstvo `appearance` sa vrednošću `none`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

Trebalo bi da dodate svojstvo `border` vašem selektoru `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

Trebalo bi da imate svojstvo `border` sa vrednošću `2px solid #f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
