---
id: 68e9a6201e2efa142d62bfe1
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Sledeće, povećaćete kutijaste za bolju vidljivost.

Počnite postavljanjem selektora za `input`, ali ciljajući specifično vaš `[type="checkbox"]`.

Unutar svog selektora, postavite `width` i `height` na `20px`. Ovo ga čini većim nego što je bilo ranije.

I da bi se uskladilo sa vašim podešavanjem `cursor`, koje je postavljeno na labelama, dodajte `cursor` i dodelite mu vrednost `pointer`. Nakon toga, kada pređete mišem preko kutijastih za čekiranje (checkboxes), pojaviće se pokazivač.

# --hints--

Trebalo bi da imate selektor u vašem CSS od `input[type="checkbox"]`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

Trebalo bi da dodate svojstvo `width` vašem selektoru `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

Trebalo bi da imate svojstvo `width` sa vrednošću `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

Trebalo bi da dodate svojstvo ``height`` vašem selektoru ``input[type="checkbox"]``.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

Trebalo bi da imate svojstvo `height` sa vrednošću `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

Trebalo bi da dodate a `cursor` svojstvo na vaš `input[type="checkbox"]` selektor.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

Trebalo bi da imate svojstvo `cursor` sa vrednošću `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
