---
id: 68e9a834dfef752dc35f6f90
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Pošto vaša polja za čekiranje menjaju boju kada ih korisnik klikne, primenićete određeno stilizovanje kako bi bilo vizuelno jasnije da je polje za čekiranje odabrano.

CSS pseudo-klase mogu se povezivati (lanacirati). Možete dodati jednu pseudo-klasu ili pseudo-element iza druge kako biste ciljali elemente koji ispunjavaju više uslova.

Kreirajte novi selektor za `input[type="checkbox"]`. Dodajte pseudo-klasu `:checked` praćenu pseudo-elementom `::after`. Ovo vam omogućava da dodate vizuelni indikator kada je polje za čekiranje označeno.

Sledeći korak, dodajte `content` sa vrednošću `"✓"`.

# --hints--

Trebalo bi da imate kombinovani tip selektor sa pseudo-class selektorem i pseudo-elemen selektorom za `input[type="checkbox"]:checked::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

Trebalo bi da dodate svojstvo `content` vašem kombinovanom tipu selektora sa pseudo-class selektorom i pseudo-element selektorom `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

Trebalo bi da imate svojstvo `content` sa vrednošću `"✓"`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```

