---
id: 68e9a7bd60c971278d2d4ffc
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

U prethodnoj lekciji ste naučili o pseudo-klasama i pseudo-elementima u CSS-u. Sada ćete primeniti to znanje kreiranjem kombinovanog tipskog selektora sa pseudo-klasnim selektorom.

Prvo, počnite sa formatom tipskog selektora postavljanjem ga kao `input[type="checkbox"]` i dodavanjem `:checked` na kraj njega.

Zatim mu date `background-color` sa vrednošću `#f1be32`.

Na kraju, postavite da `border-color` ima vrednost `#e2a60d`.

# --hints--

Trebalo bi da imate kombinovani tip i pseu-class selektor od `input[type="checkbox"]:checked`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked'));
```

Trebalo bi da dodate svojstvo `background-color` vašem kombinovanom tipu i pseudo-class selektoru od `input[type="checkbox"]:checked`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor);
```

Trebalo bi da imate svojstvo `background-color` sa vrednošću `#f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor, "rgb(241, 190, 50)");
```

Trebalo bi dodati svojstvo `border-color` vašem kombinovanom tipu i pseudo-class selektoru za `input[type="checkbox"]:checked`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor);
```

Trebalo bi da imate svojstvo `border-color` sa vrednošću `#e2a60d`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor, "rgb(226, 166, 13)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

--fcc-editable-region--

--fcc-editable-region--

```
