---
id: 68e9a3194160f8f5620ddf21
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Neposredno nakon vašeg `h1` elementa, kreirajte četiri `label` elementa.

Unutar svakog od `label` elemenata, kreirajte `input` element sa atributom `type` podešenim na `"checkbox"`.

Nakon svakog `input` elementa, dajte vašim `label` elementima sledeće tekstove, u ovom redosledu: `Sound Effects`, `Background Music`, `Hard Mode`, i `Haptic Feedback`.

# --hints--

Trebalo bi da imate četiri `label` elementa.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Vaš prvi `label` element treba da ima `input` element.

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

Vaš prvi `input` element treba da ima `type` atribut set do `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

Vaš prvi `label` element bi trebalo da sadrži tekst `Sound Effects`.

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

Vaš tekst `Sound Effects` treba da sledi nakon elementa `input`.

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

Vaš drugi `label` element bi trebalo da ima `input` element.

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

Vaš drugi `input` element treba da ima `type` atribut set do `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

Vaš drugi `label` element bi trebalo da ima tekst `Background Music`.

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

Tekst vašeg `Background Music` treba da sledi nakon elementa `input`.

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

Tvoj treći `label` element bi trebalo da ima `input` element.

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

Vaš treći `input` element bi trebalo da ima atribut `type` set do `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

Vaš treći `label` element bi trebalo da ima tekst `Hard Mode`.

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

Vaš tekst `Hard Mode` treba da sledi nakon `input` elementa.

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

Vaš četvrti `label` element bi trebalo da ima `input` element.

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

Vaš četvrti `input` element bi trebalo da ima `type` atribut set do `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

Četvrti vaš `label` element bi trebalo da ima tekst `Haptic Feedback`.

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

Tekst `Haptic Feedback` treba da dođe posle elementa `input`.

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
