---
id: 68ee86e3ee07b01626a392cc
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

U vašem selektoru `.settings-card` postavite svojstvo `margin` na `auto`.

Postavljanje svojstva `margin` na `auto` automatski prilagođava margine elementa kako bi ravnomerno rasporedilo preostali prostor u njegovom kontejneru, što se često koristi za horizontalno centriranje blok-nivo elemenata.

I na kraju, postavite svojstvo `text-align` sa vrednošću `left`. Ovo će poravnati ugrađeni sadržaj, kao što je tekst, ka levoj strani njegovog kontejner elementa.

# --hints--

Trebalo bi da dodate svojstvo `margin` na selektor `settings-card` class.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

Trebalo bi da imate svojstvo `margin` sa vrednošću `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

Trebalo bi da dodate svojstvo `text-align` na selektor `settings-card` class.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

Trebalo bi da imate svojstvo `text-align` sa vrednošću `left`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
