---
id: 68f15cd397503429f0345262
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

U deklaraciji za selektor `input[type="checkbox"]:checked::after`, postavite `display` na `block`. Postavljanje vašeg svojstva `display` kao `block` čini element blok-nivo elementom, što znači da zauzima punu širinu svog kontejnera i počinje sa nove linije, omogućavajući vam lakše kontrolisanje njegove širine, visine i razmaka.

Zatim, postavite `text-align` na `center`. Ovo će centrirati ugrađeni sadržaj (poput teksta ili inline elemenata) horizontalno unutar bloka. Blok će i dalje zauzimati punu širinu svog kontejnera, ali će sve što je unutra biti poravnato na centar.

# --hints--

Trebalo bi da dodate svojstvo `display` vašem kombinovanom tip selektoru sa pseudo-selektorom class i pseudo-element selektorom za `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

Trebalo bi da imate svojstvo `display` sa vrednošću `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

Trebalo bi da dodate `text-align` svojstvo vašem kombinovanom tip selektoru sa pseudo-class selektorom i pseudo-element selektorom za `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

Trebalo bi da imate svojstvo `text-align` sa vrednošću `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
