---
id: 68f16192e8a66f3b948d3ebd
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Da biste završili podešavanje vizuelnog efekta vaše prokupljke, postavite svojstvo `font-weight` sa vrednošću `bold`. Ovo će povećati vidljivost oznake za proveru.

Sada kada je ovo lakše vidljivo, promenite `color` na vrednost `white`. Ovo će promeniti boju oznake za proveru unutar prokupljke kada je označena.

Konačno, izmenite `line-height` na vrednost `20px`.

Završili ste kontrolnu tablu podešavanja igre!

# --hints--

Trebalo bi da dodate svojstvo `font-weight` vašem kombinovanom tipu selektora sa pseudo-class selektorom i pseudo-element selektorom `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.fontWeight);
```

Trebalo bi da imate svojstvo `font-weight` sa vrednošću `bold`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.fontWeight, "bold");
```

Trebalo bi da dodate svojstvo `color` vašem kombinovanom tip selektoru sa pseudo-class selektor i pseudo-element selektor od `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.color);
```

Trebalo bi da imate `color` svojstvo sa vrednošću od `white`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.color, ["white", "#ffffff", "rgb(255,255,255)"]);
```

Trebalo bi da dodate svojstvo `line-height` vašem kombinovanom tipu selektora sa pseudo-class selektorom i pseudo-element selektorom `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.lineHeight);
```

Trebalo bi da imate svojstvo `line-height` sa vrednošću `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.lineHeight, "20px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
  display: block;
  text-align: center;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type='checkbox']:checked:after {
  content: '✓';
  display: block;
  text-align: center;
  font-weight: bold;
  color: white;
  line-height: 20px;
}
```

