---
id: 68f1560aff1214103ab6b59b
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sada ćete početi da dekorišete stranicu nekim CSS stilizovanjem. Počnite kreiranjem selektora `body`.

Postavite da `body` ima svojstvo `height` sa vrednošću `100vh` i svojstvo `background-color` sa vrednošću `#f0f0f0`.

`height` od `100vh` čini da `body` zauzme punu visinu prozora (viewport), dok svetlo siva pozadina pruža suptilan temelj za stranicu.

Napokon, postavite svojstvo `text-align` sa vrednošću `center`. Ovo će centrirati sav inline sadržaj unutar stranice osim ako ga neki element ne nadjačava sopstvenim poravnanjem.

# --hints--

Trebalo bi da koristite tip selektora `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Trebalo bi da dodate svojstvo `height` unutar selektora tipa `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

Trebalo bi da imate svojstvo `height` sa vrednošću `100vh`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

Trebalo bi da dodate svojstvo `background-color` unutar selektora tipa `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

Trebalo bi da imate svojstvo `background-color` sa vrednošću `#f0f0f0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

Trebalo bi da dodate svojstvo `text-align` unutar tipa selektora `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

Trebalo bi da imate svojstvo `text-align` sa vrednošću `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
