---
id: 68f15ef0446464325ce1aa45
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Potrebno vam je neko razmak između polja za čekiranje i oznaka. Počnite korišćenjem selektora za `label` u vašem CSS-u i postavite svojstvo `display` sa vrednošću `block`.

Zatim, postavite svojstvo `margin` sa vrednošću `8px auto` da dodate vertikalni razmak između elemenata. Ovo će stvoriti dosledan razmak iznad i ispod svakog elementa, pomažući u razdvajanju sadržaja i poboljšavajući čitljivost.

Konačno, postavite svojstvo `cursor` na `pointer` na elementima `label`. Ovo će promeniti kursor u ikonu ruke kada korisnik pređe mišem preko `label`, signalizirajući da je element klikabilan i poboljšavajući ukupno korisničko iskustvo.

# --hints--

Trebalo bi da imate selektor tipa u vašem CSS for `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

Trebalo bi da dodate `display` svojstvo vašem `label` selektoru.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

Trebalo bi da imate svojstvo `display` sa vrednošću `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

Trebalo bi da dodate svojstvo `margin` vašem selektoru `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

Trebalo bi da imate svojstvo `margin` sa vrednošću `8px auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

Trebalo bi da dodate svojstvo `cursor` vašem selektoru `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

Trebalo bi da imate svojstvo `cursor` sa vrednošću `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
