---
id: 68f29dbf9125476b4808d6fb
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Sada je vreme da stilizujemo kontejner `settings-card`. Kreirajte selektor klase za `settings-card`.

Ovo će biti mesto gde će se nalaziti sva vaša podešavanja za kontejner. Postavite `max-width` na `250px` da biste definisali ukupnu veličinu vašeg kontejnera.

Sledeći korak, postavite vaš `padding` na `20px`, tako da je sadržaj razmaknut od ivice kontejnera.

Nakon toga, kreirajte zaobljen ugao postavljanjem vašeg `border-radius` na `10px`.

Zatim postavite `box-shadow` sa vrednostima `0 2px 6px rgba(0,0,0,0.2)`. Ovo će stvoriti suptilan „podignut“ izgled koji će dati dubinu kontejneru.

# --hints--

Trebalo bi da imate class selektor u vašem CSS for `settings-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

Trebalo bi da dodate svojstvo `max-width` na selektor `.settings-card` class.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

Trebalo bi da imate svojstvo `max-width` sa vrednošću `250px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

Trebalo bi da dodate svojstvo `padding` na `.settings-card` class selektor.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

Trebalo bi da imate svojstvo `padding` sa vrednošću `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

Trebalo bi da dodate `border-radius` svojstvo vašem `.settings-card` class selektoru.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

Trebalo bi da imate svojstvo `border-radius` sa vrednošću `10px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

Trebalo bi da dodate svojstvo `box-shadow` na vaš `.settings-card` class selektor.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

Trebalo bi da imate `box-shadow` sa vrednošću `0 2px 6px rgba(0,0,0,0.2)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
