---
id: 68f29753fd5524530c6e65c8
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Zatim, postavite svojstvo `vertical-align` sa vrednošću `middle`.

Svojstvo `vertical-align` kontroliše kako se elementi `inline` ili `inline-block` vertikalno poravnaju sa okolnim tekstom ili drugim ugrađenim elementima. Često se koristi za podešavanje vertikalne pozicije elemenata kao što su slike, ikone ili tekst unutar linije.

Zatim finalizujte svoju prokvatnicu (checkbox) svojstvom `margin` sa vrednošću `15px`.

# --hints--

Trebalo bi da dodate svojstvo `vertical-align` vašem selektoru `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign);
```

Trebalo bi da imate svojstvo `vertical-align` sa vrednošću `middle`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign, "middle");
```

Trebalo bi da dodate svojstvo `margin-right` vašem selektoru `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight);
```

Trebalo bi da imate svojstvo `margin-right` sa vrednošću `15px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight, "15px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
