---
id: 66b205e6eacba4c4e54ea434
title: Izgradnja sakrivanja e-poﾅ｡te
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

U ovoj laboratoriji, ﾄ㌃te sakriti deo korisniﾄ耕og imena u adresi e-poﾅ｡te pomoﾄ㎡ zvezdica. Sakrivanje je termin koji se koristi za skrivanje ili menjanje osetljivih informacija sa zvezdicama ili drugim slovima.

Na primer, ako je e-mail adresa bila `myEmail@email.com`, onda maskirana e-mail adresa ﾄ㌃ biti `m*****l@email.com`.

**Cilj:** Izvrﾅ｡ite korisniﾄ耕e priﾄ稿 ispod i poloﾅｾite sve testove kako biste zavrﾅ｡ili ovaj laboratorijski rad.

**Istorija korisnika:**

1. Kreirajte funkciju nazvanu `maskEmail` koja prima `email` kao argument.
2. Unutar funkcije, trebalo bi da maskujete `email` i priloﾅｾite njegov domen. Zapamtite da moﾅｾete koristiti metode poput `slice`, `repeat`, `indexOf` ili ﾄ溝k `replace`.
3. Izvan funkcije, deklariﾅ｡ite parametar nazvan `email` za skladiﾅ｡tenje e-mail adrese koju ﾅｾelite da sakrijete.
4. Pozovite funkciju `maskEmail` sa parametrom `email` i ispiﾅ｡ite rezultat u konzolu.
5. `maskEmail("apple.pie@example.com")` treba da vrati `"a*******e@example.com"`.
6. `maskEmail("freecodecamp@example.com")` treba da vrati `"f**********p@example.com"`.
7. `maskEmail("info@test.dev")` treba da vrati `"i**o@test.dev"`.
8. `maskEmail("user@domain.org")` treba da vrati `"u**r@domain.org"`.

# --hints--

Trebalo bi da objavite funkciju pod nazivom `maskEmail`.

```js
assert.isFunction(maskEmail);
```

Funkcija `maskEmail` treba da primi niz znakova, `email`, kao argument.

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

Funkcija, treba imati parametar od `email`.

```js
assert.isDefined(email);
```

Trebalo bi da unesete vaﾅｾeﾄ㎡ adresu e-poﾅ｡te u svoj parametar `email`.

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` trebalo bi da vrati `"a*******e@example.com"`.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` treba da vrati `"f**********p@example.com"`.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")` treba da vrati `"i**o@test.dev"`.

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")` treba da vrati `"u**r@domain.org"`.

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

`maskEmail` Treba da proizvede taﾄ肱e rezultate.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

Trebalo bi da napiﾅ｡ete rezultate poziva `maskEmail` sa `email` kao argumentom.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
