---
id: 56533eb9ac21ba0edf2244cf
title: Sastav kolekcije snimaka
challengeType: 26
dashedName: build-a-record-collection
---

# --description--

Kreirate funkciju koja pomaže u održavanju kolekcije muzičkih albuma. Kolekcija je kreirana kao objekat sa mnogo albuma, koji su takođe objekti. Svaki album je predstavljen u kolekciji jedinstvenim `id` kao ime atributa. Unutar svakog objekta albuma, postoje razni atributi koji definišu informacije o tom albumu. Nisu svi albumi sa potpunim informacijama.

Funkcija `updateRecords` koja prima 4 argumenta predstavljena sledećim parametrima funkcije:

-   `records` - objekat sa nekoliko privatnih albuma
-   `id` - broj koji predstavlja određeni album u objektu `records`
-   `prop` - niz znakova koji predstavlja ime atributa ažuriranog albuma
-   `value` - niz znakova koji opisuje podatke korišćene za ažuriranje atributa albuma

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratorijum.

Priče o korisniku:

1. Tvoja funkcija mora uvek da vrati kompletan objekat od `records`.
2. Ako je `value` prazan niz znakova, obriši svojstvo `prop` dobijeno iz albuma.
3. Ako `prop` nije `tracks` i ako `value` nije prazan niz znakova, postavi vrednost `value` za svojstvo `prop` tog albuma.
4. Ako `prop` je `tracks` i ako `value` nije prazan niz znakova, ali album nema svojstvo `tracks`, kreiraj prazan niz podataka i dodaj `value` unutra.
5. Ako je svojstvo `tracks` i ako `value` nije prazan niz znakova, dodaj `value` na kraj postojećeg niza podataka albuma `tracks`.

**Napomena:** Kopija objekta `recordCollection` se koristi za specifikacije. Vaša funkcija ne bi trebalo da direktno referencira objekat `recordCollection`, već samo parametar funkcije.

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

Trebalo bi da imaš funkciju `updateRecords`.

```js
assert.isFunction(updateRecords);
```

Nakon `updateRecords(recordCollection, 5439, "artist", "ABBA")`, `artist` trebalo bi da bude sekvenca slova `ABBA`

```js
assert.equal(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'],
    'ABBA'
);
```

Nakon `updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")`, `tracks` bi trebalo da ima sekvencu slova `Take a Chance on Me` kao poslednji i jedini element.

```js
updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me');
assert.lengthOf(_recordCollection[5439]['tracks'], 1);
assert.equal(_recordCollection[5439]['tracks'].pop(), 'Take a Chance on Me');
```

Nakon `updateRecords(recordCollection, 2548, "artist", "")`, `artist` treba biti postavljen

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('artist'));
```

Nakon `updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")`, `tracks` treba da ima sekvencu slova `Addicted to Love` kao poslednji element.

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop(), 'Addicted to Love'
);
```

Nakon `updateRecords(recordCollection, 2468, "tracks", "Free")`, `tracks` treba da ima niz karaktera `1999` kao prvi element.

```js
assert.equal(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0], '1999'
);
```

Nakon `updateRecords(recordCollection, 2548, "tracks", "")`, `tracks` ne bi trebalo da se postavi

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('tracks'));
```

Nakon `updateRecords(recordCollection, 1245, "albumTitle", "Riptide")`, `albumTitle` treba da bude sekvenca slova `Riptide`

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ], 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};



```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
