---
id: bad87fee1348bd9aedf08834
title: Napravi Grupu Radio Dugmica
challengeType: 0
forumTopicId: 16822
dashedName: create-a-set-of-radio-buttons
---

# --description--

Mozete koristiti <dfn>radio dugmice</dfn> za pitanja gde zelite da korisnik izabere samo jedan odgovor od vise ponudjenih.

Radio dugmad su tip `input` elementa.

Svako radio dugme moze biti ugnjezdeno unutar svog `label` elementa. Kada obuhvatite `input` element unutar `label` elementa, radio unos se automatski povezuje sa labelom koja ga okruzuje.

Sva povezana radio dugmad treba da imaju isti `name` atribut kako bi cinila jednu radio grupu. U radio grupi, izbor jednog dugmeta automatski ponistava izbor ostalih, cime se osigurava da korisnik izabere samo jedan odgovor.

Evo primera radio dugmeta:

```html
<label> 
  <input type="radio" name="indoor-outdoor">Indoor 
</label>
```

Dobra praksa je da postavite `for` atribut na `label` element, sa vrednoscu koja odgovara `id` atributu odgovarajuceg `input` elementa. To omogucava pomocnim tehnologijama da uspostave jasnu vezu izmedju labele i povezanog `input` elementa. Na primer:

```html
<input id="indoor" type="radio" name="indoor-outdoor">
<label for="indoor">Indoor</label>
```

Takodje mozemo ugnjezditi `input` element unutar `label` tagova:

```html
<label for="indoor"> 
  <input id="indoor" type="radio" name="indoor-outdoor">Indoor 
</label>
```

# --instructions--

Dodajte par radio dugmica u svoju formu, pri cemu svako treba da bude ugnjezdeno unutar svog `label` elementa. Jedno treba da ima opciju `indoor`, a drugo `outdoor`. Oba treba da dele `name` atribut sa vrednoscu `indoor-outdoor` kako bi formirali jednu radio grupu.

# --hints--

Vasa stranica treba da ima dva `radio` dugmeta.

```js
assert.lengthOf(document.querySelectorAll('input[type="radio"]'),2);
```

Vasa radio dugmad treba da imaju `name` atribut sa vrednoscu `indoor-outdoor`.

```js
assert.lengthOf([...document.querySelectorAll('input[type="radio"]')].filter(input => input.name === "indoor-outdoor"),2);
```

Svaki radio `input` element treba da bude ugnjezden unutar svog `label` elementa.

```js
assert.lengthOf(document.querySelectorAll('label > input[type="radio"]:only-child'),2);
```

Svaki `label` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length);
```

Jedno radio dugme treba da ima labelu `indoor`.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label =>  label.match(/indoor/gi));
assert.isTrue(hasMatch); 
```

Jedno radio dugme treba da ima labelu `outdoor`.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label => label.match(/outdoor/gi));
assert.isTrue(hasMatch); 
```

Svaki radio `input` element treba da bude dodat unutar `form` taga.

```js
const labelParent = document.querySelector('label').parentNode;
assert.strictEqual(labelParent.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
   <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
