---
id: bad87fee1348bd9aedf08828
title: Napravi Uredjenu Listu
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cQ3B8TM'
forumTopicId: 16824
dashedName: create-an-ordered-list
---

# --description--

HTML ima jos jedan poseban element za pravljenje <dfn>uredjenih listi</dfn>, odnosno listi sa brojevima.

Uredjene liste pocinju otvarajucim `<ol>` elementom, nakon cega sledi bilo koji broj `<li>` elemenata. Na kraju se uredjena lista zatvara sa `</ol>`.

Na primer:

```html
<ol>
  <li>Garfield</li>
  <li>Sylvester</li>
</ol>
```

sto bi napravilo numerisanu listu `Garfield` i `Sylvester`.

# --instructions--

Napravite uredjenu listu sa 3 stvari koje macke najvise mrze.

# --hints--

Treba da imate uredjenu listu za `Top 3 things cats hate:`

```js
const previousElement = document.querySelector('ol').previousElementSibling; 
assert.match(previousElement.textContent,/Top 3 things cats hate:/i);
```

Treba da imate neuredjenu listu za `Things cats love:`

```js
const previousElement = document.querySelector('ul').previousElementSibling; 
assert.match(previousElement.textContent,/Things cats love:/i);
```

Treba da imate samo jedan `ul` element.

```js
assert.lengthOf(document.querySelectorAll('ul'), 1);
```

Treba da imate samo jedan `ol` element.

```js
assert.lengthOf(document.querySelectorAll('ol'), 1);
```

Treba da imate tri `li` elementa unutar vaseg `ul` elementa.

```js
assert.lengthOf(document.querySelectorAll('ul li'), 3);
```

Treba da imate tri `li` elementa unutar vaseg `ol` elementa.

```js
assert.lengthOf(document.querySelectorAll('ol li'), 3);
```

Vas `ul` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/ul>/g);
assert.strictEqual(code.match(/<\/ul>/g).length ,code.match(/<ul>/g).length);
```

Vas `ol` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/ol>/g);
assert.strictEqual(code.match(/<\/ol>/g).length ,code.match(/<ol>/g).length);
```

Vas `li` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/li>/g);
assert.match(code,/<li>/g);
assert.strictEqual(code.match(/<\/li>/g).length ,code.match(/<li>/g).length);
```

`li` elementi u neuredjenoj listi ne smeju biti prazni.

```js
[...document.querySelectorAll('ul li')].forEach((val) =>
  assert.isNotEmpty(__helpers.removeWhiteSpace(val.textContent))
);
```

`li` elementi u uredjenoj listi ne smeju biti prazni.

```js
[...document.querySelectorAll('ol li')].forEach((val) =>
  assert.isNotEmpty(__helpers.removeWhiteSpace(val.textContent))
);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>

</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>hate 1</li>
    <li>hate 2</li>
    <li>hate 3</li>
  </ol>
</main>
```
