---
id: bad87fee1348bd9aedf08812
title: Dodaj Slike na Svoj Veb Sajt
challengeType: 0
forumTopicId: 16640
dashedName: add-images-to-your-website
---

# --description--

Mozete dodati slike na svoj sajt koristeci `img` element i navodjenjem URL adrese slike preko `src` atributa.

Primer:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg">
```

Zapamtite da je `img` prazni element.

Svi `img` elementi **moraju** imati `alt` atribut. Tekst unutar `alt` atributa koriste citaci ekrana radi bolje pristupacnosti, a prikazuje se i kada slika ne moze da se ucita.

**Napomena:** Ako je slika samo dekorativna, dobra praksa je koristiti prazan `alt` atribut.

Uopsteno, `alt` atribut ne treba da sadrzi specijalne karaktere osim ako su neophodni.

Dodajmo `alt` atribut nasem `img` primeru iznad:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg" alt="freeCodeCamp logo">
```

# --instructions--

Hajde da dodamo sliku na nas sajt:

Unutar postojeceg `main` elementa postavite `img` element pre postojecih `p` elemenata.

Zatim postavite `src` atribut da pokazuje na URL `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`.

Na kraju, ne zaboravite da svom `img` elementu dodate `alt` atribut sa prikladnim tekstom.

# --hints--

Vasa stranica treba da ima element slike.

```js
assert.exists(document.querySelector('img'));
```

Vasa slika treba da ima `src` atribut koji pokazuje na sliku macke.

```js
const url = document.querySelector('img').getAttribute('src');
assert.match(url,/^https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg$/i);
```

`alt` atribut vaseg elementa slike ne treba da bude prazan.

```js
assert.exists(document.querySelector('img').getAttribute('alt'));
assert.isNotEmpty(document.querySelector('img').getAttribute('alt'));
assert.match(__helpers.removeWhiteSpace(code),/<(?:img|IMG)\S*alt=(['"])(?!\1|>)\S+\1\S*\/?>/)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>


  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
