---
id: bad87fee1348bd9aede08835
title: Ugnjezdi Vise Elemenata unutar Jednog div Elementa
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cNW4kC3'
forumTopicId: 18246
dashedName: nest-many-elements-within-a-single-div-element
---

# --description--

`div` element, poznat i kao division element, je opsti kontejner za druge elemente.

`div` je jedan od najcesce koriscenih HTML elemenata.

Kao i svaki drugi element koji nije self-closing, `div` se otvara sa `<div>` i zatvara na drugom redu sa `</div>`.

# --instructions--

Ugnjezdite liste "Things cats love" i "Top 3 things cats hate" unutar jednog `div` elementa.

Hint: Pokusajte da stavite otvarajuci `div` tag iznad `p` elementa "Things cats love", a zatvarajuci `div` tag posle zatvarajuceg `ol` taga, kako bi obe liste bile unutar istog `div` elementa.

# --hints--

Vasi `p` elementi treba da budu ugnjezdeni unutar vaseg `div` elementa.

```js
const div = document.querySelector('div');
const children = div.querySelectorAll('p');
assert.isAbove(children.length,1);
```

Vas `ul` element treba da bude ugnjezden unutar vaseg `div` elementa.

```js
const div = document.querySelector('div');
const children = div.querySelectorAll('ul');
assert.notEmpty(children);
```

Vas `ol` element treba da bude ugnjezden unutar vaseg `div` elementa.

```js
const div = document.querySelector('div');
const children = div.querySelectorAll('ol');
assert.notEmpty(children);
```

Vas `div` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/div>/g);
assert.strictEqual(code.match(/<\/div>/g).length,code.match(/<div>/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>

  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality" value="loving" checked> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality" value="lazy"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality" value="energetic"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality" value="loving" checked> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality" value="lazy"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality" value="energetic"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
