---
id: bad87fee1348bd9aedf08827
title: Napravi Neuredjenu Listu sa Tackama
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cDKVPuv'
forumTopicId: 16814
dashedName: create-a-bulleted-unordered-list
---

# --description--

HTML ima poseban element za pravljenje <dfn>neuredjenih listi</dfn>, odnosno listi sa tackama.

Neuredjena lista pocinje otvarajucim `<ul>` elementom, zatim sledi proizvoljan broj `<li>` elemenata. Na kraju se neuredjena lista zatvara sa `</ul>`.

Na primer:

```html
<ul>
  <li>milk</li>
  <li>cheese</li>
</ul>
```

sto ce napraviti listu sa tackama `milk` i `cheese`.

# --instructions--

Uklonite poslednja dva `p` elementa i napravite neuredjenu listu od tri stvari koje macke vole pri dnu stranice.

# --hints--

Napravite `ul` element.

```js
assert.isNotEmpty(document.querySelectorAll('ul'));
```

Treba da imate tri `li` elementa unutar svog `ul` elementa.

```js
assert.lengthOf(document.querySelectorAll('ul li'),3);
```

Vas `ul` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/ul>/gi);
assert.match(code,/<ul/gi);
assert.strictEqual(code.match(/<\/ul>/gi).length,code.match(/<ul/gi).length);
```

Vasi `li` elementi treba da imaju zatvarajuce tagove.

```js
assert.match(code,/<\/li>/gi);
assert.match(code,/<li[\s>]/gi);
assert.strictEqual(code.match(/<\/li>/gi).length,code.match(/<li[\s>]/gi).length);
```

Vasi `li` elementi ne smeju biti prazni niti sadrzati samo beline.

```js
assert.isEmpty([...document.querySelectorAll('ul li')].filter((item) => item.textContent.trim() === ""));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <ul>
    <li>milk</li>
    <li>mice</li>
    <li>catnip</li>
  </ul>
</main>
```
