---
id: bad87fee1348bd9aedf08802
title: Otkomenatiraj HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBmG9T7'
forumTopicId: 18329
dashedName: uncomment-html
---

# --description--

Komentarisanje je nacin da ostavite belezke drugim developerima u svom kodu, bez uticaja na rezultat koji vidi krajnji korisnik.

Komentarisanje je takodje jednostavan nacin da privremeno onemogucite kod bez trajnog brisanja.

Komentari u HTML-u pocinju sa `<!--` i zavrsavaju se sa `-->`.

# --instructions--

Otkomenatirajte svoje `h1`, `h2` i `p` elemente.

# --hints--

Vas `h1` element treba da bude vidljiv na stranici nakon uklanjanja komentara.

```js
assert.notEmpty(document.querySelectorAll('h1'));
```

Vas `h2` element treba da bude vidljiv na stranici nakon uklanjanja komentara.

```js
assert.notEmpty(document.querySelectorAll('h2'));
```

Vas `p` element treba da bude vidljiv na stranici nakon uklanjanja komentara.

```js
assert.notEmpty(document.querySelectorAll('p'));
```

Sledeci komentarski marker ne treba da bude vidljiv na stranici (`-->`).

```js
const elements = document.querySelectorAll('*');
const potentialComments = Array.from(elements).filter(el => el.textContent.includes('-->'));
assert.notExists(potentialComments[1])
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
