---
id: bad87fee1348bd9aede08817
title: Ugnjezdi Anchor Element Unutar Paragrafa
challengeType: 0
forumTopicId: 18244
dashedName: nest-an-anchor-element-within-a-paragraph
---

# --description--

Mozete ugnjezditi linkove unutar drugih tekstualnih elemenata.

```html
<p>
  Here's a <a target="_blank" href="https://www.freecodecamp.org"> link to www.freecodecamp.org</a> for you to follow.
</p>
```

Hajde da razlozimo ovaj primer. Obican tekst je obavijen `p` elementom:

```html
<p> Here's a ... for you to follow. </p>
```

Sledece je *anchor* `<a>` element (koji zahteva zatvarajuci tag `</a>`):

```html
<a> ... </a>
```

`target` je atribut anchor elementa koji odredjuje gde ce se otvoriti link. Vrednost `_blank` znaci da se link otvara u novom tabu. `href` je atribut anchor elementa koji sadrzi URL adresu linka:

```html
<a href="https://www.freecodecamp.org" target="_blank"> ... </a>
```

Tekst `link to www.freecodecamp.org` unutar `a` elementa naziva se <dfn>anchor text</dfn> i prikazuje klikabilni link:

```html
<a href=" ... " target="...">link to freecodecamp.org</a>
```

Krajnji rezultat primera izgleda ovako:

Evo <a href="https://www.freecodecamp.org" target="_blank">linka do www.freecodecamp.org</a> koji mozete pratiti.

# --instructions--

Ugnjezdite postojeci `a` element unutar novog `p` elementa. Nemojte praviti novi anchor tag. Novi paragraf treba da ima tekst `View more cat photos`, gde je `cat photos` link, a ostatak obican tekst.

# --hints--

Treba da imate samo jedan `a` element.

```js
assert.lengthOf(document.querySelectorAll('a'), 1 );
```

`a` element treba da linkuje na "`https://www.freecatphotoapp.com`".

```js
assert.lengthOf(document.querySelectorAll('a[href="https://www.freecatphotoapp.com"]'),1);
```

Vas `a` element treba da ima anchor tekst `cat photos`.

```js
assert.match(document.querySelector('a').textContent,/cat\sphotos/gi)
```

Treba da napravite novi `p` element. U vasem HTML kodu treba da postoje najmanje 3 `p` taga.

```js
assert.lengthOf(document.querySelectorAll('p'),3)
```

Vas `a` element treba da bude ugnjezden unutar novog `p` elementa.

```js
const anchorParent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode;
assert.strictEqual(anchorParent.tagName,"P")
```

Vas `p` element treba da sadrzi tekst `View more ` (sa razmakom na kraju).

```js
const textContent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode.textContent;
assert.match(textContent,/View\smore\s/gi);
```

Vas `a` element ne treba <em>da sadrzi</em> tekst `View more`.

```js
assert.notMatch(document.querySelector('a').textContent,/View\smore/gi); 
```

Svaki vas `p` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/p>/g);
assert.match(code,/<p/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

Svaki vas `a` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>View more <a target="_blank" href="https://www.freecatphotoapp.com">cat photos</a></p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
