---
id: bad87fee1348bd9aedf08833
title: Popuni Prazninu Placeholder Tekstom
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cgR7Dc7'
forumTopicId: 18178
dashedName: fill-in-the-blank-with-placeholder-text
---

# --description--

Web developeri cesto koriste <dfn>lorem ipsum tekst</dfn> kao placeholder tekst. Lorem ipsum potice iz nasumicno izabranih delova poznatog dela Cicerona iz starog Rima.

Lorem ipsum se koristi kao placeholder tekst jos od 16. veka, a ta praksa se nastavila i na webu.

Dobro, pet vekova je sasvim dovoljno. Posto pravimo CatPhotoApp, koristite nesto sto se zove "kitty ipsum" tekst.

# --instructions--

Zamenite tekst unutar svog `p` elementa prvih nekoliko reci sledeceg kitty ipsum teksta: `Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.`

# --hints--

Vas `p` element treba da sadrzi prvih nekoliko reci datog "kitty ipsum" teksta.

```js
assert.match(document.querySelector('p').textContent,/Kitty(\s)+ipsum/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Hello Paragraph</p>
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff</p>
```
