---
id: bad87fee1348bd9aedf08801
title: Informisi Koriscenjem Paragraf Elementa
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ceZ7DtN'
forumTopicId: 18202
dashedName: inform-with-the-paragraph-element
---

# --description--

`p` element je preporucen za tekst paragrafa na sajtovima. `p` je skracenica za "paragraph".

Paragraf element mozete napraviti ovako:

```html
<p>I'm a p tag!</p>
```

# --instructions--

Napravite `p` element ispod `h2` elementa i dodelite mu tekst `Hello Paragraph`.

**Napomena:** Po konvenciji, svi HTML tagovi pisu se malim slovima, na primer `<p></p>`, a ne `<P></P>`.

# --hints--

Vas kod treba da ima validan `p` element.

```js
assert.lengthOf(document.querySelectorAll('p'),1);
```

Vas `p` element treba da sadrzi tekst `Hello Paragraph`.

```js
assert.match(document.querySelector('p').textContent,/hello(\s)+paragraph/gi);
```

Vas `p` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
```

# --solutions--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
<p>Hello Paragraph</p>
```
