---
id: bad87fee1348bd9aedf08804
title: Zakomentarisi HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cGyGbca'
forumTopicId: 16782
dashedName: comment-out-html
---

# --description--

Zapamtite da za pocetak komentara koristite `<!--`, a za kraj komentara `-->`.

Ovde je potrebno da zatvorite komentar pre nego sto pocne vas `h2` element.

# --instructions--

Zakomentarisite svoj `h1` i svoj `p` element, ali nemojte zakomentarisati `h2` element.

# --hints--

Vas `h1` element treba da bude zakomentarisan tako da ne bude vidljiv na stranici.

```js
assert.isEmpty(document.querySelectorAll('h1'));
```

Vas `h2` element ne treba da bude zakomentarisan kako bi bio vidljiv na stranici.

```js
assert.isNotEmpty(document.querySelectorAll('h2'));
```

Vas `p` element treba da bude zakomentarisan tako da ne bude vidljiv na stranici.

```js
assert.isEmpty(document.querySelectorAll('p'));
```

Svaki vas komentar treba da bude zatvoren sa `-->`.

```js
assert.isAbove(code.match(/[^fc]-->/g).length, 1);
```

Ne treba da menjate redosled `h1`, `h2` i `p` elemenata u kodu.

```js
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[0],'<h1>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[1],'<h2>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[2],'<p>');
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<!--<h1>Hello World</h1>-->
<h2>CatPhotoApp</h2> 
<!--<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p> -->
```
