---
id: bad87fed1348bd9aedf08833
title: Obrisi HTML Elemente
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ckK73C9'
forumTopicId: 17559
dashedName: delete-html-elements
---

# --description--

Nas telefon nema mnogo vertikalnog prostora.

Obrisimo nepotrebne elemente kako bismo poceli da gradimo nas CatPhotoApp.

# --instructions--

Obrisite svoj `h1` element kako bismo pojednostavili izgled.

# --hints--

Vas `h1` element treba da bude obrisan.

```js
assert.notMatch(code,/<h1>/gi);
assert.notMatch(code,/<\/h1>/gi);
```

Vas `h2` element treba da bude na stranici.

```js
assert.match(code,/<h2>[\w\W]*<\/h2>/gi);
```

Vas `p` element treba da bude na stranici.

```js
assert.match(code,/<p>[\w\W]*<\/p>/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2><p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
