---
id: bad88fee1348bd9aedf08816
title: Linkuj ka Unutrasnjim Delovima Stranice Koriscenjem Anchor Elemenata
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cyrDRUL'
forumTopicId: 301098
dashedName: link-to-internal-sections-of-a-page-with-anchor-elements
---

# --description--

`a` elementi (*anchor tagovi*) mogu se koristiti i za kreiranje unutrasnjih linkova koji skacu na razlicite delove web stranice.

Da biste napravili unutrasnji link, postavite vrednost `href` atributa linka na hash simbol `#` plus vrednost `id` atributa elementa na koji zelite da linkujete, obicno na dnu stranice. Zatim tom ciljanom elementu dodajte isti `id` atribut. `id` atribut jedinstveno identifikuje element.

Ispod je primer unutrasnjeg anchor linka i ciljanog elementa:

```html
<a href="#contacts-header">Contacts</a>
...
<h2 id="contacts-header">Contacts</h2>
```

Kada korisnici kliknu na link `Contacts`, bice preusmereni na deo stranice koji sadrzi naslov **Contacts**.

# --instructions--

Promenite spoljasnji link u unutrasnji tako sto cete promeniti `href` atribut na `#footer` i tekst sa `cat photos` na `Jump to Bottom`.

Uklonite `target="_blank"` atribut iz anchor taga, jer on otvara povezani dokument u novom tabu.

Zatim dodajte `id` atribut sa vrednoscu `footer` u `<footer>` element na dnu stranice.

# --hints--

Na vasoj stranici treba da postoji samo jedan anchor tag.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Na vasoj stranici treba da postoji samo jedan `footer` tag.

```js
assert.lengthOf(document.querySelectorAll('footer'),1);
```

`a` tag treba da ima `href` atribut postavljen na `#footer`.

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'), '#footer');
```

`a` tag ne treba da ima `target` atribut.

```js
assert.notExists(document.querySelector('a').getAttribute('target'));
```

Tekst `a` taga treba da bude "Jump to Bottom".

```js
assert.match(document.querySelector('a').textContent,/Jump to Bottom/gi);
```

`footer` tag treba da ima `id` atribut postavljen na `footer`.

```js
assert.strictEqual(document.querySelector('footer').getAttribute('id'),'footer');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer>Copyright Cat Photo App</footer>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="#footer">Jump to Bottom</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer id="footer">Copyright Cat Photo App</footer>
```
