---
id: bad87fee1348bd9aedd08830
title: Dodaj Dugme za Slanje u Formu
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cp2Nkhz'
forumTopicId: 16627
dashedName: add-a-submit-button-to-a-form
---

# --description--

Dodajmo `submit` dugme u vasu formu. Klik na ovo dugme poslace podatke iz forme na URL koji ste naveli u `action` atributu forme.

Evo primera submit dugmeta:

```html
<button type="submit">this button submits the form</button>
```

# --instructions--

Dodajte dugme kao poslednji element svog `form` elementa, sa tipom `submit` i tekstom `Submit`.

# --hints--

Vas `form` treba da sadrzi `button`.

```js
const form = document.querySelector('form');
const children = form.querySelectorAll(`:scope ${"button"}`); 
assert.isNotEmpty(children); 
```

Vase submit dugme treba da ima `type` atribut postavljen na `submit`.

```js
assert.strictEqual(document.querySelector('button').getAttribute('type'), 'submit');
```

Vase submit dugme treba da ima samo tekst `Submit`.

```js
const text = document.querySelector('button').textContent;
assert.match(text,/^\s*submit\s*$/gi);
```

Vas `button` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/button>/g);
assert.match(code,/<button/g);
assert.strictEqual(code.match(/<\/button>/g).length,code.match(/<button/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```
