---
id: bad87fee1348bd9aecf08801
title: Uvod u HTML5 Elemente
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBkZGpt7'
forumTopicId: 301097
dashedName: introduction-to-html5-elements
---

# --description--

HTML5 uvodi opisnije HTML tagove. To ukljucuje `main`, `header`, `footer`, `nav`, `video`, `article`, `section` i druge.

Ovi tagovi daju jasniju strukturu vasem HTML-u, cine ga citljivijim i pomazu u SEO-u i pristupacnosti. HTML5 `main` tag pomaze pretrazivacima i drugim developerima da pronadju glavni sadrzaj stranice.

Primer upotrebe, `main` element sa dva ugnjezdena child elementa:

```html
<main> 
  <h1>Hello World</h1>
  <p>Hello Paragraph</p>
</main>
```

**Napomena:** Mnogi novi HTML5 tagovi i njihove prednosti obradjuju se u sekciji Applied Accessibility.

# --instructions--

Napravite drugi `p` element sa sledecim kitty ipsum tekstom: `Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.`

Zatim napravite `main` element i unutar njega ugnjezdite samo ta dva `p` elementa.

# --hints--

Treba da imate dva `p` elementa sa Kitty Ipsum tekstom.

```js
assert.lengthOf(document.querySelectorAll('p'),2);
```

Svaki `p` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

Vas drugi `p` element treba da sadrzi prvih nekoliko reci dodatnog `kitty ipsum` teksta.

```js
assert.match(document.querySelectorAll('p')[1].textContent,/Purr\s+jump\s+eat/gi);
```

Vas kod treba da ima jedan `main` element.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

`main` element treba da ima dva paragraf elementa kao decu.

```js
const main = document.querySelector('main');
const children = main.querySelectorAll("p"); 
assert.lengthOf(children,2);
```

Otvarajuci `main` tag treba da bude pre prvog `p` taga.

```js
assert.match(code,(/<main>\s*?<p>/g));
```

Zatvarajuci `main` tag treba da bude posle zatvarajuceg taga drugog paragrafa.

```js
assert.match(code,(/<\/p>\s*?<\/main>/g));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
