---
id: bad87fee1348bd9aedf08830
title: Dodaj Placeholder Tekst u Tekstualno Polje
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cKdJDhg'
forumTopicId: 16647
dashedName: add-placeholder-text-to-a-text-field
---

# --description--

Placeholder tekst je ono sto se prikazuje u vasem `input` elementu pre nego sto korisnik unese bilo kakav podatak.

Placeholder tekst mozete postaviti ovako:

```html
<input type="text" placeholder="this is placeholder text">
```

**Napomena:** Zapamtite da je `input` prazan element.

# --instructions--

Postavite `placeholder` vrednost vaseg tekstualnog `input` elementa na "cat photo URL".

# --hints--

Treba da dodate `placeholder` atribut postojecem tekstualnom `input` elementu.

```js
assert.notEmpty(document.querySelectorAll('input[placeholder]'));
```

Treba da postavite vrednost `placeholder` atributa na `cat photo URL`.

```js
assert.exists(document.querySelector('input'));
assert.exists(document.querySelector('input').getAttribute('placeholder'));
const placeholder = document.querySelector('input').getAttribute('placeholder');
assert.match(placeholder,/cat\s+photo\s+URL/gi);
```

Gotov `input` element ne treba da ima zatvarajuci tag.

```js
assert.notMatch(code,/<input.*\/?>.*<\/input>/gi);
```

Gotov `input` element treba da ima validnu sintaksu.

```js
assert.notEmpty(document.querySelectorAll('input[type=text]'));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text">
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text" placeholder="cat photo URL">
</main>
```
