---
id: bad87fee1348bd9aedf08816
title: Linkuj ka Spoljnim Stranicama Koriscenjem Anchor Elemenata
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/c8EkncB'
forumTopicId: 18226
dashedName: link-to-external-pages-with-anchor-elements
---

# --description--

Mozete koristiti `a` elemente (*anchor tagove*) da linkujete ka sadrzaju izvan vase web stranice.

`a` elementima je potrebna web adresa u `href` atributu. Takodje im je potreban anchor tekst. Evo primera:

```html
<a href="https://www.freecodecamp.org">this links to freecodecamp.org</a>
```

Zatim ce browser prikazati tekst `this links to freecodecamp.org` kao klikabilan link. Taj link ce vas odvesti na web adresu `https://www.freecodecamp.org`.

# --instructions--

Napravite `a` element koji linkuje na `https://www.freecatphotoapp.com` i ima anchor tekst "cat photos".

# --hints--

Vas `a` element treba da ima anchor tekst `cat photos`.

```js
assert.match(document.querySelector('a').textContent,/cat photos/gi);
```

Potreban vam je `a` element koji linkuje na `https://www.freecatphotoapp.com`.

```js
assert.match(document.querySelector('a').getAttribute('href'),/^https?:\/\/(www\.)?freecatphotoapp\.com\/?$/i);
```

Vas `a` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/a>/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>



  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <a href="https://www.freecatphotoapp.com">cat photos</a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
