---
id: 61967e74a8e3690ab6292daa
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Ciljate na element `body` da postavite `background` na ugao kružne linije od 45 stepeni u smeru kazaljke na sat, počevši od `rgb(118, 201, 255)` i završavajući se `rgb(247, 255, 222)`.

# --hints--

Trebalo bi da koristiš ekstrakter atributa `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Trebalo bi da koristiš atribut `background` u selektoru `body`.

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

Trebalo bi da staviš `background` na mesto `linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))`.

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
