---
id: 619665c9abd72906f3ad30f9
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Počnite sa srećnim Flappy Pingvinom, a istražite i više promena za CSS i uzbudljivih crteža.

Počnite sa vašim osnovnim šablonom za HTML. Uključite deklaraciju za `DOCTYPE`, polje `html` sa jezikom postavljenim na engleski, odgovarajuće oznake za `meta`, polje `head`, `body`, i `title`. Takođe, priložite vaš stilski list na stranicu.

# --hints--

Tvoj kod mora imati prefiks `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Vaš kod treba da ima element `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Vaš element za `html` bi trebalo da ima oznaku za otvaranje sa atributom `lang` sa vrednošću `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Element vašeg `html` treba da ima zatvarajući tag.

```js
assert(code.match(/<\/html\s*>/gi));
```

Vaš kod bi trebalo da sadrži element `head` unutar elementa `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Tvoj kod mora imati element `body` unutar elementa `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Vaša komponenta `head` mora da stoji prije komponente `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Trebalo bi da ima dva elementa za `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Jedan element od `meta` mora imati `name` postavljen na `viewport`, i `content` postavljen na `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Drugi element od `meta` mora imati svojstvo `charset` postavljeno na `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Vaš kod mora imati karakteristiku `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` Vaše treba da ima tekst.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Vaš kod mora imati element `link`.

```js
assert.exists(document.querySelector('link'));
```

Tvoj element `link` treba da bude unutar tvog elementa `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

Tvoj element `link` treba da ima svojstvo `rel` sa vrednošću `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Atribut vašeg `link` mora imati svojstvo `href` sa vrednošću `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
