---
id: 5d822fd413a79914d39e98d1
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Postavite svoj element ``.background-buildings` `width` na `height` ya `100%`` ili postavite punu širinu i visinu roditelja, ``body``.

# --hints--

Trebali biste koristiti klasu objekata za `background-buildings` da izaberete pravi element.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

Vaš element za `.background-buildings` treba da ima `width` od `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

Vaš element `.background-buildings` mora imati `height` od `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--
    
```

