---
id: 5d822fd413a79914d39e98d2
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Postavite unutar `div` sa klasom objekta `bb1` u kontejner pozadnje gradnje. Otvorite svoj fajl `styles.css` i dodelite mu `.bb1` klasu objekta `width` od `10%` i `height` od `70%`. "bb" znači "background building", ovo će biti vaša prva zgrada.

# --hints--

Trebalo bi da kreirate novu komponentu za `div`.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

Trebalo bi mu dodeliti novo polje za klasu objekata `div` `bb1`.

```js
assert.exists(document.querySelector('div.bb1'));
```

Trebalo bi da koristiš klasifikator objekata `.bb1` da mu dodaš taj atribut.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

Trebalo bi mu dati atribut klase entiteta `.bb1` od `width` sa `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

Trebalo bi da mu prosledite element `.bb1`, klasu objekata `height` od `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
    
```

