---
id: 5d822fd413a79914d39e98cc
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dobrodošli u projekat CSS Variables Skyline! Počnite dodavanjem elementa `link` koji povezuje vaš fajl `styles.css` unutar elementa `head`.

# --hints--

Tvoj kod mora imati komponentu od `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Ne smete da menjate vaše postojeće oznake za `head`. Uverite se da niste izbrisali oznaku za zatvaranje.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Trebalo bi da imate jedan prazan element `link`.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Tvoj element `link` treba da bude unutar tvog elementa `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Element vašeg `link` bi trebalo da ima svojstvo elementa `rel` sa vrednošću `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Tvoj element za `link` treba da ima svojstvo elemenata `href` sa vrednošću `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
