---
id: 5d822fd413a79914d39e98d6
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Sada imate nešto što izgleda kao struktura. Vi ste spremni da kreirate svoju prvu varijablu. U prethodnim lekcijama ste naučili da deklaracije varijabli počinju sa dve crte (dashes) (`-`) i dodeljuju se ime i vrednost kao što je ovo: `--variable-name: value;`. U sekciji klase `bb1`, kreirajte varijablu pod nazivom `--building-color1` i dodelite joj vrednost `#999`.

# --hints--

Treba da kreiraš novu varijablu nazvanu `--building-color1`.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('--building-color1'));
```

Morate definisati varijablu `--building-color1` unutar `.bb1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
```

Trebalo bi mu dati vrednost `--building-color1` od `#999`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1').trim(), '#999');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}

.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
    
```
