---
id: 5d822fd413a79914d39e98d0
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Kreiraj komponentu za `div` unutar `body` sa klasom objekata `background-buildings`. Ovo će biti kontejner za grupu zgrada.

# --hints--

Trebalo bi da kreirate element za `div`.

```js
assert.exists(document.querySelector('div'));
```

Tvoj element za `div` treba da bude unutar `body`.

```js
assert.equal(document.querySelector('div')?.parentElement?.localName, 'body');
```

Vaša komponenta za `div` treba imati klasu objekata `background-buildings`

```js
assert.isTrue([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

