---
id: 5d822fd413a79914d39e98d3
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Postavite unutar četiri elementa od `div` u kontejneru `.bb1`. Dodeliti klase objekata `bb1a`, `bb1b`, `bb1c`, i `bb1d` u tom redosledu. Ova zgrada će imati četiri dela.

# --hints--

Treba da kreiraš četiri nova elementa za `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

Trebalo bi mu dati jedan od novih elemenata iz klase `div` `bb1a`.

```js
assert.exists(document.querySelector('div.bb1a'));
```

Moraš mu dati jednu od novih karakteristika klase `div` ili klase `bb1b`.

```js
assert.exists(document.querySelector('div.bb1b'));
```

Trebalo bi mu dati jedan od novih atributa klase `div` `bb1c`.

```js
assert.exists(document.querySelector('div.bb1c'));
```

Trebalo bi mu dati jedan od novih atributa klase `div` `bb1d`.

```js
assert.exists(document.querySelector('div.bb1d'));
```

Treba da umetnete nove komponente ``div`` u pravilnom redosledu.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert.isTrue(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

Trebalo bi da staviš nove elemente `div` unutar elementa `.bb1`.

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
    
```

