---
id: 5d822fd413a79914d39e98cf
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Možete videti `body` (to je kutija duboko unutar vaše stranice); okružno okvir je element za `html`. Neka `body` popuni ceo vidljivi prostor dajući mu `height` od `100vh`. Uklonite podrazumevani `margin` sa `body` postavljanjem `margin` na `0`. Na kraju, postavite svojstvo `overflow` na `hidden` da sakrije trake za skrolovanje pregledača vidljive kada nešto nadmašuje vidokrug.

# --hints--

Trebalo bi da koristiš alat `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` vaša treba da ima `height` od `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

`body` tvoje bi trebalo da ima `margin` od `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

`body` mora imati svojstvo `overflow` postavljeno na `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```

