---
id: 5d822fd413a79914d39e98ce
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Dodajte takođe `box-sizing` od `border-box` za sve. Ovo će osigurati da dodata granica ne povećava veličinu vaših elemenata.

# --hints--

Trebalo bi da koristiš atribut `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

Trebalo bi da koristiš postojeći selektor `*`.

```js
// Two selectors create two CSSStyleRule objects
assert.lengthOf(new __helpers.CSSHelp(document).getStyleDeclarations('*'), 1);
```

Svi elementi moraju imati `box-sizing` od `border-box`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```

