---
id: 68e45a2b4ab9b2f48b1cc6df
title: Napravite stranicu za podešavanje cena paketa
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Tvoja stranica treba da sadrži element `h1` sa tekstom `Pricing Plans`.
2. Tvoja stranica treba da sadrži element `div` sa klasom `pricing-container` ispod elementa `h1`.
- Selektor `.pricing-container` treba da ima svojstvo `display` sa vrednošću `flex` i svojstvo `flex-wrap` sa vrednošću `wrap`.
3. Unutar elementa `.pricing-container` treba da postoje tri `div` elementa sa klasom `pricing-card` koji predstavljaju cenovne planove.
- Jedan od elemenata `.pricing-card` treba da ima klasu `basic-plan` zajedno sa klasom `pricing-card`.
- Jedan od elemenata `.pricing-card` treba da ima klasu `pro-plan` zajedno sa klasom `pricing-card`.
- Jedan od elemenata `.pricing-card` treba da ima klasu `premium-plan` zajedno sa klasom `pricing-card`.
4. Element `.basic-plan` treba da sadrži element `h2` sa tekstom `Basic`.
5. Element `.basic-plan` treba da sadrži element `p` sa tekstom `$9/month`.
6. Element `.pro-plan` treba da sadrži element `h2` sa tekstom `Pro`.
7. Element `.pro-plan` treba da sadrži element `p` sa tekstom `$19/month`.
8. Element `.premium-plan` treba da sadrži element `h2` sa tekstom `Premium`.
9. Element `.premium-plan` treba da sadrži element `p` sa tekstom `$29/month`.
10. Svaki element `.pricing-card` treba:
- Da koristi Flexbox za raspoređivanje sadržaja u kolonu i da rasporedi prostor između dece pomoću `space-between`.
- Da postavi svojstvo `flex` na `0 0 200px` kako bi imao fiksnu širinu i sprečio širenje ili skupljanje u rasporedu.
- Da postavi svojstvo `border` na `2px solid black` kako bi se jasno videlo kako kartice zauzimaju prostor.
11. Element `.basic-plan` treba da bude prvi u rasporedu. Za to koristi svojstvo `order`.
12. Element `.pro-plan` treba da bude drugi u rasporedu. Koristi svojstvo `order` i postavi njegovo svojstvo `flex-grow` na `2` kako bi zauzimao više prostora od ostalih planova.
13. Element `.premium-plan` treba da bude poslednji u rasporedu. Za to koristi svojstvo `order`.

**Napomena:** Uverite se da ste povezali svoj stilski list na HTML i koristili CSS.

# --hints--

Trebalo bi da imaš element `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Stranica vam treba da ima element sa tekstom `Pricing Plans`, koji je `h1`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

Trebalo bi da imaš element `div` ispod elementa `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Tvoj element za `div` treba da ima klasu objekata `pricing-container`.

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

Trebalo bi da imaš birač za `.pricing-container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

Tvoj selektor za `pricing-container` mora imati svojstvo `display` sa vrednošću `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

Vaš selektor za `.pricing-container` mora imati atribut `flex-wrap` sa vrednošću `wrap`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

Treba da imaš tri `div` elementa sa klasom `pricing-card` unutar elementa `.pricing-container`.

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

Komponenta vaših `.pricing-card` treba da ima klasu objekata od `basic-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

Komponenta vaših elemenata za `.pricing-card` treba imati klasu objekata za `premium-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

Komponenta vaših `.pricing-card` treba da ima klasu od `pro-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

Trebalo bi da imate birač za `.pricing-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

Vaš selektor za `.pricing-card` mora imati svojstvo `display` sa vrednošću `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

Vaš selektor za `.pricing-card` mora imati svojstvo od `flex-direction` sa vrednošću `column`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

Tvoj selector za `.pricing-card` mora imati svojstvo `justify-content` sa vrednošću `space-between`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

Tvoj selektor za `.pricing-card` mora imati svojstvo `flex` sa vrednošću `0 0 200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

Vaš selektor za `.pricing-card` mora imati svojstvo `border` sa vrednošću `2px solid black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

Tvoj selektor za `.basic-plan` mora imati svojstvo `order` sa vrednošću `0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

Vaš element za `.basic-plan` treba da ima element `h2` sa tekstom `Basic`.

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

Komponenta za vaš `.basic-plan` treba da ima komponentu `p` sa tekstom `$9/month`.

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

Vaš element za `.pro-plan` treba da ima komponentu `h2` sa tekstom od `Pro`.

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

Vaš element od `.pro-plan` treba da ima element od `p` sa tekstom `$19/month`.

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

Vaš atribut `.pro-plan` mora imati svojstvo `order` sa vrednošću od `1`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

Vaše polje `.pro-plan` mora imati svojstvo `flex-grow` sa vrednošću `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

Vaš element od `.premium-plan` mora da ima element od `h2` sa tekstom `Premium`.

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

Vaš element od `.premium-plan` mora imati element od `p` sa tekstom `$29/month`.

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

Vaš element za `premium-plan` treba biti poslednji element u nizu i imati svojstvo `order` sa vrednošću `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```
