---
id: 5e44414f903586ffb414c950
title: Izgradite projekat kalkulatora verovatnoće
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462364
dashedName: build-a-probability-calculator-project
---

# --description--

Zamislite da imate šeširi sa 5 plavih, 4 crvenih i 2 zelenih kuglica. Koja je verovatnoća da će nasumično izvlačenje 4 kuglice sadržati najmanje 1 crvenu kruglicu i 2 zelene? Iako je moguće izračunati verovatnoću korišćenjem napredne matematike, jednostavniji način je napisati program koji će obaviti mnogo eksperimenata kako bi procenio približnu verovatnoću.

Za ovaj projekat, treba da napišete program za određivanje približne verovatnoće izvlačenja određenih kuglica nasumično iz šešira.

Prvo, kreirajte klasu `Hat` u fajlu `main.py`. Ova klasa bi trebalo da prima nepoznat broj argumenata koji definišu broj kuglica svake boje u šeširu. Na primer, objekat klase može biti kreiran na bilo koji od sledeći načini:

```py
hat1 = Hat(yellow=3, blue=2, green=6)
hat2 = Hat(red=5, orange=4)
hat3 = Hat(red=5, orange=4, black=1, blue=0, pink=2, striped=9)
```

Šešir se uvek kreira sa najmanje jednom kuglicom. Argumenti prosleđeni objektu šešira tokom inicijalizacije treba da se konvertuju u atribut pod nazivom `contents`. `contents` bi trebalo da bude lista nizova karaktera, gde svaki element predstavlja jednu kuglicu u šeširu. Svaki element u listi mora biti ime boje koje predstavlja jednu kuglicu te boje. Na primer, ako je vaš šešir `{'red': 2, 'blue': 1}`, `contents` bi trebalo da bude `['red', 'red', 'blue']`.

Klasa `Hat` treba da ima metodu `draw` koja prima argument koji definiše broj kuglica za izvlačenje iz šešira. Ova metoda bi trebalo da nasumično ukloni kuglice iz `contents` i vrati te kuglice kao listu nizova karaktera. Kuglice ne bi trebalo da se vraćaju u šešir tokom izvlačenja, kao eksperiment sa predmetima bez vraćanja. Ako je broj izvlačene kuglica veći od dostupnog, vratite sve kuglice.

Zatim, kreirajte funkciju `experiment` u fajlu `main.py` (ne unutar klase `Hat`). Ova funkcija bi trebalo da prima sledeće argumente:

- `hat`: Objekt šešira sa kuglicama koji treba simulirati unutar funkcije.
- `expected_balls`: Objekat koji predstavlja stvarni skup kuglica koje pokušavate da izvučete iz šešira za eksperimente. Na primer, da biste odredili verovatnoću izvlačenja 2 plave i 1 crvenu kruglicu iz šešira, postavite `expected_balls` na `{'blue':2, 'red':1}`.
- `num_balls_drawn`: Broj kuglica za izvlačenje iz šešira u svakom eksperimentu.
- `num_experiments`: Broj eksperimenata koje treba obaviti. (Što je veći broj eksperimenata, to je približna verovatnoća preciznija.)

Funkcija `experiment` bi trebalo da vrati verovatnoću.

Na primer, ako želite da odredite verovatnoću dobijanja najmanje dve crvene i jednu zelenu kruglicu prilikom izvlačenja pet kuglica iz šešira koji sadrži šest crnih, četiri crvene i tri zelene kuglice. Za to ćete obaviti `N` eksperimenata, prebrojavajući koliko ste puta `M` dobili najmanje dve crvene i jednu zelenu kruglicu, i procenjujući verovatnoću kao `M/N`. Svaki eksperiment uključuje početak sa šeširom sa definisanim kuglicama, izvlačenje nekoliko kuglica i proveru da li ste dobili kuglice koje ste pokušavali.

Evo kako biste napisali funkciju `experiment` na osnovu gornjeg primera i 2000 eksperimenata:

```py
hat = Hat(black=6, red=4, green=3)
probability = experiment(hat=hat,
                  expected_balls={'red':2,'green':1},
                  num_balls_drawn=5,
                  num_experiments=2000)
```

Rezultat će biti sličan ovome:

```bash
0.356
```

Pošto zavisi od nasumičnog izvlačenja, verovatnoća će varirati svaki put kada se kod pokrene.

_Napomena: Razmislite o korišćenju modula koji su već uvezeni na vrhu. Nemojte inicijalizovati random seed unutar fajla._

Zapamti: otvorite konzolu pretraživača pritiskom tastera F12 da vidite detaljne rezultate eksperimenta.

# --hints--

Kreiranje objekta `hat` bi trebalo da doda tačan sadržaj.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_class_contents(self):
        hat = probability_calculator.Hat(red=3,blue=2)
        actual = hat.contents
        expected = ["red","red","red","blue","blue"]
        self.assertEqual(actual, expected, 'Expected creation of hat object to add correct contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Metoda `draw` u klasi `hat` bi trebalo da smanji broj elemenata u sadržaju.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)

class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw(self):
        hat = probability_calculator.Hat(red=5,blue=2)
        actual = hat.draw(2)
        expected = ['blue', 'red']
        self.assertEqual(actual, expected, 'Expected hat draw to return two random items from hat contents.')
        actual = len(hat.contents)
        expected = 5
        self.assertEqual(actual, expected, 'Expected hat draw to reduce number of items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Metoda `draw` bi trebalo da radi ispravno kada je broj izvlačenih kuglica veći od broja kuglica u šeširu.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw_2(self):
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        actual = sorted(hat.draw(20))
        expected = sorted(['yellow', 'yellow', 'yellow', 'yellow', 'yellow', 'red', 'green', 'green', 'green', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'test'])
        self.assertEqual(actual, expected, 'Expected hat draw to return all items from hat contents.')
        actual = len(hat.contents)
        expected = 0
        self.assertEqual(actual, expected, 'Expected hat draw to leave no items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Metoda `experiment` bi trebalo da vrati različitu verovatnoću.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_prob_experiment(self):
        hat = probability_calculator.Hat(blue=3,red=2,green=6)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"blue":2,"green":1}, num_balls_drawn=4, num_experiments=1000)
        actual = probability
        expected = 0.272
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected