---
id: 66fbcf750a62784cf11f562d
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Kreiraj objekat nazvan `recipe1`. Unutar objekta `recipe1`, kreiraj svojstvo `name` sa vrednošću `"Spaghetti Carbonara"`.

I unutar objekta `recipe1`, kreiraj svojstvo `ingredients` sa vrednošću niza podataka. Taj niz podataka treba da sadrži `spaghetti`, `Parmesan cheese`, `pancetta` i `black pepper`.

# --hints--

Trebalo bi da napraviš nešto što se zove `recipe1`.

```js
assert.isObject(recipe1);
```

Vaš `recipe1` mora imati svojstvo `name`.

```js
assert.property(recipe1, 'name');
```

Tvoja vrednost za `name` treba da bude postavljena na `Spaghetti Carbonara`.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

Vaš element sa `recipe1` treba da ima svojstvo od `ingredients`.

```js
assert.property(recipe1, 'ingredients');
```

Tvoja vrednost za `ingredients` mora imati rednu vrednost podataka.

```js
assert.isArray(recipe1.ingredients);
```

Vrednost niza podataka vašeg atributa `ingredients` treba da sadrži `spaghetti`, `Parmesan cheese`, `pancetta`, unutar `black pepper`.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
