---
id: 66fbcf750a62784cf11f562e
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Kreiraj objekat za `recipe2` sa sledećim svojstvima i vrednostima:

| Ključ | Vrednost |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

Trebalo bi da kreirate nešto što se zove `recipe2`.

```js
assert.isObject(recipe2);
```

Vas deo za `recipe2` treba da ima osobinu od `name`.

```js
assert.property(recipe2, "name");
```

Tvoja vrijednost za `name` treba da bude postavljena na `Chicken Curry`.

```js
assert.equal(recipe2.name, "Chicken Curry");
```

Tvoj `recipe2` mora imati karakteristiku od `ingredients`.

```js
assert.property(recipe2, "ingredients");
```

Vaše svojstvo `ingredients` mora imati vrednost podatkovnog niza.

```js
assert.isArray(recipe2.ingredients);
```

Vrednost podataka vašeg atributa za `ingredients` treba da sadrži `chicken breast`, `coconut milk`, `curry powder`, `onion` i `garlic` unutar njega.

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

Vaš predmet sa `recipe2` trebao bi da ima svojstvo `cookingTime`.

```js
assert.property(recipe2, "cookingTime");
```

Vrednost vašeg atributa `cookingTime` mora biti broj.

```js
assert.isNumber(recipe2.cookingTime);
```

Tvoja vrednost za `cookingTime` bi trebalo da bude postavljena na `42`.

```js
assert.equal(recipe2.cookingTime, 42);
```

Vaš predmet od `recipe2` treba da ima svojstvo `totalIngredients`.

```js
assert.property(recipe2, "totalIngredients");
```

Tvoja vrednost za `totalIngredients` treba da bude postavljena na `null`.

```js
assert.isNull(recipe2.totalIngredients);
```

Твој део `recipe2` мора имати карактеристику `difficultyLevel`.

```js
assert.property(recipe2, "difficultyLevel");
```

Vaš atribut `difficultyLevel` treba da bude postavljen na prazan niz.

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
  };

--fcc-editable-region--

--fcc-editable-region--
```
