---
id: 66fbcf750a62784cf11f5630
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Sada treba da gurate tri stavke u niz podataka `recipes`. Da biste to uradili, možete koristiti metodu `push()`.

Koristite putanju `push()` za pushovanje svih recepata u niz podataka `recipes`. Uverite se da pushujete `recipe1`, `recipe2` i `recipe3` u tom redosledu.

Takođe obrišite parametre `recipe1Name`, `recipe2CookingTime` na `recipe3Ingredients`, zajedno sa informacijama o `console.log` koje prikazuju te parametre.

# --hints--

Trebalo bi da ubaciš sve sastojke za recepte unutar niza podataka `recipes`. Osiguraj da ih ubacuješ po redosledu objavljivanja.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

Trebalo bi da obrišete parametar `recipe1Name`.

```js
assert.throws(() => recipe1Name, ReferenceError);
```

Trebalo bi da obrišeš informaciju `console.log(recipe1Name);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

Trebalo bi da obrišete parametar `recipe2CookingTime`.

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

Trebalo bi da obrišeš informaciju `console.log(recipe2CookingTime);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

Treba da obrišeš parametar `recipe3Ingredients`.

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

Treba da obrišeš informaciju `console.log(recipe3Ingredients);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
