---
id: 66fbcdfed6848e48947dbbc9
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, kreiraćete pratioca recepata koristeći stvari od JavaScript.

Počnite kreiranjem praznog reda podataka pod nazivom `recipes`. Ovo je red podataka gde ćete kasnije uneti stavke recepta.

# --hints--

Trebalo bi da kreirate tabelu podataka nazvanu `recipes`.

```js
assert.isArray(recipes);
```

Vaš red podataka za `recipes` treba da bude prazan.

```js
assert.isEmpty(recipes);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
