---
id: 66fbcf750a62784cf11f5632
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Kreiraj funkciju `getTotalIngredients` koja prima jedan argument, koji predstavlja niz poveznih podataka, i vraća broj poveznih elemenata iz niza podataka poslatog funkciji.

# --hints--

Trebalo bi da kreirate funkciju za `getTotalIngredients`.

```js
assert.isFunction(getTotalIngredients);
```

Твоја функција за `getTotalIngredients` мора имати један параметар.

```js
assert.lengthOf(getTotalIngredients, 1);
```

Tvoja funkcija za `getTotalIngredients` treba da vrati broj.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients));
```

`getTotalIngredients(recipe1.ingredients)` treba da vrati `4`.

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)` treba vratiti `5`.

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)` trebalo bi da vrati `3`.

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
