---
id: 66fbcf750a62784cf11f5634
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Nije vreme za testiranje svake funkcije. Možete koristiti bilo koji recept za ovo, ali će ova obuka početi sa `recipe1`.

Napravite dva nova parametra: `recipe1TotalIngredients` i `recipe1DifficultyLevel`. Postavite imanje vrednosti pozivanjem odgovarajuće funkcije za svaki parametar i prosleđujući neophodni atribut `recipe1`.

Na kraju, napišite svaki parametar u konzolu da biste videli rezultate.

# --hints--

Treba da kreiraš parametar za `recipe1TotalIngredients`.

```js
assert.isNotNull(recipe1TotalIngredients);
```

Vaš parametar za `recipe1TotalIngredients` treba biti postavljen na vrednost `getTotalIngredients(recipe1.ingredients)`.

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

Treba da ukucate `recipe1TotalIngredients` u konzolu.

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/);
```

Treba da kreiraš parametar `recipe1DifficultyLevel`.

```js
assert.isNotNull(recipe1DifficultyLevel);
```

Vaš parametar za `recipe1DifficultyLevel` mora biti postavljen na vrednost `getDifficultyLevel(recipe1.cookingTime)`.

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

Trebalo bi da napišete `recipe1DifficultyLevel` u konzoli.

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
