---
id: 66fbcf750a62784cf11f5633
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Kreira funkciju koja `getDifficultyLevel` prima broj koji određuje vreme kuvanja kao argument.

Ako je vreme kuvanja manje ili jednako od `30`, funkcija treba da vrati `"easy"`. Ako je manje ili jednako od `60`, funkcija treba da vrati `"medium"`. Inače, funkcija treba da vrati `"hard"`.

# --hints--

Treba da kreiraš funkciju za `getDifficultyLevel`.

```js
assert.isFunction(getDifficultyLevel);
```

Tvoja funkcija za `getDifficultyLevel` bi trebalo da ima jedan parametar.

```js
assert.lengthOf(getDifficultyLevel, 1);
```

Vaš uređaj za `getDifficultyLevel` treba da vrati `"easy"` kada je vreme kuvanja manje ili jednako od `30`.

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

Tvoja funkcija od `getDifficultyLevel` treba da vrati `"medium"` kada je vreme kuvanja veće od `31` i manje ili jednako `60`.

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

Tvoja funkcija za `getDifficultyLevel` mora da vrati `"hard"` kada je vreme kuvanja veće od `60`.

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
