---
id: 66fbcf750a62784cf11f5635
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Sada možete popuniti svaki element kolone podataka `recipes` sa vrednostima atributa `totalIngredients` i `difficultyLevel`.

Trenutno, dobiti pristup `totalIngredients` i `difficultyLevel` za `recipe1`, pa postaviti njihove vrednosti za odgovarajuće rezultate funkcionalskih poziva i argumenata.

# --hints--

Trebaće vam pristup atributu `totalIngredients` od `recipe1`.

```js
assert.isNotNull(recipe1.totalIngredients);
```

Trebalo bi postaviti vrednost rezultata poziva `getTotalIngredients` sa `recipe1.ingredients` za svojstvo `totalIngredients` od `recipe1`.

```js
assert.strictEqual(recipe1.totalIngredients, getTotalIngredients(recipe1.ingredients));
```

Trebalo bi da dobijete pristup atributu `difficultyLevel` od `recipe1`.

```js
assert.isNotEmpty(recipe1.difficultyLevel);
```

Morate postaviti vrednost rezultata poziva `getDifficultyLevel` sa `recipe1.cookingTime` za svojstvo `cookingTime` od `recipe1`.

```js
assert.strictEqual(recipe1.difficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

--fcc-editable-region--

--fcc-editable-region--
```
