---
id: 66ff02792b3042a2b811f475
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Dodajte sledeće atribute na objekat `recipe1`:

| Funkcija | Vrednost |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

Atributi koji imaju `null` i vrednosti za praznu sekvencu karaktera biće ažurirani kasnije nakon brojanja.

# --hints--

Tvoj `recipe1` mora imati svojstvo `cookingTime`.

```js
assert.property(recipe1, "cookingTime");
```

Tvoja vrednost za `cookingTime` treba da bude postavljena na `22`.

```js
assert.equal(recipe1.cookingTime, 22);
```

Vaš element za `recipe1` treba da ima svojstvo od `totalIngredients`.

```js
assert.property(recipe1, "totalIngredients");
```

Tvoja vrednost za `totalIngredients` treba da bude postavljena na `null`.

```js
assert.isNull(recipe1.totalIngredients);
```

Vaš element za `recipe1` mora imati svojstvo od `difficultyLevel`.

```js
assert.property(recipe1, "difficultyLevel");
```

Vaš atribut `difficultyLevel` treba da bude postavljen na prazan niz.

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
