---
id: 670e4f45f7116c0f216a5177
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Pre nego što nastaviš, treba da vežbaš kako da dobiješ karakteristike od nečega.

Možete koristiti tačku (`.`) ili znak zagrade (`[]`) da ovo uradite. Evo primera:

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Dobij svojstvo `name` od `recipe1`, i postavi njegovu vrednost u izmene `recipe1Name`.

Zatim, dobiti svojstvo `cookingTime` od `recipe2` i postaviti njegovu vrednost u izmene `recipe2CookingTime`.

Konačno, dobiti svojstvo `ingredients` od `recipe3` i postaviti njegovu vrednost u izmene `recipe3Ingredients`.

Uverite se da su sve promene koje ste napravili/kreirali zapisane na konzoli.

# --hints--

Trebalo bi da napraviš izmene `recipe1Name`.

```js
assert.isDefined(recipe1Name);
```

Trebalo bi da postaviš vrednost atributa `name` na `recipe1` u promenljivu `recipe1name`.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

Treba da napišete `recipe1Name` u konzoli.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

Trebalo bi da napravite promene `recipe2CookingTime`.

```js
assert.isDefined(recipe2CookingTime);
```

Trebalo bi da postaviš vrednost atributa `cookingTime` za `recipe2` u promenljivu `recipe2CookingTime`.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

Morate napisati `recipe2CookingTime` u konzolu.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

Trebalo bi da napraviš izmene `recipe3Ingredients`.

```js
assert.isDefined(recipe3Ingredients);
```

Trebalo bi da postaviš vrednost atributa `ingredients` na `recipe3` u promenljivoj `recipe3Ingredients`.

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

Trebalo bi da ukucate `recipe3Ingredients` u konzolu.

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
