---
id: 69161f4c0c9c8c48b3c843c5
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Kreiraj predmet za `recipe3` sa sledećim svojstvima i vrednostima:

|Key|Value|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

Trebalo bi da kreirate nešto nazvano `recipe3`.

```js
assert.isObject(recipe3);
```

Vaš predmet sa `recipe3` mora imati svojstvo od `name`.

```js
assert.property(recipe3, "name");
```

Tvoja vrednost za `name` mora biti postavljena na `Vegetable Stir Fry`.

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

Tvoj predmet za `recipe3` mora imati atribut od `ingredients`.

```js
assert.property(recipe3, "ingredients");
```

Tvoja vrednost za `ingredients` mora imati rednu vrednost podataka.

```js
assert.isArray(recipe3.ingredients);
```

Vrednost reda podataka vašeg atributa `ingredients` mora da sadrži `broccoli`, `carrot` i `bell pepper`.

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

Vaš `recipe3` mora imati svojstvo od `cookingTime`.

```js
assert.property(recipe3, "cookingTime");
```

Vrednost vašeg atributa `cookingTime` mora biti broj.

```js
assert.isNumber(recipe3.cookingTime);
```

Tvoja vrednost za `cookingTime` treba da bude postavljena na `15`.

```js
assert.equal(recipe3.cookingTime, 15);
```

Tvoj element sa `recipe3` mora imati svojstvo `totalIngredients`.

```js
assert.property(recipe3, "totalIngredients");
```

Tvoja vrednost za `totalIngredients` treba da bude postavljena na `null`.

```js
assert.isNull(recipe3.totalIngredients);
```

Твој садржај `recipe3` треба да има карактеристику `difficultyLevel`.

```js
assert.property(recipe3, "difficultyLevel");
```

Vaš atribut `difficultyLevel` treba da bude postavljen na prazan niz.

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
