---
id: 66ada3f46945763dd97f43f8
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Ranije ste kreirali varijable `bot` i `botLocation`. Sada ćete ih koristiti za ispis novih poruka na konzolu.

U prethodnim lekcijama, naučili ste kako da radite sa konkatenacijom stringova koristeći operator `+` za spajanje stringova zajedno na ovaj način:

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

Zapamtite da morate biti pažljivi sa razmakama prilikom konkatenacije stringova sa varijablama.

Kreirajte varijablu nazvanu `botIntroduction`.

Zatim koristite konkatenaciju stringova sa operatorom `+` da spojite string `"My name is "`, praćen varijablom `bot`, a zatim tačkom (`.`).

Dodelite ovu vrednost varijabli `botIntroduction`.

Zatim, logujte (ili: Ispišite) varijablu `botIntroduction` na konzoli.

# --hints--

Trebalo bi da imate varijablu pod nazivom `botIntroduction`.

```js
assert.isNotNull(botIntroduction);
```

Vaša promenljiva `botIntroduction` bi trebalo da drži vrednost stringa.

```js
assert.isString(botIntroduction);
```

Trebalo bi da koristite konkatenaciju niza sa `+` operaterom kako biste spojili niz `"My name is "` sa varijablom `bot`, praćen tačkom (`.`). Pazite na razmake. Dodelite spojeni niz this varijabli `botIntroduction`.

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

Trebalo bi da imate četvrti `console.log()` izraz u svom kodu.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

Vaša izjava `console` treba da izbaci varijablu `botIntroduction`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
