---
id: 66adc6046acc18453decf577
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

Da biste videli novi nadimak bota, moraćete da logujete poruku u konzolu.

Kreirajte novu varijablu nazvanu `newNicknameGreeting`.

Nakon toga koristite konkatenaciju niza da spojite niz `"I love my nickname but I wish people would call me "` sa varijablom `bot`, praćen tačkom.

Dodeli rezultat varijable `newNicknameGreeting`.

Zatim, logujte vrednost `newNicknameGreeting` u konzolu.

# --hints--

Trebalo bi da imate varijablu nazvanu `newNicknameGreeting`.

```js
assert.isNotNull(newNicknameGreeting);
```

Vaša promenljiva `newNicknameGreeting` treba da bude string.

```js
assert.isString(newNicknameGreeting);
```

Trebalo bi da koristite konkatenaciju nizova za spajanje niza `"I love my nickname but I wish people would call me "` sa varijablom `bot`, a zatim tačkom.

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

Trebalo bi da imate izjavu `console` u svom kodu.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

Trebalo bi da se loguje vrednost `newNicknameGreeting` na console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
