---
id: 66adcf383276814776aba3ca
title: Korak 14
challengeType: 1
dashedName: step-14
---

# --description--

Zatim, kreirajte varijablu nazvanu `favoriteSubjectSentence`.

Koristite konkatenaciju stringova da spojite string `"My favorite subject is "` sa varijablom `favoriteSubject`, praćenim tačkom.

Dodeliti rezultat varijable `favoriteSubjectSentence`.

Zatim, logujte vrednost `favoriteSubjectSentence` u konzolu.

# --hints--

Trebalo bi da imate varijablu pod nazivom `favoriteSubjectSentence`.

```js
assert.isNotNull(favoriteSubjectSentence);
```

Vaša varijabla `favoriteSubjectSentence` treba da bude niz karaktera (string).

```js
assert.isString(favoriteSubjectSentence);
```

Trebalo bi da koristite konkatenaciju stringova za spajanje stringa `"My favorite subject is "` sa varijablom `favoriteSubject`, a zatim tačkom. Pazite na razmake.

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

Trebalo bi da imate izjavu `console` u svom kodu.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

Trebalo bi da zapisujete vrednost `favoriteSubjectSentence` u console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
