---
id: 66ad984e92ce213a66f7e51c
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada je vreme da dodelite vrednost vašoj promenljivoj `bot`.

U prethodnim lekcijama, naučili ste kako da dodeljujete vrednosti promenljivama na ovaj način:

```js
variableName = "Here is the value";
```

Zapamtite da je ono što se nalazi sa desne strane operatera dodeljivanja `=` vrednost koju dodeljujete varijabli na levoj strani.

Dodelite string `"teacherBot"` promenljivoj `bot`.

# --hints--

Trebalo bi da dodelite string vašoj varijabli `bot`.

```js
assert.isString(bot);
```

Vaša varijabla `bot` treba da dobije niz `"teacherBot"`.

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
