---
id: 66adc42868cab843ccee87d9
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Sada je vreme da vidimo novu vrednost `bot`.

Počnite kreiranjem nove varijable nazvane `nicknameIntroduction`.

Koristite konkatenaciju nizova da spojite niz `"My nickname is "` sa varijablom `bot`, praćenim tačkom (`.`).

Dodelite rezultatni niz varijabli `nicknameIntroduction`.

Zatim, logujte vrednost `nicknameIntroduction` u konzolu.

# --hints--

Trebalo bi da imate varijablu nazvanu `nicknameIntroduction`.

```js
assert.isNotNull(nicknameIntroduction);
```

Vaša promenljiva `nicknameIntroduction` bi trebalo da bude niz karaktera (string).

```js
assert.isString(nicknameIntroduction);
```

Trebalo bi da koristite konkatenaciju stringa za spajanje stringa `"My nickname is "` sa varijablom `bot`, a zatim tačkom (`.`). Obratite pažnju na razmake.

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

Trebalo bi da imate izjavu `console` u svom kodu.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

Trebalo bi da snimite vrednost `nicknameIntroduction` u console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
