---
id: 66ad998285df023b23dacdd3
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Sada je vreme da se inicijalizuje promenljiva `botLocation`.

Kada morate deklarisati varijable sa više reči, možete koristiti konvenciju imenovanja <dfn>camelCase</dfn>.

Kada koristite `camelCase`, prva reč je u malim slovima, a prvi slovo svake naknadne reči je veliko.

Evo primera:

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

Deklarišite i dodelite string `"the universe"` promenljivoj `botLocation` na istoj liniji koristeći ključnu reč `let`.

# --hints--

Koristite ključnu reč `let` za deklarisanje varijable.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Vaša varijabla bi trebalo da se zove `botLocation`.

```js
assert.isNotNull(botLocation);
```

Trebalo bi da dodelite string varijabli vašeg `botLocation`.

```js
assert.isString(botLocation);
```

Vaša promenljiva `botLocation` treba da dobije niz `"the universe"`.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
