---
id: 66ad9c633fa26d3c1475eae3
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Sada je vreme da dodamo poruku od bota.

Dodajte još jednu izjavu ``console`` u kod koji loguje poruku ``"Allow me to introduce myself."``.

# --hints--

Trebalo bi da imate treći `console.log()` izraz u vašem kodu.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

Vaša `console` izjava bi trebalo da ima poruku `"Allow me to introduce myself."`. Nemojte zaboraviti navodnike.

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
