---
id: 66ad8294a0ad902f1b31b612
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

Na ovom radionici ćete naučiti kako da radite sa osnovama JavaScript-a kreiranjem bota za pozdrave.

U ovom prvom koraku, želite da izbacite poruku u konzolu sa pozdravnog bota.

Zapamtite da ste u prethodnim lekcijama naučili o `console.log()` i stringovima.

Evo podsetnika kako koristiti `console.log()` sa stringovima:

```js
console.log("Hello, World!");
```

Dodajte izjavu `console.log()` koja ispisuje string `"Hi there!"` u konzolu. Ne zaboravite navodnike oko poruke!

# --hints--

Trebalo bi da imate izjavu `console.log()` u svom kodu.

```js
assert.match(__helpers.removeJSComments(code), /console\.log(.*)/);
```

Tvoj ``console`` izjava treba da ima poruku ``"Hi there!"``. Ne zaboravi navodnike.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((['"])Hi\s+there!\1\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
