---
id: 66ad8ab945d266383f318cbf
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Sada bi trebalo da vidite prvu poruku od bota u konzoli.

Vrijeme je da se doda druga poruka od bota.

Dodajte još jednu izjavu ``console.log`` da bi se poruka ``"I am excited to talk to you."`` ispisala u konzolu.

# --hints--

Trebalo bi da imate drugu izjavu `console.log()` u svom kodu.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

Vaša izjava `console` treba da ima poruku `"I am excited to talk to you."`. Ne zaboravite navodnike.

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
