---
id: 66ad8d150264db3926eccfeb
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

U prethodnim časovima, naučili ste o ključnoj riječi `let` i kako da deklarišete varijable u JavaScriptu.

Evo podsjetnika kako da deklarišete varijablu koristeći ključnu riječ `let`:

```js
let greeting;
```

Koristite ključnu riječ `let` da deklarišete varijablu nazvanu `bot`.

**NAPOMENA**: Koristite `let` ovdje jer kasnije u radionici mijenjate vrijednost varijable `bot`.

# --hints--

Trebalo bi da imate ključnu reč `let` u vašem kodu.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

Trebalo bi da deklariš promenljivu nazvanu `bot`.

```js
assert.isNotNull(bot);
```

Trebalo bi da koristite ključnu reč `let` za deklarisanje varijable.

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
