---
id: 66adb844118ba74107ce771f
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Sledeća poruka od bota će se tičati lokacije bota.

Kreirajte varijablu nazvanu `botLocationSentence`.

Zatim koristite konkatenaciju niza sa operaterom `+` da spojite niz `"I live in "` sa varijablom `botLocation` praćenom tačkom (`.`).

Dodelite ovu vrednost varijabli `botLocationSentence`.

Zatim, logujte vrednost `botLocationSentence` u konzolu.

# --hints--

Trebalo bi da imate varijablu `botLocationSentence`.

```js
assert.isNotNull(botLocationSentence);
```

Vaša `botLocationSentence` varijabla treba da sadrži vrednost stringa.

```js
assert.isString(botLocationSentence);
```

Trebalo bi da koristite konkatenaciju stringa sa operatorom `+` za spajanje stringa `"I live in "` sa varijablom `botLocation`, praćenog tačkom (`.`). Obratite pažnju na razmake.

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

Trebalo bi da imate izjavu `console` u svom kodu.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

Vaš ``console`` izraz bi trebalo da izbaci varijablu ``botLocationSentence``.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
