---
id: 61537c5f81f0cf325b4a854c
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Dodaj komponentu `header` unutar komponente `body` i stavi klasu elemenata `header` u nju.

Unutar `header`, generiši `h1` sa `css flexbox photo gallery` kao tekst.

# --hints--

Trebalo bi da imaš element `header` unutar tvog elementa `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Vaše polje za `header` treba da ima `class` sa vrednošću `header`.

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Tvoj element `header` bi trebalo da sadrži element `h1` unutar sebe.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

Element vašeg `h1` treba da sadrži tekst od `css flexbox photo gallery` unutar sebe.

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
