---
id: 6153893900438b4643590367
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Poravnajte svoj tekst `.header` po sredini. Postavite tekst u velika slova koristeći svojstvo `text-transform` i vrednost `uppercase`.

Postavi unutrašnji prostor za `32px` sa svih strana. Postavi pozadinsku boju na `#0a0a23` i boju teksta na `#fff` kao vrednosti boja.

Iznos `border-bottom` sa vrednošću od `4px solid #fdb347`.

# --hints--

Trebalo bi da imaš birač za `.header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

Vaš selektor za `.header` mora imati svojstvo `text-align` postavljeno na `center` kao vrednost.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

Vaš selektor za `.header` mora imati svojstvo `text-transform`, postavljeno na `uppercase` kao vrednost.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

Vaš selektor za `.header` mora imati svojstvo `padding` postavljeno na `32px` kao vrednost.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

Tvoj selektor za `.header` mora imati atribut `background-color` postavljen na `#0a0a23` kao vrednost.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

Tvoj selektor za `.header` mora imati svojstvo `color` postavljeno na vrednost `#fff`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

Vaš selektor za `.header` mora imati atribut `border-bottom`, postavljen na `4px solid #fdb347` kao vrednost.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
