---
id: 61537a8054753e2f1f2a1574
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Unutar vašeg elementa `head`, dodajte oznaku/etiketu `meta` sa svojstvom elementa `name` postavljeno na `viewport` i svojstvo elementa `content` postavljeno na `width=device-width, initial-scale=1`.

Takođe dodaj oznaku za `meta` sa atributom elementa `charset` postavljenu na `UTF-8`.

# --hints--

Trebalo bi da ima dva elementa za `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Jedna komponenta sa `meta` treba da ima atribut komponente `name` postavljen na `viewport`, i `content` postavljen na `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

Drugi element od `meta` mora imati svojstvo elementa `charset` postavljeno na `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
