---
id: 61539f32a206bd53ec116465
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Primećujete kako su neke od vaših slika počele da se degradiraju. To je zato što slike imaju različite omore lica. Umesto da postavljate omor lica za svaku sliku pojedinačno, možete koristiti atribut `object-fit` da odredite kako bi slike trebalo da izgledaju.

Postavi svoj selektor za `.gallery img`, svojstvo `object-fit` i postavi ga na `cover`. Ovo će osigurati da slika popuni kontejner `img`, zadržavajući omjer stranica i uzrokujući rezanje kako bi odgovaralo.

# --hints--

Tvoj selektor za `.gallery img` mora imati svojstvo `object-fit` postavljeno na `cover` kao vrednost.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.objectFit === 'cover');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}

--fcc-editable-region--
.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
--fcc-editable-region--
```
