---
id: 61537c9eecea6a335db6da79
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Ispod vašeg elementa sa `.header`, kreirajte novi element `div` i postavite vrednost za `class` na `gallery`. Ovaj `div` će se koristiti kao kontejner za slike biblioteke.

Unutar tog elementa od `.gallery`, kreiraj devet elemenata od `img`.

# --hints--

Treba da napraviš element `div` unutar svog elementa `body`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

Vaš novi element za `div` mora imati `class` sa vrednošću od `gallery`.

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

Vaš novi element za `div` treba da dođe nakon vašeg elementa za `.header`.

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

Vaš deo `.gallery` treba da sadrži devet elemenata `img`.

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
