---
id: 61537bb9b1a29430ac15ad38
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Unutar vašeg elementa `head`, dodajte element `title` sa tekstom postavljenim na `Photo Gallery`, i element `link` za učitavanje vašeg fajla `styles.css` na stranicu.

# --hints--

Vaš element za `link` mora imati atribut elementa `href` sa vrednošću `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

Vaš kod mora imati karakteristiku `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Vaš projekat bi trebalo da ima zaglavlje `Photo Gallery`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

Zapamtite, velika i mala slova zajedno sa pravopisom su važne za naslov.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
