---
id: 61537d86bdc3dd343688fceb
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Sledeći, dodeli svakom `img` svojstvo elementa `src` prema njegovom redosledu u dokumentu. `img` prvi treba da ima `src` od `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`. Preostali neka budu isti, osim što promeniš `1` sa brojem koji `img` u dokumentu.

# --hints--

Sve devet elemenata od `img` moraju imati svojstvo elementa `src`.

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

Vaš prvi element za `img` treba da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` postavljen kao vrednost atributa elementa za `src`.

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Drugi element vašeg `img` mora imati `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` postavljen kao vrednost svojstva elementa za `src`.

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

Tvoj treći element od `img` treba da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` postavljen kao vrednost atributa elemenata `src`.

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

Tvoj četvrti element od `img` treba da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` postavljeno kao vrednost svojstva polja `src`.

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

Tvoj peti element ``img`` treba da ima ``https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg`` postavljen kao vrednost atributa elementa ``src``.

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

Tvoje šesto polje od `img` mora imati `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` postavljeno kao vrednost svojstva za `src`.

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

Element vašeg sedmog `img` mora imati `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` postavljeno kao vrednost atributa elementa za `src`.

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

Tvoj osmi element od `img` mora imati `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` postavljeno kao vrednost svojstva elemenata `src`.

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

Tvoj deveti element za ``img`` treba da ima ``https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg`` postavljen kao vrednost atributa elemenata ``src``.

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
