---
id: 61537485c4f2a624f18d7794
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Počni sa svojom osnovnom strukturom od HTML. Dodaj deklaraciju `DOCTYPE`, element od `html`, element od `head`, i element od `body`.

Dodaj atribut `lang` otvarajućoj oznaci `<html>` i postavi njegovu vrednost na `en`.

# --hints--

Vaš kod treba da ima referencu na `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Treba da staviš razmak nakon reference `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Trebalo bi da definišete `html` kao vrstu dokumenta.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Trebalo bi da zatvorite deklaraciju `DOCTYPE` sa etikete za zatvaranje `>` nakon tipa.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Vaš element `html` bi trebalo da ima otvarajući labelu.

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Element vašeg `html` treba da ima zatvarajući tag.

```js
assert(code.match(/<\/html\s*>/));
```

Otvarajući deo za `<html>` mora imati svojstvo atributa `lang` sa vrednošću `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

Tvoje saopštenje za `DOCTYPE` trebalo bi biti na početku tvog HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Trebalo bi da imate etiketu za otvaranje `head`.

```js
assert(code.match(/<head\s*>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Trebalo bi da imaš oznaku za otvaranje `body`.

```js
assert(code.match(/<body\s*>/i));
```

Trebalo bi da imaš oznaku za zatvaranje `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Elementi za `head` i `body` bi trebalo da budu sestrinski elementi.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Element `head` treba da bude unutar elementa `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Element za `body` treba da bude unutar elementa `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
