---
id: 615392916d83fa4f02f7e2cf
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--


Svojstvo `flex-wrap` određuje kako će vaše premeštive komponente da se ponašaju kada je pokretni uređaj veoma mali. Postavljanje njegove vrednosti na `wrap` će omogućiti delovima da se rotiraju u sledeći ili naredni red. `nowrap` (podrazumevana vrednost) će sprečiti vaše delove da se rotiraju i oslabave ako je to potrebno.

Neka tvoji pokretni elementi se premešte u sledeći red kada im nedostaje mesta.

# --hints--

Vaš selektor za `.gallery` mora imati svojstvo `flex-wrap` postavljeno na `wrap` kao vrednost.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexWrap === 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--
.gallery {
  display: flex;
  flex-direction: row;
}
--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
