---
id: 6153908a366afb4d57185c8d
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Flexbox ima glavnu osobinu i poprečnu osobinu. Glavna osa je definisana svojstvom `flex-direction`, koje ima četiri moguća vrednosti:

- `row` (podrazumevano): horizontalna osa i prilagodljivi delovi od leve ka desne
- `row-reverse`: horizontalna osa i prilagodljivi delovi od desne ka levo
- `column`: vertikalna osa i prilagodljivi delovi od gore ka dole
- `column-reverse`: vertikalna osa i prilagodljivi delovi od dole ka gore

**Napomena**: Osa i orijentacija će biti različiti u zavisnosti od pravca pisanja. Prikazane vrednosti su za smer pisanja od leve ka desnoj.

Isprobajte različite vrednosti da biste videli kako utiču na raspored.

Po završetku, stavi `flex-direction` vidljivi deo od `row` na elementu `.gallery`.

# --hints--

Tvoj selektor za `.gallery` mora imati svojstvo `flex-direction` postavljeno na `row` kao vrednost.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection === 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--
.gallery {
  display: flex;
}
--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
