---
id: 6494da0daf5df5197963671d
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Pogledajte kako plava granica slike se proteže preko crvene granice biblioteke. To je zbog načina na koji pretraživači mere dimenzije kontejnerskih elemenata.

Atribut za `box-sizing` se koristi za postavljanje ove karakteristike. Podrazumevano, koristi se primer `content-box`. U ovom primeru, kada element ima određenu širinu, ta širina se meri uzimajući u obzir samo sadržaj elementa. Vrednosti za unutrašnji razmakotvore i ivicu dodaju se na ukupnu širinu, tako da element raste kako bi odgovarao ovim vrednostima.

Pokušajte da postavite `box-sizing` na `content-box` eksplicitno, koristeći međunarodni selektor za `*`. Trenutno nećete videti nikakve promene jer koristite podrazumevanu vrednost.

# --hints--

Trebalo bi da imate birač za `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Vaš selektor za `*` mora imati svojstvo od `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

Vaš selektor za `*` bi trebalo da ima svojstvo od `box-sizing` ako je postavljeno na `content-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
