---
id: 64dd076a1b33c86d84b4232f
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Da biste bolje videli kako su vam komponente merene, dodavanje granice moÅ¾e pomoÄi.

Postavi `.gallery` tvoju karakteristiku `width` od `50%` na `border` koja je postavljena na `5px solid red`.

Zatim ukloni tvoje komponente od `img` `width` iz `100%`, `padding` koje su postavljene na `5px`, i `border` koje je postavljeno na `5px solid blue`.

# --hints--

Trebalo bi da imaÅ” biraÄ za `.gallery`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Tvoj selektor za `.gallery` mora imati svojstvo `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

VaÅ” selektor za `.gallery` mora imati svojstvo od `width`, postavljeno na `50%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

VaÅ” selektor za `.gallery` mora imati svojstvo od `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

Tvoj selektor za `.gallery` treba da ima svojstvo od `border` postavljeno na `5px solid red`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

Trebalo bi da imate biraÄ za `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

Tvoj selektor za `img` mora imati svojstvo od `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

Tvoj selektor za `img` mora imati svojstvo `width` postavljeno na `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

Tvoj selektor za `img` mora da ima svojstvo `padding`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

VaÅ” izbor za `img` mora imati atribut od `padding`, postavljen na `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

Tvoj selektor za `img` mora imati svojstvo od `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

Tvoj selektor za `img` mora imati svojstvo `border` postavljeno na `5px solid blue`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
