---
id: 669e2f60e83c011754f711f9
title: Narođite veb stranicu za tranzitnu agenciju
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Mora da ima deklaraciju `DOCTYPE`.
1. Mora da ima element `html` sa `lang` podešenim na `en`.
1. Mora da ima element `head` sa praznim elementom `meta`, i `charset` koji ima tekst `Travel Agency Page` podešen na `utf-8` i `title`.
1. Mora da ima oznaku `meta` u vašem elementu `head` sa kratkim opisom vaše veb stranice za SEO.
1. Mora da ima element `h1` za prikaz destinacija putovanja.
1. Mora da ima odlomak ispod elementa `h1` koji predstavlja mogućnosti putovanja.
1. Mora da ima element `h2` sa tekstom `Packages`.
1. Mora da ima element `p` koji ukratko predstavlja različite pakete.
1. Mora da ima neuređenu listu sa dve stavke liste. Ove dve stavke liste moraju imati tekst `Group Travels` i `Private Tours`, respektivno. Tekst svake stavke liste mora biti pokriven elementom sidra (anchor).
1. Mora da ima element `h2` sa tekstom `Top Itineraries`.
1. Mora da ima najmanje tri elementa `figure`, svaki sa elementom sidra i elementom `figcaption`.
1. Tri elementa sidra moraju imati atribut elementa `img` sa odgovarajućim atributom elementa `alt` i atributom elementa `src` podešenim kao validna slika u njihovom sadržaju. Možete koristiti `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` i `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg` kako god želite.
1. Sve vaše pet oznaka sidra moraju imati atribut elementa `href` sa vrednošću `https://www.freecodecamp.org/learn` i atribut elementa `target` sa vrednošću `_blank`.

# --hints--

Tvoja stranica agenta za putovanja treba da ima deklaraciju `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Trebalo bi da ima element sa `html` koji je podešen na `lang` i postavljen na `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/i);
```

Trebalo bi da imaš element `head` unutar elementa `html`.

```js
assert.match(
  code,
  /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i
);
```

Trebalo bi da imaš dva elementa `meta` unutar svog elementa `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*(<\s*meta[\s\S]*>[\s\S]*){2,}<\/\s*head\s*>/i);
```

Jedan element iz `meta` treba da ima svojstvo elementa/polja `name` sa vrednošću `description` i svojstvo elementa/polja `content` koje nije prazno.

 ```js
 assert.match(code,
  /<\s*meta\s+[^>]*name\s*=\s*["']description["'][^>]*content\s*=\s*["'][^"']+?["'][^>]*>/i,
);
```

Jedan element `meta` treba da ima atribut `charset` set na `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Trebalo bi da imate element `title` unutar vašeg elementa `head`.

```js
assert.match(
  code,
  /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i
);
```

Vaš `title` element bi trebalo da ima ime vaše putničke agencije.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Trebalo bi da imate element `body` unutar vašeg elementa `html`.

```js
assert.match(
  code,
  /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i
);
```

Trebalo bi da imate element `h1` za prikaz destinacija putovanja.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Trebalo bi da imate samo jedan `h1` element.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Trebalo bi da imate element `p` ispod elementa `h1`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

Vaš prvi `p` element bi trebalo da uvede mogućnosti putovanja.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Vaš prvi `h2` element bi trebalo da sadrži tekst `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

Trebalo bi da imate element `p` ispod vašeg prvog `h2` elementa.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

Vaš drugi `p` element bi trebalo da kratko uvede razne pakete.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

Trebalo bi da imate element neuređenog spiska ispod vašeg drugog `p` elementa.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

Trebalo bi da imate dva elementa u vašoj neuređenoj listi.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

Obje vaše stavke liste treba da sadrže jangkorni element.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

Anker element vaše prve stavke liste treba da obuhvati tekst `Group Travels`.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

Element jangkor (anchor) drugog stavke liste treba da obuhvati tekst `Private Tours`.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

Trebalo bi da imate element `h2` nakon vaše neuređene liste.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

Vaš drugi element `h2` treba da ima tekst `Top Itineraries`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

Trebalo bi da imate najmanje tri `figure` elementa.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

Svaki element `figure` treba da sadrži jangkor element as kao svoj prvi ugnježdeni element.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

Each `figure` element should contain a `figcaption` element as its second nested element.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

Svaki `figcaption` bi trebalo da sadrži neki tekst.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

Svaki element iz `a` koji je ugnježdjen unutar vaših elemenata `figure` treba da sadrži sliku.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

Svaki element `img` mora imati validan atribut `src`.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

Svaki element `img` treba da ima atribut `alt` sa odgovarajućom vrednošću.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Svaki element `a` treba da ima atribut `href` sa vrednošću `https://www.freecodecamp.org/learn`. Ne zaboravite linkove u stavkama liste.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Svaki element `a` treba da ima atribut `target` sa vrednošću `_blank`. Ne zaboravite linkove u listama stavki.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
