---
id: 68a47dcdbce12193caf77012
title: Izvršiće pronaći konfigurisani model
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

U ovom okruženju ćete kreirati funkciju koja vraća najmanji indeks gde treba uneti vrednost u niz podataka nakon sortiranja po rastućem redosledu.

**Cilj**: Popuni priče korisnika ispod i prođi sve testove kako bi se završio laboratorijum.

**Priče korisnika:**

1. Treba da ima funkciju `getIndexToIns` koja prima dva argumenta: niz podataka i broj.
2. Treba koristiti metodu `sort` za sortiranje niza podataka po uzlaznom redosledu.
3. Tvoja funkcija `getIndexToIns` treba da vrati najnižu vrednost gde bi trebalo uneti broj koristeći metodu `findIndex`.
4. Tvoja funkcija `getIndexToIns` treba da uvek vrati broj.

**Snipet:**
Metoda za `findIndex` je metoda napravljena unutar niza podataka u JavaScript. Prima *callback* funkciju i vraća instancu prvog elementa koji zadovoljava uslov. Sve `findIndex` i `sort` su funkcije visokog nivoa.

**Primjeri:**

- `getIndexToIns([1, 2, 3, 4], 1.5)` trebalo bi da vrati `1` jer `1.5` je veće od `1` (primer `0`) i manje od `2` (primer `1`).
- `getIndexToIns([20, 3, 5], 19)` trebalo bi da vrati `2` jer nakon organizovanja u `[3, 5, 20]`, `19` je manje od `20` (primer `2`) i veće od `5` (primer `1`).


# --hints--

Trebalo bi da imaš funkciju `getIndexToIns`.

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns` Treba da uvek vrati broj.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

Tvoja funkcija mora da koristi metodu `sort`.

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```

Tvoja funkcija treba da koristi `findIndex`.

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` treba da vrati `3`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` treba da vrati `2`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` trebalo bi da vrati `1`.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` trebalo bi da vrati `0`.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` trebalo bi vratiti `2`.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` trebalo bi da vrati `2`.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)` treba da vrati `3`

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)` trebalo bi da vrati `0`

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
```
