---
id: a0b5010f579e69b815e7c5d6
title: Pronađi i Zameni
challengeType: 1
forumTopicId: 16045
dashedName: search-and-replace
---

# --description--

Izvršite pretragu i zamenu u rečenici koristeći navedene izraze i vratite novu rečenicu.

Prva stranica je rečenica koju ćeš koristiti za pretragu i zamenu.

Drugi argument je reč koju menjate (pre).

Treća tačka je ono što ćeš promeniti drugo reč sa njim (nakon).

**Napomena:** Sačuvajte prvo slovo originalne reči koju menjate. Na primer, ako nameravate da promenite reč `Book` u reč `dog`, trebalo bi da se promeni u `Dog`

# --hints--

`myReplace("Let us go to the store", "store", "mall")` trebalo bi da vrati sekvencu znakova `Let us go to the mall`.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` treba da vrati sekvencu karaktera `He is Sitting on the couch`.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` treba da vrati niz znakova `I think we should look down there`.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` treba da vrati niz znakova `This has a spelling error`.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` trebalo bi da vrati sekvencu karaktera `His name is John`.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` trebalo bi da vrati niz znakova `Let us get back to more Algorithms`.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js
function myReplace(str, before, after) {
  return str;
}

myReplace("A quick brown fox jumped over the lazy dog", "jumped", "leaped");
```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
