---
id: a105e963526e7de52b219be9
title: Planirano spajanje
challengeType: 1
forumTopicId: 16077
dashedName: sorted-union
---

# --description--

Napišite funkciju koja prima dva ili više nizova podataka i vraća novi niz sa jedinstvenim vrednostima, zadržavajući redosled prvobitno unetih nizova.

Drugim rečima, sve vrednosti koje se nalaze u svim redovima podataka treba da budu uključene u njihov originalni redosled, ali bez ponavljanja u poslednjem redu.

Jedinstveni brojevi treba da se sortiraju prema njihovom prvobitnom redosledu, ali konačni red ne bi trebalo da se sortira po numeričkom redosledu.

Proverite testove validacije za primere.

# --hints--

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` trebalo bi da vrati `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` trebalo bi da vrati `[1, 2, 3, 5]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` trebalo bi vratiti `[1, 2, 3, 5, 4, 6, 7, 8]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` trebao bi da vrati `[1, 3, 2, 5, 4, 6]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` trebalo bi da vrati `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js
function uniteUnique(arr) {
  return arr;
}

uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]);
```

# --solutions--

```js
function uniteUnique(arr) {
  return [].slice.call(arguments).reduce(function(a, b) {
    return [].concat(
      a, 
      b.filter(function(e, currentIndex) {
        return b.indexOf(e) === currentIndex && a.indexOf(e) === -1;
      }));
  }, []);
}
```
