---
id: a2f1d72d9b908d0bd72bb9f6
title: Napravi osobu
challengeType: 1
forumTopicId: 16020
dashedName: make-a-person
---

# --description--

Popunite komponentu stavke i sledeće metode ispod:

```js
getFirstName()
getLastName()
getFullName()
setFirstName(first)
setLastName(last)
setFullName(first, last)
```

Pokrenite testove da biste videli očekivane rezultate za svaku metodu. Ove su jedine dostupne metode za interakciju sa tim objektom. Svaki test će prikazati novu instancu od `Person` kao `new Person('Bob', 'Ross')`.

# --hints--

Ne smijete mijenjati potpis administratora.

```js
assert.match(code, /const\s+Person\s*=\s*function\s*\(\s*first\s*,\s*last\s*\)\s*{/);
```

Ne smete da postavite novu vrednost za parametar `first`.

```js
assert.notMatch(code, /\bfirst\s*=\s*/);
```

Nije moguće postaviti novu vrednost za parametar `last`.

```js
assert.notMatch(code, /\blast\s*=\s*/);
```

Nema atributa koji treba dodati. `Object.keys(Person).length` bi trebalo uvek da vrati 6.

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
_person.setLastName('Curry');
_person.setFullName('John', 'Smith');
assert.lengthOf(Object.keys(_person), 6);
```

Trebalo bi da budeš u stanju da napraviš primer svog `Person`.

```js
const _person = new Person('Bob', 'Ross');
assert.instanceOf(_person, Person);
```

Vaš predmet sa `Person` ne bi trebalo da ima svojstvo od `firstName`.

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'firstName');
```

Objekat koji je `Person` ne bi trebalo da ima svojstvo `lastName`.

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'lastName');
```

Putanja za `.getFirstName()` treba da vrati niz karaktera `Bob`.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFirstName(), 'Bob');
```

`.getLastName()` trebalo bi da vrati niz karaktera `Ross`.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getLastName(), 'Ross');
```

Putanja za `.getFullName()` treba da vrati niz znakova `Bob Ross`.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFullName(), 'Bob Ross');
```

Putanja za `.getFullName()` treba da vrati niz znakova `Haskell Ross` nakon pozivanja `.setFirstName('Haskell')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
assert.strictEqual(_person.getFullName(), 'Haskell Ross');
```

Putanja za `.getFullName()` treba da vrati niz znakova `Bob Curry` nakon pozivanja `.setLastName('Curry')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setLastName('Curry');
assert.strictEqual(_person.getFullName(), 'Bob Curry');
```

Putanja za `.getFullName()` treba da vrati sekvencu karaktera `Haskell Curry` nakon pozivanja `.setFullName('Haskell', 'Curry')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFullName(), 'Haskell Curry');
```

Funkcija ``.getFirstName()`` treba da vrati niz karaktera ``Haskell`` nakon pozivanja ``.setFullName('Haskell', 'Curry')``.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFirstName(), 'Haskell');
```

Put `.getLastName()` trebalo bi da vrati niz karaktera `Curry` nakon pozivanja `.setFullName('Haskell', 'Curry')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getLastName(), 'Curry');
```

Metoda `.getFullName()` treba da vrati niz karaktera `Emily Martinez de la Rosa` nakon pozivanja `.setFullName('Emily Martinez', 'de la Rosa')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFullName(), 'Emily Martinez de la Rosa');
```

Svojstvo `.getFirstName()` trebalo bi vratiti niz znakova `Emily Martinez` nakon pozivanja `.setFullName('Emily Martinez', 'de la Rosa')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFirstName(), 'Emily Martinez');
```

Svojstvo `.getLastName()` treba da vrati niz karaktera `de la Rosa` nakon pozivanja `.setFullName('Emily Martinez', 'de la Rosa')`.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getLastName(), 'de la Rosa');
```

# --seed--

## --seed-contents--

```js
const Person = function(first, last) {
  this.getFullName = function() {
    return "";
  };
  return "";
};
```

# --solutions--

```js
const Person = function(first, last) {
  let firstName = first;
  let lastName = last;

  this.getFirstName = function(){
    return firstName;
  };

  this.getLastName = function(){
    return lastName;
  };

  this.getFullName = function(){
    return firstName + " " + lastName;
  };

  this.setFirstName = function(str){
    firstName = str;
  };


  this.setLastName = function(str){
    lastName = str;
  };

  this.setFullName = function(first, last){
    firstName = first;
    lastName = last;
  };
};
```
