---
id: a3566b1109230028080c9345
title: Izračunaj zbir svih brojeva u koloni
challengeType: 1
forumTopicId: 16083
dashedName: sum-all-numbers-in-a-range
---

# --description--

Dati ćemo vam red podataka sa dve brojeve. Vratite zbir tih dve broja zajedno sa zbirom svih brojeva koji su između njih. Manji broj neće uvek biti prvi.

Na primer, `sumAll([4,1])` treba da vrati `10` jer je suma svih brojeva između 1 i 4 (uključujući sve) `10`.

# --hints--

`sumAll([1, 4])` Trebalo bi da vrati broj.

```js
assert(typeof sumAll([1, 4]) === 'number');
```

`sumAll([1, 4])` Trebalo bi da vrati 10.

```js
assert.deepEqual(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` trebalo bi da vrati 10.

```js
assert.deepEqual(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` trebalo bi da vrati 45.

```js
assert.deepEqual(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` treba da vrati 45.

```js
assert.deepEqual(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
function sumAll(arr) {
  return 1;
}

sumAll([1, 4]);
```

# --solutions--

```js
function sumAll(arr) {
  var sum = 0;
  arr.sort(function(a,b) {return a-b;});
  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }
  return sum;
}
```
