---
id: a5229172f011153519423690
title: Saberi sve nepoznate brojeve za Fibonacci
challengeType: 1
forumTopicId: 16084
dashedName: sum-all-odd-fibonacci-numbers
---

# --description--

Ako uzmete pozitivan savršeni broj `num`, vratite zbir svih nepoznatih brojeva Fibonaccija koji su manji ili jednaki `num`.

Prve dve brojeve u Fibonaccijevoj sekvenci su 0 i 1. Svaki sledeći broj u sekvenci je zbir prethodnih dva broja. Prvih sedam brojeva u Fibonaccijevoj sekvenci su 0, 1, 1, 2, 3, 5 i 8.

Na primer, `sumFibs(10)` treba da vrati `10` jer su svi Fibonacci brojevi manji ili jednaki `10` broj 1, 1, 3 i 5.

# --hints--

`sumFibs(1)` trebalo bi da vrati broj.

```js
assert(typeof sumFibs(1) === 'number');
```

`sumFibs(1000)` treba da vrati 1785.

```js
assert(sumFibs(1000) === 1785);
```

`sumFibs(4000000)` trebalo bi da vrati 4613732.

```js
assert(sumFibs(4000000) === 4613732);
```

`sumFibs(4)` trebalo bi da vrati 5.

```js
assert(sumFibs(4) === 5);
```

`sumFibs(75024)` trebalo bi da vrati 60696.

```js
assert(sumFibs(75024) === 60696);
```

`sumFibs(75025)` treba da vrati 135721.

```js
assert(sumFibs(75025) === 135721);
```

# --seed--

## --seed-contents--

```js
function sumFibs(num) {
  return num;
}

sumFibs(4);
```

# --solutions--

```js
function sumFibs(num) {
  var a = 1;
  var b = 1;
  var s = 0;
  while (a <= num) {
    if (a % 2 !== 0) {
      s += a;
    }
    a = [b, b=b+a][0];
  }
  return s;
}
```
