---
id: a97fd23d9b809dac9921074f
title: Nije neophodno
challengeType: 1
forumTopicId: 14271
dashedName: arguments-optional
---

# --description--

Kreiraj funkciju koja sabira dva argumenta zajedno. Ako se pruži samo jedan argument, onda vrati funkciju koja očekuje jedan argument i vraća zbir.

Na primer, `addTogether(2, 3)` treba da vrati `5`, i `addTogether(2)` treba da vrati funkciju.

Pozivanje funkcije koja vraća jedan argument dodaje i vraća zbir:

```js
var sumTwoAnd = addTogether(2);
```

`sumTwoAnd(3)` vraća `5`.

Ako je bilo koji argument nevažeći broj, vratite undefined.

# --hints--

`addTogether(2, 3)` trebalo bi da vrati 5.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` treba da vrati 53.4.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` trebalo bi vratiti `undefined`.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` treba da vrati `undefined`.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` trebalo bi da vrati `undefined`.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` Trebalo bi da vrati funkciju.

```js
assert.deepEqual(typeof(addTogether(5)), 'function');
```

`addTogether(5)(7)` Trebalo bi da vrati 12.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` trebalo bi vratiti `undefined`.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` trebalo bi da vrati `undefined`.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js
function addTogether() {
  return false;
}

addTogether(2,3);
```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
