---
id: aa7697ea2477d1316795783b
title: Pig Latin
challengeType: 1
forumTopicId: 16039
dashedName: pig-latin
---

# --description--

Pig Latin je način za menjanje engleskih reči. Pravila su sledeća:

\- Ako reč počinje sa konsonantom, uzmi prvu konsonantu ili grupu konsonanti, premesti je na kraj reči, a zatim dodaj `ay`.

\- Ako reč počinje samoglasnikom, samo dodaj `way` na kraju.

# --instructions--

Prevedite dati niz slova u Pig Latina. Ulazni niz slova je garantovan da su sve mala engleska reči.

# --hints--

`translatePigLatin("california")` treba da vrati sekvencu slova `aliforniacay`.

```js
assert.deepEqual(translatePigLatin('california'), 'aliforniacay');
```

`translatePigLatin("paragraphs")` bi trebalo da vrati sekvencu slova `aragraphspay`.

```js
assert.deepEqual(translatePigLatin('paragraphs'), 'aragraphspay');
```

`translatePigLatin("glove")` treba vratiti sekvencu karaktera `oveglay`.

```js
assert.deepEqual(translatePigLatin('glove'), 'oveglay');
```

`translatePigLatin("algorithm")` mora vratiti sekvencu karaktera `algorithmway`.

```js
assert.deepEqual(translatePigLatin('algorithm'), 'algorithmway');
```

`translatePigLatin("eight")` treba da vrati sekvencu karaktera `eightway`.

```js
assert.deepEqual(translatePigLatin('eight'), 'eightway');
```

Trebalo bi da se obrade reči gde je prva vokala na sredini reči. `translatePigLatin("schwartz")` treba da vrati niz znakova `artzschway`.

```js
assert.deepEqual(translatePigLatin('schwartz'), 'artzschway');
```

Trebalo bi da se obrađuju reči bez vokala. `translatePigLatin("rhythm")` treba da vrati nizu slova `rhythmay`.

```js
assert.deepEqual(translatePigLatin('rhythm'), 'rhythmay');
```

# --seed--

## --seed-contents--

```js
function translatePigLatin(str) {
  return str;
}

translatePigLatin("consonant");
```

# --solutions--

```js
function translatePigLatin(str) {
  if (isVowel(str.charAt(0))) return str + "way";
  var front = [];
  str = str.split('');
  while (str.length && !isVowel(str[0])) {
    front.push(str.shift());
  }
  return [].concat(str, front).join('') + 'ay';
}

function isVowel(c) {
  return ['a', 'e', 'i', 'o', 'u'].indexOf(c.toLowerCase()) !== -1;
}
```
