---
id: af4afb223120f7348cdfc9fd
title: Organizuj stvari za svemir
challengeType: 1
forumTopicId: 16021
dashedName: map-the-debris
---

# --description--

Prema Trećem zakonu Keplera, period kruženja $T$ za dva objekta koji kruže u krugu ili elipsi je:

$$
T = 2 \pi \sqrt{\frac{a^{3}}{\mu}}
$$

- $a$ je polumajor osa orbite
- $μ = GM$ je gravitacioni parametar
- $G$ je gravitačna konstanta,
- $M$ je masa tela sa najvećom masom.

Vrati novi red podataka koji konvertuje prosečnu dužinu stavki u njihove ciklusne periode (u sekundama).

Podatni red će imati stavke u strukturi `{name: 'name', avgAlt: avgAlt}`.

Vrednosti treba zaokružiti na najbliži ceo broj. Telo koje se zaokružuje je Zemlja.

Oko Zemlje je kilometara 6367.4447, a vrednost GM Zemlje je 398600.4418 km<sup>3</sup>s<sup>-2</sup>.

# --hints--

`orbitalPeriod([{name : "sputnik", avgAlt : 35873.5553}])` trebalo bi vratiti `[{name: "sputnik", orbitalPeriod: 86400}]`.

```js
assert.deepEqual(orbitalPeriod([{ name: 'sputnik', avgAlt: 35873.5553 }]), [
  { name: 'sputnik', orbitalPeriod: 86400 }
]);
```

`orbitalPeriod([{name: "iss", avgAlt: 413.6}, {name: "hubble", avgAlt: 556.7}, {name: "moon", avgAlt: 378632.553}])` treba da vrati `[{name: "iss", orbitalPeriod: 5557}, {name: "hubble", orbitalPeriod: 5734}, {name: "moon", orbitalPeriod: 2377399}]`.

```js
assert.deepEqual(
  orbitalPeriod([
    { name: 'iss', avgAlt: 413.6 },
    { name: 'hubble', avgAlt: 556.7 },
    { name: 'moon', avgAlt: 378632.553 }
  ]),
  [
    { name: 'iss', orbitalPeriod: 5557 },
    { name: 'hubble', orbitalPeriod: 5734 },
    { name: 'moon', orbitalPeriod: 2377399 }
  ]
);
```

# --seed--

## --seed-contents--

```js
function orbitalPeriod(arr) {
  const GM = 398600.4418;
  const earthRadius = 6367.4447;
  return arr;
}

orbitalPeriod([{name : "sputnik", avgAlt : 35873.5553}]);
```

# --solutions--

```js
function orbitalPeriod(arr) {
  const GM = 398600.4418;
  const earthRadius = 6367.4447;
  const TAU = 2 * Math.PI;
  return arr.map(function(obj) {
    return {
      name: obj.name,
      orbitalPeriod: Math.round(TAU * Math.sqrt(Math.pow(obj.avgAlt+earthRadius, 3)/GM))
    };
  });
}
```
