---
id: af7588ade1100bde429baf20
title: Nedostajuća slova
challengeType: 1
forumTopicId: 16023
dashedName: missing-letters
---

# --description--

Pronađi nedostajućicu slovo u nizu slova koje je dato i vrati ga.

Kada su sva slova u redu, obnovite `undefined`.

# --hints--

`fearNotLetter("abce")` trebalo bi da vrati niz znakova `d`.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` mora vratiti sekvencu karaktera `i`.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` trebalo bi da vrati sekvencu karaktera `u`.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` treba da vrati niz karaktera `e`.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` trebalo bi da vrati `undefined`.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js
function fearNotLetter(str) {
  return str;
}

fearNotLetter("abce");
```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    var letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
