---
id: 587d78af367417b2b2512b04
title: Izgradite stranicu za prodaju proizvoda
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301144
dashedName: build-a-product-landing-page
---

# --description--

**Cilj:** Napraviti aplikaciju koja funkcionalno liči na <a href="https://product-landing-page.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://product-landing-page.freecodecamp.rocks</a>. **Ne kopirajte ovaj demo projekat**.

**Priče korisnika:**

1. Vaša stranica za prodaju proizvoda mora imati element `header` sa `id="header"` koji odgovara
1. Možete videti sliku unutar elementa `header` sa `id="header-img"` koja odgovara (Logo će biti lepa slika ovde)
1. Unutar elementa `#header`, možete videti element `nav` sa `id="nav-bar"` koji odgovara
1. Možete videti najmanje tri klikabilna elementa unutar elementa `nav`, svaki sa klasom `nav-link`
1. Kada kliknete na dugme `.nav-link` unutar elementa `nav`, preusmeravate se na odgovarajući deo stranice za prodaju
1. Možete pogledati video proizvoda sa `id="video"`
1. Vaša stranica za prodaju ima element `form` sa `id="form"` koji odgovara
1. Unutar forme, postoji polje za unos `input` sa `id="email"` gde možete uneti e-mail adresu
1. Polje za unos `#email` mora imati placeholder tekst da obavesti korisnike za šta je to polje
1. Polje za unos `#email` koristi HTML5 validaciju da bi proverilo da li su unesovani podaci e-mail adresa
1. Unutar forme, postoji dugme za slanje `input` sa `id="submit"` koje odgovara
1. Kada kliknete na element `#submit`, e-mail se šalje na statičnu stranicu (koristite ovaj primer URL: `https://www.freecodecamp.com/email-submit`)
1. Navbar mora uvek biti na vrhu ekrana
1. Vaša stranica za prodaju proizvoda mora imati najmanje jednu media query
1. Vaša stranica za prodaju proizvoda mora koristiti CSS flexbox barem jednom

Završen je priča korisnika i sve prođite testove ispod da završite ovaj projekat. Dodajte svoj stil. Srećno kodiranje!

**Napomena:** Uverite se da dodate `<link rel="stylesheet" href="styles.css">` u vaš HTML kako biste povezali vašu list stilova i koristili CSS

# --hints--

Trebalo bi imati element `header` sa `id` na `header`.

```js
const el = document.getElementById('header')
assert(!!el && el.tagName === 'HEADER')
```

Trebalo bi imati element `img` sa `id` na `header-img`.

```js
const el = document.getElementById('header-img')
assert(!!el && el.tagName === 'IMG')
```

Vaš `#header-img` treba da bude potomak `#header`.

```js
const els = document.querySelectorAll('#header #header-img')
assert(els.length > 0)
```

Vaš `#header-img` treba da imati atribut `src` elementa.

```js
const el = document.getElementById('header-img')
assert(!!el && !!el.src)
```

Vrednost `src` za `#header-img` treba da bude važeća URL (počinje sa `http`).

```js
const el = document.getElementById('header-img')
assert(!!el && /^http/.test(el.src))
```

Trebalo bi imati element `nav` sa `id` na `nav-bar`.

```js
const el = document.getElementById('nav-bar')
assert(!!el && el.tagName === 'NAV')
```

Vaš `#nav-bar` treba da bude potomak `#header`.

```js
const els = document.querySelectorAll('#header #nav-bar')
assert(els.length > 0)
```

Trebalo bi imati najmanje tri elementa `.nav-link` unutar `#nav-bar`.

```js
const els = document.querySelectorAll('#nav-bar .nav-link')
assert(els.length >= 3)
```

Svaki element `.nav-link` treba da ima atribut `href` elementa.

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  if (!el.href) assert(false)
})
assert(els.length > 0)
```

Svaki element `.nav-link` treba da se veže za odgovarajući deo stranice (koji ima `id` vrednost, npr. `#footer`).

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  const linkDestination = el.getAttribute('href').slice(1)
  if (!document.getElementById(linkDestination)) assert(false)
})
assert(els.length > 0)
```

Trebalo bi imati element `video` ili `iframe` sa `id` na `video`.

```js
const el = document.getElementById('video')
assert(!!el && (el.tagName === 'VIDEO' || el.tagName === 'IFRAME'))
```

Vaš `#video` treba da imati atribut `src` elementa.

```js
let el = document.getElementById('video')
const sourceNode = el.children;
let sourceElement = null;
if (sourceNode.length) {
  sourceElement = [...video.children].filter(el => el.localName === 'source')[0];
}
if (sourceElement) {
  el = sourceElement;
}
assert(el.hasAttribute('src'));
```

Trebalo bi imati element `form` sa `id` na `form`.

```js
const el = document.getElementById('form')
assert(!!el && el.tagName === 'FORM')
```

Trebalo bi imati element `input` sa `id` na `email`.

```js
const el = document.getElementById('email')
assert(!!el && el.tagName === 'INPUT')
```

Vaš `#email` treba da biti potomak `#form`.

```js
const els = document.querySelectorAll('#form #email')
assert(els.length > 0)
```

Vaš `#email` treba da imati atribut `placeholder` sa placeholder tekstom.

```js
const el = document.getElementById('email')
assert(!!el && !!el.placeholder && el.placeholder.length > 0)
```

Vaš `#email` treba da koristi HTML5 validaciju postavljanjem `type` na `email`.

```js
const el = document.getElementById('email')
assert(!!el && el.type === 'email')
```

Trebalo bi imati element `input` sa `id` na `submit`.

```js
const el = document.getElementById('submit')
assert(!!el && el.tagName === 'INPUT')
```

Vaš `#submit` treba da biti potomak `#form`.

```js
const els = document.querySelectorAll('#form #submit')
assert(els.length > 0)
```

Vaš `#submit` treba da imati `type` na `submit`.

```js
const el = document.getElementById('submit')
assert(!!el && el.type === 'submit')
```

Vaš `#form` treba da imati atribut `action` sa vrednošću `https://www.freecodecamp.com/email-submit`.

```js
const el = document.getElementById('form')
assert(!!el && el.action === 'https://www.freecodecamp.com/email-submit')
```

Vaš `#email` treba da imati atribut `name` sa vrednošću `email`.

```js
const el = document.getElementById('email')
assert(!!el && el.name === 'email')
```

Vaš `#nav-bar` mora biti uvek na vrhu ekrana.

```js
  const timeout = (milliseconds) => new Promise((resolve) => setTimeout(resolve, milliseconds));

  const header = document.getElementById('header');
  const headerChildren = header.children;
  const navbarCandidates = [header, ...headerChildren];
  
  // Return smallest top position of all navbar candidates
  const getNavbarPosition = (candidates = []) => {
    return candidates.reduce(
      (min, candidate) =>
        Math.min(min, Math.abs(candidate?.getBoundingClientRect().top)),
      Infinity
    );
  };
  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the viewport '
  );

  window.scroll(0, 500);
  await timeout(1);

  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the ' +
      'viewport even after scrolling '
  );
    
  window.scroll(0, 0);
```

Vaša stranica za prodaju proizvoda mora koristiti najmanje jednu media query.

```js
const htmlSourceAttr = Array.from(document.querySelectorAll('source')).map(el => el.getAttribute('media'))
const cssCheck = new __helpers.CSSHelp(document).getCSSRules('media')
assert(cssCheck.length > 0 || htmlSourceAttr.length > 0);
```

Vaša stranica za prodaju proizvoda mora koristiti CSS Flexbox barem jednom.

```js
const hasFlex = (rule) => ["flex", "inline-flex"].includes(rule.style?.display)
const stylesheet = new __helpers.CSSHelp(document).getStyleSheet()
const cssRules = new __helpers.CSSHelp(document).styleSheetToCssRulesArray(stylesheet)
const mediaRules = new __helpers.CSSHelp(document).getCSSRules('media')
const usesFlex = cssRules.find(rule => hasFlex(rule))
const usesFlexMedia = mediaRules.find(mediaRule => {
  return [...mediaRule.cssRules].find(rule => hasFlex(rule))
})
assert(usesFlex || usesFlexMedia)
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Product Landing Page</title>
  </head>
  <body>
    <header id="header">
      <nav id="nav-bar">
        <img
          id="header-img"
          src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
          max-height="50px"
        />
        <a href="#Features" class="nav-link">Features</a> |
        <a href="#Video" class="nav-link">See our facility!</a> |
        <a href="#Pricing" class="nav-link">Purchase</a>
        <hr />
      </nav>
    </header>
    <main>
      <h1>
        Pokemon Daycare Service
      </h1>
      <section id="Features">
        <h2>What we offer</h2>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">Guaranteed friendly and loving staff!</div>
        </div>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">
            Comfortable environment for Pokemon to explore and play!
          </div>
        </div>
        <div class="flex-here">
          <div class="flex-left">
            <img
              id="bullet"
              src="https://upload.wikimedia.org/wikipedia/commons/3/39/Pokeball.PNG"
              max-height="25px"
            />
          </div>
          <div class="flex-right">
            Multiple membership