---
id: 6718d2d59337c822ecb697ff
title: Napravi program za upravljanje bankovnim računima
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba da definišete klasu objekata nazvanu `BankAccount` koja ima komponentu. Komponenta treba da postavi osnovnu vrednost `balance` na `0` i da inicijalizuje prazan niz podataka nazvan `transactions` za skladištenje transakcionih zapisa kao objekata.

2. Svaka transakcija sačuvana u red podataka za `transactions` mora biti stavka sa dve karakteristike: `type` i `amount`. Karakteristika za `type` mora biti `deposit` ili `withdraw`, a karakteristika za `amount` mora biti uneseni ili uzet iznos.

3. Trebalo bi definisati metodu pod nazivom `deposit` koja prima iznos depozita kao parametar. Kada je iznos depozita veći od `0`, treba da:

- Dodajte novi red u transakcionoj bazi podataka sa tipom `deposit` i navedenim iznosom.
   - Ažuriraj saldo.   
   - Vrati `"Successfully deposited $[amount]. New balance: $[balance]"`.

4. Ako je iznos mali ili jednak `0`, metoda za `deposit` bi trebalo da vrati `"Deposit amount must be greater than zero."`.

5. Trebalo bi da definišete funkciju pod nazivom `withdraw` koja prima iznos kao parametar. Ova funkcija treba da ažurira trenutni saldo na osnovu oduzimanja. Kada je iznos koji se želi uzeti veći od `0` i manji ili jednak trenutnom saldu, trebalo bi:

- Dodaj novi unos u red podatke o transakciji sa tipom `withdraw` i iznosom koji je uzet.
   - Ažuriraj saldo.
   - Vrati `"Successfully withdrew $[amount]. New balance: $[balance]"`.

6. Ako je iznos koji se želi povući manji ili jednak `0` ili veći od trenutnog salda, funkcija `withdraw` treba da vrati `"Insufficient balance or invalid amount."`.

7. Treba da definišete metodu pod nazivom `checkBalance` koja vraća trenutni saldo u formatu `"Current balance: $[balance]"`.

8. Treba da definišete metodu pod nazivom `listAllDeposits` koja iterira kroz niz podataka transakcija i vraća sve depozite u formatu od `"Deposits: amount,amount,..."`.

9. Treba da definišete funkciju pod nazivom `listAllWithdrawals` koja iterira kroz niz podataka o transakcijama i vraća sve brisanje u formatu `"Withdrawals: amount,amount,..."`.

10. Trebalo bi kreirati novi primer `BankAccount`, nazvan `myAccount`.

11. Vaš bankovni račun `myAccount` mora imati najmanje pet transakcija.

12. Vaš bankovni račun `myAccount` mora imati najmanje dva depozita.

13. Vaš bankovni račun `myAccount` mora imati najmanje dva povlačenja.

14. Vaš bankovni račun za `myAccount` mora imati stanje veće od `$100`.

# --hints--

Treba da definišete klasu stvari nazvanu `BankAccount`.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

Element `BankAccount` treba da ima `balance` od `0` na praznom redu podataka `transactions` za skladištenje transakcionih zapisa.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

Treba vam metoda `deposit` koja prima iznos depozita kao parametar.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

Trebalo bi da imate putanju `withdraw` koja prima količinu uklanjanja kao parametar.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

Trebalo bi da imaš putanju `checkBalance` koja proverava trenutni saldo.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

Treba vam način `listAllDeposits` koji navodi sve depozite.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

Trebalo bi da imate putanju `listAllWithdrawals` koja navodi sve uklanjanje.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)` trebalo bi da vrati `"Successfully deposited $100. New balance: $100"`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` treba vratiti `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` treba da vrati `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

Kada je saldo računa `100`, `BankAccount.withdraw(150)` treba da vrati `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` treba da vrati `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` treba da vrati `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

Kada je stanje računa `200`, `BankAccount.withdraw(150)` treba da vrati `"Successfully withdrew $150. New balance: $50"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

Kada je saldo računa `200`, `BankAccount.checkBalance()` bi trebalo da vrati `"Current balance: $200"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

Kada uplaćujete depozit `10`, `35`, `90`, putem `listAllDeposits` bi trebalo da se vrati `"Deposits: 10,35,90"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

Kada povlačite depozit `20`, `50`, `100`, putem od `listAllWithdrawals` treba da vrati `"Withdrawals: 20,50,100"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

Trebalo bi da imaš primer `BankAccount` nazvan `myAccount`.

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

Račun vaše banke `myAccount` mora imati najmanje pet transakcija.

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

Vaš bankovni račun `myAccount` mora imati najmanje dva depozita.

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

Račun vaše banke `myAccount` mora imati najmanje dva povlačenja.

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```

Račun vaše banke `myAccount` mora imati saldo veći od `$100`.

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
