---
id: af7588ade1100bde429baf20
title: Izgradnja praćenja nedostajućeg slova
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Treba da postoji funkcija nazvana `fearNotLetter`.
1. Funkcija za `fearNotLetter` treba da primi jedan argument: sekvencu karaktera koja predstavlja opseg slova u abecednom redosledu i koji može imati jedno nedostajuće slovo.
1. Funkcija treba da pronađe nedostajuće slovo u poslatom opsegu slova i vrati ga.
1. Ako su sva slova prisutna u opsegu, funkcija treba da vrati `undefined`.

# --hints--

Trebalo bi da imaš funkciju `fearNotLetter`.

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` trebalo bi da vrati niz znakova `d`.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` mora vratiti sekvencu karaktera `i`.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` trebalo bi da vrati sekvencu karaktera `u`.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` treba da vrati niz karaktera `e`.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` trebalo bi da vrati `undefined`.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
