---
id: 587d7b7c367417b2b2512b19
title: Izmeni Ugnjezdeni Objekat unutar Objekta
challengeType: 1
forumTopicId: 301164
dashedName: modify-an-object-nested-within-an-object
---

# --description--

Hajde sada da pogledamo malo slozeniji objekat. Svojstva objekta mogu biti ugnjezdena na bilo kom nivou, a njihove vrednosti mogu biti bilo koji tip podatka koji JavaScript podrzava, ukljucujuci nizove i druge objekte. Pogledajte sledece:

```js
let nestedObject = {
  id: 28802695164,
  date: 'December 31, 2016',
  data: {
    totalUsers: 99,
    online: 80,
    onlineStatus: {
      active: 67,
      away: 13,
      busy: 8
    }
  }
};
```

`nestedObject` ima tri svojstva: `id` (vrednost je broj), `date` (vrednost je string) i `data` (vrednost je objekat sa svojom ugnjezdenom strukturom). Iako strukture brzo mogu postati slozene, i dalje mozemo koristiti istu notaciju da pristupimo podacima koji nam trebaju. Da bismo postavili vrednost `10` za svojstvo `busy` u ugnjezdenom objektu `onlineStatus`, koristimo tackastu notaciju:

```js
nestedObject.data.onlineStatus.busy = 10;
```

# --instructions--

Ovde smo napravili objekat `userActivity`, koji sadrzi jos jedan ugnjezdeni objekat. Postavite vrednost kljuca `online` na `45`.

# --hints--

`userActivity` treba da ima svojstva `id`, `date` i `data`.

```js
assert(
  'id' in userActivity && 'date' in userActivity && 'data' in userActivity
);
```

`userActivity` treba da ima kljuc `data` postavljen na objekat sa kljucevima `totalUsers` i `online`.

```js
assert('totalUsers' in userActivity.data && 'online' in userActivity.data);
```

Svojstvo `online` ugnjezdeno unutar kljuca `data` objekta `userActivity` treba da bude postavljeno na `45`

```js
assert(userActivity.data.online === 45);
```

Svojstvo `online` treba postaviti koristeci tackastu notaciju ili notaciju sa uglastim zagradama.

```js
assert.strictEqual(__helpers.removeJSComments(code).search(/online: 45/), -1);
```

# --seed--

## --seed-contents--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

// Only change code below this line

// Only change code above this line

console.log(userActivity);
```

# --solutions--

```js
let userActivity = {
  id: 23894201352,
  date: 'January 1, 2017',
  data: {
    totalUsers: 51,
    online: 42
  }
};

userActivity.data.online = 45;
```
