---
id: 587d7b7b367417b2b2512b17
title: Spoji Nizove Koriscenjem Spread Operatora
challengeType: 1
forumTopicId: 301156
dashedName: combine-arrays-with-the-spread-operator
---

# --description--

Jos jedna velika prednost <dfn>spread</dfn> operatora je mogucnost spajanja nizova, odnosno ubacivanja svih elemenata jednog niza u drugi, na bilo kom mestu po indeksu. Standardnom sintaksom takodje mozemo spajati nizove, ali to uglavnom omogucava dodavanje na kraj jednog i pocetak drugog. Spread sintaksa cini sledecu operaciju veoma jednostavnom:

```js
let thisArray = ['sage', 'rosemary', 'parsley', 'thyme'];

let thatArray = ['basil', 'cilantro', ...thisArray, 'coriander'];
```

`thatArray` bi imao vrednost `['basil', 'cilantro', 'sage', 'rosemary', 'parsley', 'thyme', 'coriander']`.

Koriscenjem spread sintakse postigli smo operaciju koja bi bila slozenija i opseznija da smo koristili standardne metode.

# --instructions--

Napravili smo funkciju `spreadOut` koja vraca promenljivu `sentence`. Izmenite funkciju koristeci <dfn>spread</dfn> operator tako da vrati niz `['learning', 'to', 'code', 'is', 'fun']`.

# --hints--

`spreadOut` treba da vrati `["learning", "to", "code", "is", "fun"]`

```js
assert.deepEqual(spreadOut(), ['learning', 'to', 'code', 'is', 'fun']);
```

Funkcija `spreadOut` treba da koristi spread sintaksu

```js
assert.notStrictEqual(spreadOut.toString().search(/[...]/), -1);
```

# --seed--

## --seed-contents--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence; // Change this line
  return sentence;
}

console.log(spreadOut());
```

# --solutions--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence = ['learning', ...fragment, 'is', 'fun'];
  return sentence;
}
```
