---
id: 587d78b2367417b2b2512b10
title: Ukloni Stavke Koriscenjem splice()
challengeType: 1
forumTopicId: 301166
dashedName: remove-items-using-splice
---

# --description--

Dobro, naucili smo kako da uklonimo elemente sa pocetka i kraja niza koristeci `shift()` i `pop()`, ali sta kada treba ukloniti element iz sredine? Ili ukloniti vise elemenata odjednom? Tu nastupa `splice()`. `splice()` nam omogucava upravo to: **uklanjanje bilo kog broja uzastopnih elemenata** sa bilo koje pozicije u nizu.

`splice()` moze da primi do 3 argumenta, ali za sada cemo se fokusirati samo na prva dva. Prva dva argumenta `splice()` su celi brojevi koji predstavljaju indeks, odnosno poziciju stavki u nizu nad kojim je `splice()` pozvan. I zapamtite, nizovi imaju *indeksiranje od nule*, pa za prvi element koristimo `0`. Prvi argument `splice()` predstavlja indeks od kog pocinje uklanjanje, a drugi broj elemenata koji se uklanja. Na primer:

```js
let array = ['today', 'was', 'not', 'so', 'great'];

array.splice(2, 2);
```

Ovde uklanjamo 2 elementa, pocevsi od treceg elementa (indeks 2). `array` bi imao vrednost `['today', 'was', 'great']`.

`splice()` ne menja samo niz nad kojim je pozvan, vec i vraca novi niz sa vrednostima uklonjenih elemenata:

```js
let array = ['I', 'am', 'feeling', 'really', 'happy'];

let newArray = array.splice(3, 2);
```

`newArray` ima vrednost `['really', 'happy']`.

# --instructions--

Inicijalizovali smo niz `arr`. Koristite `splice()` da uklonite elemente iz `arr`, tako da u nizu ostanu samo elementi ciji je zbir `10`.

# --hints--

Ne treba da menjate originalnu liniju `const arr = [2, 4, 5, 1, 7, 5, 2, 1];`.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/constarr=\[2,4,5,1,7,5,2,1\];?/)
);
```

`arr` treba da sadrzi samo elemente ciji je zbir `10`.

```js
assert.strictEqual(
  arr.reduce((a, b) => a + b),
  10
);
```

Vas kod treba da koristi `splice()` metod nad `arr`.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(/));
```

`splice` treba samo da ukloni elemente iz `arr`, bez dodavanja novih elemenata u `arr`.

```js
assert(
  !__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(\d+,\d+,\d+.*\)/g)
);
```

# --seed--

## --seed-contents--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
// Only change code below this line

// Only change code above this line
console.log(arr);
```

# --solutions--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
arr.splice(1, 4);
```
