---
id: 587d7b7c367417b2b2512b1a
title: Pristupi Nazivima Svojstava Pomocu Notacije sa Uglastim Zagradama
challengeType: 1
forumTopicId: 301150
dashedName: access-property-names-with-bracket-notation
---

# --description--

U prvom izazovu o objektima pomenuli smo upotrebu notacije sa uglastim zagradama kao nacin pristupa vrednostima svojstava preko evaluacije promenljive. Na primer, zamislite da se nas `foods` objekat koristi u aplikaciji kase supermarketa. Imamo funkciju koja postavlja `selectedFood` i zelimo da proverimo `foods` objekat za prisustvo te namirnice. To moze izgledati ovako:

```js
let selectedFood = getCurrentFood(scannedItem);
let inventory = foods[selectedFood];
```

Ovaj kod ce evaluirati vrednost sacuvanu u promenljivoj `selectedFood` i vratiti vrednost tog kljuca u objektu `foods`, ili `undefined` ako ne postoji. Notacija sa uglastim zagradama je korisna jer ponekad svojstva objekta nisu poznata pre izvrsavanja ili nam treba fleksibilniji pristup.

# --instructions--

Napravili smo funkciju `checkInventory`, koja prima skeniranu stavku kao argument. Vratite trenutnu vrednost kljuca `scannedItem` u objektu `foods`. Mozete pretpostaviti da ce se funkciji `checkInventory` prosledjivati samo validni kljucevi.

# --hints--

`checkInventory` treba da bude funkcija.

```js
assert.strictEqual(typeof checkInventory, 'function');
```

Objekat `foods` treba da ima samo sledece parove kljuc-vrednost: `apples: 25`, `oranges: 32`, `plums: 28`, `bananas: 13`, `grapes: 35`, `strawberries: 27`.

```js
assert.deepEqual(foods, {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
});
```

`checkInventory("apples")` treba da vrati `25`.

```js
assert.strictEqual(checkInventory('apples'), 25);
```

`checkInventory("bananas")` treba da vrati `13`.

```js
assert.strictEqual(checkInventory('bananas'), 13);
```

`checkInventory("strawberries")` treba da vrati `27`.

```js
assert.strictEqual(checkInventory('strawberries'), 27);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(checkInventory("apples"));
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  return foods[scannedItem];
}
```
