---
id: 587d7b7d367417b2b2512b1c
title: Proveri da li Objekat Ima Svojstvo
challengeType: 1
forumTopicId: 301155
dashedName: check-if-an-object-has-a-property
---

# --description--

Sada mozemo dodavati, menjati i uklanjati kljuceve iz objekata. Ali sta ako zelimo samo da proverimo da li objekat ima odredjeno svojstvo? JavaScript nam daje dva razlicita nacina za to. Jedan koristi metod `hasOwnProperty()`, a drugi kljucnu rec `in`. Ako imamo objekat `users` koji ima svojstvo `Alan`, njegovo prisustvo mozemo proveriti na jedan od sledecih nacina:

```js
users.hasOwnProperty('Alan');
'Alan' in users;
```

Oba izraza ce vratiti `true`.

# --instructions--

Dovrsite funkciju tako da vraca `true` ako prosledjeni objekat ima sva cetiri imena, `Alan`, `Jeff`, `Sarah` i `Ryan`, a u suprotnom `false`.

# --hints--

Objekat `users` ne treba direktno da se koristi unutar funkcije

```js 

assert(__helpers.removeJSComments(code).match(/users/gm).length <= 2)

```

Objekat `users` treba da ima samo kljuceve `Alan`, `Jeff`, `Sarah` i `Ryan`

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

Funkcija `isEveryoneHere` treba da vrati `true` ako su `Alan`, `Jeff`, `Sarah` i `Ryan` svojstva u prosledjenom objektu.

```js
assert(isEveryoneHere(users) === true);
```

Funkcija `isEveryoneHere` treba da vrati `false` ako `Alan` nije svojstvo prosledjenog objekta.

```js
assert(
  (function () {
    delete users.Alan;
    return isEveryoneHere(users);
  })() === false
);
```

Funkcija `isEveryoneHere` treba da vrati `false` ako `Jeff` nije svojstvo prosledjenog objekta.

```js
assert(
  (function () {
    delete users.Jeff;
    return isEveryoneHere(users);
  })() === false
);
```

Funkcija `isEveryoneHere` treba da vrati `false` ako `Sarah` nije svojstvo prosledjenog objekta.

```js
assert(
  (function () {
    delete users.Sarah;
    return isEveryoneHere(users);
  })() === false
);
```

Funkcija `isEveryoneHere` treba da vrati `false` ako `Ryan` nije svojstvo prosledjenog objekta.

```js
assert(
  (function () {
    delete users.Ryan;
    return isEveryoneHere(users);
  })() === false
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  // Only change code below this line
  
  // Only change code above this line
}

console.log(isEveryoneHere(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  return [
    'Alan',
    'Jeff',
    'Sarah',
    'Ryan'
  ].every(user => userObj.hasOwnProperty(user));
}

console.log(isEveryoneHere(users));
```
