---
id: 587d7b7e367417b2b2512b20
title: Koristi Niz za Cuvanje Kolekcije Podataka
challengeType: 1
forumTopicId: 301167
dashedName: use-an-array-to-store-a-collection-of-data
---

# --description--

Ispod je primer vrlo jednostavne implementacije strukture podataka niza. Ovo se naziva <dfn>jednodimenzionalni niz</dfn>, sto znaci da ima samo jedan nivo, odnosno da ne sadrzi druge ugnjezdene nizove. Primetite da sadrzi <dfn>bool vrednosti</dfn>, <dfn>stringove</dfn> i <dfn>brojeve</dfn>, zajedno sa drugim validnim JavaScript tipovima podataka:

```js
let simpleArray = ['one', 2, 'three', true, false, undefined, null];
console.log(simpleArray.length);
```

Poziv `console.log` ispisuje `7`.

Svi nizovi imaju svojstvo duzine, koje se, kao sto je prikazano iznad, lako dobija sintaksom `Array.length`. Slozenija implementacija niza prikazana je ispod. To je <dfn>ugnjezdeni niz</dfn>, odnosno niz koji cuva druge nizove. Primetite da ovaj niz sadrzi i JavaScript <dfn>objekte</dfn>, koje cemo detaljnije obraditi u sledecoj sekciji, ali za sada je dovoljno da znate da nizovi mogu cuvati i slozene strukture.

```js
let complexArray = [
  [
    {
      one: 1,
      two: 2
    },
    {
      three: 3,
      four: 4
    }
  ],
  [
    {
      a: "a",
      b: "b"
    },
    {
      c: "c",
      d: "d"
    }
  ]
];
```

# --instructions--

Postavili smo promenljivu pod imenom `yourArray`. Dovrsite deklaraciju tako sto cete promenljivoj `yourArray` dodeliti niz sa najmanje 5 elemenata. Vas niz treba da sadrzi najmanje jedan string, jedan broj i jednu bool vrednost.

# --hints--

`yourArray` treba da bude niz.

```js
assert.strictEqual(Array.isArray(yourArray), true);
```

`yourArray` treba da ima najmanje 5 elemenata.

```js
assert.isAtLeast(yourArray.length, 5);
```

`yourArray` treba da ima bar jedan `boolean`.

```js
assert(yourArray.filter((el) => typeof el === 'boolean').length >= 1);
```

`yourArray` treba da ima bar jedan `number`.

```js
assert(yourArray.filter((el) => typeof el === 'number').length >= 1);
```

`yourArray` treba da ima bar jedan `string`.

```js
assert(yourArray.filter((el) => typeof el === 'string').length >= 1);
```

# --seed--

## --seed-contents--

```js
let yourArray; // Change this line
```

# --solutions--

```js
let yourArray = ['a string', 100, true, ['one', 2], 'another string'];
```
