---
id: 587d7b7c367417b2b2512b1b
title: Koristi delete Kljucnu Rec za Uklanjanje Svojstava Objekta
challengeType: 1
forumTopicId: 301168
dashedName: use-the-delete-keyword-to-remove-object-properties
---

# --description--

Sada znate sta su objekti, koje su njihove osnovne osobine i prednosti. Ukratko, to su skladista parova kljuc-vrednost koja nude jednostavan i jasan nacin organizacije podataka, ***a*** pritom omogucavaju vrlo brzo pretrazivanje. U ovim izazovima prolazimo kroz nekoliko uobicajenih operacija koje mozete raditi sa objektima kako biste se lakse snalazili sa ovim vaznim strukturama podataka u svojim programima.

U prethodnim izazovima dodavali smo i menjali parove kljuc-vrednost u objektu. Ovde cemo videti kako da *uklonimo* par kljuc-vrednost iz objekta.

Vratimo se jos jednom na nas primer objekta `foods`. Ako zelimo da uklonimo kljuc `apples`, to mozemo uraditi koristeci kljucnu rec `delete` ovako:

```js
delete foods.apples;
```

# --instructions--

Koristite kljucnu rec `delete` da uklonite kljuceve `oranges`, `plums` i `strawberries` iz objekta `foods`.

# --hints--

Objekat `foods` treba da ima samo tri kljuca: `apples`, `grapes` i `bananas`.

```js
assert(
  !foods.hasOwnProperty('oranges') &&
    !foods.hasOwnProperty('plums') &&
    !foods.hasOwnProperty('strawberries') &&
    Object.keys(foods).length === 3
);
```

Kljucevi `oranges`, `plums` i `strawberries` treba da budu uklonjeni koristeci `delete`.

```js
assert(
  __helpers.removeJSComments(code).search(/oranges:/) !== -1 &&
    __helpers.removeJSComments(code).search(/plums:/) !== -1 &&
    __helpers.removeJSComments(code).search(/strawberries:/) !== -1
);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

delete foods.oranges;
delete foods.plums;
delete foods.strawberries;

console.log(foods);
```
