---
id: 587d7b7d367417b2b2512b1d
title: Iteriraj Kroz Kljuceve Objekta Pomocu for...in Naredbe
challengeType: 1
forumTopicId: 301162
dashedName: iterate-through-the-keys-of-an-object-with-a-for---in-statement
---

# --description--

Ponekad je potrebno da prodjete kroz sve kljuceve u objektu. Za to mozete koristiti <dfn>for...in</dfn> petlju. for...in petlja izgleda ovako:

```javascript
const refrigerator = {
  'milk': 1,
  'eggs': 12,
};

for (const food in refrigerator) {
  console.log(food, refrigerator[food]);
}
```

Ovaj kod ispisuje `milk 1` i `eggs 12`, pri cemu je svaki par kljuc-vrednost u svom redu.

Definisali smo promenljivu `food` u zaglavlju petlje i ta promenljiva dobija svaki kljuc objekta pri svakoj iteraciji, sto dovodi do ispisa svakog naziva namirnice u konzoli.

**NAPOMENA:** Objekti ne cuvaju redosled kljuceva kao nizovi; zato pozicija kljuca u objektu, odnosno redosled pojavljivanja, nije bitan pri referenciranju ili pristupu tom kljucu.

# --instructions--

Definisali smo funkciju `countOnline` koja prima jedan argument, `allUsers`. Koristite <dfn>for...in</dfn> naredbu unutar ove funkcije da iterirate kroz objekat `allUsers` i vratite broj korisnika kojima je svojstvo `online` postavljeno na `true`. Primer objekta koji moze biti prosledjen funkciji `countOnline` prikazan je ispod. Svaki korisnik ima svojstvo `online` postavljeno na `true` ili `false`.

```js
{
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
```

# --hints--

Funkcija `countOnline` treba da koristi `for in` naredbu za iteraciju kroz kljuceve prosledjenog objekta.

```js
assert.match(
  __helpers.removeJSComments(code),
    /for\s*\(\s*(var|let|const)\s+[a-zA-Z_$]\w*\s+in\s+[a-zA-Z_$]\w*\s*\)/
);
```

Funkcija `countOnline` treba da vrati `1` kada joj se prosledi objekat `{ Alan: { online: false }, Jeff: { online: true }, Sarah: { online: false } }`

```js
const usersObj1 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj1), 1);
```

Funkcija `countOnline` treba da vrati `2` kada joj se prosledi objekat `{ Alan: { online: true }, Jeff: { online: false }, Sarah: { online: true } }`

```js
const usersObj2 = {
  Alan: {
    online: true
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: true
  }
}
assert.equal(countOnline(usersObj2), 2);
```

Funkcija `countOnline` treba da vrati `0` kada joj se prosledi objekat `{ Alan: { online: false }, Jeff: { online: false }, Sarah: { online: false } }`

```js
const usersObj3 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj3), 0);
```

# --seed--

## --seed-contents--

```js
const users = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}

function countOnline(allUsers) {
  // Only change code below this line

  // Only change code above this line
}

console.log(countOnline(users));
```

# --solutions--

```js
function countOnline(allUsers) {
  let numOnline = 0;
  for(const user in allUsers){
    if(allUsers[user].online) {
      numOnline++;
    }
  }
  return numOnline;
}
```
