---
id: 587d7b7d367417b2b2512b1e
title: Generisi Niz Svih Kljuceva Objekta Pomocu Object.keys()
challengeType: 1
forumTopicId: 301160
dashedName: generate-an-array-of-all-object-keys-with-object-keys
---

# --description--

Takodje mozemo napraviti niz koji sadrzi sve kljuceve sacuvane u objektu koristeci metod `Object.keys()`. Ovaj metod prima objekat kao argument i vraca niz stringova koji predstavljaju svako svojstvo u tom objektu. Ponovo, redosled stavki u tom nizu nije garantovan.

# --instructions--

Dovrsite funkciju `getArrayOfUsers` tako da vraca niz koji sadrzi sva svojstva objekta primljenog kao argument.

# --hints--

Objekat `users` treba da ima samo kljuceve `Alan`, `Jeff`, `Sarah` i `Ryan`

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

Funkcija `getArrayOfUsers` treba da vrati niz koji sadrzi sve kljuceve u objektu `users`

```js
assert(
  (function () {
    users.Sam = {};
    users.Lewis = {};
    let R = getArrayOfUsers(users);
    return (
      R.indexOf('Alan') !== -1 &&
      R.indexOf('Jeff') !== -1 &&
      R.indexOf('Sarah') !== -1 &&
      R.indexOf('Ryan') !== -1 &&
      R.indexOf('Sam') !== -1 &&
      R.indexOf('Lewis') !== -1
    );
  })() === true
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  // Only change code below this line

  // Only change code above this line
}

console.log(getArrayOfUsers(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  return Object.keys(obj);
}

console.log(getArrayOfUsers(users));
```
