---
id: 587d7b7b367417b2b2512b14
title: Proveri Prisustvo Elementa Pomocu indexOf()
challengeType: 1
forumTopicId: 301154
dashedName: check-for-the-presence-of-an-element-with-indexof
---

# --description--

Posto nizovi mogu da se menjaju u bilo kom trenutku, nije sigurno na kojoj se poziciji nalazi odredjeni element, niti da li je i dalje prisutan. Srecom, JavaScript ima ugradjeni metod `indexOf()` koji nam omogucava da brzo i jednostavno proverimo prisustvo elementa u nizu. `indexOf()` prima element kao argument i, kada se pozove, vraca poziciju odnosno indeks tog elementa, ili `-1` ako element nije u nizu.

Na primer:

```js
let fruits = ['apples', 'pears', 'oranges', 'peaches', 'pears'];

fruits.indexOf('dates');
fruits.indexOf('oranges');
fruits.indexOf('pears');
```

`indexOf('dates')` vraca `-1`, `indexOf('oranges')` vraca `2`, a `indexOf('pears')` vraca `1` (prvi indeks na kom se element pojavljuje).

# --instructions--

`indexOf()` je veoma koristan za brzu proveru prisustva elementa u nizu. Deklarisali smo funkciju `quickCheck`, koja prima niz i element kao argumente. Izmenite funkciju koristeci `indexOf()` tako da vrati `true` ako je prosledjeni element prisutan u nizu, a `false` ako nije.

# --hints--

Funkcija `quickCheck` treba da vraca bool vrednost (`true` ili `false`), a ne string (`"true"` ili `"false"`)

```js
assert.isBoolean(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

`quickCheck(["squash", "onions", "shallots"], "mushrooms")` treba da vrati `false`

```js
assert.strictEqual(
  quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'),
  false
);
```

`quickCheck(["onions", "squash", "shallots"], "onions")` treba da vrati `true`

```js
assert.strictEqual(
  quickCheck(['onions', 'squash', 'shallots'], 'onions'),
  true
);
```

`quickCheck([3, 5, 9, 125, 45, 2], 125)` treba da vrati `true`

```js
assert.strictEqual(quickCheck([3, 5, 9, 125, 45, 2], 125), true);
```

`quickCheck([true, false, false], undefined)` treba da vrati `false`

```js
assert.strictEqual(quickCheck([true, false, false], undefined), false);
```

Funkcija `quickCheck` treba da koristi metod `indexOf()`

```js
assert.notStrictEqual(quickCheck.toString().search(/\.indexOf\(/), -1);
```

# --seed--

## --seed-contents--

```js
function quickCheck(arr, elem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

# --solutions--

```js
function quickCheck(arr, elem) {
  return arr.indexOf(elem) >= 0; 
}
```
