---
id: 587d7b7c367417b2b2512b18
title: Dodaj Parove Kljuc-Vrednost u JavaScript Objekte
challengeType: 1
forumTopicId: 301153
dashedName: add-key-value-pairs-to-javascript-objects
---

# --description--

U svojoj osnovi, objekti su samo kolekcije <dfn>kljuc-vrednost</dfn> parova. Drugim recima, to su delovi podataka (<dfn>vrednosti</dfn>) organizovani jedinstvenim identifikatorima koji se zovu <dfn>svojstva</dfn> (<dfn>kljucevi</dfn>). Pogledajte ovaj primer:

```js
const tekkenCharacter = {
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true
};
```

Kod iznad opisuje objekat igraca iz video igre Tekken pod imenom `tekkenCharacter`. Ima tri svojstva, svako sa odredjenom vrednoscu. Ako zelite da dodate jos jedno svojstvo, na primer "origin", to mozete uraditi tako sto cete ga dodeliti objektu:

```js
tekkenCharacter.origin = 'South Korea';
```

Ovde se koristi tackasta notacija. Ako pogledate objekat `tekkenCharacter`, sada ce sadrzati i svojstvo `origin`. Hwoarang je takodje imao narandzastu kosu. Ovo svojstvo mozete dodati koristeci notaciju sa uglastim zagradama:

```js
tekkenCharacter['hair color'] = 'dyed orange';
```

Notacija sa uglastim zagradama je potrebna ako naziv svojstva sadrzi razmak ili ako zelite da koristite promenljivu kao naziv svojstva. U ovom primeru, naziv svojstva je pod navodnicima da bi se pokazalo da je to string i da ce biti dodat bas tako kako je napisan. Bez navodnika, tumacilo bi se kao promenljiva, pa bi naziv svojstva postao vrednost te promenljive. Evo primera sa promenljivom:

```js
const eyes = 'eye color';

tekkenCharacter[eyes] = 'brown';
```

Posle svih ovih dodataka, objekat ce izgledati ovako:

```js
{
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true,
  origin: 'South Korea',
  'hair color': 'dyed orange',
  'eye color': 'brown'
};
```

# --instructions--

Objekat `foods` je inicijalno napravljen sa tri unosa. Koristeci sintaksu po izboru, dodajte jos tri unosa: `bananas` sa vrednoscu `13`, `grapes` sa vrednoscu `35` i `strawberries` sa vrednoscu `27`.

# --hints--

`foods` treba da bude objekat.

```js
assert(typeof foods === 'object');
```

Objekat `foods` treba da ima kljuc `bananas` sa vrednoscu `13`.

```js
assert(foods.bananas === 13);
```

Objekat `foods` treba da ima kljuc `grapes` sa vrednoscu `35`.

```js
assert(foods.grapes === 35);
```

Objekat `foods` treba da ima kljuc `strawberries` sa vrednoscu `27`.

```js
assert(foods.strawberries === 27);
```

Definicija objekta `foods` ne treba da se menja.

```js
assert(
  __helpers.removeJSComments(code).search(/let foods/) === -1 &&
  __helpers.removeJSComments(code).search(/const\s+foods\s*=\s*{\s*apples:\s*25,\s*oranges:\s*32,\s*plums:\s*28\s*};/
) !== -1
);
```

# --seed--

## --seed-contents--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

foods['bananas'] = 13;
foods['grapes']  = 35;
foods['strawberries'] = 27;
```
