---
id: 587d78b3367417b2b2512b11
title: Dodaj Stavke Koriscenjem splice()
challengeType: 1
forumTopicId: 301152
dashedName: add-items-using-splice
---

# --description--

Secate li se da smo u prethodnom izazovu pomenuli da `splice()` moze da primi tri argumenta? Treci argument, koji moze sadrzati jedan ili vise elemenata, koristimo za dodavanje u niz. Ovo je veoma korisno za brzo menjanje jednog elementa ili grupe elemenata drugim vrednostima.

```js
const numbers = [10, 11, 12, 12, 15];
const startIndex = 3;
const amountToDelete = 1;

numbers.splice(startIndex, amountToDelete, 13, 14);
console.log(numbers);
```

Druga pojava `12` se uklanja, a dodajemo `13` i `14` na isti indeks. Niz `numbers` sada ce biti `[ 10, 11, 12, 13, 14, 15 ]`.

Ovde pocinjemo sa nizom brojeva. Zatim `splice()` prosledjujemo sledece: indeks od kog pocinje uklanjanje (3), broj elemenata za uklanjanje (1), a preostali argumenti (13, 14) ubacuju se pocevsi od tog istog indeksa. Imajte na umu da nakon `amountToDelete` moze stajati bilo koji broj elemenata (odvojenih zarezima), koji ce biti ubaceni.

# --instructions--

Napravili smo funkciju `htmlColorNames`, koja prima niz HTML boja kao argument. Izmenite funkciju koristeci `splice()` da ukloni prva dva elementa niza i na njihovo mesto doda `'DarkSalmon'` i `'BlanchedAlmond'`.

# --hints--

`htmlColorNames` treba da vrati `["DarkSalmon", "BlanchedAlmond", "LavenderBlush", "PaleTurquoise", "FireBrick"]`

```js
assert.deepEqual(
  htmlColorNames([
    'DarkGoldenRod',
    'WhiteSmoke',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]),
  [
    'DarkSalmon',
    'BlanchedAlmond',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]
);
```

Funkcija `htmlColorNames` treba da koristi metod `splice()`

```js
assert(/.splice/.test(__helpers.removeJSComments(code)));
```

Ne treba da koristite `shift()` niti `unshift()`.

```js
assert(!/shift|unshift/.test(__helpers.removeJSComments(code)));
```

Ne treba da koristite indeksnu notaciju niza.

```js
assert(!/\[\d\]\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function htmlColorNames(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(htmlColorNames(['DarkGoldenRod', 'WhiteSmoke', 'LavenderBlush', 'PaleTurquoise', 'FireBrick']));
```

# --solutions--

```js
function htmlColorNames(arr) {
  arr.splice(0,2,'DarkSalmon', 'BlanchedAlmond');
  return arr;
}
```
