---
id: 587d7b7b367417b2b2512b15
title: Prodji Kroz Sve Stavke Niza Koriscenjem for Petlji
challengeType: 1
forumTopicId: 301161
dashedName: iterate-through-all-an-arrays-items-using-for-loops
---

# --description--

Ponekad je pri radu sa nizovima veoma korisno proci kroz svaki element da biste pronasli jedan ili vise elemenata koji su vam potrebni, ili da biste izmenili niz na osnovu elemenata koji ispunjavaju odredjene uslove. JavaScript nudi vise ugradjenih metoda koji iteriraju kroz nizove na malo razlicite nacine radi razlicitih rezultata (kao `every()`, `forEach()`, `map()`, itd.), ali pristup koji daje najvise slobode i kontrole je jednostavna `for` petlja.

Fikiria yafuatayo:

```js
function greaterThanTen(arr) {
  let newArr = [];
  for (let i = 0; i < arr.length; i++) {
    if (arr[i] > 10) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}

greaterThanTen([2, 12, 8, 14, 80, 0, 1]);
```

Koriscenjem `for` petlje, ova funkcija prolazi kroz svaki element niza i testira ga jednostavnim uslovom. Na ovaj nacin lako utvrdjujemo koji elementi su veci od `10` i vracamo novi niz `[12, 14, 80]` koji ih sadrzi.

# --instructions--

Deklarisali smo funkciju `filteredArray`, koja prima `arr`, niz ugnjezdenih nizova, i `elem` kao argumente i vraca novi niz. `elem` predstavlja element koji moze, ali ne mora, da postoji u jednom ili vise ugnjezdenih nizova unutar `arr`. Izmenite funkciju koristeci `for` petlju tako da vrati filtriranu verziju prosledjenog niza, gde je uklonjen svaki ugnjezdeni niz iz `arr` koji sadrzi `elem`.

# --hints--

`filteredArray([[10, 8, 3], [14, 6, 23], [3, 18, 6]], 18)` treba da vrati `[[10, 8, 3], [14, 6, 23]]`

```js
assert.deepEqual(
  filteredArray(
    [
      [10, 8, 3],
      [14, 6, 23],
      [3, 18, 6]
    ],
    18
  ),
  [
    [10, 8, 3],
    [14, 6, 23]
  ]
);
```

`filteredArray([["trumpets", 2], ["flutes", 4], ["saxophones", 2]], 2)` treba da vrati `[["flutes", 4]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['trumpets', 2],
      ['flutes', 4],
      ['saxophones', 2]
    ],
    2
  ),
  [['flutes', 4]]
);
```

`filteredArray([["amy", "beth", "sam"], ["dave", "sean", "peter"]], "peter")` treba da vrati `[["amy", "beth", "sam"]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['amy', 'beth', 'sam'],
      ['dave', 'sean', 'peter']
    ],
    'peter'
  ),
  [['amy', 'beth', 'sam']]
);
```

`filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3)` treba da vrati `[]`

```js
assert.deepEqual(
  filteredArray(
    [
      [3, 2, 3],
      [1, 6, 3],
      [3, 13, 26],
      [19, 3, 9]
    ],
    3
  ),
  []
);
```

Funkcija `filteredArray` treba da koristi `for` petlju

```js
assert.notStrictEqual(filteredArray.toString().search(/for/), -1);
```

# --seed--

## --seed-contents--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  // Only change code below this line

  // Only change code above this line
  return newArr;
}

console.log(filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3));
```

# --solutions--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  for (let i = 0; i<arr.length; i++) {
    if (arr[i].indexOf(elem) < 0) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}
```
