---
id: 587d7b7b367417b2b2512b16
title: Kreiraj Slozene Visedimenzionalne Nizove
challengeType: 1
forumTopicId: 301159
dashedName: create-complex-multi-dimensional-arrays
---

# --description--

Odlicno! Naucili ste mnogo o nizovima! Ovo je bio pregled na visokom nivou, a i dalje ima mnogo toga da se nauci o radu sa nizovima, sto cete videti u narednim sekcijama. Ali pre nego sto predjemo na <dfn>Objekte</dfn>, hajde da pogledamo kako nizovi mogu postati slozeniji nego u prethodnim izazovima.

Jedna od najjacih osobina nizova kada ih posmatramo kao strukture podataka jeste to sto mogu sadrzati, ili cak biti potpuno sastavljeni od drugih nizova. U prethodnim izazovima smo vec videli ugnjezdene nizove, ali su bili jednostavni. Medjutim, nizovi mogu imati neogranicenu dubinu ugnjezdenih nizova, gde svaki nivo moze sadrzati nove nizove, i tako dalje. Na taj nacin niz brzo moze postati veoma slozena struktura podataka, poznata kao <dfn>visedimenzionalni</dfn> ili ugnjezdeni niz. Pogledajte sledeci primer:

```js
let nestedArray = [
  ['deep'],
  [
    ['deeper'], ['deeper'] 
  ],
  [
    [
      ['deepest'], ['deepest']
    ],
    [
      [
        ['deepest-est?']
      ]
    ]
  ]
];
```

Niz `deep` je ugnjezden na nivou dubine 2. Nizovi `deeper` su na nivou 3. Nizovi `deepest` su na nivou 4, a `deepest-est?` je na nivou 5.

Iako ovaj primer moze delovati slozeno, ovaj nivo slozenosti nije neuobicajen kada radite sa velikim kolicinama podataka. Ipak, i dalje mozemo lako pristupiti i najdubljim delovima ovog slozenog niza koristeci indeksnu notaciju:

```js
console.log(nestedArray[2][1][0][0][0]);
```

Ovo ispisuje string `deepest-est?`. A sada kada znamo gde se taj podatak nalazi, mozemo ga promeniti ako je potrebno:

```js
nestedArray[2][1][0][0][0] = 'deeper still';

console.log(nestedArray[2][1][0][0][0]);
```

Sada ispisuje `deeper still`.

# --instructions--

Napravili smo promenljivu `myNestedArray`, postavljenu na niz. Izmenite `myNestedArray`, koristeci bilo koju kombinaciju <dfn>stringova</dfn>, <dfn>brojeva</dfn> i <dfn>bool vrednosti</dfn> kao elemente, tako da ima tacno pet nivoa dubine (spoljasnji niz je nivo 1). Negde na trecem nivou ubacite string `deep`, na cetvrtom nivou string `deeper`, a na petom nivou string `deepest`.

# --hints--

`myNestedArray` treba da sadrzi samo brojeve, bool vrednosti i stringove kao elemente

```js
assert.strictEqual(
  (function (arr) {
    let flattened = (function flatten(arr) {
      const flat = [].concat(...arr);
      return flat.some(Array.isArray) ? flatten(flat) : flat;
    })(arr);
    for (let i = 0; i < flattened.length; i++) {
      if (
        typeof flattened[i] !== 'number' &&
        typeof flattened[i] !== 'string' &&
        typeof flattened[i] !== 'boolean'
      ) {
        return false;
      }
    }
    return true;
  })(myNestedArray),
  true
);
```

`myNestedArray` treba da ima tacno 5 nivoa dubine

```js
assert.strictEqual(
  (function (arr) {
    let depth = 0;
    function arrayDepth(array, i, d) {
      if (Array.isArray(array[i])) {
        arrayDepth(array[i], 0, d + 1);
      } else {
        depth = d > depth ? d : depth;
      }
      if (i < array.length) {
        arrayDepth(array, i + 1, d);
      }
    }
    arrayDepth(arr, 0, 0);
    return depth;
  })(myNestedArray),
  4
);
```

`myNestedArray` treba da ima tacno jedno pojavljivanje stringa `deep` u ugnjezdenom nizu na nivou dubine 3

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deep').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deep')[0] === 2
);
```

`myNestedArray` treba da ima tacno jedno pojavljivanje stringa `deeper` u ugnjezdenom nizu na nivou dubine 4

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deeper').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deeper')[0] === 3
);
```

`myNestedArray` treba da ima tacno jedno pojavljivanje stringa `deepest` u ugnjezdenom nizu na nivou dubine 5

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deepest').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deepest')[0] === 4
);
```

# --seed--

## --seed-contents--

```js
let myNestedArray = [
  // Only change code below this line
  ['unshift', false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
  // Only change code above this line
];
```

# --solutions--

```js
let myNestedArray = [
  ['unshift', ['deep', ['deeper', ['deepest']]],false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
];
```
