---
id: 587d7b7b367417b2b2512b13
title: Kopiraj Niz Koriscenjem Spread Operatora
challengeType: 1
forumTopicId: 301157
dashedName: copy-an-array-with-the-spread-operator
---

# --description--

Dok nam `slice()` omogucava da izaberemo koje elemente niza zelimo da kopiramo, novi ES6 spread operator nam omogucava da lako kopiramo *sve elemente* niza redom, uz jednostavnu i citljivu sintaksu. Spread sintaksa izgleda ovako: `...`

U praksi, spread operator mozemo koristiti za kopiranje niza ovako:

```js
let thisArray = [true, true, undefined, false, null];
let thatArray = [...thisArray];
```

`thatArray` je jednak `[true, true, undefined, false, null]`. `thisArray` ostaje nepromenjen, a `thatArray` sadrzi iste elemente kao `thisArray`.

# --instructions--

Napravili smo funkciju `copyMachine` koja prima `arr` (niz) i `num` (broj) kao argumente. Funkcija treba da vrati novi niz sastavljen od `num` kopija niza `arr`. Vecina posla je uradjena, ali jos ne radi ispravno. Izmenite funkciju koristeci spread sintaksu da radi kako treba (hint: jos jedan metod koji smo vec obradili moze pomoci ovde).

# --hints--

`copyMachine([true, false, true], 2)` treba da vrati `[[true, false, true], [true, false, true]]`

```js
assert.deepEqual(copyMachine([true, false, true], 2), [
  [true, false, true],
  [true, false, true]
]);
```

`copyMachine([1, 2, 3], 5)` treba da vrati `[[1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3]]`

```js
assert.deepEqual(copyMachine([1, 2, 3], 5), [
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3]
]);
```

`copyMachine([true, true, null], 1)` treba da vrati `[[true, true, null]]`

```js
assert.deepEqual(copyMachine([true, true, null], 1), [[true, true, null]]);
```

`copyMachine(["it works"], 3)` treba da vrati `[["it works"], ["it works"], ["it works"]]`

```js
assert.deepEqual(copyMachine(['it works'], 3), [
  ['it works'],
  ['it works'],
  ['it works']
]);
```

Funkcija `copyMachine` treba da koristi `spread operator` sa nizom `arr`

```js
assert(__helpers.removeJSComments(__helpers.removeJSComments(code)).match(/\.\.\.\s*arr/));
```

# --seed--

## --seed-contents--

```js
function copyMachine(arr, num) {
  let newArr = [];
  while (num >= 1) {
    // Only change code below this line

    // Only change code above this line
    num--;
  }
  return newArr;
}

console.log(copyMachine([true, false, true], 2));
```

# --solutions--

```js
function copyMachine(arr,num){
    let newArr=[];
    while(num >=1){
    newArr.push([...arr]);
    num--;
    }
    return newArr;
}
console.log(copyMachine([true, false, true], 2));
```
