---
id: 587d7b7d367417b2b2512b1f
title: Izmeni Niz Sacuvan u Objektu
challengeType: 1
forumTopicId: 301163
dashedName: modify-an-array-stored-in-an-object
---

# --description--

Sada ste videli sve osnovne operacije nad JavaScript objektima. Možete dodavati, menjati i uklanjati parove ključ-vrednost, proveravati da li ključevi postoje i prolaziti kroz sve ključeve objekta. Kako budete dalje učili JavaScript, videćete još efikasnijih načina upotrebe objekata. Pored toga, lekcije o strukturama podataka u delu za pripremu za intervju takođe obrađuju ES6 objekte <dfn>Map</dfn> i <dfn>Set</dfn>, koji su slični običnim objektima, ali nude dodatne funkcionalnosti. Sada kada ste naučili osnove nizova i objekata, spremni ste da krenete na zahtevnije JavaScript izazove.

# --instructions--

Pogledajte objekat koji smo vam dali u editoru koda. Objekat `user` ima tri kljuca. Kljuc `data` ima pet kljuceva, od kojih jedan sadrzi niz `friends`. Ovde mozete videti koliko su objekti fleksibilni kao strukture podataka. Poceli smo da pisemo funkciju `addFriend`. Dovrsite je tako da primi objekat `user` i doda ime iz argumenta `friend` u niz sacuvan u `user.data.friends`, a zatim vrati taj niz.

# --hints--

Objekat `user` treba da ima kljuceve `name`, `age` i `data`.

```js
assert('name' in user && 'age' in user && 'data' in user);
```

Funkcija `addFriend` treba da primi objekat `user` i string `friend` kao argumente i doda prijatelja u niz `friends` unutar objekta `user`.

```js
assert(
  (function () {
    let L1 = user.data.friends.length;
    addFriend(user, 'Sean');
    let L2 = user.data.friends.length;
    return L2 === L1 + 1;
  })()
);
```

`addFriend(user, "Pete")` treba da vrati `["Sam", "Kira", "Tomo", "Pete"]`.

```js
assert.deepEqual(
  (function () {
    delete user.data.friends;
    user.data.friends = ['Sam', 'Kira', 'Tomo'];
    return addFriend(user, 'Pete');
  })(),
  ['Sam', 'Kira', 'Tomo', 'Pete']
);
```

# --seed--

## --seed-contents--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  // Only change code below this line

  // Only change code above this line
}

console.log(addFriend(user, 'Pete'));
```

# --solutions--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  userObj.data.friends.push(friend);
  return userObj.data.friends;
}
```
