---
id: 587d78b2367417b2b2512b0f
title: Ukloni Stavke iz Niza Pomocu pop() i shift()
challengeType: 1
forumTopicId: 301165
dashedName: remove-items-from-an-array-with-pop-and-shift
---

# --description--

I `push()` i `unshift()` imaju odgovarajuce metode koje su im skoro suprotne: `pop()` i `shift()`. Kao sto ste verovatno vec primetili, umesto dodavanja, `pop()` *uklanja* element sa kraja niza, dok `shift()` uklanja element sa pocetka. Glavna razlika izmedju `pop()` i `shift()` i njihovih parnjaka `push()` i `unshift()` je to sto ne primaju argumente i svaki od njih menja niz za samo jedan element odjednom.

Hebu tuchunguze:

```js
let greetings = ['whats up?', 'hello', 'see ya!'];

greetings.pop();
```

`greetings` bi imao vrednost `['whats up?', 'hello']`.

```js
greetings.shift();
```

`greetings` bi imao vrednost `['hello']`.

Takodje mozemo sacuvati vrednost uklonjenog elementa iz bilo kog metoda ovako:

```js
let popped = greetings.pop();
```

`greetings` bi imao vrednost `[]`, a `popped` bi imao vrednost `hello`.

# --instructions--

Napravili smo funkciju `popShift`, koja prima niz kao argument i vraca novi niz. Izmenite ovu funkciju koristeci `pop()` i `shift()` da uklonite prvi i poslednji element argumenta niza, a uklonjene elemente smestite u odgovarajuce promenljive, tako da vraceni niz sadrzi njihove vrednosti.

# --hints--

`popShift(["challenge", "is", "not", "complete"])` treba da vrati `["challenge", "complete"]`

```js
assert.deepEqual(popShift(['challenge', 'is', 'not', 'complete']), [
  'challenge',
  'complete'
]);
```

Funkcija `popShift` treba da koristi metod `pop()`

```js
assert.notStrictEqual(popShift.toString().search(/\.pop\(/), -1);
```

Funkcija `popShift` treba da koristi metod `shift()`

```js
assert.notStrictEqual(popShift.toString().search(/\.shift\(/), -1);
```

# --seed--

## --seed-contents--

```js
function popShift(arr) {
  let popped; // Change this line
  let shifted; // Change this line
  return [shifted, popped];
}

console.log(popShift(['challenge', 'is', 'not', 'complete']));
```

# --solutions--

```js
function popShift(arr) {
  let popped = arr.pop(); // Change this line
  let shifted = arr.shift(); // Change this line
  return [shifted, popped];
}
```
