---
id: 587d7b7a367417b2b2512b12
title: Kopiraj Stavke Niza Koriscenjem slice()
challengeType: 1
forumTopicId: 301158
dashedName: copy-array-items-using-slice
---

# --description--

Sledeci metod koji cemo nauciti je `slice()`. Umesto menjanja niza, `slice()` kopira ili *izdvaja* odredjen broj elemenata i smesta ih u novi niz, dok originalni niz ostaje nepromenjen. `slice()` prima samo dva argumenta - prvi je indeks od kog pocinje izdvajanje, a drugi indeks na kom se zavrsava (izdvajanje ide do tog indeksa, ali bez ukljucivanja elementa na njemu). Pogledajte:

```js
let weatherConditions = ['rain', 'snow', 'sleet', 'hail', 'clear'];

let todaysWeather = weatherConditions.slice(1, 3);
```

`todaysWeather` ce imati vrednost `['snow', 'sleet']`, dok ce `weatherConditions` i dalje biti `['rain', 'snow', 'sleet', 'hail', 'clear']`.

Na ovaj nacin kreirali smo novi niz izdvajanjem elemenata iz postojeceg niza.

# --instructions--

Napravili smo funkciju `forecast`, koja prima niz kao argument. Izmenite ovu funkciju koristeci `slice()` da izdvoji informacije iz argumenta niza i vrati novi niz koji sadrzi elemente `warm` i `sunny`.

# --hints--

`forecast` treba da vrati `["warm", "sunny"]`

```js
assert.deepEqual(
  forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']),
  ['warm', 'sunny']
);
```

Funkcija `forecast` treba da koristi metod `slice()`

```js
assert(/\.slice\(/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function forecast(arr) {
  // Only change code below this line

  return arr;
}

// Only change code above this line
console.log(forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']));
```

# --solutions--

```js
function forecast(arr) {
  return arr.slice(2,4);
}
```
