---
id: 587d78b2367417b2b2512b0e
title: Dodaj Stavke u Niz Pomocu push() i unshift()
challengeType: 1
forumTopicId: 301151
dashedName: add-items-to-an-array-with-push-and-unshift
---

# --description--

Duzina niza, kao i tipovi podataka koje moze da sadrzi, nije fiksna. Nizovi mogu biti definisani sa bilo kojim brojem elemenata, a elementi se mogu dodavati ili uklanjati tokom vremena; drugim recima, nizovi su <dfn>promenljivi</dfn>. U ovom izazovu cemo pogledati dva metoda koja mozemo koristiti za programsko menjanje niza: `Array.push()` i `Array.unshift()`.

Oba metoda primaju jedan ili vise elemenata kao argumente i dodaju ih u niz nad kojim su pozvani; metod `push()` dodaje elemente na kraj niza, a `unshift()` ih dodaje na pocetak. Pogledajte sledece:

```js
let twentyThree = 'XXIII';
let romanNumerals = ['XXI', 'XXII'];

romanNumerals.unshift('XIX', 'XX');
```

`romanNumerals` bi imao vrednost `['XIX', 'XX', 'XXI', 'XXII']`.

```js
romanNumerals.push(twentyThree);
```

`romanNumerals` bi imao vrednost `['XIX', 'XX', 'XXI', 'XXII', 'XXIII']`. Primetite da takodje mozemo prosledjivati promenljive, sto nam daje vecu fleksibilnost pri izmeni podataka niza u realnom vremenu.

# --instructions--

Napravili smo funkciju `mixedNumbers`, kojoj prosledjujemo niz kao argument. Izmenite funkciju koristeci `push()` i `unshift()` tako da doda `'I', 2, 'three'` na pocetak niza i `7, 'VIII', 9` na kraj, kako bi vraceni niz sadrzao predstavljanje brojeva 1-9 po redu.

# --hints--

`mixedNumbers(["IV", 5, "six"])` sada treba da vrati `["I", 2, "three", "IV", 5, "six", 7, "VIII", 9]`

```js
assert.deepEqual(mixedNumbers(['IV', 5, 'six']), [
  'I',
  2,
  'three',
  'IV',
  5,
  'six',
  7,
  'VIII',
  9
]);
```

Funkcija `mixedNumbers` treba da koristi metod `push()`

```js
assert(mixedNumbers.toString().match(/\.push/));
```

Funkcija `mixedNumbers` treba da koristi metod `unshift()`

```js
assert(mixedNumbers.toString().match(/\.unshift/));
```

# --seed--

## --seed-contents--

```js
function mixedNumbers(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(mixedNumbers(['IV', 5, 'six']));
```

# --solutions--

```js
function mixedNumbers(arr) {
  arr.push(7,'VIII',9);
  arr.unshift('I',2,'three');
  return arr;
}
```
