---
id: 5a661e0f1068aca922b3ef17
title: Pristupi Sadrzaju Niza Koriscenjem Notacije sa Uglastim Zagradama
challengeType: 1
forumTopicId: 301149
dashedName: access-an-arrays-contents-using-bracket-notation
---

# --description--

Osnovna osobina svake strukture podataka je, naravno, sposobnost ne samo da cuva podatke vec i da omoguci pristup tim podacima po zahtevu. Zato, sada kada ucimo kako da pravimo niz, hajde da razmislimo kako da pristupimo informacijama u njemu.

Kada definisemo jednostavan niz kao ispod, on sadrzi 3 elementa:

```js
let ourArray = ["a", "b", "c"];
```

U nizu svaki element ima svoj <dfn>indeks</dfn>. Taj indeks predstavlja i poziciju elementa u nizu i nacin na koji se na njega referise. Medjutim, vazno je napomenuti da su JavaScript nizovi <dfn>indeksirani od nule</dfn>, sto znaci da je prvi element na poziciji ***nula***, a ne jedan. Da bismo pristupili elementu niza, stavljamo indeks u uglaste zagrade iza niza, odnosno najcesce iza promenljive koja referise na taj niz. Ovo se zove <dfn>notacija sa uglastim zagradama</dfn>. Na primer, ako zelimo da uzmemo `a` iz `ourArray` i smestimo tu vrednost u promenljivu, mozemo uraditi sledece:

```js
let ourVariable = ourArray[0];
```

Sada `ourVariable` ima vrednost `a`.

Pored citanja vrednosti sa indeksa, mozete i *postaviti* vrednost na indeks koristeci istu notaciju:

```js
ourArray[1] = "not b anymore";
```

Koriscenjem notacije sa uglastim zagradama, promenili smo element na indeksu 1 iz stringa `b` u `not b anymore`. Sada je `ourArray` `["a", "not b anymore", "c"]`.

# --instructions--

Da biste zavrsili ovaj izazov, postavite bilo koju vrednost koju zelite na drugu poziciju (indeks `1`) u `myArray`, osim slova `b`.

# --hints--

`myArray[0]` treba da bude jednako slovu `a`

```js
assert.strictEqual(myArray[0], 'a');
```

`myArray[1]` ne treba da bude jednako slovu `b`

```js
assert.notStrictEqual(myArray[1], 'b');
```

`myArray[2]` treba da bude jednako slovu `c`

```js
assert.strictEqual(myArray[2], 'c');
```

`myArray[3]` treba da bude jednako slovu `d`

```js
assert.strictEqual(myArray[3], 'd');
```

# --seed--

## --seed-contents--

```js
let myArray = ["a", "b", "c", "d"];
// Only change code below this line

// Only change code above this line
console.log(myArray);
```

# --solutions--

```js
let myArray = ["a", "b", "c", "d"];
myArray[1] = "e";
```
