---
id: 699d1f8e51fcdc8230f779fb
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, kreiraćete kataloški sistem za biblioteku nasleđa. Parsiraće se sirovi podaci o knjigama, pretraživati i grupisati unose, validisati kvalitet podataka i eksportovati rezultate u različite formate.

Svaka katalog kartica je sačuvana kao pojedinačni niz u formatu `"Title | Author | Year | Location"`, sa poljima razdvojenim znakom pipe (|) (`|`).

Koristićete seriju ovih nizova tokom radionice za izradu i testiranje vaših funkcija kataloga. Kreirajte niz sa nazivom `rawCatalogCards` i dodajte mu sledeće dve linije:

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

Trebalo bi da deklarišete varijablu pod nazivom `rawCatalogCards`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` bi trebalo da bude niz.

```js
assert.isArray(rawCatalogCards);
```

Prvi element `rawCatalogCards` treba da bude `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

Drugi element `rawCatalogCards` bi trebalo da bude `"The Shining | King, Stephen | 1977 | Shelf K1"`.

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
