---
id: 699d20623e57e8ffee885456
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Sada kada su delovi čisti, možete ih pretvoriti u strukturirani objekat. U JavaScriptu, objekti koriste vitičaste zagrade sa parovima "ključ: vrednost". Primer je, `{ title: "Dune", year: 1965 }`.

Deklarišete četiri promenljive: `title`, `author`, `year`, i `location`. Dodelite svakoj od njih odgovarajući element iz `trimmedParts`. Zatim ažurirajte deklaraciju `return` da vrati objekat sa ovim četiri svojstva.

# --hints--

`parseCard` bi trebalo da return objekat sa `title`, `author`, `year`, i `location` svojstava.

```js
const result = parseCard("The Stand | King, Stephen | 1978 | Shelf K2");
assert.isObject(result);
assert.property(result, "title");
assert.property(result, "author");
assert.property(result, "year");
assert.property(result, "location");
```

Svojstvo `title` objekta vraćenog od strane `parseCard("The Stand | King, Stephen | 1978 | Shelf K2")` treba da ima vrednost `"The Stand"`.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").title, "The Stand");
```

Svojstvo ``author`` objekta koji je vraćen od ``parseCard("The Stand | King, Stephen | 1978 | Shelf K2")`` trebalo bi da ima vrednost ``"King, Stephen"``.

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").author, "King, Stephen");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
--fcc-editable-region--

  return trimmedParts;
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
