---
id: 699d20623e57e8ffee885454
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Prvi korak u parsiranju kataloške kartice je razdvajanje četiri polja. Metoda `.split()` deli niz na niz podnizova pri svakom pojavljivanju razdelioca koji pružite.

Primer:

```js
"The Shining | King, Stephen | 1977 | Shelf K1".split("|")
// ["The Shining ", " King, Stephen ", " 1977 ", " Shelf K1"]
```

Unutar `parseCard`, pozovite `.split("|")` na `rawString` i vratite rezultat. Nakon ovog koraka, konzola bi trebalo da prikaže niz od četiri niza.

**Napomena:** Možda primetite dodatni razmak oko vrednosti. To ćete očistiti u sledećem koraku.

# --hints--

Trebalo bi da pozovete `.split("|")` na `rawString` unutar vaše `parseCard` funkcije.

```js
assert.match(parseCard.toString(), /rawString\s*\.\s*split\s*\(\s*('|"|`)\|\1\s*\)\s*;?/);
```

`parseCard` bi trebalo da return niz.

```js
assert.isArray(parseCard("A | B | C | D"));
```

Kada se pozove sa `"A | B | C | D"`, `parseCard` treba da return niz sa četiri elementa.

```js
assert.strictEqual(parseCard("A | B | C | D").length, 4);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
--fcc-editable-region--
  
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
