---
id: 642df9df4b5216350de7b0d2
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

BiÄ‡e vam potrebna funkcija za kreiranje niza brojeva.

DeklariĹˇite praznu funkciju za `range` koja prima parametar(e) od `start` i `end`. Koristite komponentu `Array()` i vratite prazan red indirektno.

# --hints--

Trebalo bi da objaviĹˇ promenljivi deo `range`.

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

Trebalo bi da koristiĹˇ `const` za oglaĹˇavanje svog promenljivog sadrĹľaja od `range`.

```js
assert.match(code, /const\s+range/);
```

Tvoj dinamiÄŤki deo `range` mora biti funkcionalan.

```js
assert.isFunction(range);
```

VaĹˇa funkcija `range` treba da koristi strelicastu sintaksu.

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

Tvoja funkcija `range` mora da primi parametar `start` prvo.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

Tvoja funkcija za `range` treba da primi drugi parametar `end`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

VaĹˇa funkcija `range` bi trebalo da koristi indirektan povrat. Zapamtite, ovo znaÄŤi da neÄ‡ete koristiti vitiÄŤaste zagrade.

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

Tvoja funkcija za `range` treba da koristi komponentu `Array()`. Komponenta osnovnih objekata ne zahteva glavnu kljuÄŤnu reÄŤ od `new`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

Ne Ĺľelite da prenesete niĹˇta za komponentu `Array()`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
