---
id: 642db8c409d9991d0b3b2f0d
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Tvoj projekat počinje sa osnovnim kontejnerom HTML i nekim odgovarajućim CSS. Tvoja prva vežba će biti programiranje ćelija u svojoj kalkulacionoj tablici.

Globalni element `window` predstavlja prozor pretraživača (ili karticu/tab). Ima svojstvo `onload` koje mu omogućava da definiše ponašanje kada prozor učita celu stranicu, uključujući listove stilova i skripte programa.

Počnite postavljanjem svojstva `onload` od `window` kao pointer na funkciju bez parametara. U toj funkciji, deklarišite parametar `container` i postavite mu vrednost za dobijanje elementa za `id` od `"container"`.

# --hints--

Trebalo bi da dobijete pristup atributu `onload` elementa `window`.

```js
assert.match(code, /window\.onload/);
```

Trebalo bi da postaviš atribut `onload` kao funkciju/metodu.

```js
assert.isFunction(window.onload);
```

Trebalo bi da koristiš sintaksu strelice.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

Funkcija `onload` ne bi trebalo da uzima nikakve parametre.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

Treba da deklariraš parametar `container` u svojoj funkciji `onload`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

Tvoj parametar od `container` treba da se objavi za `const`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

Trebalo bi da koristite `document.getElementById()`

```js
assert.match(code, /document\.getElementById\(/);
```

Trebalo bi da dobiješ element za `id` od `"container"`.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

Trebalo bi da staviš element `#container` na `container`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
