---
id: 642ddfdea4200e313f80a4b6
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Zapamtite da ima nešto od `document` putanju `.createElement()` koja vam omogućava kreiranje novih komponenti za HTML putem modifikacija.

Unutar vašeg funkcionera od `createLabel`, deklariši parametar za `label` i postavi vrednost novog polja za `div`.

# --hints--

Treba da deklariraš parametar za `label` u svojoj funkciji `createLabel`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*(?:const|let|var)\s+label/);
```

Tvoj parametar za `label` mora biti objavljen na `const`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*const\s+label/);
```

Trebalo bi da koristiš način `.createElement()` objekta `document`.

```js
assert.match(code, /document\.createElement\(/);
```

Treba da prođeš niz slova `"div"` putem `.createElement()`.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)div\1\s*\)/);
```

Trebalo bi da staviš svoj novi element `div` na `label`.

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {

  }
}
--fcc-editable-region--
```
