---
id: 642dccb78549c9285835ebc2
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Funkcije su najbolje za logiku koja se može ponovo koristiti. Kada funkcija sama treba da ponovno koristi logiku, možete deklarisati ugnježdenu funkciju koja će se baviti tom logikom. Evo primera ugnježdenoj funkciji:

```js
const outer = () => {
  const inner = () => {

  };
};
```

Objavi funkciju definisanu unutar `createLabel` koristeći sintaksu strelice. Trebalo bi da uzme argument `name`.

# --hints--

Trebalo bi da objaviš razliku od `createLabel` u svojoj funkciji `onload`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?)?\s*(?:let|var|const)\s+createLabel/);
```

Tvoja razlika (ili: Tvoj *commit*) sa `createLabel` treba da se objavi nakon tvoje razlike (ili: tvog *commita*) sa `container`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?\s*(?:let|var|const)\s+createLabel/);
```

Vaš razlika od `createLabel` treba da se objavi na `const`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?\s*const\s+createLabel/);
```

Tvoja razlika od `createLabel` bi trebalo da bude strelna funkcija.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?\s*const\s+createLabel\s*=\s*(\(.*\)|[^\s()]+)\s*=>/);
```

Vaša funkcija sa `createLabel` mora imati parametar za `name`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?\s*const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>/);
```

Funkcija za `createLabel` mora biti prazna.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)\s*;?\s*const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");

}
--fcc-editable-region--
```
