---
id: 642def66e6a60432c9a0371e
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Postavi `className` za element `label` na `"label"`, i postavi `textContent` kao parametar za `name`.

# --hints--

Trebalo bi da dobijete atribut za element `className` sa elementa `label`.

```js
assert.match(code, /label\.className/);
```

Treba da postavite atribut za `className` na niz znakova `"label"`.

```js
assert.match(code, /label\.className\s*=\s*('|"|`)label\1/);
```

Trebalo bi da dobiješ atribut `textContent` za element `label`.

```js
assert.match(code, /label\.textContent/);
```

Treba da postavite atribut `textContent` na vrednost `name`.

```js
assert.match(code, /label\.textContent\s*=\s*name/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");

  }
}
--fcc-editable-region--
```
