---
id: 642e02be7845f13b014cd2b0
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Sada, nakon Ĺˇto imate funkciju `range`, moĹľete ga koristiti i za kreiranje niza znakova.

DeklariĹˇite funkciju `charRange` koristeÄ‡i `const` i strelicastu sintaksu. Ta funkcija mora da primi parametar `start` i `end`. Funkcija mora da vrati rezultat poziva `range()` za `start` i `end` kao argument.

# --hints--

Trebalo bi da objaviĹˇ parametar `charRange`.

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

VaĹˇ parametar za `charRange` mora se objaviti za `const`.

```js
assert.match(code, /const\s+charRange/);
```

Parametar za `charRange` treba da bude funkcionalan.

```js
assert.isFunction(charRange);
```

Tvoja funkcija za `charRange` mora da koristi strelnu sintaksu.

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

Funkcija ``charRange`` mora da uzme ``start`` kao prvi parametar.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

Tvoja funkcija za `charRange` bi trebalo da uzme `end` kao drugi parametar.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

VaĹˇa funkcija `charRange` treba da koristi nejasnu/neprozirnu povratnu vrednost.

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

Tvoja funkcija za `charRange` mora pozvati tvoju funkciju za `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

Trebalo bi da prosledite `start` i `end` kao argumente u vaĹˇem pozivu `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
