---
id: 643475e13dc727231acd0f72
title: Korak 17
challengeType: 0
dashedName: step-17
---

# --description--

Sada, unutar vašeg definisanog poziva u `.forEach()`, deklarišite parametar za `input`. Koristite metodu `.createElement()` elementa/svojstva `document` za kreiranje elementa/svojstva `input`. Postavite svojstvo elementa `type` na `"text"` i svojstvo elementa `id` na `letter + number`.

# --hints--

Trebalo bi da objaviš parametar `input`.

```js
assert.match(code, /(?:var|let|const)\s+input/)
```

Trebalo bi da koristiš `const` za deklarisanje svog parametra `input`.

```js
assert.match(code, /const\s+input/)
```

Trebalo bi da pozoveš putanju `.createElement()` elementa `document`.

```js
assert.lengthOf(code.match(/document\.createElement\(/g), 2)
```

Trebalo bi da prođete niz slova `"input"` putem `.createElement()`.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Treba da staviš svoj novi element `input` u svoj atribut `input`.

```js
assert.match(code, /const\s+input\s*=\s*document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Trebalo bi da dobiješ pristup atributu elementa `type` tvog elementa `input`.

```js
assert.match(code, /input\.type/);
```

Treba da postavite atribut elementa za `type` vašeg elementa `input` na `"text"`.

```js
assert.match(code, /input\.type\s*=\s*('|"|`)text\1/)
```

Trebaće vam pristup atributu elementa `id` vašeg elementa `input`.

```js
assert.match(code, /input\.id/);
```

Trebalo bi da postaviš atribut elementa `id` tvog elementa `input` na `letter + number`.

```js
assert.match(code, /input\.id\s*=\s*letter\s\+\snumber/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {

    })
--fcc-editable-region--
  })
}
```
