---
id: 6434750c53db16218f41e6e1
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Zapamtite da `range()` vraća red podataka, pa možete povezati putanje redova podataka direktno na poziv funkcije.

Izvršite poziv `range()` sa `1` i `99` kao argumentima, i povežite rute `.forEach()`. Prođite kroz rutu `.forEach()` sa praznim handlerom koji prima `number` kao parametar.

# --hints--

Trebalo bi da pozoveš svoju funkciju `range()`.

```js
assert.lengthOf(code.match(/range\(/g), 2);
```

Trebalo bi da prosledite `1` kao prvi argument u vašem pozivu `range()`.

```js
assert.match(code, /range\(\s*1/);
```

Treba da proslediš `99` kao drugi argument u tvom pozivu `range()`.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)/);
```

Trebalo bi da povežete putanju `.forEach()` u vaš poziv `range()`.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(/);
```

Trebalo bi da prosledite funkciju povratnika na `.forEach()` koristeći sintaksu strelice.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

Tvoja funkcija povratnog poziva mora imati ``number`` kao jedini parametar.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--

--fcc-editable-region--
}
```
