---
id: 6434759f78ec812264ff8f34
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Kada ga pozovete nazad, moraćete izvršiti dva funkcionalna poziva. Počnite pozivanjem `createLabel()` i prosleđivanjem `number` kao argumenta. Trebalo bi da vidite neke brojeve na vašem kalkulacionom listu.

Zatim prati put `.forEach()` na redu podataka za `letters`. Pozovi praznu funkciju povratnog poziva koja prima parametar sa `letter`.

# --hints--

Trebalo bi da pozoveš svoju funkciju `createLabel()`.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*createLabel\(/);
```

Treba da prođeš kroz `number` za svoj poziv od `createLabel()`.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*createLabel\(/)
```

Treba ti da pozoveš putanju `.forEach()` na red podataka `letters`.

```js
assert.lengthOf(code.match(/letters\.forEach\(/g), 2)
```

Trebalo bi da prođeš kroz funkciju povratne veze sa sintaksom strelice na način vašeg `.forEach()`.

```js
assert.match(code, /letters\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>\s*\{/)
```

Vaša funkcija za povratni poziv treba da ima parametar `letter`.

```js
assert.match(code, /letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{/)
```

Funkcija tvog odgovora mora biti prazna.

```js
assert.match(code, /letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{\s*\}/)
```

Vaša funkcija `letters.forEach()` treba da bude smeštena unutar funkcije `range(1, 99).forEach(number => {}`.

```js
assert.match(code, /range\s*\(\s*1\s*,\s*99\s*\)\s*.forEach\s*\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*[^}]*letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{\s*\}\s*\)\s*;?\s*\}\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--
  range(1, 99).forEach(number => {

  })
--fcc-editable-region--
}
```
