---
id: 6434552bcc0a951a0a99df3b
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Tvoja funkcija za `range` očekuje broj, ali tvoje vrednosti od `start` i `end` će biti niz slova (konkretno, biće jedno slovo kao `A`).

Promenite vaše vrednosti za `start` i `end` u vašem pozivu `range()` u broj koristeći metodu `.charCodeAt()` na njima, prosljeđujući broj `0` kao argument toj metodi.

# --hints--

Trebalo bi da koristiš putanju `.charCodeAt()`.

```js
assert.match(code, /\.charCodeAt\(/);
```

Trebalo bi da pozoveš putanju `.charCodeAt()` na `start`.

```js
assert.match(code, /start\.charCodeAt\(/);
```

Morate da prosledite `0` putem `.charCodeAt()` od `start`.

```js
assert.match(code, /start\.charCodeAt\(\s*0\s*\)/);
```

Trebalo bi da pozoveš putanju `.charCodeAt()` na `end`.

```js
assert.match(code, /end\.charCodeAt\(/);
```

Trebalo bi da prođeš `0` putem `.charCodeAt()` od `end`.

```js
assert.match(code, /end\.charCodeAt\(\s*0\s*\)/);
```

Trebalo bi da koristiš metode `.charCodeAt()` direktno u svom pozivu `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--
const charRange = (start, end) => range(start, end);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
