---
id: 643498755d54c6279ba09078
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Mnogi programi za rad sa tabelama imaju ugrađene funkcije za izračunavanja.

Deklariši funkciju `sum` koja prima parametar `nums`, koji će biti niz podataka brojeva. Trebalo bi da vrati rezultat poziva `reduce` u taj niz kako bi sabrao sve brojeve.

# --hints--

Trebalo bi da objavite parametar `sum`.

```js
assert.match(code, /(?:let|const|var)\s+sum/);
```

Morate koristiti `const` da deklarišete vaš parametar za `sum`.

```js
assert.match(code, /const\s+sum/);
```

Vaš kriterijum za `sum` mora biti funkcionalan.

```js
assert.isFunction(sum);
```

Tvoja funkcija za `sum` treba da koristi sinaksu sa strelicom.

```js
assert.match(code, /const\s+sum\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

Tvoja funkcija za `sum` treba da ima parametar `nums`.

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>/);
```

Vaša funkcija `sum` treba da koristi da direktno vrati rezultate.

```js
assert.notMatch(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

Tvoja funkcija za `sum` trebalo bi da vrati rezultat poziva na/u `.reduce()` u `nums`.

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*nums\.reduce\(/);
```

Tvoja funkcija ``sum`` treba da vrati zbir svih brojeva u ``nums``.

```js
const numbers = [1, 2, 3, 4, 5];
assert.equal(sum(numbers), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
