---
id: 643715013330824ecaa70442
title: Korak 23
challengeType: 0
dashedName: step-23
---

# --description--

Sledeća tvoja funkcija će izračunati vrednost `median` brojne sekvence. Počni deklarisanjem strelične funkcije sa parametrom `nums`.

Unutar funkcije, deklariraj parametar `sorted` i postavi mu vrednost na rezultat sortiranja kopije reda `nums`.

Trebalo bi da koristiš metodu `slice()` za kreiranje duboke kopije tog reda.

# --hints--

Trebalo bi da objaviš parametar `median`.

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

Trebalo bi da koristiš `const` za deklarisanje svog parametra `median`.

```js
assert.match(code, /const\s+median/);
```

Vaš parametar za `median` bi trebalo da bude funkcionalan.

```js
assert.isFunction(median);
```

Tvoja funkcija za `median` mora koristiti sintaksu sa strelicom.

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

Vaša funkcija za `median` mora imati parametar za `nums`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

Vaša funkcija `median` ne bi trebalo da koristi implicitno vraćanje vrednosti.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

Tvoja funkcija za `median` treba da ima parametar `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

Trebalo bi da koristiš `const` za deklarisanje svog parametra `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

Trebalo bi da koristiš `.slice()` za umetanje kopije rednog broja `nums` u `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

Treba da povežete putanju `.sort()` i vašu putanju `.slice()`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

Trebalo bi da prosledite funkciju povratnog poziva vašoj metodi od `sort` kako biste tačno sortirali brojeve u uzlaznom redosledu. Koristite vraćanje vrednosti eksplicitnim načinom radi jasnoće.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
