---
id: 6437133052eaf04d7300e622
title: Korak 22
challengeType: 0
dashedName: step-22
---

# --description--

Definišite funkciju `average` koja prima niz brojeva kao argument za `nums`. Ta funkcija bi trebala vratiti aritmetičku sredinu svih brojeva u nizu podataka.

Prosečna vrednost se može izračunati deljenjem zbira svih brojeva u nizu podataka dužinom tog niza. Zapamtite da imate funkciju `sum` koju možete koristiti.

# --hints--

Treba da objaviš parametar `average`.

```js
assert.match(code, /(?:let|const|var)\s+average/);
```

Trebalo bi da koristiš `const` za deklarisanje svog parametra `average`.

```js
assert.match(code, /const\s+average/);
```

Vaš parametar za `average` mora biti funkcionalan.

```js
assert.isFunction(average);
```

Funkcija vam za `average` mora koristiti streličnu sintaksu.

```js
assert.match(code, /const\s+average\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

Tvoja funkcija za `average` mora imati parametar `nums`.

```js
assert.match(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)/);
```

Tvoja funkcija za `average` mora da koristi vraćanje implicitne vrednosti (implicit return).

```js
assert.notMatch(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

Tvoja funkcija od `average` treba da vrati proseka red podataka za `nums`.

```js
assert.equal(average([1,2,3]), 2);
assert.equal(average([1,2,3,4,5]), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
