---
id: 64496df724dd3716a71fe971
title: Korak 26
challengeType: 0
dashedName: step-26
---

# --description--

Svojstva objekata uključuju parove ključ-vrednost. Možete koristiti skraćena imena svojstava kada deklarišete objekat tipa literala. Kada koristite sintaksu sa skraćenim svojstvima, ime parametra postaje ključ svojstva, a njegova vrednost postaje vrednost svojstva.

Sledeći primer prikazuje stavku `user` sa atributima `userId`, `firstName`, na `loggedIn`.

```js
const userId = 1;
const firstName = "John";
const loggedIn = true;

const user = {
  userId,
  firstName,
  loggedIn,
};

console.log(user); // { userId: 1, firstName: 'John', loggedIn: true }
```

Da biste pratili sve funkcije vaše kalkulacione tabele, deklarište objekat `spreadsheetFunctions`. Uz korišćenje sintakse skraćivanja atributa, postavi `sum`, `average` i `median` kao atribute na objekat `spreadsheetFunctions`.

# --hints--

Trebalo bi da objaviš parametar `spreadsheetFunctions`.

```js
assert.match(code, /(?:const|let|var)\s+spreadsheetFunctions/);
```

Trebalo bi da koristiš `const` za deklarisanje svog parametra `spreadsheetFunctions`.

```js
assert.match(code, /const\s+spreadsheetFunctions/);
```

Vaš parametar za `spreadsheetFunctions` trebao bi da bude nešto.

```js
assert.isObject(spreadsheetFunctions);
```

Vaš `spreadsheetFunctions` mora imati svojstvo od `sum`.

```js
assert.property(spreadsheetFunctions, "sum");
```

Tvoja vrednost za `sum` bi trebalo da bude tvoj `sum` atribut.

```js
assert.equal(spreadsheetFunctions?.sum, sum);
```

Tvoj deo `spreadsheetFunctions` treba da ima svojstvo `average`.

```js
assert.property(spreadsheetFunctions, "average");
```

Tvoja vrednost za `average` bi trebalo da bude tvoja funkcija `average`.

```js
assert.equal(spreadsheetFunctions?.average, average);
```

Tvoj deo `spreadsheetFunctions` treba da ima karakteristiku `median`.

```js
assert.property(spreadsheetFunctions, "median");
```

Tvoja vrednost za `median` bi trebalo da bude tvoj atribut `median`.

```js
assert.equal(spreadsheetFunctions?.median, median);
```

Trebalo bi da koristiš sintaksu rešavanja strukture za postavljanje vrednosti svojih atributa.

```js
const objectText = code.replace(/.*const\s+spreadsheetFunctions\s*=\s*\{([^}]*)}.*/s, "$1");
assert.include(objectText, "sum");
assert.include(objectText, "average");
assert.include(objectText, "median");
assert.notInclude(objectText, ":");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
