---
id: 67c562286b29447da020d407
title: Generišite jednokratni tajni broj
challengeType: 25
dashedName: build-a-one-time-password-generator
demoType: onClick
---

# --description--

U ovom laboratoriju, ćete generisati jednokratni tajni broj sa 6 cifara (OTP) i prikazati ga korisniku. OTP će imati rok trajanja nakon 5 sekundi, a novi OTP će biti generisan kada korisnik klikne na dugme `"Generate New OTP"`.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da koristite hook `useEffect` za upravljanje indikatorom brojača dole.
2. Vaš deo `OTPGenerator` treba da vrati element `div` sa imenom klase `container`.
3. `div` sa klasom `container` treba da uključuje sledeće elemente:

- Element sa `h1`, ID-om `otp-title` i tekstom `"OTP Generator"`.
- Element sa `h2`, ID-om `otp-display`, koji prikazuje poruku `"Click 'Generate OTP' to get a code"` ili prikaže generisani OTP ako je dostupan.
- Element sa `p`, ID-om `otp-timer` i atributom `aria-live` postavljen na važeću vrednost koja:
    - Počinje prazan.
    - Prikazuje `"Expires in: X seconds"` nakon klika na dugme, gde `X` predstavlja preostalo vreme do isteka OTP-a.
    - Prikazuje poruku `"OTP expired. Click the button to generate a new OTP."` kada indikator broja unatrag dostigne `0`.
- Element sa `button`, ID-om `generate-otp-button` i oznakom `"Generate OTP"`. Kada se klikne, treba da generiše novi OTP i pokrene indikator broja unatrag od 5 sekundi.
- Dugme `"Generate OTP"` bi trebalo da bude onemogućeno dok indikator broja unatrag radi.

4. Treba da osigurate da se spuštač brojač automatski zaustavi kada dostigne sekund `0` ili spreči nepotrebna ažuriranja.
5. Generisani OTP treba da ima 6 cifara.

# --before-all--

```js
let clock = __FakeTimers.install();
```

# --after-all--

```js
clock.uninstall();
```

# --hints--

Ne smete ukloniti placeholder vrednosti `const { useState, useEffect, useRef } = React;` koji je u kodu.

```js
assert.match(code, /{\s*(useState\s*,\s*useEffect\s*,\s*useRef|useState\s*,\s*useRef\s*,\s*useEffect|useEffect\s*,\s*useState\s*,\s*useRef|useEffect\s*,\s*useRef\s*,\s*useState|useRef\s*,\s*useState\s*,\s*useEffect|useRef\s*,\s*useEffect\s*,\s*useState|useState|useEffect|useRef)\s*}\s*=\s*React;/);
```

Morate prikazati element `div` koji ima klasu `.container`.

```js
const containerDiv = document.querySelector('.container');
assert.exists(containerDiv);
```

`.container` treba da ima karakteristiku `h1` sa identifikatorom `otp-title` i tekstualnim sadržajem `"OTP Generator"`.

```js
const h1 = document.querySelector(".container h1#otp-title");
assert.exists(h1);
assert.equal(h1.textContent, "OTP Generator");
```

`.container` trebalo bi imati polje `h2` sa identifikatorom `otp-display` koje prikazuje OTP kada se generiše.

```js
const h2 = document.querySelector(".container h2#otp-display");
assert.exists(h2);
```

Početak, `h2` unutar `.container` treba da prikaže poruku `"Click 'Generate OTP' to get a code"`.

```js
assert.strictEqual(document.querySelector('.container h2#otp-display').textContent.trim(), "Click 'Generate OTP' to get a code");
```

Element za `.container` mora da ima element `p` sa identifikatorom `otp-timer`.

```js
const container = document.querySelector('.container');
assert.exists(container, "The container element should exist");
const pElements = container.querySelectorAll('p');
assert.strictEqual(pElements.length, 1);
assert.strictEqual(pElements[0].id, "otp-timer", "The p element should have the id otp-timer");
```

Element za `p` mora imati svojstvo `aria-live` postavljeno na `assertive` ili `polite`.

```js
const pElement = document.querySelector('.container p');
assert.exists(pElement?.ariaLive);
assert.oneOf(pElement?.ariaLive, ["assertive", "polite"]);
```

Početak, polje za `p` trebalo bi da bude prazno.

```js
const pElement = document.querySelector('.container p');
assert.strictEqual(pElement?.innerText, "", "The `p` element should be empty");
```

`.container` mora imati element `button` sa identifikatorom `generate-otp-button` i tekst od `"Generate OTP"`.

```js
const button = document.querySelector('.container button#generate-otp-button').textContent;
assert.exists(button);
assert.strictEqual(button, 'Generate OTP');
```

Kada se dugme klikne, trebalo bi da prikaže novi OTP u polju ``h2`` sa identifikatorom ``otp-display``.

```js
try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)

assert.match(h2.textContent?.trim(), /[0-9]/, "OTP should be a number");
} finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

Generisani OTP mora imati 6 cifara.

```js
  try {
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1);
const h2 = document.querySelector(".container h2#otp-display");

assert.match(h2.textContent.trim(), /^[0-9]{6}$/, "OTP should be a 6-digit number");

 } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

Kada se klikne na dugme, element `p` koji ima identifikator `otp-timer` treba da prikaže tajmer broja do nule od 5 sekundi.

```js
  try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)
const p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/, "Countdown should start at 5 seconds");

  } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

Poruka u elementu `p` sa identifikatorom `otp-timer` treba da se ažurira svake sekunde kako bi prikazala preostalo vreme.

```js
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1)
let p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/);
await clock.tickAsync(1000)
p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 4 seconds/, "Countdown should update to 4 seconds after 1 second");
```

Dugme `"Generate OTP"` treba biti onemogućeno kada je brojač odbrojavanje aktivan.

```js
try {
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();
  clock.tick(1);
  
  assert.isTrue(button.disabled);
} finally {
  await clock.runAllAsync();
}
```

Dugme `"Generate OTP"` treba da bude omogućeno kada indikator broja dođe do nule.

```js

const button = document.querySelector(".container button#generate-otp-button");
button.click();
await clock.tickAsync(6000);
  
assert.isFalse(button.disabled);
```

Kada brojač dostigne `0`, treba da prikaže poruku `OTP expired. Click the button to generate a new OTP.`.

```js
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();

  await clock.tickAsync(6000);

  const p = document.querySelector(".container p#otp-timer");
  assert.equal(p.textContent.trim(), "OTP expired. Click the button to generate a new OTP.");
```

Trebalo bi da prenesete deo `OTPGenerator`.

```js
assert.isFunction(window.index.OTPGenerator);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:disabled {
  background-color:#57585c;
  cursor: not-allowed;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {
  const intervalId = useRef(null);
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
    intervalId.current = setInterval(() => {
      setSecondsLeft(prevSeconds => prevSeconds - 1);
    }, 1000);
  };

  useEffect(() => {
    if (secondsLeft === 0) {
      setIsCounting(false);
      clearInterval(intervalId.current);
    }
  }, [secondsLeft]);

  useEffect(
    () => () => {
      clearInterval(intervalId.current);
    },
    []
  );

  return (
    <div className='container'>
      <h1 id='otp-title'>OTP Generator</h1>
      <h2 id='otp-display'>
        {otp ? otp : "Click 'Generate OTP' to get a code"}
      </h2>
      <p id='otp-timer' aria-live='assertive'>
        {isCounting
          ? `Expires in: ${secondsLeft} seconds`
          : hasGeneratedOtp && 
            "OTP expired. Click the button to generate a new OTP."}
      </p>
      <button
        id='generate-otp-button'
        onClick={handleGenerateOtp}
        disabled={isCounting}
      >
        Generate OTP
      </button>
    </div>
  );
};
```
