---
id: 56533eb9ac21ba0edf2244b3
title: Pretvori Celsius u Fahrenheit
challengeType: 1
forumTopicId: 16806
dashedName: convert-celsius-to-fahrenheit
---

# --description--

Formula za konverziju iz Celsiusa u Fahrenheit je temperatura u Celsijusima pomnozena sa `9/5`, pa zatim dodajte `32`.

Data vam je promenljiva `celsius` koja predstavlja temperaturu u Celsijusima. Koristite vec deklarisanu promenljivu `fahrenheit` i postavite njenu vrednost na odgovarajucu Fahrenheit temperaturu za zadatu Celsius temperaturu. Koristite gorenavedenu formulu da biste izvrsili konverziju.

# --hints--

`convertCtoF(0)` treba da vrati broj

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` treba da vrati vrednost `-22`

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` treba da vrati vrednost `14`

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` treba da vrati vrednost `32`

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` treba da vrati vrednost `68`

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` treba da vrati vrednost `86`

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js
function convertCtoF(celsius) {
  let fahrenheit;
  return fahrenheit;
}

convertCtoF(30);
```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
