---
id: a202eed8fc186c8434cb6d61
title: Obrni String
challengeType: 1
forumTopicId: 16043
dashedName: reverse-a-string
---

# --description--

Obrnite dati string i vratite obrnuti string.

Na primer, `"hello"` treba da bude `"olleh"`.

# --hints--

`reverseString("hello")` treba da vrati string.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` treba da vrati string `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` treba da vrati string `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` treba da vrati string `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js
function reverseString(str) {
  return str;
}

reverseString('hello');
```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
