---
id: a24c1a4622e3c05097f71d67
title: Gde Pripadam
challengeType: 1
forumTopicId: 16094
dashedName: where-do-i-belong
---

# --description--

Vratite najnizi indeks na koji vrednost (drugi argument) treba da bude ubacena u niz (prvi argument) nakon sto se niz sortira. Povratna vrednost treba da bude broj.

Na primer, `getIndexToIns([1,2,3,4], 1.5)` treba da vrati `1` zato sto je veci od `1` (indeks 0), ali manji od `2` (indeks 1).

Slicno tome, `getIndexToIns([20,3,5], 19)` treba da vrati `2` zato sto ce niz nakon sortiranja izgledati kao `[3,5,20]`, a `19` je manje od `20` (indeks 2) i vece od `5` (indeks 1).

# --hints--

`getIndexToIns([10, 20, 30, 40, 50], 35)` treba da vrati broj.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` treba da vrati `3`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` treba da vrati broj.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 30));
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` treba da vrati `2`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` treba da vrati broj.

```js
assert.isNumber(getIndexToIns([40, 60], 50));
```

`getIndexToIns([40, 60], 50)` treba da vrati `1`.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` treba da vrati broj.

```js
assert.isNumber(getIndexToIns([3, 10, 5], 3));
```

`getIndexToIns([3, 10, 5], 3)` treba da vrati `0`.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` treba da vrati broj.

```js
assert.isNumber(getIndexToIns([5, 3, 20, 3], 5));
```

`getIndexToIns([5, 3, 20, 3], 5)` treba da vrati `2`.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` treba da vrati `2`.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([2, 20, 10], 19)` treba da vrati broj.

```js
assert.isNumber(getIndexToIns([2, 20, 10], 19));
```

`getIndexToIns([2, 5, 10], 15)` treba da vrati `3`.

```js
assert.strictEqual(getIndexToIns([2, 5, 10], 15), 3);
```

`getIndexToIns([2, 5, 10], 15)` treba da vrati broj.

```js
assert.isNumber(getIndexToIns([2, 5, 10], 15));
```

`getIndexToIns([], 1)` treba da vrati broj.

```js
assert.isNumber(getIndexToIns([], 1));
```

`getIndexToIns([], 1)` treba da vrati `0`.

```js
assert.strictEqual(getIndexToIns([], 1), 0);
```

# --seed--

## --seed-contents--

```js
function getIndexToIns(arr, num) {
  return num;
}

getIndexToIns([40, 60], 50);
```

# --solutions--

```js
function getIndexToIns(arr, num) {
  arr = arr.sort((a, b) => a - b);

  for (let i = 0; i < arr.length; i++) {
    if (arr[i] >= num) {
      return i;
    }
  }

  return arr.length;
}

getIndexToIns([40, 60], 50);
```
