---
id: a302f7aae1aa3152a5b413bc
title: Izracunaj Faktorijel Broja
challengeType: 1
forumTopicId: 16013
dashedName: factorialize-a-number
---

# --description--

Vratite faktorijel datog celog broja.

Ako je ceo broj predstavljen slovom `n`, faktorijel je proizvod svih pozitivnih celih brojeva manjih ili jednakih od `n`.

Faktorijel se cesto zapisuje skraceno kao `n!`

Na primer: `5! = 1 * 2 * 3 * 4 * 5 = 120`

Celi brojevi prosledjeni funkciji bice samo veci ili jednaki nuli.

# --hints--

`factorialize(5)` treba da vrati broj.

```js
assert.isNumber(factorialize(5));
```

`factorialize(5)` treba da vrati `120`.

```js
assert.strictEqual(factorialize(5), 120);
```

`factorialize(10)` treba da vrati `3628800`.

```js
assert.strictEqual(factorialize(10), 3628800);
```

`factorialize(20)` treba da vrati `2432902008176640000`.

```js
assert.strictEqual(factorialize(20), 2432902008176640000);
```

`factorialize(0)` treba da vrati `1`.

```js
assert.strictEqual(factorialize(0), 1);
```

# --seed--

## --seed-contents--

```js
function factorialize(num) {
  return num;
}

factorialize(5);
```

# --solutions--

```js
function factorialize(num) {
  return num < 1 ? 1 : num * factorialize(num - 1);
}

factorialize(5);
```
