---
id: a6e40f1041b06c996f7b2406
title: Ko Nadje Njegovo Je
challengeType: 1
forumTopicId: 16016
dashedName: finders-keepers
---

# --description--

Napravite funkciju koja pretrazuje niz `arr` i vraca prvi element koji prolazi "truth test". To znaci da za element `x` test prolazi ako je `func(x)` jednako `true`. Ako nijedan element ne prodje test, vratite `undefined`.

# --hints--

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` treba da vrati `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` treba da vrati `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js
function findElement(arr, func) {
  let num = 0;
  return num;
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
