---
id: a77dbc43c33f39daa4429b4f
title: Boo Ko
challengeType: 1
forumTopicId: 16000
dashedName: boo-who
---

# --description--

Proverite da li je vrednost klasifikovana kao osnovna boolean vrednost true ili false. Vratite `true` ili `false`.

Osnovne boolean vrednosti su `true` i `false`.

# --hints--

`booWho(true)` treba da vrati `true`.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` treba da vrati `true`.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` treba da vrati `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` treba da vrati `false`.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` treba da vrati `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` treba da vrati `false`.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` treba da vrati `false`.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` treba da vrati `false`.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` treba da vrati `false`.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` treba da vrati `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js
function booWho(bool) {
  return bool;
}

booWho(null);
```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

booWho(null);
```
