---
id: ab6137d4e35944e21037b769
title: Pretvori Svaku Rec u Naslovni Oblik
challengeType: 1
forumTopicId: 16088
dashedName: title-case-a-sentence
---

# --description--

Vratite dati string tako da prvo slovo svake reci bude veliko. Ostatak svake reci treba da bude malim slovima.

Za potrebe ovog zadatka, takodje treba da pisete velikim pocetnim slovom veznicke reci kao sto su `the` i `of`.

# --hints--

`titleCase("I'm a little tea pot")` treba da vrati string.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` treba da vrati string `I'm A Little Tea Pot`.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` treba da vrati string `Short And Stout`.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` treba da vrati string `Here Is My Handle Here Is My Spout`.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js
function titleCase(str) {
  return str;
}

titleCase("I'm a little tea pot");
```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}

titleCase("I'm a little tea pot");
```
