---
id: acda2fb1324d9b0fa741e6b5
title: Potvrdi Zavrsetak
challengeType: 1
forumTopicId: 16006
dashedName: confirm-the-ending
---

# --description--

Proverite da li string (prvi argument, `str`) zavrsava zadatim ciljnim stringom (drugi argument, `target`).

Ovaj izazov _moze_ da se resi metodom `.endsWith()`, uvedenom u ES2015. Ali za potrebe ovog izazova, potrebno je da umesto toga koristite neki od JavaScript substring metoda.

# --hints--

`confirmEnding("Bastian", "n")` treba da vrati `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` treba da vrati `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` treba da vrati `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` treba da vrati `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` treba da vrati `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` treba da vrati `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` treba da vrati `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` treba da vrati `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing", "mountain")` treba da vrati `false`.

```js
assert.isFalse(
  confirmEnding(
    'If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing',
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` treba da vrati `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Vas kod ne treba da koristi ugradjeni `.endsWith()` metod za resavanje ovog izazova.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js
function confirmEnding(str, target) {
  return str;
}

confirmEnding('Bastian', 'n');
```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
