---
id: a77dbc43c33f39daa4429b4f
title: Kreiraj funkciju za proveru tačne ili netačne vrednosti
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

U ovoj laboratoriji, kreiraćeš funkciju koja proverava da li je vrednost definisana kao tačno ili netačno.

Osnovne vrednosti koje su tačne ili netačne su `true` i `false`.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da postoji funkcija nazvana `booWho` koja prima jedan argument.
1. Ako je primljeni argument osnovna booleanska vrednost (tačno ili netačno), funkcija bi trebalo da vrati `true`.
1. Ako je argument bilo koja druga vrednost, funkcija bi trebalo da vrati `false`.

# --hints--

Trebalo bi da imate funkciju `booWho`.

```js
assert.isFunction(booWho);
```

`booWho(true)` treba da vrati `true`.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` trebalo bi da vrati `true`.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` treba da vrati `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` treba da vrati `false`.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` treba da vrati `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` trebalo bi da vrati `false`.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` treba vratiti `false`.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` trebalo bi da vrati `false`.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` treba da vrati `false`.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` trebalo bi da vrati `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
