---
id: 68e4c81577243e4db1676491
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, veštaćete se sa radom sa semantičnim HTML-om izgradnjom web stranice koja uključuje neke od saveta Kvinsija Larsona za dobijanje posla developera. Osnovni HTML šablon je pripremljen za vas.

Počnite kreiranjem elementa `h1` sa tekstom `Quincy's Tips for Getting a Developer Job`.

# --hints--

You should have an `h1` element.

```js
assert.exists(document.querySelector('h1'));
```

Your `h1` element should have the text `Quincy's Tips for Getting a Developer Job`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
