---
id: 68e964b84f954c322376e592
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Ako je izvor citata veb stranica, možete ga navesti pomoću atributa `cite`. Vrednost ovog atributa treba da bude važeća URL adresa. Iako ovaj atribut ne menja prikaz citata u bloku, veoma je koristan jer pruža čitačima ekrana i pretraživačima više informacija o citatu.

Evo primera inline elementa za citat sa atributom `cite`:

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

Dodajte atribut `cite` na inline element za citat sa ovom URL adresom:

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

Your `q` element should have a `cite` attribute.

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

The `cite` attribute of your `q` element should have the value `https://www.freecodecamp.org/news/learn-to-code-book/`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
