---
id: 68ea7e7b916e8270483689b6
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Ispod elementa `h2`, dodaćete još jedan citat Kvinsija. Ovog puta, citat neće biti deo većeg pasusa. Umesto toga, ceo pasus će biti citat. Da biste razlikovali ovakvo citirano teksto, trebalo bi da koristite element za blok citate: `blockquote`. U pretraživaču ćete videti da je tekst blago uvlačen.

Here is an example of a block quotation element with quoted text:

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

Sada, unutar prvog odeljka, dodajte element za blok citate ispod elementa `h2` sa tekstom `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

# --hints--

Inside the first `section` element, you should have a `blockquote` element below the `h2` element.

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

Your `blockquote` element should have the text `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
