---
id: 68e4276483b0c71658aab1d1
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

Ako je potrebno kreirati novu sekvencu karaktera gde su svi karakteri u malim slovima, možete koristiti metodu `toLowerCase()`.

Počnite kreiranjem parametra nazvanog `lowerCaseInput` i postavite vrednost rezultata koristeći metodu `toLowerCase()` u parametar `cleanedInput`.

Zatim ukucajte niz znakova `"Result of using the toLowerCase() method:"` u konzolu.

Na kraju, dodaj još jedan poziv za `console.log()` za parametar `lowerCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da imaš kriterijum za `lowerCaseInput`.

```js
assert.exists(lowerCaseInput);
```

Vaš parametar za `lowerCaseInput` mora biti niz slova.

```js
assert.isString(lowerCaseInput);
```

Trebalo bi da postaviš vrednost `cleanedInput.toLowerCase()` na svoj parametar `lowerCaseInput`.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

Treba da ukucaš niz karaktera `"Result of using the toLowerCase() method:"` u konzolu.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

Treba da unesete parametar za `lowerCaseInput` na konzoli.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
