---
id: 68e4276483b0c71658aab1cc
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Tvoja vrednost `userInput` ima dodatne razmake na početku i na kraju nizu karaktera. Da bi kreirao novu nizu karaktera bez tih dodatnih razmaka, možeš koristiti metodu `trim()`.

Evo primera:

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

Kreiraj parametar nazvan `cleanedInput` i postavi mu vrednost rezultata korišćenjem metode `trim()` na parametar `userInput`.

# --hints--

Trebalo bi da imaš kriterijum za `cleanedInput`.

```js
assert.exists(cleanedInput);
```

Tvoj parametar za `cleanedInput` treba da bude niz slova.

```js
assert.isString(cleanedInput);
```

Trebalo bi da postaviš vrednost `userInput.trim()` na svoj parametar `cleanedInput`.

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
