---
id: 68e4276483b0c71658aab1ce
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Biće trenutaka kada samo želite da obrišete prazan razmak na početku niza slova. Tamo je metoda `trimStart()` korisna.

Evo primera:

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

Zapamtite da sekvenca slova ne može se promeniti. Ovaj način će stvoriti novu sekvencu slova sa rezultatima. Neće promeniti originalnu.

Kreiraj parametar `trimmedStart` i postavi mu vrednost rezultata korišćenja metode `trimStart()` u parametru `userInput`.

# --hints--

Trebalo bi da imaš kriterijum `trimmedStart`.

```js
assert.exists(trimmedStart);
```

Vaš parametar za `trimmedStart` mora biti niz slova.

```js
assert.isString(trimmedStart);
```

Treba da postaviš vrednost `userInput.trimStart()` u svoj parametar `trimmedStart`.

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
