---
id: 68e4276483b0c71658aab1d0
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Biće vremena kada ćete morati da kreirate sekvencu novih karaktera koji sadrže znake iz različitih tipova slova.

Da biste stvorili niz novih karaktera sa svim slovima konvertovanim u VELIKA SLOVA, možete koristiti metod `toUpperCase()`.

Evo primera:

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

Definši parametar nazvan `upperCaseInput` i postavi mu vrednost kao rezultat korišćenja metode `toUpperCase()` na parametru `cleanedInput`.

# --hints--

Trebalo bi da imate kriterijum za `upperCaseInput`.

```js
assert.exists(upperCaseInput);
```

Parametar vašeg `upperCaseInput` mora biti niz slova.

```js
assert.isString(upperCaseInput);
```

Treba da staviš vrednost `cleanedInput.toUpperCase()` u svoj parametar `upperCaseInput`.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
