---
id: 68e4276483b0c71658aab1ca
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, veštaćete korišćenje različitih metoda sekvenciranja znakova koje se koriste za sortiranje stringova.

Počnite kreiranjem parametra pod nazivom `userInput` i postavite mu sledeću vrednost: `"   Hello World!   "`.

# --hints--

Trebalo bi da imaš kriterijum `userInput`.

```js
assert.exists(userInput);
```

Vaše polje za `userInput` treba da bude niz slova.

```js
assert.isString(userInput);
```

Trebalo bi da postavite vrednost `"   Hello World!   "` u parametar `userInput`. Proverite da ima tri razmaka pre `H` i tri razmaka posle `!`.

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
