---
id: 68e4276483b0c71658aab1cd
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Sada je vreme da vidimo rezultate.

Počnite sa dodavanjem `console.log()` i sekvencu karaktera `"Result of trimming whitespace from both ends:"`.

Zatim ispod `console.log()`, dodaj `console.log()` još jedan sa varijablom `cleanedInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da ukucate niz znakova `"Result of trimming whitespace from both ends:"` u konzolu.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

Treba da napišete variable `cleanedInput` na konzoli.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
