---
id: 68e4276483b0c71658aab1d3
title: Korak 14
challengeType: 1
dashedName: step-14
---

# --description--

Kao što se sećate iz prethodnih lekcija i radionica, možete izdvojiti delove niza znakova koristeći metod `slice()`.

Ovo je podsetnik:

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

Trebalo bi da promenite mapping vrednosti postojećeg `camelCasedVersion` menjanjem praznog stringa sa `lowercaseWord.slice(0, 5)`.

Sada bi trebalo da vidiš reč `"camel"` na konzoli.

# --hints--

Vaše polje za `camelCasedVersion` ne sme biti niz praznih znakova.

```js
assert.isNotEmpty(camelCasedVersion);
```

Trebalo bi da postaviš vrednost `lowercaseWord.slice(0, 5)` za parametar `camelCasedVersion`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

Morate uneti tačnu sekvencu slova `"camel"` za `camelCasedVersion`. Molimo ponovo pročitajte uputstva na ispravan način.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
