---
id: 68e4276483b0c71658aab1d2
title: Korak 13
challengeType: 1
dashedName: step-13
---

# --description--

Biće vremena kada ćete morati da koristite više metoda za kreiranje nove sekvence slova. U narednim koracima, koristićete neke od metoda koje ste naučili do sada kako biste napravili novu sekvencu slova koja koristi camelCase konvenciju.

Počnite kreiranjem parametra nazvanog `lowercaseWord` i postavite mu vrednost tekstualnog niza `"camelcase"`.

Ispod tog parametra `lowercaseWord`, kreiraj još jedan parametar nazvan `camelCasedVersion` i privremeno mu dodeli prazan niz znakova.

Pod parametrom `camelCasedVersion`, dodaj `console.log()` koji ima sekvencu slova `"Camel cased version:"` praćenu/nakon čega sledi `console.log()` za parametar `camelCasedVersion`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da imaš parametar koji se zove `lowercaseWord`.

```js
assert.exists(lowercaseWord);
```

`lowercaseWord` Treba da bude niz (ili: sekvenca) slova.

```js
assert.isString(lowercaseWord);
```

Trebalo bi da unesete vrednost sekvence slova `"camelcase"` za vaš parametar `lowercaseWord`.

```js
assert.equal(lowercaseWord, "camelcase");
```

Trebalo bi da imaš kriterijum za `camelCasedVersion`.

```js
assert.exists(camelCasedVersion);
```

Vaš parametar za `camelCasedVersion` treba da bude sekvenca slova.

```js
assert.isString(camelCasedVersion);
```

Trebalo bi da postavite praznu vrednost za niz karaktera za vaš parametar `camelCasedVersion`.

```js
assert.isEmpty(camelCasedVersion);
```

Trebaće ti da ukucaš niz znakova `"Camel cased version:"` u konzolu.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

Trebalo bi da napišete parametar `camelCasedVersion` u konzoli.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
