---
id: 68f1179eaea3bb59ba434386
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Da biste kreirali prazan prostor na kraju niza karaktera, možete koristiti metodu `trimEnd()`.

Evo primera:

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

Kreiraj parametar nazvan `trimmedEnd` i postavi mu vrednost kao rezultat korišćenja metode `trimEnd()` na parametar `userInput`.

# --hints--

Trebalo bi da imaš kriterijum `trimmedEnd`.

```js
assert.exists(trimmedEnd);
```

Vaš parametar za `trimmedEnd` treba da bude niz slova.

```js
assert.isString(trimmedEnd);
```

Treba da postavite vrednost `userInput.trimEnd()` u vaš parametar/polje `trimmedEnd`.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
