---
id: 68f3f890010c870e17bbe754
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Sada je vreme da vidimo rezultate.

Ispod tvoje varijable za `upperCaseInput`, dodaj `console.log()` sa sekvencom slova `"Result of using the toUpperCase() method:"`.

Zatim ispod `console.log()`, dodaj `console.log()` drugi sa varijablom `upperCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Treba da ukucaš niz karaktera `"Result of using the toUpperCase() method:"` u konzolu.

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

Trebalo bi da napišete varijablu `upperCaseInput` u konzoli.

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
