---
id: 68f3f5ca1d2a490c97f14f03
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Sada je vreme da vidimo rezultate uklanjanja neželjenih razmaka na kraju niza karaktera.

Ispod vaše varijable `trimmedEnd`, dodajte `console.log()` sa sekvencom karaktera `"After using the trimEnd() method, trailing spaces removed:"`.

Zatim ispod `console.log()`, ponovo dodaj `console.log()` sa varijablom `trimmedEnd`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da ukucate niz znakova `"After using the trimEnd() method, trailing spaces removed:"` u konzolu.

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

Treba da napišete varijablu `trimmedEnd` u konzoli.

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
