---
id: 68f3eeaa65614107ecd5f1f2
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Ispod vašeg prvog `console.log()`, dodajte drugi `console.log()` da zapamti vrednost promene `userInput` u konzolu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Treba da snimiš promene za `userInput` u konzoli.

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
