---
id: 68f3f3d65d51d80ac0b412fc
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Sada je vreme da vidimo rezultate uklanjanja neželjenih razmaka sa početka niza karaktera.

Ispod vaše varijable `trimmedStart`, dodajte `console.log()` sa nizom karaktera `"After using the trimStart() method, leading spaces removed:"`.

Zatim ispod `console.log()`, dodaj `console.log()` drugu koja ima varijablu od `trimmedStart`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Treba da napišete niz znakova `"After using the trimStart() method, leading spaces removed:"` u konzolu.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

Treba da napišete varijablu `trimmedStart` u konzolu.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
