---
id: 68f405c8c9df44132c8fe1bb
title: Korak 15
challengeType: 1
dashedName: step-15
---

# --description--

Kao što se sećate iz ranijih lekcija, konvencija imenovanja u stilu camelCase počinje malim slovom za prvu reč, a svaka sledeća reč počinje velikim slovom praćen malim slovima.

Evo primera da vas podsetim:

```md
freeCodeCamp
camelCase
learningIsFun
```

Druga reč u promeni `lowercaseWord` je `"case"`. Da biste dobili dostupnost `c` u toj reči, možete koristiti `lowercaseWord[5]`.

Ispravite opseg vrednosti od `camelCasedVersion` ili koristite operater za `+` da spojite `lowercaseWord.slice(0, 5)` sa rezultatima korišćenjem ispravne metode za konverziju sekvence slova u velika slova na `lowercaseWord[5]`.

Sada bi trebalo da vidiš `camelC` na konzoli.

# --hints--

Trebalo bi da povežete `lowercaseWord.slice(0, 5)` sa `lowercaseWord[5].toUpperCase()`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

Ne smete da postavite vrednost sekvence karaktera `"camelC"` direktno na `camelCasedVersion`. Molimo vas da ponovo pročitate uputstva pravilno.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
