---
id: 68f40925d810ae14b2d5a169
title: Korak 16
challengeType: 1
dashedName: step-16
---

# --description--

Trenutno, tvoja varijabla `camelCasedVersion` ima vrednost `camelC`. Ali rezultat koji želiš je `camelCase`.

Da bi se dobili preostali znakovi za `ase`, možete koristiti metodu `slice()` sa negativnom indeksnom vrednošću. Zapamtite, korišćenje negativnih indeksnih vrednosti govori metodi `slice()` da izvuku od kraja.

Evo primera podsećanja:

```js
const str = "abcd";
str.slice(-2); // cd
```

Validiraj opsege `camelCasedVersion` i poveži `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` sa rezultatima korišćenjem metode `slice()` na `lowercaseWord`. Unesite broj `-3` putem `slice()`.

Sa ovim poslednjim izmenama, vaša radionica će biti završena!

# --hints--

Trebalo bi da povežete `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` sa `lowercaseWord.slice(-3)`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

Treba da umetnete sekvencu karaktera za `"camelCase"` direktno u `camelCasedVersion`. Molimo, ponovo pročitajte uputstva na ispravan način.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
