---
id: 66c057041df6394ca796bf33
title: Napravi generator rečenica
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

U ovoj laboratoriji, kreiraćeš dve različite priče koristeći šablon rečenice. Koristićeš parametre da sačuvaš različite delove priče, a zatim ćeš te priče izbaciti na konzolu.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratorijum.

Priče o korisniku:

1. Trebalo bi da deklarišete sledeće parametre koristeći `let`:
   
    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

Trebalo bi da unesete vrednosti sekvence slova za navedene parametre u skladu sa vašim izborom.

1. Morate deklarisati parametar ``firstStory``.
   
1. Morate koristiti šablon priče ispod da biste kreirali prvu priču i postavili je u parametar ``firstStory``: ``"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."``;

Trebalo bi da izneseš svoju prvu priču na konzolu koristeći poruku `"First story: [firstStory]"`.

1. Morate da postavite novu vrednost za vaše parametre `adjective`, `noun`, `verb`, `place`, `adjective2`, i `noun2`.

1. Morate deklarisati parametar `secondStory`.

1. Kreiraj drugu priču koristeći isti šablon i stavi je u parametar `secondStory`.

Trebalo bi ti da izneseš svoju drugu priču na konzolu koristeći poruku `"Second story: [secondStory]"`.


# --hints--

Morate definisati parametar za `adjective`.

```js
assert.isNotNull(adjective);
```

Trebalo bi da objaviš parametar `noun`.

```js
assert.isNotNull(noun);
```

Treba da deklariš parametar za `verb`.

```js
assert.isNotNull(verb);
```

Treba da objaviš parametar `place`.

```js
assert.isNotNull(place);
```

Trebalo bi da objaviš parametar `adjective2`.

```js
assert.isNotNull(adjective2);
```

Trebalo bi da deklariš parametar `noun2`.

```js
assert.isNotNull(noun2);
```

Treba da unesete vrednost niza karaktera za parametar `adjective`.

```js
assert.isString(adjective);
```

Morate postaviti vrednost tekstualnog niza za parametar `noun`.

```js
assert.isString(noun);
```

Trebalo bi da postavite vrednost sekvence karaktera za parametar `verb`.

```js
assert.isString(verb);
```

Trebalo bi da unesete vrednost niza karaktera za parametar `place`.

```js
assert.isString(place);
```

Trebalo bi da unesete vrednost niza karaktera za parametar `adjective2`.

```js
assert.isString(adjective2);
```

Treba da postavite vrednost niza karaktera za parametar `noun2`.

```js
assert.isString(noun2);
```

Trebalo bi da objavite parametar `firstStory`.

```js
assert.isNotNull(firstStory);
```

Trebalo bi da koristiš ispravan format priče za prvu priču: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Obrati pažnju na razmak.

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

Trebalo bi da povežete svoju prvu priču koristeći parametre koje ste naveli u ispravnom redosledu.

```js
assert.match(
  __helpers.removeJSComments(code),
  /firstStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

Trebalo bi ti da napišeš svoju prvu priču koristeći poruku `"First story: [firstStory]"`.

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

Trebalo bi da deklariš polje `secondStory`.

```js
assert.isNotNull(secondStory);
```

Trebalo bi da postaviš novu vrednost za parametar `adjective` za drugu priču.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

Trebalo bi da postaviš novu vrednost za parametar `noun` za drugu priču.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

Trebalo bi da postaviš novu vrednost za parametar `verb` za drugu priču.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

Treba da postavite novu vrednost za parametar `place` za drugu priču.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

Treba da postaviš novu vrednost za parametar `adjective2` za drugu priču.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

Morate postaviti novu vrednost za parametar `noun2` za drugu priču.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

Treba da koristiš ispravan format priče za drugu priču: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Obrati pažnju na razmak.

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

Morate povezati svoju drugu priču korišćenjem parametara koje ste naveli u pravilnom redosledu.

```js
assert.match(
  __helpers.removeJSComments(code),
  /secondStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

Trebalo bi ti da napišeš svoju drugu priču koristeći format `"Second story: [secondStory]"`.

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

`firstStory` ne sme biti jednako `secondStory`.

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```
