---
id: 66629f407d679d3105e8317f
title: Jenga mchezo wa jiwe karatasi mkono
challengeType: 25
dashedName: top-build-a-rock-paper-scissors-game
---

# --description--
Tvoja igra će se igrati protiv računara. Napisaćeš funkciju koja nasumično vraća `"rock"`, `"paper"` ili `"scissors"`.

Ne moraš da brineš o korisničkom interfejsu igre. Pisaćeš samo logiku igre. Otvori `script.js` i počni da pratiš uputstva.

**Hadithi za mtumizi:**

1. Treba da imaš funkciju pod nazivom `getComputerChoice`.

1. Funkcija `getComputerChoice` treba nasumično da vraća `"rock"`, `"paper"` ili `"scissors"`.

**Savet:** Metoda [Math.random](https://developer.mozilla.org/en-US/docs/Web/JavaScript/Reference/Global_Objects/Math/random) vraća nasumičan broj veći ili jednak 0, a manji od 1. Razmisli kako možeš da je iskoristiš da nasumično vrati jednu od više mogućih opcija.

Igrom će upravljati ljudski igrač. Napisaćeš funkciju koja prihvata izbor korisnika i vraća ga.

1. Tengeneza kitendakazi kinachoitwa `getHumanChoice`.

1. Napiši kod tako da `getHumanChoice` vrati jednu od važećih opcija na osnovu korisnikovog izbora.

**Savet:** Koristi metodu [prompt](https://developer.mozilla.org/en-US/docs/Web/API/Window/prompt) da dobiješ korisnički unos.

Tvoja igra će pratiti rezultat igrača. Napisaćeš promenljive za čuvanje rezultata.

1. Napravi dve nove globalne promenljive pod nazivom `humanScore` i `computerScore`.

1. Inicijalizuj obe promenljive vrednošću `0`.

Igra će se igrati rundu po rundu. Napisaćeš funkciju koja prima izbor igrača i računara kao argumente, odigrava jednu rundu, uvećava rezultat pobednika i ispisuje poruku o pobedniku.

1. Tengeneza kitendakazi kipya kinachoitwa `playRound`.

1. Definiši dva parametra za `playRound`. Prvi parametar neka bude `humanChoice`, a drugi `computerChoice`. Koristi ih da primiš izbore igrača i računara kao argumente.

1. Omogući da parametar `humanChoice` u funkciji `playRound` ne razlikuje velika i mala slova, tako da igrači mogu uneti `"rock"`, `"ROCK"`, `"RocK"` ili bilo koju drugu kombinaciju.

1. Napiši kod za funkciju `playRound` tako da vraća tekstualnu poruku koja predstavlja pobednika runde.

- Ako je nerešeno, treba da vrati `"It's a tie!"`.

- Ako igrač pobedi, treba da vrati `"You win! [player choice] beats [computer choice]"`.

- Ako računar pobedi, treba da vrati `"You lose! [computer choice] beats [player choice]"`.

1. Uvećaj vrednost promenljive `humanScore` ili `computerScore` u zavisnosti od pobednika runde.

Igra će trajati 3 runde. Napisaćeš funkciju `playGame` koja poziva `playRound` tri puta, prati rezultat i na kraju proglašava pobednika.

1. Tengeneza kitendakazi kipya kinachoitwa `playGame`.

1. Napravi petlju koja igra 3 runde i pri svakom prolazu poziva funkciju `playRound` sa izborom igrača i izborom računara kao argumentima.

1. Mwishoni mwa mchezo, rudisha mshindi wa mchezo kulingana na aliyeshinda raundi nyingi zaidi.

- Ako ljudski igrač pobedi u više rundi od računara, vrati poruku `"You win the game!"`.

- Ako računar pobedi u više rundi od ljudskog igrača, vrati poruku `"You lose the game!"`.

# --hints--

Treba da postoji funkcija pod nazivom `getComputerChoice`.

```js
assert.isFunction(getComputerChoice);
```

Vaša funkcija `getComputerChoice` treba nasumično da vraća `"rock"`, `"paper"` ili `"scissors"`.

```js
const counts = {}
for (let i = 0; i < 100; i++) {
  const result = getComputerChoice();
  counts[result] = (counts[result] ?? 0) + 1;
}
assert.lengthOf(Object.keys(counts), 3);
assert.include(Object.keys(counts), "rock");
assert.include(Object.keys(counts), "paper"); 
assert.include(Object.keys(counts), "scissors");
```

Treba da postoji funkcija pod nazivom `getHumanChoice`.

```js
assert.isFunction(getHumanChoice);
```

Treba da postoje dve promenljive pod nazivom `humanScore` i `computerScore` u globalnom opsegu.

```js
assert.exists(humanScore);
assert.exists(computerScore);
```

Treba da postoji funkcija pod nazivom `playRound`.

```js
assert.isFunction(playRound);
```

Vaša funkcija `playRound` treba da prihvata izbor igrača i izbor računara kao argumente putem parametara `humanChoice` i `computerChoice`.

```js
assert.match(playRound.toString(), /\s*(?:\bhumanChoice\b\s*,\s*\bcomputerChoice\b)/);
```

Vaša funkcija `playRound` ne treba da razlikuje velika i mala slova, tako da igrači mogu da unesu `"rock"`, `"ROCK"`, `"RocK"` ili bilo koju drugu varijaciju.

```js
assert.match(playRound.toString(), /\bhumanChoice\s*\.toLowerCase\(\)/);
```

Vaša funkcija `playRound` treba da vrati string `"It's a tie!"` ako su izbor igrača i računara isti.

```js
assert.equal(playRound("rock", "rock"), "It's a tie!");
```

Vaša funkcija `playRound` treba da vrati string `"You win! [player choice] beats [computer choice]"` ako igrač pobedi.

```js
assert.equal(playRound("rock", "scissors"), "You win! rock beats scissors");
```

Vaša funkcija `playRound` treba da vrati string `"You lose! [computer choice] beats [player choice]"` ako računar pobedi.

```js
assert.equal(playRound("rock", "paper"), "You lose! paper beats rock");
```

Vaša funkcija `playRound` treba da poveća vrednost promenljive `humanScore` ili `computerScore` u zavisnosti od pobednika runde.

```js
humanScore = 0;
computerScore = 0;
playRound("rock", "scissors");
assert.equal(humanScore, 1);
```

Treba da postoji funkcija pod nazivom `playGame`.

```js
assert.isFunction(playGame);
```

Treba da koristite petlju za odigravanje 3 runde.

```js
assert.match(playGame.toString(), /\bfor\s*\(/);
```

Treba da vratite pobednika igre na osnovu toga ko je pobedio u više rundi.

```js
window.prompt = () => "rock";
assert.match(playGame(), /You (win|lose) the game!/);
```

# --seed--

## --seed-contents--

```js
const hand = ['rock', 'paper', 'scissors'];
```

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

```js
const hand = ['rock', 'paper', 'scissors'];


let computerScore = 0;
let humanScore = 0;

const getComputerChoice = () => {
  return hand[Math.floor(Math.random() * hand.length)];
}

const getHumanChoice = () => {
  return prompt();
}

const playRound = (humanChoice, computerChoice) => {
  humanChoice = humanChoice.toLowerCase()

  const tie = "It's a tie!"
  const win = `You win! ${humanChoice} beats ${computerChoice}`
  const lose = `You lose! ${computerChoice} beats ${humanChoice}`


  if (humanChoice == 'rock') {
    if (computerChoice == 'rock') {
      computerScore++;
      humanScore++;
      return tie;
    } else if (computerChoice == 'paper') {
      computerScore++;
      return lose;
    } else {
      humanScore++;
      return win;
    }
  } else if (humanChoice == 'paper') {
    if (computerChoice == 'rock') {
      humanScore++;
      return win;
    } else if (computerChoice == 'paper') {
      computerScore++;
      humanScore++;
      return tie;
    } else {
      computerScore++;
      return lose;
    }
  } else if (humanChoice == 'scissors') {
    if (computerChoice == 'rock') {
      computerScore++;
      return lose;
    } else if (computerChoice == 'paper') {
      humanScore++;
      return win;
    } else {
      computerScore++;
      humanScore++;
      return tie;
    }
  }
}

const playGame = () => {
  for(let i = 0; i < 3; i++){
    playRound(getHumanChoice(), getComputerChoice())
  }

  return humanScore > computerScore ? "You win the game!" : "You lose the game!"
}

```
