---
id: 587d8254367417b2b2512c71
title: Ukloni nit (iličnik) iz seta u ES6
challengeType: 1
forumTopicId: 301713
dashedName: remove-items-from-a-set-in-es6
---

# --description--
Pokušaćemo ukloniti elemente iz ES6 Set-a koristeći metod `delete`.

Prvo, kreiraj ES6 Set:

```js
var set = new Set([1,2,3]);
```

Sada uklonite element iz vašeg Skupa koristeći `delete`.

```js
set.delete(1);
console.log([...set]) // should return [ 2, 3 ]
```

# --instructions--
Sada, kreiraj skup sa kompletnim brojevima 1, 2, 3, 4 i 5.

Uklonite vrednost od 2 do 5, a zatim vratite taj set.

# --hints--
Tvoj set treba da ima vrednosti 1, 3 i 4

```js
assert(
  (function () {
    var test = checkSet();
    return test.has(1) && test.has(3) && test.has(4) && test.size === 3;
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet(){
  // Only change code below this line
  var set = null;

  // Only change code above this line
  return set;   
}
```

# --solutions--

```js
function checkSet(){
var set = new Set([1,2,3,4,5]);
set.delete(2);
set.delete(5);
return set;}
```
