---
id: 587d825b367417b2b2512c8c
title: Izvrši sortiranje po najnižoj vrednosti
challengeType: 1
forumTopicId: 301643
dashedName: implement-heap-sort-with-a-min-heap
---

# --description--
Sada, kada možemo dodavati i uklanjati elemente, istražimo neke programe u kojima se mogu koristiti rondi (heapovi). Heapovi se široko koriste za implementaciju cola sa prioritetom jer uvijek čuvaju element s najvećom ili najmanjom vrijednošću na prvoj poziciji. Pored toga, koriste se za implementaciju algoritma sortiranja koji se naziva heap sort. Pokazaće nam kako to učiniti ovdje. Heap sort koristi minimum heap (minimalni rondo), što je suprotno od maximum heap-a (maksimalnog ronda). Minimalni rondo uvijek čuva element s najmanjom vrijednošću na poziciji korijena.

Heap sort radi uzimanjem neorganizovanog niza podataka, dodavanjem svakog elementa u minimalni heap (min-heap), a zatim uklanjanjem svakog elementa iz tog heapa i postavljanjem ga u novi niz. Struktura minimalnog heapa osigurava da će novi niz imati elemente poređane od najmanjeg do najvećeg. Ovo je jedan od efikasnih algoritama sortiranja, sa prosečnom i najgorom performansom od O(nlog(n)).

# --instructions--
Implementiraj heap sort koristeći min-heap. Molimo vas da ovde promenite svoj max-heap kod. Kreirajte objekat `MinHeap` sa metodama `insert`, `remove` i `sort`. Metoda `sort` treba da vrati niz svih elemenata iz min-heap-a, sortirovanih od najmanjeg do najvećeg.

# --hints--
Struktura podataka `MinHeap` mora postojati.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MinHeap` treba imati metod nazvanu `insert`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MinHeap` treba da ima putanju nazvanu `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

`MinHeap` treba da ima metodu nazvanu `sort`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.sort == 'function';
  })()
);
```

Njia `sort` treba da vrati red sa svim dodanim elementima u najnižem steku po redosledu.

```js
assert(
  (() => {
    if (typeof MinHeap === 'undefined') {
      return false;
    }

    const heap = new MinHeap();
    const arr = createRandomArray(25);

    for (let i of arr) {
      heap.insert(i);
    }

    const result = heap.sort();
    arr.sort((a, b) => a - b);

    for (let i = 0; i < arr.length; i++) {
      if (arr[i] !== result[i]) {
        return false;
      }
    }
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
// Generate a randomly filled array
function createRandomArray(size = 5){
  let a = new Array(size);
  for(let i = 0; i < size; i++)
    a[i] = Math.floor(Math.random() * 100);
  
  return a;
}
const array = createRandomArray(25);

var MinHeap = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
