---
id: 587d8256367417b2b2512c77
title: Lista odnosa nodesa
challengeType: 1
forumTopicId: 301620
dashedName: adjacency-list
---

# --description--
Grafovi se mogu predstaviti na različite načine. Ovde objašnjavamo jedan način, poznat kao <dfn>lista odnosa</dfn>. Lista odnosa je lista sa oznakama gde je levi deo čvor, a desni deo navodi sve ostale čvorove povezane s njim. Ispod je reprezentacija liste odnosa.

<blockquote>Čvor1: Čvor2, Čvor3<br>Čvor2: Čvor1<br>Čvor3: Čvor1</blockquote>

Ovo je nedirektovani graf jer su `Node1` povezani sa `Node2` i `Node3`, a ta informacija odgovara vezi prikazanoj na `Node2` i `Node3`. Lista veza za direktovan graf znači da svaki red u listi pokazuje smer. Da je ono gore direktovano, onda bi `Node2: Node1` značilo da postoji tok/smer koji teče od `Node2` ka `Node1`. Nedirektovani graf gore možemo predstaviti kao listu veza stavivanjem u objekat JavaScript.

```js
var undirectedG = {
  Node1: ["Node2", "Node3"],
  Node2: ["Node1"],
  Node3: ["Node1"]
};
```

Ovo se takođe može predstaviti kao podatkovni niz gde čvorovi imaju brojeve umesto alfabetskih oznaka.

```js
var undirectedGArr = [
  [1, 2], // Node1
  [0],    // Node2
  [0]     // Node3
];
```

# --instructions--
Izgradite društvenu mrežu kao nepoumereni graf sa 4 čvora/spomenute osobe: `James`, `Jill`, `Jenny`, i `Jeff`. Postoje veze/odnosi između Jamesa i Jeffa, Jill i Jenny, i Jeffa i Jenny.

# --hints--

`undirectedAdjList` treba da sadrži samo četiri čvora.

```js
assert(Object.keys(undirectedAdjList).length === 4);
```

Mora postojati veza između `Jeff` i `James`.

```js
assert(
  undirectedAdjList.James.indexOf('Jeff') !== -1 &&
    undirectedAdjList.Jeff.indexOf('James') !== -1
);
```

Treba da postoji konekcija između `Jill` i `Jenny`.

```js
assert(
  undirectedAdjList.Jill.indexOf('Jenny') !== -1 &&
    undirectedAdjList.Jenny.indexOf('Jill') !== -1
);
```

Treba da postoji konekcija između `Jeff` i `Jenny`.

```js
assert(
  undirectedAdjList.Jeff.indexOf('Jenny') !== -1 &&
    undirectedAdjList.Jenny.indexOf('Jeff') !== -1
);
```

# --seed--

## --seed-contents--

```js
var undirectedAdjList = {};
```

# --solutions--

```js
var undirectedAdjList = {
  James: ['Jeff'],
  Jill: ['Jenny'],
  Jenny: ['Jill', 'Jeff'],
  Jeff: ['James', 'Jenny']
};
```
