---
id: 587d8251367417b2b2512c61
title: Radi sa čvorom u povezanoj listi
challengeType: 1
forumTopicId: 301721
dashedName: work-with-nodes-in-a-linked-list
---

# --description--
Još jedna struktura podataka koja se često susreće u računarstvu je <dfn>spojena lista</dfn>. Spojena lista je linearni skup elemenata podataka, koji se nazivaju 'čvorovi', a svaki ukazuje na sledeći. Svaki <dfn>čvor</dfn> u spojenoj listi ima dva ključna elementa informacija: `element` sam, i referencu na sledeći `node`.

Zamislite da ste u conga liniji. Imate ruke na osobi koja vas prati u liniji, a osoba iza vas ima svoje ruke na vama. Možete videti samo osobu direktno ispred sebe, ali oni blokiraju pogled na ostale ljude koji su dalje u liniji. Node je kao osoba u conga liniji: znaju ko su i mogu videti samo osobu koja ih prati, ali ne znaju ništa o drugim ljudima ispred ili iza njih.

# --instructions--
U našem kod editoru, kreirali smo dve čvora, `Kitten` i `Puppy`, i spojili smo čvor `Kitten` sa čvorom `Puppy` ručno.

Kreiraj čvor za `Cat` na `Dog` ručno na liniji.

# --hints--
Tvoj čvor (node) `Puppy` treba da ima referencu na čvor (node) `Cat`.

```js
assert(Puppy.next.element === 'Cat');
```

Tvoj čvor za `Cat` mora da ima referencu ka čvoru `Dog`.

```js
assert(Cat.next.element === 'Dog');
```

# --seed--

## --seed-contents--

```js
var Node = function(element) {
  this.element = element;
  this.next = null;
};
var Kitten = new Node('Kitten');
var Puppy = new Node('Puppy');

Kitten.next = Puppy;
// Only change code below this line
```

# --solutions--

```js
// solution required
```
