---
id: 587d8255367417b2b2512c72
title: Koristi .has i .size u ES6 Setu
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--
Istražite metode `has` i `size` dostupne na objektu Set iz ES6.

Prvo, kreiraj Set za ES6

```js
var set = new Set([1,2,3]);
```

Način za .has proveriću da li je vrednost unutar skupa.

```js
var hasTwo = set.has(2);
```

Metoda .size će vratiti kompletan broj koji predstavlja veličinu Set

```js
var howBig = set.size;
```

# --instructions--
U ovoj vežbi ćemo pregledati niz podataka i vrednosti za funkciju `checkSet()`. Tvoja funkcija treba da kreira ES6 `Set` iz niza podataka. Pronađi da li `Set` sadrži vrednost argumenta. Pronađi veličinu `Set`-a. I vrati te dve vrednosti u niz podataka.

# --hints--

`checkSet([4, 5, 6], 3)` inapaswa kurudisha [ false, 3 ]

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
