---
id: 587d8258367417b2b2512c7f
title: "Koristite široki pretragu prvo u binarnoj tranci (ili: binarnom stablu)"
challengeType: 1
forumTopicId: 301718
dashedName: use-breadth-first-search-in-a-binary-search-tree
---

# --description--
Ovde ćemo predstaviti drugi način za pretraživanje stabljika: pretraživanje po širini. Za razliku od metoda pretraživanja po dubini iz prethodnog izazova, pretraživanje po širini pronalazi sve čvorove na određenom nivou unutar stabljika pre nego što pređe na sledeći nivo. Obično se koristi red kao pomoćna struktura podataka u implementaciji algoritama pretraživanja po širini.

Na ovaj način, počinjemo dodavanjem korešnog čvora u red. Zatim pokrećemo petlju gde izbacujemo prvi element iz reda, dodajemo ga u novu listu, a zatim istražujemo sva njegova poddrveća. Ako mu deca nisu null, svaki se dodaje u red. Ovaj proces nastavlja dok red nije prazan.

# --instructions--
Implementirajmo metodu pretraživanja po širini u našem stablu pod nazivom `levelOrder`. Ova metoda treba da vrati niz podataka koji sadrži sve vrednosti čvorova stabla, pretražene po širini. Uverite se da vraćate samo vrednosti u nizu, a ne same čvorove. Nivo treba da se prolazi od leve ka desnoj. Zatim, napišite sličnu metodu pod nazivom `reverseLevelOrder` koja obavlja istu pretragu, ali u suprotnom smeru (od desnog ka levom) na svakom nivou.

# --hints--
Struktura podataka `BinarySearchTree` mora da postoji.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Binarno stablje pretrage treba da ima putanju nazvanu `levelOrder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.levelOrder == 'function';
  })()
);
```

Binarno stablo pretrage mora imati putanju nazvanu `reverseLevelOrder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.reverseLevelOrder == 'function';
  })()
);
```

Njia za `levelOrder` bi trebalo da vrati niz vrednosti pronađenih čvorova drveta po nivou redu.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.levelOrder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.levelOrder().join('') == '719038102546';
  })()
);
```

Putanja za `reverseLevelOrder` treba da vrati niz vrednosti pronađenih čvorova drveta u obrnutom nivo redu.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.reverseLevelOrder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.reverseLevelOrder().join('') == '791108305264';
  })()
);
```

Putanja za `levelOrder` treba da vrati `null` na prazno stablo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.levelOrder !== 'function') {
      return false;
    }
    return test.levelOrder() == null;
  })()
);
```

Putanja za `reverseLevelOrder` treba da vrati `null` na prazno stablo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.reverseLevelOrder !== 'function') {
      return false;
    }
    return test.reverseLevelOrder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  this.levelOrder = (root = this.root) => {
    if(!root) return null;
    let queue = [root];
    let results = [];
    while(queue.length > 0) {
      let node = queue.shift();
      results.push(node.value);
      if(node.left) queue.push(node.left);
      if(node.right) queue.push(node.right);
    }
    return results;
  }

  this.reverseLevelOrder = (root = this.root) => {
    if(!root) return null;
    let queue = [root];
    let results = [] ;
    while ( queue.length > 0) {
      let node = queue.shift();
      results.push(node.value);
      if(node.right) queue.push(node.right);
      if(node.left ) queue.push(node.left);
    }
    return results;
  }
  // Only change code above this line
}
```
