---
id: 587d8259367417b2b2512c84
title: Tengeneza mti wa utafutaji wa mti wa maneno
challengeType: 1
forumTopicId: 301634
dashedName: create-a-trie-search-tree
---

# --description--
Ovde ćemo otići sa binarnih stabla pretrage i pogledati drugi tip strukture stabla koji se naziva trij (ili stablo reči). Trie je strukturirano stablo za pretragu koje se često koristi za skladištenje sekvenci karaktera, ili generalnije, povezanih redova podataka ili dinamičnih skupova podataka gde su ključevi sekvence karaktera. Izuzetno je efikasno za skladištenje skupova podataka kada mnogi ključevi imaju slične prefikse, na primer, sve reči u rječniku. Za razliku od binarnog stabla, čvorovi nisu povezani sa stvarnom vrednošću. Umesto toga, put do čvora predstavlja specifični ključ. Na primer, ako bismo hteli da sačuvamo sekvencu karaktera `code` u triju, imali bismo četiri čvora, po jedan za svaki karakter: c — o — d — e. Praćenje tog puta kroz sve ove čvorove stvoriće `code` kao sekvencu karaktera — taj put je ključ koji smo sačuvali. Zatim, ako bismo hteli da dodamo sekvencu karaktera `coding`, podijelili bi tri prvih čvora od `code` pre nego što se razlikuju nakon d. Na ovaj način, veliki skupovi podataka mogu biti sačuvani veoma kompaktno. Štaviše, pretraga može biti izuzetno brza jer je efikasnost ograničena dužinom sekvence karaktera koja se skladišti. Pored toga, za razliku od binarnih stabala, čvor može da sadrži bilo koji broj čvorova dece. Kao što možete zaključiti iz gore navedenog primera, određeni metadata se obično skladišti na čvorima koji označavaju kraj ključa kako bi u budućim pregledima taj ključ i dalje bio dostupan. Na primer, ako bismo dodali `codes` u naš gornji primer, morali bismo imati način da znamo da `e` u `code` predstavlja kraj unesenog ključa. Inače, te informacije bi efikasno nestale kada dodajemo `codes`.

# --instructions--
Napravićemo stablo reči za čuvanje reči. Primaćeš reči putem ``add`` i skladištiće ih u strukturu podataka stabla reči. Takođe nam omogućava da pitamo da li je zadana sekvenca karaktera reč putem ``isWord``, i da dobijemo sve reči unete u stablo reči putem ``print``. ``isWord`` bi trebalo da vrati booleanski vrednost, a ``print`` bi trebalo da vrati niz svih reči kao sekvencu karaktera. Da bismo potvrdili da je ova struktura podataka implementirana ispravno, obezbedili smo strukturu ``Node`` za svaki čvor u stablu. Svaki čvor će biti objekat sa svojstvom ``keys``, koje je Mapa objekat (`Map`) tipa `JavaScript`. Ovo će čuvati pojedinačne karaktere koji su validni ključevi za svaki čvor. Takođe smo kreirali svojstvo ``end`` na čvorovima koje se može postaviti na vrednost ``true`` ako čvor predstavlja kraj reči.

# --hints--
`Trie` trebalo bi da ima puta do `add`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

`Trie` treba da ima način za `print`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

`Trie` treba da ima način za `isWord`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.isWord == 'function';
  })()
);
```

Njia ya `print` treba da vrati sve elemente dodate u drvo reči kao sekvencu karaktera u nizu podataka.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    test.add('jump');
    test.add('jumps');
    test.add('jumped');
    test.add('house');
    test.add('mouse');
    var added = test.print();
    return (
      added.indexOf('jump') != -1 &&
      added.indexOf('jumps') != -1 &&
      added.indexOf('jumped') != -1 &&
      added.indexOf('house') != -1 &&
      added.indexOf('mouse') != -1 &&
      added.length == 5
    );
  })()
);
```

Putanja za `isWord` treba da vrati `true` samo za reči unete u stablo reči i `false` za sve ostale reči.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    test.add('hop');
    test.add('hops');
    test.add('hopped');
    test.add('hoppy');
    test.add('hope');
    return (
      test.isWord('hop') &&
      !test.isWord('ho') &&
      test.isWord('hopped') &&
      !test.isWord('hopp') &&
      test.isWord('hoppy') &&
      !test.isWord('hoping')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
var Node = function() {
  this.keys = new Map();
  this.end = false;
  this.setEnd = function() {
    this.end = true;
  };
  this.isEnd = function() {
    return this.end;
  };
};
var Trie = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
