---
id: 587d825b367417b2b2512c8d
title: Unda ramani ya ES6 JavaScript
challengeType: 1
forumTopicId: 301635
dashedName: create-an-es6-javascript-map
---

# --description--
Nova verzija JavaScript daje nam ugrađeni objekat Map (Map) koji pruža veliki deo funkcionalnosti koju smo ručno napisali u prethodnom izazovu. Ovaj Map objekat, iako liči na standardne stvari od JavaScript, pruža ključnu funkcionalnost koju standardne stvari nemaju. Na primer, ES6 Map prati redosled unosa dodatih elemenata. Evo kompletnog pregleda njegovih metoda:

- `.has(key)` vraća da li je tačno ili netačno zavisno od postojanja ključa
- `.get(key)` vraća vrednost povezana sa ključem
- `.set(key, value)` postavlja novu par ključ-vrednost
- `.delete(key)` uklanja par ključ-vrednost
- `.clear()` uklanja sve parove ključ-vrednost
- `.keys()` vraća novi iterator mape sa svim ključevima u redosledu unosa
- `.values()` vraća novi iterator mape sa svim vrednostima u redosledu unosa
- `.entries()` vraća novi iterator mape sa svim parovima ključ-vrednost u redosledu unosa

# --instructions--
Definisati objekat mape za JavaScript i postavi njegovu vrednost u parametar nazvan myMap. Dodaj par ključ-vrednost, `freeCodeCamp`, `Awesome!` unutar njega.

# --hints--

Objekat `myMap` treba da postoji.

```js
assert(typeof myMap === 'object');
```

`myMap` trebao bi imati par ključ-vrednost `freeCodeCamp`, `Awesome!`.

```js
assert(myMap.get('freeCodeCamp') === 'Awesome!');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const myMap = new Map();

myMap.set("freeCodeCamp", "Awesome!");
```
