---
id: 587d8250367417b2b2512c60
title: Kreiraj klasu za redove
challengeType: 1
forumTopicId: 301631
dashedName: create-a-queue-class
---

# --description--
Kao i staci, redovi su kolekcija elemenata. Ali za razliku od staka, redovi prate princip FIFO (First In First Out ã Prvo uéÀlo, prvo izaáe). Elementi koji se dodaju u red ubacuju se na zadnji deo reda, a uklanjanju je dozvoljeno samo element koji se nalazi na prednjem kraju reda.

Moéƒemo koristiti strukturu podataka da predstavimo redove, ali kao i stike, éƒelimo smanjiti nivo kontrole koji moéƒemo imati na vrhu naéÀih redova.

Dve glavne metode klase reda su `enqueue` i `dequeue`. Metoda `enqueue` ubacuje element na zadnji kraj reda, dok metoda `dequeue` uklanja i vraáa element koji se nalazi na prednjem kraju reda. Ostale vaéƒne metode su `front`, `size` i `isEmpty`.

# --instructions--
NapiéÀite metodu za `enqueue` koja ubacuje element na zadnji kraj reda, metodu za `dequeue` koja uklanja i vraáa prednji element, metodu za `front` koja nam omoguáava da vidimo prednji element, metodu za `size` koja prikazuje veliáinu, i metodu za `isEmpty` koja proverava da li je red prazan.

# --hints--
VaéÀa klasa za `Queue` treba da ima metodu `enqueue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.enqueue === 'function';
  })()
);
```

Darasa lako la `Queue` bi trebalo da ima metodu za `dequeue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.dequeue === 'function';
  })()
);
```

VaéÀa klasa za `Queue` treba da ima metodu `front`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.front === 'function';
  })()
);
```

VaéÀa klasa za `Queue` treba da ima metodu `size`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.size === 'function';
  })()
);
```

VaéÀa klasa za `Queue` bi trebalo da ima metodu `isEmpty`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.isEmpty === 'function';
  })()
);
```

Putanja `dequeue` treba da ukloni i vrati prednji element reda

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.dequeue() === 'Smith';
  })()
);
```

Putanja `front` treba da vrati vrednost prednjeg elementa redova

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.front() === 'Smith';
  })()
);
```

Putanja `size` bi trebalo da vrati duéƒinu reda

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return test.size() === 1;
  })()
);
```

Putanja `isEmpty` mora vratiti `false` ako postoje elementi u folderu

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return !test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Queue() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function Queue () { 
    var collection = [];
    this.print = function() {
        console.log(collection);
    };
    // Only change code below this line
    this.enqueue = function(item) {
        collection.push(item);
    }

    this.dequeue = function() {
        return collection.shift();
    }

    this.front = function() {
        return collection[0];
    }

    this.size = function(){
        return collection.length;
    }

    this.isEmpty = function() {
        return collection.length === 0 ? true : false;
    }
    // Only change code above this line
}
```
