---
id: 587d825a367417b2b2512c8a
title: Unesite element u najviši sloj
challengeType: 1
forumTopicId: 301703
dashedName: insert-an-element-into-a-max-heap
---

# --description--
Sada ćemo nastaviti sa drugom strukturom podataka baziranom na drvetu, binarnim kupom. Binarni kup je binarno drvo sa rasporedom elemenata koji zadovoljava svojstvo kupa. Svojstvo kupa opisuje odnos između roditeljskog čvora i njegovih detiljih čvorova. Možete imati maksimalni kup, gde su svi roditeljski čvorovi veći ili jednaki svojim detiljim čvorovima, ili minimalni kup, gde je obrnuto tačno. Binarni kupovi su takođe i potpuno binarna drva. To znači da su svi nivoi drveta potpuno popunjeni, a ako je poslednji nivo delimično popunjen, on se popunjava od leve ka desnoj strani.

Iako se binarni stablići mogu implementirati kao strukture stabla sa čvorovima koji drže reference levo i desno, sortiranje elemenata prema svojstvu stablaca omogućava nam predstavljanje stabla korišćenjem niza podataka. Odnos roditelj-dete je ono što nam je cilj, i jednostavnim izračunavanjem možemo izračunati decu bilo kog roditelja i roditelja bilo kog čvorice deteta.

Na primer, zamislite ovu reprezentaciju niza podataka najnižeg binarnog paketa:

```js
[ 6, 22, 30, 37, 63, 48, 42, 76 ]
```

Korenski čvor je prvi element, `6`. Njegovi deca su `22` i `30`. Ako pogledamo odnos između indeksa niza podataka ovih vrednosti, za primer `i`, deca su `2 * i + 1` i `2 * i + 2`. Slično tome, element na primeru `0` je roditelj ovih dvaju dečaka koji se nalaze na indeksima `1` i `2`. Generalno, možemo dobiti roditelja bilo kog čvora za bilo koji primer koristeći sledeće: `Math.floor((i - 1) / 2)`. Ovi sistemi će biti tačni bez obzira na veličinu binarnog staba. Na kraju, možemo napraviti malu korekciju kako bismo ovaj proračun olakšali preskočenjem prvog elementa u nizu podataka. To stvara sledeći odnos za bilo koji element na primeru `i`:

Primer predstavljanja niza podataka:

```js
[ null, 6, 22, 30, 37, 63, 48, 42, 76 ]
```

Levo dete elementa: `i * 2`

Desni potomak elementa: `i * 2 + 1`

Roditeljski element: `Math.floor(i / 2)`

Čim se razumevanje ove formule postigne, korišćenje reprezentacije niza podataka je veoma važno jer lokacije čvorova se mogu brzo odrediti korišćenjem ovog računanja i smanjuje se potrošnja memorije jer nije potrebno skladištiti reference na čvorove potomaka.

# --instructions--
Uputstva: Ovde ćemo kreirati maksimalni heap. Počnite kreiranjem metode ``insert`` koja dodaje elemente u naš heap. Prilikom ubacivanja, važno je uvek održavati svojstvo heapa. Za maksimalni heap, to znači da element korena mora imati najveću vrednost u drvetu i da sve roditeljske čvorove moraju biti veći od svojih potomaka. Za implementaciju heapa korišćenjem niza podataka, ovo se obično dešava u tri koraka:

<ol>
  <li>Dodajte novi element na kraj niza podataka.</li>
  <li>Ako je element veći od svog roditelja, zamijenite njihove pozicije.</li>
  <li>Nastavite sa zamjenama dok novi element nije manji od svog roditelja ili dosegnete korijen stabla.</li>
</ol>

Na kraju, dodaj metod za `print` koja vraća niz svih stavki dodanih u stak.

# --hints--
Struktura podataka za `MaxHeap` treba da postoji.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap` treba imati metod nazvan `insert`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap` treba da ima putanju nazvanu `print`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

Putanja `insert` treba da doda funkcionalnosti prema svojstvu najvišeg steka.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    test.insert(50);
    test.insert(100);
    test.insert(700);
    test.insert(32);
    test.insert(51);
    test.insert(800);
    const result = test.print();
    const solution = JSON.stringify([null,800,51,700,32,50,100]);
    const solutionWithoutNull = JSON.stringify([800,51,700,32,50,100]);

    return (result.length == 6) ? (JSON.stringify(result) == solutionWithoutNull) : (JSON.stringify(result) == solution);
  })()
);
```

# --seed--

## --seed-contents--

```js
var MaxHeap = function() {
  // Only change code below this line
    
  // Only change code above this line
};
```

# --solutions--

```js
var MaxHeap = function() {
    // Only change code below this line
    this.heap = [];
    this.parent = index => {
      return Math.floor((index - 1) / 2);
    }
    this.insert = element => {
      this.heap.push(element);
      this.heapifyUp(this.heap.length - 1);
    }
    this.heapifyUp = index => {
      let currentIndex = index,
      parentIndex = this.parent(currentIndex);
      while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
        this.swap(currentIndex, parentIndex);
        currentIndex = parentIndex;
        parentIndex = this.parent(parentIndex);
      }
    }
    this.swap = (index1, index2) => {
      [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
    }
    this.print = () => {
      return this.heap;
    }
    // Only change code above this line
};
```
