---
id: 587d8255367417b2b2512c74
title: Kreiraj klasu za red prioriteta
challengeType: 1
forumTopicId: 301630
dashedName: create-a-priority-queue-class
---

# --description--
U ovom izazovu ćete kreirati prioritetnu kole (priority queue). Prioritetna cola je poseban tip kole u kojoj elementi mogu imati dodatne informacije koje određuju njihov prioritet. Ovo se lako može predstaviti celim brojem. Prioritet elementa nadjačava redosled unosa pri određivanju sekvence elemenata koji se izvlače iz kole. Ako se element sa visokim prioritetom unese nakon elemenata sa niskim prioritetom, element sa visokim prioritetom će biti uklonjen prvi, pre svih ostalih.

Na primer, zamislite da imamo prioritetnu cola sa tri elementa:

```js
[['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Ovde druga vrednost (kompletni broj) predstavlja prioritet elementa. Ako postavimo `['human', 1]` i prioritet od `1` (uzimajući u obzir da niži prioritet dobija viši prioritet), onda će biti prvi element uklonjen iz redova. Kolekcija bi izgledala ovako:

```js
[['human', 1], ['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Počeli smo da pišemo `PriorityQueue` u kodni editor. Biće vam potrebno dodati metodu `enqueue` za dodavanje stavki sa prioritetom, metodu `dequeue` za uklanjanje i vraćanje stavki, metodu `size` za povratak broja stavki u red, metodu `front` za vraćanje elementa na početku reda, i napokon metodu `isEmpty` koja će vratiti `true` ako je red prazan ili `false` ako nije prazan.

`enqueue` treba da prima stavke sa strukturom prikazanom iznad (`['human', 1]`) gde `1` predstavlja prioritet. `dequeue` i `front` bi trebalo da vrate samo ime stavke, a ne njen prioritet.

# --hints--
Darasa lako za `PriorityQueue` mora imati metod od/za `enqueue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.enqueue === 'function';
  })()
);
```

Vaša klasa za `PriorityQueue` mora imati metodu za `dequeue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.dequeue === 'function';
  })()
);
```

Klasa za `PriorityQueue` mora imati metodu `size`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.size === 'function';
  })()
);
```

Vaša klasa za `PriorityQueue` mora imati metodu `front`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.front === 'function';
  })()
);
```

Vaša klasa za `PriorityQueue` mora imati metodu `isEmpty`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.isEmpty === 'function';
  })()
);
```

Vaša klasa za `PriorityQueue` treba da tačno prati trenutni broj predmeta koristeći metodu `size` kada se predmeti dodaju i uklanjaju iz redova.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    test.enqueue(['Jon Snow', 1]);
    var size1 = test.size();
    test.dequeue();
    var size2 = test.size();
    test.enqueue(['A', 3]);
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    return size1 === 2 && size2 === 1 && test.size() === 4;
  })()
);
```

Putanja za `front` bi trebalo da vrati ispravan element koji se nalazi pre rednog niza kada se elementi dodaju i uklanjaju.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    var front1 = test.front();
    test.enqueue(['Jon Snow', 1]);
    var front2 = test.front();
    test.dequeue();
    test.enqueue(['A', 3]);
    var front3 = test.front();
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    test.dequeue();
    var front4 = test.front();
    return (
      front1 === 'David Brown' &&
      front2 === 'Jon Snow' &&
      front3 === 'David Brown' &&
      front4 === 'A'
    );
  })()
);
```

Putanja za `isEmpty` bi trebalo da vrati `true` kada je red prazan.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 1]);
    test.enqueue(['B', 1]);
    test.dequeue();
    var first = test.isEmpty();
    test.dequeue();
    return !first && test.isEmpty();
  })()
);
```

Prioritetni red treba da vraća elemente sa visokim prioritetom pre elemenata sa niskim prioritetom, a takođe i da ih vraća u FIFO (First-In, First-Out) redosledu.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 5]);
    test.enqueue(['B', 5]);
    test.enqueue(['C', 5]);
    test.enqueue(['D', 3]);
    test.enqueue(['E', 1]);
    test.enqueue(['F', 7]);
    var result = [];
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    return result.join('') === 'EDABCF';
  })()
);
```

# --seed--

## --seed-contents--

```js
function PriorityQueue () {
  this.collection = [];
  this.printCollection = function() {
    console.log(this.collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function PriorityQueue() {
  this.collection = [];
  this.printCollection = function () {
    console.log(this.collection);
  };
  // Only change code below this line
  this.enqueue = function (newitem) {
    if (this.isEmpty()) {
      return this.collection.push(newitem);
    }

    this.collection = this.collection.reverse();
    var found_index = this.collection.findIndex(function (item) {
      return newitem[1] >= item[1];
    });
    if (found_index === -1) {
      this.collection.push(newitem);
    } else {
      this.collection.splice(found_index, 0, newitem);
    }
    this.collection = this.collection.reverse();
  };
  this.dequeue = function () {
    if (!this.isEmpty()) {
      return this.collection.shift()[0];
    } else {
      return "The queue is empty.";
    }
  };
  this.size = function () {
    return this.collection.length;
  };
  this.front = function () {
    return this.collection[0][0];
  };
  this.isEmpty = function () {
    return this.size() > 0 ? false : true;
  };
  // Only change code above this line
}
```
