---
id: 587d8255367417b2b2512c75
title: "Napravi krug (ili: napravi ciklus)"
challengeType: 1
forumTopicId: 301625
dashedName: create-a-circular-queue
---

# --description--
U ovom izazovu ćete kreirati kružni buffer (circular buffer). Kružni buffer je niz koji dodaje podatke na kraj kontejnera, a zatim počinje da prepisuje podatke sa samog početka tog kontejnera. Ova vrsta strukture podataka važna je u određenim situacijama. Na primer, kružni buffer može se koristiti za reprodukciju strimovanih medija. Kada se bafer popuni, novi podaci o medijima će prepisati stare podatke.

Dobar način da se ovaj koncept objasni je korišćenjem reda podataka sa dužinom `5`:

```js
[null, null, null, null, null]
 ^Read @ 0
 ^Write @ 0
```

Ovde su sve zaglavlja za čitanje i pisanje u lokaciji `0`. Sada red prima 3 novih zapisa: `a`, `b` i `c`. Naš red sada izgleda ovako:

```js
[a, b, c, null, null]
 ^Read @ 0
          ^Write @ 3
```

Kada uređaj za čitanje pročita, može da ukloni ili sapremi vrednosti:

```js
[null, null, null, null, null]
                   ^Read @ 3
                   ^Write @ 3
```

Sada pišemo vrednosti `d`, `e`, i `f` u redu. Kada pisanje dođe do kraja reda, ponavlja se na početku:

```js
[f, null, null, d, e]
                ^Read @ 3
    ^Write @ 1
```

Ovaj način zahteva fiksnu količinu memorije, ali omogućava obradu većih fajlova.

# --instructions--
U ovom zadatku ćemo implementirati kružni niz. Kružni niz treba da pruži metode za `enqueue` i `dequeue` koje će omogućiti čitanje i pisanje u niz. Sama klasa takođe treba da primi argument celog broja koji možete koristiti za određivanje veličine niza prilikom kreiranja. Napisali smo početnu verziju ove klase za vas u kodnom editoru.

Kada dodajete stavke u red, kursor bi trebalo da se pomera napred i vrati na početak kada dođe do kraja reda. Funkcija `enqueue` bi trebalo da vrati element koji ste dodali ako je uspešno; inače bi trebalo da vrati `null`.

Isto tako, savet za čitanje mora da se pomakne napred kada se elementi uklone iz liste. Kada uklonite element, taj element bi trebalo da se vrati. Ako ne možete ukloniti element, morate vratiti `null`.

Kursor za pisanje ne bi trebalo da se dozvoli pomerati dalje od kursora za čitanje (naša klasa neće dozvoliti pisanje na podatke koje još niste pročitali) i kursor za čitanje ne bi trebalo da se pomeri dalje od podataka koje ste napisali.

# --hints--
Putanja `enqueue` bi trebalo da doda stavke u ciklički niz.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Ne dodajte ništa više od saveta za čitanje.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    test.enqueue(13);
    test.enqueue(25);
    test.enqueue(59);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Putanja `dequeue` bi trebalo da vrati element iz niza.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 && test.dequeue() === 32 && test.dequeue() === 591
    );
  })()
);
```

Nakon uklanjanja elementa, njegova pozicija u redu bi trebalo da se ponovo postavi na `null`.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(672);
    test.dequeue();
    test.dequeue();
    var print = test.print();
    return print[0] === null && print[1] === null && print[2] === 672;
  })()
);
```

Pokušaj da se izvuče više od saveta za pisanje treba vratiti `null` i nije pomerio savet za pisanje.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 &&
      test.dequeue() === 32 &&
      test.dequeue() === 591 &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.enqueue(100) === 100 &&
      test.dequeue() === 100
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
class CircularQueue {
  constructor(size) {

    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line

    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line

    // Only change code above this line
  }
}
```

# --solutions--

```js
class CircularQueue {
  constructor(size) {
    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line
    console.log(this.write, this.max);
    if (this.queue[this.write] === null) {
      this.queue[this.write++] = item;

      if (this.write > this.max) {
        this.write = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line
    if (this.queue[this.read] !== null) {
      let item = this.queue[this.read];
      this.queue[this.read++] = null;
      if (this.read > this.max) {
        this.read = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }
}
```
