---
id: 587d8250367417b2b2512c5f
title: Kreiraj klasu za stack
challengeType: 1
forumTopicId: 301633
dashedName: create-a-stack-class
---

# --description--
U poslednjem delu, razgovarali smo o tome šta je stack i kako možemo da koristimo niz podataka za predstavljanje steka. U ovom delu, kreiraćemo našu klasu za stack. Iako možete koristiti nizove podataka za kreiranje steka, ponekad je bolje smanjiti nivo kontrole koji imamo nad našim stekovima. Pored metoda `push` i `pop`, stakovi imaju druge važne metode. Dodajmo metode `peek`, `isEmpty` i `clear` u našu klasu za stack.

# --instructions--
Izdvojiti metodu `push` koja dodaje element na stak, metodu `pop` koja izvlači i vraća element sa vrha staka, metodu `peek` koja proverava element na vrhu staka, metodu `isEmpty` koja proverava da li je stak prazan, i metodu `clear` koja uklanja sve elemente sa staka. Stakovi obično nemaju ovo, ali dodali smo pomoćnu metodu `print` koja piše u konsolu sklopnje.

# --hints--
Tvoja klasa od `Stack` treba da ima metodu za `push`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.push === 'function';
  })()
);
```

Tvoja klasa za `Stack` treba da ima metodu `pop`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.pop === 'function';
  })()
);
```

Vaša klasa za `Stack` bi trebalo da ima metodu za `peek`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.peek === 'function';
  })()
);
```

Tvoja klasa od `Stack` mora da ima metodu za `isEmpty`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.isEmpty === 'function';
  })()
);
```

Vaša klasa `Stack` treba da ima metodu za/od `clear`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.clear === 'function';
  })()
);
```

Putanja za `peek` bi trebalo da vrati najgornji element steka

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.peek() === 'CS50' && test.peek() === 'CS50';
  })()
);
```

Njia za `pop` treba da ukloni i vrati najgornji element steka

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.pop() === 'CS50' && test.pop() === 'CS61';
  })()
);
```

Putanja za `isEmpty` bi trebalo da vrati istinu ako stos nema elemenata

```js
assert(
  (function () {
    var test = new Stack();
    return test.isEmpty();
  })()
);
```

Metoda `clear` treba da ukloni sve elemente sa steka

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    test.clear();
    return test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Stack() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
class Stack {
  constructor() {
    this.collection = [];
  }
  print() {
    console.log(this.collection);
  }
  push(val) {
    this.collection.push(val);
  }
  pop() {
    return this.collection.pop();
  }
  peek() {
    return this.collection[this.collection.length - 1];
  }
  isEmpty() {
    return this.collection.length === 0;
  }
  clear() {
    return (this.collection.length = 0);
  }
}
```
