---
id: 587d8251367417b2b2512c64
title: Tražite u povevenoj listi
challengeType: 1
forumTopicId: 301715
dashedName: search-within-a-linked-list
---

# --description--
Dodajmo više korisnih metoda u našu klasu povezanog popisa. Zar ne bi bilo korisno da možemo reći da li je naš popis prazan ili ne, kao u našim klasama za `Stack` i `Queue`?

Takođe bi trebalo da imamo mogućnost dohvaćanja specifičnih elemenata iz naše povezane liste. Kroz strukturu podataka je nešto sa čime ćete želeti mnogo da vežbate! Napravimo funkciju `indexOf` koja uzima `element` kao argument, i vraća `index` tog elementa iz povezane liste. Ako element nije pronađen u povezanoj listi, vrati `-1`.

Takođe napravimo funkciju koja radi suprotno: funkcija za `elementAt` koja uzima `index` kao argument i vraća `element` u datoem `index`. Ako nije dostupan `element`, vrati `undefined`.

# --instructions--
Napiši metodu za `isEmpty` koja proverava da li je povezana lista prazna, metodu `indexOf` koja vraća `index` nekog elementa, i metodu `elementAt` koja vraća `element` u dostavljenom `index`.

# --hints--
Darasa lako za `LinkedList` treba da ima metodu od `isEmpty`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.isEmpty === 'function';
  })()
);
```

Tvoja putanja za `isEmpty` bi trebalo da vrati `false` kada postoji barem jedan element u povezanoj listi.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.isEmpty() === false;
  })()
);
```

Tvoj put za `isEmpty` bi trebalo da vrati `true` kada nema stavki u povezanoj listi.

```js
assert(
  (function () {
    var test = new LinkedList();
    return test.isEmpty() === true;
  })()
);
```

Vaša klasa za `LinkedList` treba da ima metodu od `indexOf`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.indexOf === 'function';
  })()
);
```

Tvoja putanja ``indexOf`` bi trebalo da vrati model nekog atributa pronađenog u povezanoj listi.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('cat') === 0;
  })()
);
```

Tvoja putanja ``indexOf`` bi trebalo da vrati ``-1`` ako element nije pronađen u povezanoj listi.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('pony') === -1;
  })()
);
```

Vaša klasa za `LinkedList` mora imati metodu za `elementAt`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.elementAt === 'function';
  })()
);
```

Vaša putanja za `elementAt` treba da vrati element pronađen u određenom modelu/šemi u povezanoj listi.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(1) === 'dog';
  })()
);
```

Tvoja putanja za `elementAt` treba da vrati `undefined` ako dostavljeni element nije pronađen u određenom modelu u povezanoj listi.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(5) === undefined;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
      head = currentNode.next;
    } else {
      while(currentNode.element !== element) {
        previousNode = currentNode;
        currentNode = currentNode.next;
      }

      previousNode.next = currentNode.next;
    }

    length --;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
        head = currentNode.next;
    } else {
        while(currentNode.element !== element) {
            previousNode = currentNode;
            currentNode = currentNode.next;
        }

        previousNode.next = currentNode.next;
    }

    length --;
  };

  this.indexOf = function(element) {
    if (head === null) return -1

    let current = head;
    let index = 0;

    while (current.element !== element && current.next !== null) {
      current = current.next;
      index++
    }

    if (current.element !== element && current.next === null) {
      return -1
    }

    return index;
  }

  this.elementAt = function(index) {
    if (head === null) return undefined;

    let current = head;
    let currentIndex = 0;

    while (currentIndex !== index && current.next !== null) {
      current = current.next;
      currentIndex++
    }

    if (currentIndex !== index && current.next === null) {
      return undefined;
    }

    return current.element;
  }

  this.isEmpty = function() {
    return length === 0;
  }
}
```
