---
id: 587d825d367417b2b2512c96
title: Prvo duboka pretraga
challengeType: 1
forumTopicId: 301640
dashedName: depth-first-search
---

# --description--
Kao što je sa <dfn>široka pretraga prvo</dfn>, ovde ćemo naučiti o drugom algoritmu za pretraživanje grafa koji se zove <dfn>duboka pretraga prvo</dfn>.

Kada širi prvojne pretraživanje traži po koracima dužine ivica od izvornog čvora, <dfn>dubinsko prvo pretraživanje</dfn> počinje spuštanjem po putu ivica dok ne stigne.

Kada se vrati do kraja putanje, pretraga se vraća nazad do poslednjeg čvora sa ivicom koja nije posećena i nastavlja pretraživanje.

Shema dole prikazuje kako algoritam funkcioniše. Algoritam počinje sa gornjim čvorom i poseti čvorove numeričkim redosledom.

<img alt="animation showing the depth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/depth-first-search.gif'>

Posmatraj kako, za razliku od pretrage u širinu (BFS), svaki put kada se čvor posjeti, on ne prolazi kroz sve svoje susjedne čvorove. Umjesto toga, počinje posjećivanjem jednog susjednog čvora i nastavlja duž tog puta sve dok ne ostane nijedan drugi čvor koji treba posjetiti na tom putu.

Da biste implementirali ovaj algoritam, moraćete da koristite stakove. Stak je struktura podataka gde se poslednji dodati element prvi uklanja. Ovo je takođe poznato kao struktura podataka <dfn>Last-In-First-Out</dfn>. Stakovi su važni u algoritmima pretraživanja dubine jer kada dodajemo susjednike na stak, želimo da prvo posjetimo najnovije dodate susjednike i uklonimo ih sa staka.

Jednostavan rezultat ovog algoritma je lista čvorova koji se mogu dostignuti od određenog čvora. Stoga, takođe ćete morati da pratite čvorove koje posjećujete.

# --instructions--
Definisati funkciju `dfs()` koja prima nedirekcionalnu matricu susjednosti `graph` i etiketu čvora `root` kao parametre. Etiketa čvora će biti samo vrijednost broja čvora između `0` i `n - 1`, gdje je `n` ukupan broj čvorova u grafu.

Tvoja funkcija treba da vrati niz podataka svih dostupnih čvorova od `root`.

# --hints--
Ulazni graf `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `1` bi trebalo da vrati red podataka koji sadrži `0`, `1`, `2` i `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })(),
  [0, 1, 2, 3]
);
```

Grafik ulaza `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `3` treba da vrati niz podataka sa `3`, `2`, `1`, na `0`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3, 2, 1, 0]
);
```

Ulazna grafika `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `1` treba da vrati niz podataka sa četiri elementa.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })().length === 4
);
```

Grafik ulaza `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` sa početnim čvorom `3` bi trebalo da vrati niz podataka koji sadrži `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3]
);
```

Graf ulaza `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` sa početnim čvorom `3` treba da vrati red podataka sa jednim elementom.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })().length === 1
);
```

Ulazni graf `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `3` treba da vrati red podataka `2` i `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [2, 3]
);
```

Ulazni graf `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `3` treba da vrati red podataka sa dva elementa.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })().length === 2
);
```

Ulazni grafikon `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `0` mora vratiti red podataka koji sadrži `0` i `1`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })(),
  [0, 1]
);
```

Ulazni graf `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `0` treba da vrati red podataka sa dve komponente.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })().length === 2
);
```

# --seed--

## --seed-contents--

```js
function dfs(graph, root) {

}

var exDFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(dfs(exDFSGraph, 3));
```

# --solutions--

```js
function dfs(graph, root) {
    var stack = [];
    var tempV;
    var visited = [];
    var tempVNeighbors = [];
    stack.push(root);
    while (stack.length > 0) {
        tempV = stack.pop();
        if (visited.indexOf(tempV) == -1) {
            visited.push(tempV);
            tempVNeighbors = graph[tempV];
            for (var i = 0; i < tempVNeighbors.length; i++) {
                if (tempVNeighbors[i] == 1) {
                    stack.push(i);
                }
            }
        }
    }
    return visited;
}
```
