---
id: 587d8257367417b2b2512c7c
title: Proveri da li čvor postoji u binarnom stablu pretrage
challengeType: 1
forumTopicId: 301623
dashedName: check-if-an-element-is-present-in-a-binary-search-tree
---

# --description--
Kada imamo opšte razumevanje binarnog stabla pretrage, hajde da pričamo o tome detaljnije. Binarna stablo pretrage pruža logaritamsko vreme za uobičajene operacije pretraživanja, dodavanja i brisanja u prosečnom slučaju, a linearno vreme u najgorem slučaju. Zašto je to tako? Svaka od ovih osnovnih operacija zahteva da pronađemo element u stablu (ili u slučaju dodavanja, da pronađemo gde bi trebalo da bude postavljen) i zbog strukture stabla na svakom roditeljskom čvoru, delimo levu ili desnu stranu, efikasno uklanjajući polovinu veličine preostalog stabla. To čini pretraživanje ekvivalentnim logaritmu broja čvorova u stablu, što stvara logaritamsko vreme za ove operacije u prosečnom slučaju. Dobro, ali šta je sa najgorem slučaju? Pa, zamislite da gradimo stablo iz sledećih vrednosti, dodajući ih od leve ka desnoj: `10`, `12`, `17`, `25`. Prateći naše pravilo binarnog stabla pretrage, dodali bismo `12` na desnu stranu `10`, `17` na desnu stranu ovog, i `25` na desnu stranu ovog. Sada naše stablo izgleda kao spajana lista, a prolazak kroz njega za pretraživanje `25` zahteva prolazak kroz sve elemente u linearnom redosledu. Stoga, linearno vreme u najgorem slučaju. Problem ovde je što stablo nije uravnoteženo. Detaljnije ćemo pogledati značenje toga u sledećim izazovima.

# --instructions--
U ovom izazovu, kreiraćemo alat za naše drvo. Napišite funkciju ``isPresent`` koja uzima ceo broj kao ulaz i vraća boolean vrednost (tačno ili netačno) za postojanje ili nepostojanje te vrednosti u binarnom stablu pretrage.

# --hints--
Struktura podataka `BinarySearchTree` mora da postoji.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Binarni stab za pretragu mora imati putanju nazvanu `isPresent`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isPresent == 'function';
  })()
);
```

Putanja `isPresent` mora tačno da proveri prisustvo ili neprisutnost dodatih elemenata na drvetu.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    test.add(4);
    test.add(7);
    test.add(411);
    test.add(452);
    return (
      test.isPresent(452) &&
      test.isPresent(411) &&
      test.isPresent(7) &&
      !test.isPresent(100)
    );
  })()
);
```

`isPresent` Trebalo bi da obradi slučaj gde je stablo prazno.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    return test.isPresent(5) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.isPresent = function (value) {
    var current = this.root
    while (current) {
      if (value === current.value) {
        return true;
      }
      current = value < current.value ? current.left : current.right;
    }
    return false;
  }
}
```
