---
id: 587d8257367417b2b2512c7b
title: Dodaj novu karakteristiku u binarno stablo pretrage
challengeType: 1
forumTopicId: 301618
dashedName: add-a-new-element-to-a-binary-search-tree
---

# --description--
Ova serija izazova će uvesti strukturu podataka drvo. Drveće je važna i široko korišćena struktura u računarskoj nauci. Naravno, njegovo ime potiče od činjenice da kada pogledaš njegovu sliku, izgleda kao drveće koje poznajemo u stvarnom svetu. Struktura podataka drvo počinje sa jednim čvorom, obično nazvanim koren, i odatle se širi na druge čvorove, svaki sa više potomaka čvorova, i tako dalje. Ova struktura podataka je obično prikazana sa korenom na vrhu; možeš je zamisliti kao prirodno drvo okrenuto naopako.

Prvo, objasnićemo neke uobičajene termine koje ćemo naići kada je reč o drveću. Koren (root) je na vrhu drveća. Tačke podataka unutar drveća nazivaju se čvorovi (nodes). Čvorovi koji imaju grane koje vode ka drugom čvoru nazivaju se roditelji tog grana (grana koja vodi do čvora, odnosno čvor deteta). Ostali terminološki pojmovi koriste se kao što očekujete. Poddrvo odnosi se na sve potomke određenog čvora; grane mogu se zvati ivice (edges), a listni čvorovi su konačni čvorovi drveća koji nemaju potomke. Na kraju, zapamtite da su drveći strukture podataka sa rekurentnom (rekurzivnom) strukturom. To znači da su deca čvora roditelji njihovog poddrveća, i tako dalje. Razumevanje rekurzivne prirode drveća je ključno kada dizajnirate algoritme za standardne operacije na drveću.

Početi ćemo sa diskusijom o specifičnom tipu drveća, binarnom drvetu. Zapravo, govorićemo o binarnom stablu za pretragu. Hajde da objasnimo šta to znači. Iako struktura podataka drvo može imati bilo koje grane za jedan čvor, binarno stablo može imati samo dve grane po čvoru. Štaviše, binarno stablo za pretragu je organizovano prema svojim podstablima; tako da je vrednost svakog čvora u levom podstablu manja ili jednaka vrednosti roditeljskog čvora, a vrednost svakog čvora u desnom podstablu je veća ili jednaka vrednosti roditeljskog čvora. Veoma je korisno videti ovu vezu kako bi se bolje razumelo:

<div style='width: 100%; display: flex; justify-content: center; align-items: center;'><img alt="an example of a binary search tree" style='width: 100%; max-width: 350px; background-color: var(--gray-05);' src='https://user-images.githubusercontent.com/18563015/32136009-1e665d98-bbd6-11e7-9133-63184f9f9182.png'></div>

Sada je ovaj organizovani odnos veoma lako vidljiv. Zapamtite da je svaka vrednost sa leve strane čvora korena 8 manja od 8, a svaka vrednost sa desne strane veća od 8. Takođe pogledajte da se ovaj odnos primenjuje i na svako poddrvo. Na primer, prvi levi detilji je poddrvo. 3 je roditeljski čvor i ima samo dva detilja: po pravilima binarnog stabla pretrage, znamo bez čak i proveravanja da će levi detilji ovog čvora (i sva njegova detilja) biti manji od 3, a desni detilji (i sva njihova detilja) bit će veći od 3 (ali takođe manji od vrednosti korena strukture), i tako dalje.

Binarni stabljici za pretragu su strukture podataka koje su veoma dostupne i korisne jer pružaju logaritamsko vreme u proseku za nekoliko standardnih operacija kao što su pretraživanje, ubacivanje i brisanje.

# --instructions--
Počećemo lako. Ovde smo objasnili osnovnu strukturu binarnog stabla pretrage, zajedno sa funkcijom za kreiranje čvorova našeg stabla. Napomena: Svaki čvor može imati levu i desnu vrednost. Oni će dobiti podstabla dece, ako postoje. U našem binarnom stablu pretrage, kreiraćete metodu za dodavanje nove vrednosti u stablo. Ova metoda treba da se zove `add` i mora primiti celobrojnu vrednost za dodavanje u stablo. Pazite da održavate pravilo binarnog stabla pretrage: vrednost u svakom levom detetu mora biti manja ili jednaka vrednosti roditelja, a vrednost u svakom desnom detetu mora biti veća ili jednaka vrednosti roditelja. Ovde, učini da naše stablo ne može da sadrži duplikovane vrednosti. Ako pokušamo dodati vrednost koja već postoji, metoda treba da vrati `null`. Inače, ako je dodavanje uspešno, treba da se vrati `undefined`.

**Napomena:** Drvo je struktura podataka sa prirodnom rekurzivnom strukturom!

# --hints--
Struktura podataka za `BinarySearchTree` mora da postoji.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Stablo binarnog pretraživanja treba da ima metodu nazvanu `add`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

Metoda ubacivanja treba da unosi elemente prema pravilima binarnog stablastog traženja.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    const expectedResult = [1, 4, 7, 8, 34, 45, 73, 87];
    const result = test.inOrder();
    return expectedResult.toString() === result.toString();
  })()
);
```

Dodavanje funkcije koja već postoji treba da vrati `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    return test.add(4) == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    inOrder() {
      if (!this.root) {
        return null;
      }
      var result = new Array();
      function traverseInOrder(node) {
        node.left && traverseInOrder(node.left);
        result.push(node.value);
        node.right && traverseInOrder(node.right);
      }
      traverseInOrder(this.root);
      return result;
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.add = function(element) {
    let current = this.root;
    if (!current) {
      this.root = new Node(element);
      return;
    } else {
      const searchTree = function(current) {
        if (current.value > element) {
          if (current.left) {
            return searchTree(current.left);
          } else {
            current.left = new Node(element);
            return;
          }
        } else if (current.value < element) {
          if (current.right) {
            return searchTree(current.right);
          } else {
            current.right = new Node(element);
            return;
          }
        } else {
          return null;
        }
      };
      return searchTree(current);
    }
  };
}
```
