---
id: 587d8254367417b2b2512c70
title: Kreiranje i dodavanje u Sets u ES6
challengeType: 1
forumTopicId: 301636
dashedName: create-and-add-to-sets-in-es6
---

# --description--
Sada nakon rada sa ES5-om, uradićeš nešto slično u ES6-u. Ovo će biti mnogo lakše. ES6 ima ugrađenu strukturu podataka `Set` pa su mnoge operacije koje si ručno kreirao sada spremne za tebe. Istražimo:

Kreiranje novog praznog skupa:

```js
var set = new Set();
```

Možete kreirati skup sa vrednostima:

```js
var set = new Set(1);
```

Možete kreirati skup sa redovima podataka:

```js
var set = new Set([1, 2, 3]);
```

Nakon kreiranja skupa, možete dodati vrednost koju želite koristeći `add`:

```js
var set = new Set([1, 2, 3]);
set.add([4, 5, 6]);
```

Napomena, skup je struktura podataka koja nema duplikovane vrednosti:

```js
var set = new Set([1, 2, 3, 1, 2, 3]);
// set contains [1, 2, 3] only
```

# --instructions--
Za ovu vežbu, vrati set sa sledećim vrednostima: `1, 2, 3, 'Taco', 'Cat', 'Awesome'`

# --hints--
`Set` vaše treba da ima samo vrednost `1, 2, 3, Taco, Cat, Awesome`.

```js
assert(
  (function () {
    var test = checkSet();
    return (
      test.size == 6 &&
      test.has(1) &&
      test.has(2) &&
      test.has(3) &&
      test.has('Taco') &&
      test.has('Cat') &&
      test.has('Awesome')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet() {
  var set = new Set([1, 2, 3, 3, 2, 1, 2, 3, 1]);
  // Only change code below this line

  // Only change code above this line
  console.log(Array.from(set));
  return set;
}

checkSet();
```

# --solutions--

```js
function checkSet(){var set = new Set([1,2,3,'Taco','Cat','Awesome']);
return set;}
```
