---
id: 587d8257367417b2b2512c7d
title: Tafuta kiwango cha chini na cha juu cha mti wa utafutaji wa kibainari
challengeType: 1
forumTopicId: 301641
dashedName: find-the-minimum-and-maximum-height-of-a-binary-search-tree
---

# --description--
U poslednjem izazovu napisali smo situaciju u kojoj drvo može biti neizbalansirano. Da bismo razumeli koncept balansa, istražujemo drugu osobinu drveća: visinu. Visina u drvetu predstavlja rastojanje od korenovog čvora do bilo kog listnog čvora. Različiti putevi u strukturi graneastog drveća mogu imati različite visine, ali za određeno drvo će postojati minimalna i maksimalna visina. Ako je drvo balansirano, ove vrednosti će se razlikovati najviše za jedan nivo. To znači da u balansiranom drvetu su svi listni čvorovi na istom nivou, ili ako nisu na istom nivou, onda su barem jedan nivo udaljeni.

Svojstvo ravnoteže je važno za drveće jer određuje efikasnost operacija na drveću. Kao što smo objasnili u poslednjem izazovu, susrećemo se sa problemom najgoreg slučaja kod potpuno neravnoteženog drveća. Samoravnotežna drveća često se koriste za rešavanje ovog problema u drvećima sa promenljivim skupovima podataka. Uobičajeni primeri su AVL drveće, crveno-crno drvo i B drvo. Sva ova drveća imaju dodatnu internu logiku koja ispravlja ravnotežu drveća kada ubacivanje ili brisanje uzrokuje neravnoteženost.

**Napomena:** Atribut sličan dužini je dubina, koja se odnosi na udaljenost određenog čvora od korenovog čvora.

# --instructions--
Napišite dve metode za naš binarni stablo: `findMinHeight` i `findMaxHeight`. Ove metode treba da vraćaju celobrojnu vrednost minimalne i maksimalne visine unutar datih binarnog stabla, respektivno. Ako je čvor prazan (null), molim vas postavite vrednost visine `-1` (to je osnovni slučaj). Na kraju, dodajte treću metodu `isBalanced` koja vraća `true` ili `false` u zavisnosti od toga da li je stablo balansirano. Možete koristiti dve metode koje ste ranije napisali za određivanje ovoga.

# --hints--
Struktura podataka `BinarySearchTree` treba da postoji.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Binarno stablo pretrage treba da ima put nazvan `findMinHeight`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMinHeight == 'function';
  })()
);
```

Binarni stab za pretraživanje mora imati metodu nazvanu `findMaxHeight`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMaxHeight == 'function';
  })()
);
```

Binarno stabljiko za pretraživanje mora da ima putanju nazvanu `isBalanced`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isBalanced == 'function';
  })()
);
```

Ruta za `findMinHeight` bi trebalo da vrati minimalnu visinu drveta.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMinHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMinHeight() == 1;
  })()
);
```

Putanja za `findMaxHeight` treba da vrati maksimalnu visinu drveta.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMaxHeight() == 5;
  })()
);
```

Prazno stablo treba da vrati visinu `-1`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    return test.findMaxHeight() == -1;
  })()
);
```

Put za `isBalanced` bi trebalo da vrati `false` ako je drvo neravnotežno binarno stablo pretrage.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.isBalanced() === false;
  })()
);
```

Putanja za `isBalanced` bi trebalo da vrati `true` ako je stablo balansirano binarno stablo pretrage.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(10);
    test.add(3);
    test.add(22);
    test.add(1);
    test.add(4);
    test.add(17);
    test.add(32);
    return test.isBalanced() === true;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
  this.findMinHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMinHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMinHeight(root.left) + 1;
    }
    const lHeight = this.findMinHeight(root.left);
    const rHeight = this.findMinHeight(root.right);
    return Math.min(lHeight, rHeight) + 1;
  };
  this.findMaxHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMaxHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMaxHeight(root.left) + 1;
    }
    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    return Math.max(lHeight, rHeight) + 1;
  };
  this.isBalanced = function(root = this.root) {
    if (root === null) {
      return true;
    }

    if (root.left === null && root.right === null) {
      return true;
    }

    if (root.left === null) {
      return this.findMaxHeight(root.right) <= 0;
    }

    if (root.right === null) {
      return this.findMaxHeight(root.left) <= 0;
    }

    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    if (Math.abs(lHeight - rHeight) > 1) {
      return false;
    }
    return this.isBalanced(root.left) && this.isBalanced(root.right);
  };
}
```
