---
id: 587d8257367417b2b2512c7e
title: Koristite dubinsku pretragu prvo u binarnoj stablo pretrage
challengeType: 1
forumTopicId: 301719
dashedName: use-depth-first-search-in-a-binary-search-tree
---

# --description--
Znamo kako pronaći određenu vrednost u binarnom stablu pretrage. Ali šta ako želimo samo da istražimo celo stablo? Ili šta ako nemamo organizovano stablo i samo nam je potrebno pronaći vrednost? Ovde ćemo predstaviti neke metode pretraživanja stabla koje se mogu koristiti za istraživanje struktura podataka tipa stablo. Prvo je dubinska pretraga (Depth-First Search). U dubinskoj pretrazi, izvučeno podstablo istražuje se što je moguće dublje pre nego što se pretraga nastavi na drugo podstablo. Postoje tri načina da se to uradi:

In-order (Porredno): Počnite pretragu od najlevšeg čvora i završite na najdesnjem čvoru.
Pre-order (Predredno): Istražite sve korene pre listova.
Post-order (Nakonredno): Istražite sve listove pre korena.

Kao što možete da zamislite, možete izabrati različite metode pretraživanja zavisno od vrste podataka koje vaše stablo skladišti i šta tražite. Za binarno stablo pretrage, *in-order* (porredno) pretraživanje vraća čvorove u sortiranom redosledu.

# --instructions--
Ovde ćemo kreirati ove tri metode pretraživanja na našoj binarnoj stablu za pretraživanje. Prethodno je pretraživanje po dubini inherentna rekurzivna operacija koja nastavlja da istražuje najdublje grane sve dok postoje čvorovi dece. Kada razumete ovaj osnovni koncept, možete jednostavno promeniti redosled posete čvorovima i podstabljama kako biste proizveli jedno od tri navedena pretraživanja. Na primer, pri post-order pretrazi želimo da rekurzivno idemo do lista čvora pre nego što počnemo vraćati bilo koji čvor, dok u pre-order pretrazi želimo da vratimo čvor prvo, a zatim nastavimo sa rekursijom niz stablo. Tafsirajte metode `inorder`, `preorder` i `postorder` na našem stablu. Svaka metoda treba da vrati listu elemenata koja predstavlja pretraživanje stabla. Uverite se da vraćate celobrojne vrednosti za svaki čvor u listi, a ne sam čvor. Na kraju, vratite `null` ako je stablo prazno.

# --hints--
Struktura podataka `BinarySearchTree` mora da postoji.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Binarno stablo pretrage treba da ima metodu nazvanu `inorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.inorder == 'function';
  })()
);
```

Binarno stablo pretrage treba da ima metodu nazvanu `preorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.preorder == 'function';
  })()
);
```

Binarni stab za pretraživanje mora imati metodu nazvanu `postorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.postorder == 'function';
  })()
);
```

Putanja za `inorder` trebalo bi da vrati niz vrednosti čvorova dobijen rednim obilaskom.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.inorder().join('') == '012345678910';
  })()
);
```

Put za `preorder` bi trebalo da vrati niz vrednosti čvorova koji proističu iz traversiranja pre sortiranja.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.preorder().join('') == '710325469810';
  })()
);
```

Metoda `postorder` treba da vrati niz vrednosti čvorova koji rezultuju traversovanjem nakon sortiranja.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.postorder().join('') == '024653181097';
  })()
);
```

Put od `inorder` bi trebalo da vrati `null` na prazno stablo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    return test.inorder() == null;
  })()
);
```

Putanja za `preorder` treba da vrati `null` na prazno stablo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    return test.preorder() == null;
  })()
);
```

Putanja za `postorder` treba da vrati `null` na prazno stablo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    return test.postorder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.result = [];

  this.inorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.inorder(node.left);
    this.result.push(node.value);
    if (node.right) this.inorder(node.right);
    return this.result;
  };
  this.preorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    this.result.push(node.value);
    if (node.left) this.preorder(node.left);
    if (node.right) this.preorder(node.right);
    return this.result;
  };
  this.postorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.postorder(node.left);
    if (node.right) this.postorder(node.right);
    this.result.push(node.value);

    return this.result;
  };
}
```
