---
id: 587d825c367417b2b2512c90
title: Pretraga širine prvo
challengeType: 1
forumTopicId: 301622
dashedName: breadth-first-search
---

# --description--
Dakle, učili smo različite načine kreiranja reprezentacija grafova. Zar ne? Originalno pitanje je: koja je udaljenost između bilo koje dve čvorove u grafu? Ovde ulaze algoritmi za <dfn>prolazak kroz graf </dfn>.

<dfn>Algoritmi prolaska</dfn> su algoritmi za prolazak ili posetu čvorova u grafu. Jedan od tipova algoritama prolaska je algoritam široko-prve pretrage.

Ovaj algoritam počinje od jednog čvora i poseti sve njegove komšije koje su udaljene jednom ivicom. Zatim nastavlja da poseti komšije svakog od njih i tako dalje, sve dok se ne dosegnu svi čvorovi.

Esencijalna struktura podataka koja pomaže u implementaciji algoritma pretraživanja po širini je red. Ovo je linearna struktura podataka gde možete dodavati elemente na jednu stranu i uklanjati elemente sa druge strane. To se takođe naziva strukturom podataka <dfn>FIFO</dfn> ili <dfn>First-In-First-Out</dfn>.

Prikazaće kako algoritam funkcioniše:
<img alt="animation showing the breadth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/breadth-first-search.gif'>

Siva boja ukazuje čvor koji se dodaje u redovnik, a crna boja ukazuje čvor koji se uklanja iz redovnika. Pogledajte kako se svaki put kada čvor izađe iz redovnika (čvor postaje crn), svi njegovi susedi dodaju u redovnik (čvor postaje siv).

Da biste izvršili ovaj algoritam, moraćete da unesete strukturu grafa i oznaku čvora sa kojeg želite početi.

Prvo, biće vam potrebno razumevanje udaljenosti od, ili broja ivica od, početnog čvora. Želite da sve udaljenosti inicijalizujete velikim brojem, kao što je `Infinity`. Ovo sprečava probleme sa računanjem kada do čvora nije moguće stići iz početnog čvora. Zatim, želite da krenete od početnog čvora ka njegovim susjedima. Ovi susjedi su na udaljenosti jedne ivice i u ovom trenutku bi trebalo da dodate jednu jedinicu udaljenosti na udaljenost koju sačuvavate.

# --instructions--
Definisati funkciju `bfs()` koja prima graf matrice susednosti (dvodimenzionalni niz podataka) i oznaku korenovog čvora kao parametre. Oznaka čvora će biti samo vrednost broja čvora između `0` i `n - 1`, gde je `n` ukupan broj čvorova u grafu.

Vaša funkcija će vratiti parove ključ-vrednost za JavaScript koji sadrže čvorove i njihovu udaljenost od korena. Ako je čvor nedostupan, trebalo bi da ima udaljenost od `Infinity`.

# --hints--

Ulazni graf `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `1` treba da vrati `{0: 1, 1: 0, 2: 1, 3: 2}`.

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 1);
assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: 2 });
```

Ulazni graf `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` sa početnim čvorom `1` treba da vrati `{0: 1, 1: 0, 2: 1, 3: Infinity}`.

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 0],
  [0, 0, 0, 0]
];
var results = bfs(graph, 1);
 assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: Infinity });
```

Ulazni graf `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `0` treba da vrati `{0: 0, 1: 1, 2: 2, 3: 3}`.

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 0);
 assert.deepEqual(results, { 0: 0, 1: 1, 2: 2, 3: 3 });
```

Ulazni graf `[[0, 1], [1, 0]]` sa početnim čvorom `0` treba da vrati `{0: 0, 1: 1}`.

```js
var graph = [
  [0, 1],
  [1, 0]
];
var results = bfs(graph, 0);
assert.deepEqual(results, { 0: 0, 1: 1 });
```

# --seed--

## --seed-contents--

```js
function bfs(graph, root) {
  var nodesLen = {};

  return nodesLen;
};

var exBFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(bfs(exBFSGraph, 3));
```

# --solutions--

```js
function bfs(graph, root) {
  var nodesLen = {};
  // Set all distances to infinity
  for (var i = 0; i < graph.length; i++) {
    nodesLen[i] = Infinity;
  }
  nodesLen[root] = 0; // ...except root node
  var queue = [root]; // Keep track of nodes to visit
  var current; // Current node traversing
  // Keep on going until no more nodes to traverse
  while (queue.length !== 0) {
    current = queue.shift();
    // Get adjacent nodes from current node
    var curConnected = graph[current]; // Get layer of edges from current
    var neighborIdx = []; // List of nodes with edges
    var idx = curConnected.indexOf(1); // Get first edge connection
    while (idx !== -1) {
      neighborIdx.push(idx); // Add to list of neighbors
      idx = curConnected.indexOf(1, idx + 1); // Keep on searching
    }
    // Loop through neighbors and get lengths
    for (var j = 0; j < neighborIdx.length; j++) {
      // Increment distance for nodes traversed
      if (nodesLen[neighborIdx[j]] === Infinity) {
        nodesLen[neighborIdx[j]] = nodesLen[current] + 1;
        queue.push(neighborIdx[j]); // Add new neighbors to queue
      }
    }
  }
  return nodesLen;
}
```
