---
id: 587d8258367417b2b2512c80
title: Izbriši listni čvor u binarnom stablu pretrage
challengeType: 1
forumTopicId: 301637
dashedName: delete-a-leaf-node-in-a-binary-search-tree
---

# --description--
Ovo je prvi od tri izazova gde ćemo implementirati najtežu operaciju u binarnim stablama pretrage: brisanje. Brisanje je teško jer brisanje čvora prekida veze unutar stabla. Ove veze moraju pažljivo ponovo povezati kako bi se osigurao integritet strukture binarnog stabla. Za neko brisanje, to znači da stablo mora biti restrukturirano. Generalno, naići ćete na jedan od tri slučaja kada želite obrisati čvor: Listni čvor: Cilj za brisanje nema detilje. Jedan detilji: Cilj za brisanje ima samo jednog detilja. Dva detilja: Cilj za brisanje ima dva detilja. Brisanje listnog čvora je jednostavno, samo ga brišemo. Brisanje čvora sa jednim detiljem takođe je lako, samo ga brišemo i povezujemo njegovog roditelja sa čvorom tog detilja koji smo obrisali. Brisanje čvora sa dva detilja je najteže, jer to dovodi do dva detilja koja moraju biti ponovo povezana sa roditeljskim stablom. Videćemo kako da se reši ovaj slučaj u trećem izazovu. Štaviše, morate biti pažljivi sa nekim posebnim slučajevima kada upravljate brisanjem. Da li je stablo prazno? Je li čvor za brisanje korenovi čvor? Ima li stablo samo dva elementa? Za sada, hajde da uzmemo prvi slučaj u kojem brišemo listni čvor.

# --instructions--
Implementirajte funkciju na našem binarnom stablu pod nazivom `remove`. Ovde ćemo izgraditi logiku za našu operaciju brisanja. Prvo, želite da napravite funkciju unutar `remove` koja traži čvor koji želimo obrisati u trenutnom stablu. Ako čvor nije u stablu, `remove` bi trebalo da vrati `null`. Sada, ako je ciljni čvor listni i nema potomaka, onda bi referenca njegovog roditelja na taj čvor trebalo da se postavi na `null`. Ovo efikasno briše taj čvor iz stabla. Da biste to uradili, takođe ćete morati da pratite roditelja čvora koji želimo obrisati. Takođe je važno napraviti način za praćenje broja potomaka ciljnog čvora, jer će to odrediti koji slučaj brisanja obrađujemo. Drugi i treći slučaj ćemo obraditi u budućim izazovima. Srećno!

# --hints--
Struktura podataka za `BinarySearchTree` mora da postoji.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Binarno stablo pretrage treba da ima metodu nazvanu `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Pokušaj brisanja elementa iz praznog stabla bi trebalo da vrati `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    return test.remove(100) == null;
  })()
);
```

Pokušaj brisanja elementa koji ne postoji bi trebalo da vrati `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(15);
    test.add(30);
    return test.remove(100) == null;
  })()
);
```

Ako koren čvor nema dece, brisanje treba da postavi koren na `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(500);
    test.remove(500);
    return test.inorder() == null;
  })()
);
```

Putanja `remove` treba da obriše listne čvorove iz drveta.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(5);
    test.add(3);
    test.add(7);
    test.add(6);
    test.add(10);
    test.add(12);
    test.remove(3);
    test.remove(12);
    test.remove(10);
    return test.inorder().join('') == '567';
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    },
    inorder: function() {
      if (this.root == null) {
        return null;
      } else {
        var result = new Array();
        function traverseInOrder(node) {
          if (node.left != null) {
            traverseInOrder(node.left);
          }
          result.push(node.value);
          if (node.right != null) {
            traverseInOrder(node.right);
          }
        }
        traverseInOrder(this.root);
        return result;
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}

function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
}
```

# --solutions--

```js
// solution required
```
