---
id: 587d825a367417b2b2512c87
title: Kreiraj dvostruko povezanu listu
challengeType: 1
forumTopicId: 301626
dashedName: create-a-doubly-linked-list
---

# --description--
Sve spojisne liste koje smo do sada kreirali su jednostručne spojisne liste. Ovde ćemo kreirati <dfn>dvostručno spojisnu listu</dfn>. Kao što ime kaže, čvorovi u dvostručno spojisnoj listi imaju reference ka sledećem i prethodnom čvoru u listi.

Ovo nam omogućava da traversiramo listu u oba pravca, ali zahteva i dodatnu memoriju jer svaki čvor mora imati dodatnu referencu na čvor koji ga prethodi u listi.

# --instructions--
Uklonili smo stavke za `Node` i pokrenuli našu `DoublyLinkedList`. Dodali smo dve metode u naš dvosmerno povezani spisak koje nazivaju `add` i `remove`. Metoda `add` treba da doda uklonenu funkciju u spisak, dok metoda `remove` treba da ukloni sve instance uklonene funkcije iz spiska.

Budite pažljivi prilikom rukovanja bilo kojim ivicama dok pišete ove linije, kao što je brisanje prvog ili poslednjeg elementa. Takođe, uklanjanje bilo kog elementa iz prazne liste bi trebalo da vrati `null`.

# --hints--
Struktura podataka za `DoublyLinkedList` treba da postoji.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return typeof test == 'object';
  })()
);
```

`DoublyLinkedList` trebalo bi da ima metodu nazvanu `add`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.add == undefined) {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

`DoublyLinkedList` trebalo bi da ima metodu nazvanu `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.remove == undefined) {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Uklanjanje elementa iz prazne liste trebalo bi da vrati `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return test.remove(100) == null;
  })()
);
```

Putanja za `add` bi trebalo da doda stavke na listu.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(5);
    test.add(6);
    test.add(723);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '56723';
  })()
);
```

Svaki čvor mora pratiti prethodni čvor.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(50);
    test.add(68);
    test.add(73);
    var result = [];
    var curr = test.tail;
    while (curr) {
      result.push(curr.data);
      curr = curr.prev;
    }
    return result.join('') == '736850';
  })()
);
```

Prvi element treba ukloniti sa liste.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(25);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '3560';
  })()
);
```

Poslednji element treba da se ukloni sa liste.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(60);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '2535';
  })()
);
```


# --seed--

## --seed-contents--

```js
var Node = function(data, prev) {
  this.data = data;
  this.prev = prev;
  this.next = null;
};
var DoublyLinkedList = function() {
  this.head = null;
  this.tail = null;
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
