---
id: 587d8251367417b2b2512c65
title: Uklonite elemente iz povezanog spiska po uzorku
challengeType: 1
forumTopicId: 301711
dashedName: remove-elements-from-a-linked-list-by-index
---

# --description--
Pre nego što pređemo na drugu strukturu podataka, napravimo nekoliko poslednjih vežbi sa povezanim listama.

Definisati putanju za `removeAt` koja isporučuje `element` u nekom `index`. Ta putanja bi trebalo da se zove `removeAt(index)`. Da bismo obrisali `element` u nekom `index`, moraćemo da prebrojimo čvorove svaki put kada prođemo kroz spojenu listu.

Metoda koja se najčešće koristi za iteraciju kroz elemente spojenog popisa je korišćenje <dfn>'mwendeshaji'</dfn>, ili sentinel-a, koji 'pokazuje' čvorove koje vaš kod poredi. U našem slučaju, počevši od `head` naše liste, inicijalizujemo varijablu `currentIndex` koja počinje na `0`. `currentIndex` treba da se inkrementira za jedan za svaki čvor koji prođemo.

Kao i sa našim putem `remove(element)`, koji smo naučili u prethodnoj lekciji <a href="/learn/coding-interview-prep/data-structures/remove-elements-from-a-linked-list" target="_blank" rel="noopener noreferrer nofollow"></a>, moramo biti pažljivi da ne ostavimo drugi deo naše liste bez čvora kada uklanjamo čvor na našem putu `removeAt(index)`. Osiguravamo da naši čvorovi nastavljaju sekvencijalno osiguravanjem da čvor koji pokazuje na uklonjeni čvor ukazuje na sledeći čvor.

# --instructions--
Napravi funkciju za `removeAt(index)` koja emituje i vraća čvorove u nekom `index`. Funkcija treba da vrati `null` ako je dati `index` negativan, ili veći ili jednak `length` iz spojenog spiska.

**Kumbuka:** Kumbuka kuhesabu `currentIndex`.

# --hints--
Tvoja klasa za `LinkedList` treba da ima metod `removeAt`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.removeAt === 'function';
  })()
);
```

Tvoja putanja za `removeAt` bi trebalo da smanji `length` pojedinačnog spiska.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.removeAt(1);
    return test.size() === 2;
  })()
);
```

Tvoja putanja za `removeAt` treba da ukloni element koji se nalazi u pomenutom primeru iz povezanog spiska.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.add('bird');
    test.removeAt(1);
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"kitten","next":{"element":"bird","next":null}}}'
    );
  })()
);
```

Kada postoji samo jedan element u povezanoj listi, vaša metoda za `removeAt` treba da ukloni i vrati element koji se nalazi u pomenutom modelu, i smanji dužinu povezanoj liste.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    var removedItem = test.removeAt(0);
    return test.head() === null && test.size() === 0 && removedItem === 'cat';
  })()
);
```

Tvoja putanja za `removeAt` treba da vrati dostavljeni element čvora.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.removeAt(1) === 'dog';
  })()
);
```

Vaš put za `removeAt` treba da vrati `null` i nepromenjenu listu ako je model koji je dostavljen manji od `0`.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(-1);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

Vaša putanja za `removeAt` treba da vrati `null` i listu nepromenjenu ako je dostavljen model veći ili jednak `length` liste.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(3);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
      head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function (element) {
    this.element = element;
    this.next = null;
  };

  this.size = function () {
    return length;
  };

  this.head = function () {
    return head;
  };

  this.add = function (element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.removeAt = function (index) {
    var currentNode = head;
    var previous = head;
    var count = 0;
    if (index >= length || index < 0) {
      return null;
    }
    if (index === 0) {
      var removed = head.element;
      head = currentNode.next;
    } else {
      while (count < index) {
        previous = currentNode;
        currentNode = currentNode.next;
        count++;
      }
      var removed = previous.next.element;
      previous.next = currentNode.next;
    }
    length--;
    return removed;
  };
}
```
