---
id: 587d8252367417b2b2512c67
title: Dodaj elemente u specifičan šablon u povezanoj listi
challengeType: 1
forumTopicId: 301619
dashedName: add-elements-at-a-specific-index-in-a-linked-list
---

# --description--
Napravimo metodu `addAt(index,element)` koja dodaje element u zadatu strukturu. Kao što koristimo za uklanjanje elemenata iz određene strukture, moramo pratiti `currentIndex` dok prolazimo kroz spojenu listu. Kada se `currentIndex` poklapa sa zadatom strukturom, moramo postaviti novu vrednost za svojstvo `next` prethodnog čvora (node) tako da on referencira novi dodati čvor. A novi čvor treba da ima referencu na sledeći čvor na poziciji `currentIndex`. Vraćajući se na primer sa konge linije, nova osoba želi da se pridruži liniji, ali želi da se pridruži usred. Ti si usred linije, pa protežeš ruke osobi koja je bila ispred tebe. Nova osoba dolazi i stavlja svoje ruke na osobu čije su ruke bile u tvojim rukama, a sada ti imaš ruke sa novom osobom.

# --instructions--
Napravi put `addAt(index,element)` koji dodaje element na dati model. Vrati false ako ne može da se doda element. **Zapamti:** Proveri da li je dati model negativan ili da je prešao dužinu povezanog popisa.

# --hints--
Tvoja putanja za `addAt` treba da postavi vrednost `head` na novi čvor kada je dato nulti.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'fish');
    return test.head().element === 'fish' && test.head().next.element === 'cat';
  })()
);
```

Tvoja putanja ``addAt`` bi trebalo da poveća dužinu spojenog popisa za svaki novi čvor koji se dodaje.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'cat');
    return test.size() === 3;
  })()
);
```

Tvoja putanja ``addAt`` bi trebalo da vrati ``false`` ako čvor nije mogao biti dodan.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.addAt(4, 'cat') === false;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element) {
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function() {
    return head;
  };

  this.add = function(element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if (head === null){
        head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };
  this.addAt = function (index, element) {
    if (index > length || index < 0) {
      return false;
    }
    var newNode = new Node(element);
    var currentNode = head;
    if (index === 0) {
      head = newNode;
    } else {
      var previousNode = null;
      var i = 0;
      while (currentNode && i < index) {
        previousNode = currentNode;
        currentNode = currentNode.next;
        i++;
      }
      previousNode.next = newNode;
    }
    newNode.next = currentNode;
    length++;
  }
}
```
