---
id: 587d8251367417b2b2512c62
title: Kreiraj klasu za povezanu listu
challengeType: 1
forumTopicId: 301628
dashedName: create-a-linked-list-class
---

# --description--
Početićemo kreiranjem klase za `linked list`. Svaka povezana lista treba da počinje sa nekoliko osnovnih svojstava: `head` (prvi element u vašoj listi) i `length` (broj elemenata u vašoj listi). Ponekad ćete videti implementaciju povezanih lista koja uključuje `tail` za poslednji element liste, ali za sada ćemo se držati samo ovih dve. Svaki put kada dodamo element u povezanu listu, naše svojstvo `length` treba da se poveća za jedan.

Biće nam potreban način za dodavanje stavki u našu povezanu listu, pa je prvi metod koji želimo da kreiramo je `add`.

Ako je naša lista prazna, dodavanje elementa u povezanu listu je lako: rotiramo taj element u klasu `Node` i postavljamo taj čvor kao `head` naše povezane liste.

Međutim, šta ako naša lista već ima jednog ili više članova? Kako dodajemo element u listu? Zapamtite da svaki čvor u spojenoj listi ima atribut `next`. Da biste dodali čvor u listu, pronađite poslednji čvor u listi i usmerite atribut `next` tog poslednjeg čvora ka našem novom čvoru. (Napomena: znate da ste dostigli kraj spojenoj liste kada je atribut `next` čvora `null`.)

# --instructions--
Napisati metodu za dodavanje koja postavlja prvi čvor koji dodaješ u spojenoj listi kao `head`; nakon toga, svaki put kada dodaš čvor, svaki čvor bi trebalo da ima referencu na svojstvo `next` prethodnog čvora.

Kumbuka

`length` Indeks vaše liste treba da se poveća za jedan svaki put kada se dodaje element u spojenoj listi.

# --hints--
Klasa za `LinkedList` mora da ima metod za `add`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.add === 'function';
  })()
);
```

Vaša klasa za `LinkedList` treba da postavi `head` za prvi dodati čvor.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    return test.head().element === 'cat';
  })()
);
```

`node` koji prethodi u vašoj klasi od `LinkedList` mora imati referencu na najnovije kreiran čvor.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('fish');
    return test.head().next.element === 'dog' && test.head().next.next.element === 'fish';
  })()
);
```

`size` za vašu klasu od `LinkedList` mora biti jednako broju čvorova u spojenoj listi.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.size() === 2;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() { 
  var length = 0; 
  var head = null; 

  var Node = function(element){
    this.element = element; 
    this.next = null; 
  }; 

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line
    if (head == null) {
      head = new Node(element);
    } 
    else {
      let currentNode = head;
      while (currentNode.next != null) {
        // currentNode.next will be last node of linked list after loop
        currentNode = currentNode.next;
      }
      currentNode.next = new Node(element);
    }
    length++;
    // Only change code above this line
  };
}
```
