---
id: 587d825b367417b2b2512c8b
title: Ukloni element iz najgornjeg paketa/skupa.
challengeType: 1
forumTopicId: 301710
dashedName: remove-an-element-from-a-max-heap
---

# --description--
Sada kada možemo dodati elemente na naš stak, istražimo kako da uklonimo elemente. Uklanjanje i ubacivanje svih elemenata zahteva sličnu logiku. U najdubljem staku, obično želite da uklonite najveću vrednost, pa ovo uključuje samo preuzimanje iz korena našeg stabla. Ovo će narušiti svojstvo staka našeg stabla, pa moramo ga na neki način popraviti. Obično, za najdublji stak, ovo se radi na sledeći način:

<ol>
  <li>Premesti poslednji element iz steka na poziciju korena.</li>
  <li>Ako bilo koji dete korena ima veću vrednost od samog korena, zameni koren detetom koje ima najveću vrednost.</li>
  <li>Nastavi sa zamjenom dok roditelj nije veći od oba dete ili dođeš do poslednjeg nivoa u drvetu.</li>
</ol>

# --instructions--
Upute: Dodaj metodu na našem maksimalnom heap-u nazvanom `remove`. Ova metoda treba da vrati najveću vrednost koja je dodata u naš maksimalni heap i ukloni je iz heap-a. Takođe, treba da reorganizuje heap kako bi svojstvo heap-a ostalo održano. Nakon uklanjanja elementa, najveći preostali element u heap-u treba da bude koren.

# --hints--
Struktura podataka za `MaxHeap` mora da postoji.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap` trebalo bi da ima putanju nazvanu `print`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

`MaxHeap` trebala bi da ima metodu nazvanu `insert`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap` trebalo bi da ima putanju nazvanu `remove`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Metoda `remove` treba da ukloni veliki element iz najgornjeg stoga, uz održavanje svojstva najgornjeg stoga.

```js
function isHeap(arr, i, n) {
    if( arr[i] < arr[2 * i + 1] || arr[i] < arr[2 * i + 2] ){
        return false;
    }
    if (i > (n - 1) / 2) {
        return true;
    }
    if (isHeap(arr, 2 * i + 1, n) && isHeap(arr, 2 * i + 2, n)) {
        return true;
    }
    return false;
}


assert(
  (function () {
    let test = false;

    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    
    let max = Infinity;
    const [result, vals] = [[], [9, 3, 5, 2, 15, 3, 7, 12, 7, 10, 90]];

    vals.forEach((val) => test.insert(val));

    for (let i = 0; i < vals.length; i++) {
      const curHeap = test.print();
      const arr = curHeap[0] === null ? curHeap.slice(1) : curHeap;
      
      if (!isHeap(arr, 0, arr.length - 1)) {
        return false;
      }

      const removed = test.remove();
      if (!vals.includes(removed)) return false;
      if (removed > max) return false
      max = removed;
      result.push(removed);
    }
    
    for (let i = 0; i < vals.length; i++) {
      if (!result.includes(vals[i])) {
        return false;
      }
    }
    
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
const MaxHeap = function () {
  this.heap = [];
  this.parent = index => {
    return Math.floor((index - 1) / 2);
  }
  this.insert = element => {
    this.heap.push(element);
    this.heapifyUp(this.heap.length - 1);
  }
  this.heapifyUp = index => {
    let currentIndex = index,
    parentIndex = this.parent(currentIndex);
    while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
      this.swap(currentIndex, parentIndex);
      currentIndex = parentIndex;
      parentIndex = this.parent(parentIndex);
    }
  }
  this.swap = (index1, index2) => {
    [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
  }
  this.print = () => {
    return this.heap;
  }
  // Only change code below this line

  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
