---
id: 587d8251367417b2b2512c63
title: Uklonite elemente iz povezanog spiska
challengeType: 1
forumTopicId: 301712
dashedName: remove-elements-from-a-linked-list
---

# --description--
Sledeća važna metoda koju bilo koja implementacija spojenog spiska će zahtevati je metoda `remove`. Ova metoda bi trebalo da uzme element koji želimo da uklonimo kao argument, a zatim da pretraži listu kako bi pronašla i uklonila čvor koji sadrži taj element.

Svaki put kada uklonimo čvor iz spojenog popisa, važno je osigurati da ne ostavimo drugi deo popisa bez veze. Zapamtite da atribut ``next`` svakog čvora ukazuje na sledeći čvor u popisu. Ako uklonimo srednji element, na primer, želimo da osiguramo vezu od atributa ``next`` prethodnog čvora do atributa ``next`` srednjeg elementa (koji je sledeći čvor u popisu!).

Ovo može izgledati veoma zbunjujuće, pa da se vratimo na primer konge linije kako bismo imali dobar uvid koncepta. Stani u konge liniju, i osoba koja je ispred tebe napušta liniju. Osoba koja je napustila liniju više nema ruke na nikoga u liniji — i ti više nemate ruke na osobi koja je otišla. Napreduješ napred i stavljaš ruke na osobu koja te prati.

Ako je element koji želimo ukloniti element `head`, menjamo `head` u drugi čvor spojenog popisa.

# --instructions--
Napiši putanju za `remove` koja uzima element i uklanja ga iz povezanog popisa.

**Napomena:** Broj `length` liste treba da se smanji za jedan svaki put kada se element ukloni iz povezane liste.

# --hints--
Vaša klasa za `LinkedList` bi trebalo da ima metodu za `remove`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.remove === 'function';
  })()
);
```

Tvoja putanja za `remove` treba da ponovo postavi `head` na drugi čvor kada se prvi čvor ukloni.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.remove('cat');
    return test.head().element === 'dog';
  })()
);
```

Vaša putanja od `remove` bi trebalo da smanji `length` povezanog spiska za svaki uklonjeni čvor.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('hamster');
    test.remove('cat');
    test.remove('fish');
    return test.size() === 2;
  })()
);
```

Vaša putanja za `remove` treba da resetuje referencu prethodnog čvora za referencu `next` uklonjenog čvora.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('snake');
    test.add('kitten');
    test.remove('snake');
    return test.head().next.next.element === 'kitten';
  })()
);
```

Tvoja putanja za `remove` ne bi trebalo da modifikuje povezanu listu ako element nije u povezanoj listi.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.remove('elephant');
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    if (head === null) {
      return;
    }
    var previous;
    var currentNode = head;

    while (currentNode.next !== null && currentNode.element !== element) {
      previous = currentNode;
      currentNode = currentNode.next;
    }
    
    if (currentNode.next === null && currentNode.element !== element) {
      return;
    }
    else if (previous) {
      previous.next = currentNode.next;
    } else {
      head = currentNode.next;
    }

    length--;
  };
} 
```
