---
id: 587d8255367417b2b2512c73
title: Koristite operater za raspršivanje i opise za konekciju ES5 Set()
challengeType: 1
forumTopicId: 301720
dashedName: use-spread-and-notes-for-es5-set-integration
---

# --description--

Je, unakumbuka opereta ya spread ya ES6 `...`?

`...` može uzeti iterabilne vrednosti u ES6 i pretvoriti ih u nizove podataka.

Tufanye Set, na tuchunguze kitendakazi cha spread.

```js
var set = new Set([1,2,3]);
var setToArr = [...set]
console.log(setToArr) // returns [ 1, 2, 3 ]
```

# --instructions--
U ovoj vežbi ćemo proslediti objekat set funkciji `checkSet`. Treba da vrati red podataka sa vrednostima Set.

Sada si uspešno naučio kako da koristiš nešto od `Set()` ES6, odlično!

# --hints--

`checkSet(new Set([1,2,3,4,5,6,7])` inapaswa kurudisha `[1, 2, 3, 4, 5, 6, 7]`.

```js
var test = checkSet(new Set([1, 2, 3, 4, 5, 6, 7]));
assert.deepEqual(test, [1, 2, 3, 4, 5, 6, 7]);
```

# --seed--

## --seed-contents--

```js
function checkSet(set){
   // Only change code below this line

   // Only change code above this line
}
```

# --solutions--

```js
function checkSet(set){
return [...set];}
```
