---
id: 587d825b367417b2b2512c8e
title: Unda jedwali la hash
challengeType: 1
forumTopicId: 301627
dashedName: create-a-hash-table
---

# --description--
U ovom izazovu ćemo naučiti o hash tabelama. Hash tabela se koristi za implementaciju asocijativnih struktura podataka, ili mapa parova ključ-vrednost, kao što su objekti i Mapovi koje smo nedavno učili. Objekat JavaScript može se implementirati kao hash tabela, na primer (njegova stvarna implementacija će zavisiti od okruženja u kojem se izvršava). Način na koji hash tabela radi je da uzme ulazni ključ i pouzdano ga hashuje do određene numeričke vrednosti. Ova numerička vrednost koristi se kao stvarni ključ kojim se čuva pripadajuća vrednost. Zatim, ako pokušate da dobijete pristup istom ključu ponovo, hash funkcija će obradovati ključ, vraćajući isti numerički rezultat, koji će zatim biti korišćen za pretraživanje pripadajuće vrednosti. Ovo pruža prosečno vreme pristupa O(1) sa velikom efikasnošću.

Hash mape mogu se implementirati kao nizovi podataka sa funkcijama hasovanja koje generišu indekse niza unutar definisanog opsega. Na ovaj način, izbor veličine niza je važan, posebno u odnosu na funkciju hasovanja. Na primer, da li će funkcija hasovanja generisati istu vrednost za dva različita ključa? To se naziva kolizija. Jedan način rešavanja kolizija je skladištenje svih parova ključ-vrednost u toj strukturi. Zatim, prilikom dohvaćanja jednog od njih, moraćete da iterirate kroz grupu elemenata kako biste pronašli traženi ključ. Dobra funkcija hasovanja smanjujeće kolizije radi održavanja dobrog vremena pretrage.

Ovde, nećemo se fokusirati na detalje heširanja ili implementaciju hash tabela; samo ćemo pokušati da uhvatimo opšti osećaj kako funkcionišu.

# --instructions--
Osnovna funkcionalnost hash tabele. Napravili smo standardnu hash funkciju koju možete koristiti. Možete proslediti vrednost niza funkciji `hash`, a ona će vam vratiti hash vrednost koju možete koristiti kao ključ za skladištenje. Skladištite stavke na osnovu ove hash vrednosti u objekt `this.collection`. Kreirajte ove tri metode: `add`, `remove`, i `lookup`. Prva treba da primi par ključ-vrednost za dodavanje u hash tabelu. Druga treba da ukloni par ključ-vrednost kada mu se prosledi ključ. Treća treba da primi ključ i vrati pripadajuću vrednost ili `null` ako ključ ne postoji.

Hakikisha unaandika msimbo wako kuzingatia migongano!

**Napomena:** Test za putanju `remove` neće proći do putanja `add` i `lookup` ako su ispravno implementirane.

# --hints--

Funkcija `hash` treba da bude ispravna.

```js
let calls = 0;
const ourHash = string => {
  calls++;
  let hashCode = 0;
  for (let i = 0; i < string.length; i++) {
    hashCode += string.charCodeAt(i);
  }
  return hashCode;
};

assert.strictEqual(hash('yek'),ourHash('yek'));

assert.strictEqual(hash('key'),ourHash('key'));
assert.strictEqual(hash('key1'),ourHash('key1'));
assert.strictEqual(hash('key2'),ourHash('key2'));
assert.strictEqual(hash('key3'),ourHash('key3'));

assert.strictEqual(hash('1key'),ourHash('1key'));
assert.strictEqual(hash('ke1y'),ourHash('ke1y'));
assert.strictEqual(hash('altKey'),ourHash('altKey'));

assert.strictEqual(called,calls); 
```

Struktura podataka za `HashTable` mora da postoji.

```js
assert.isDefined(HashTable); 
let newHashTable = new HashTable(); 
assert.isObject(newHashTable); 
```

`HashTable` trebalo bi da ima način za `add`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.add); 
```

`HashTable` treba da ima puta do `lookup`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.lookup); 
```

`HashTable` treba da ima način za `remove`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.remove); 
```

Metoda `add` treba da doda parove ključ-vrednost, a metoda `lookup` treba da vrati vrednosti povezane sa određenim ključem.

```js
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
assert.strictEqual(newHashTable.lookup('key'),'value'); 
```

Metoda za `remove` treba da primi ključ kao ulaz i ukloni pripadajući par ključ-vrednost.

```js
let hashValue = hash('key');
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
newHashTable.remove('key'); 
assert.notProperty(newHashTable.collection,hashValue); 
```

Metoda `remove` treba da ukloni samo ispravan par ključ-vrednost.

```js
let newHashTable = new HashTable(); 
let hashValue = hash('key');
newHashTable.add('key','value');
newHashTable.add('yek','value');
newHashTable.add('altKey','value');
newHashTable.remove('yek');

assert.notExists(newHashTable.lookup('yek'));
assert.exists(newHashTable.lookup('altKey'));
assert.exists(newHashTable.lookup('key'));

newHashTable.remove('key');

assert.notProperty(newHashTable.collection,hashValue); 
assert.exists(newHashTable.lookup('altKey'));

```

Podaci moraju biti dodati korišćenjem hash funkcije.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1', 'value1')
newHashTable.add('key2', 'value2');
newHashTable.add('key3', 'value3');
assert.strictEqual(called,3); 
```

Tabela haša mora da reši kolizije.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1','value1');
newHashTable.add('1key', 'value2'); 
newHashTable.add('ke1y', 'value3'); 
assert.strictEqual(newHashTable.lookup('key1'),'value1');
assert.strictEqual(newHashTable.lookup('1key'),'value2');
assert.strictEqual(newHashTable.lookup('ke1y'),'value3');
```

# --seed--

## --seed-contents--

```js
var called = 0;
var hash = string => {
  called++;
  var hashed = 0;
  for (var i = 0; i < string.length; i++) {
    hashed += string.charCodeAt(i);
  }
  return hashed;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
var called = 0;
var hash = (string) => {
  called++;
  var hash = 0;
  for (var i = 0; i < string.length; i++) { hash += string.charCodeAt(i); }
  return hash;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line

  this.add = function(key, val) {
    var theHash = hash(key);
    if (!this.collection.hasOwnProperty(theHash)) {
      this.collection[theHash] = {};
    }
    this.collection[theHash][key] = val;
  }

  this.remove = function(key) {
    var theHash = hash(key);
    var hashedObj = this.collection[theHash];
    if (hashedObj.hasOwnProperty(key)) {
      delete hashedObj[key];
    }
    if (!Object.keys(hashedObj).length) {
      delete this.collection[theHash];
    }
  }

  this.lookup = function(key) {
    var theHash = hash(key);
    if (this.collection.hasOwnProperty(theHash)) {
      return this.collection[theHash][key];
    }
    return null
  }
  // Only change code above this line
};
```
