---
id: 5cc0c1b32479e176caf3b422
title: Proveri da li je drvo binarno drvo pretrage
challengeType: 1
forumTopicId: 301624
dashedName: check-if-tree-is-binary-search-tree
---

# --description--
Pošto već znate šta je binarni stab za pretraživanje, ovaj izazov će razjasniti kako možete utvrditi da li je stab binarni stab za pretraživanje ili ne.

Glavna razlika binarnog stabla pretrage je što su čvorovi organizovani po usklađenom redosledu. Čvor ima maksimalno dva ćeraka (desni i/ili levi) u zavisnosti od toga da li je vrednost ćerakovog čvora veća ili jednaka (desna) ili manja od (leva) čvora roditelja.

# --instructions--
U ovom izazovu, kreiraćeš pomoćnu funkciju za svoj stabao. Napiši metodu JavaScript `isBinarySearchTree` koja prima stabao kao ulaz i vraća vrednost tačno ili netačno da li je stabao binarno stabalo pretrage ili ne. Koristi rekurziju (recursion) kada je to moguće.

# --hints--

Tvoje binarno stablo pretrage treba da vrati `true` kada se proveri funkcijom `isBinarySearchTree()`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    test.push(1);
    test.push(5);
    test.push(3);
    test.push(2);
    test.push(4);
    return isBinarySearchTree(test) == true;
  })()
);
```

`isBinarySearchTree()` trebalo bi da vrati lažno kada se proveri stabljuka koje nije binarno stabljuka za pretragu.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    test.push(1);
    test.root.left = new Node(1);
    return isBinarySearchTree(test) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype.push = function(val) {
  var root = this.root;

  if (!root) {
    this.root = new Node(val);
    return;
  }

  var currentNode = root;
  var newNode = new Node(val);

  while (currentNode) {
    if (val < currentNode.value) {
      if (!currentNode.left) {
        currentNode.left = newNode;
        break;
      } else {
        currentNode = currentNode.left;
      }
    } else {
      if (!currentNode.right) {
        currentNode.right = newNode;
        break;
      } else {
        currentNode = currentNode.right;
      }
    }
  }
};
```

# --seed--

## --seed-contents--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
}
function isBinarySearchTree(tree) {
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
}
function isBinarySearchTree(tree) {
  if (tree.root == null) {
    return null;
  } else {
    let isBST = true;
    function checkTree(node) {
      if (node.left != null) {
        const left = node.left;
        if (left.value >= node.value) {
          isBST = false;
        } else {
          checkTree(left);
        }
      }
      if (node.right != null) {
        const right = node.right;
        if (right.value < node.value) {
          isBST = false;
        } else {
          checkTree(right);
        }
      }
    }
    checkTree(tree.root);
    return isBST;
  }
};
```
