---
id: 66e45c8140f9fda5c105ae26
title: Kreiraj stranicu sa letakama za događaj
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Trebalo bi da imaš element `header` unutar tela.
1. Tvoj element `header` treba da sadrži sliku za tvoj događaj, a unutrašnjost mora sadržati `h1` sa naslovom tvog događaja, u tom redosledu. Možeš koristiti ovu sliku ako želiš: `https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. Trebalo bi da imaš element `main` unutar tela.
1. Tvoj element `main` treba da sadrži najmanje dva elementa `section` koji prikazuju detalje događaja.
1. Svaki tvoj element `section` treba da sadrži `h2` unutar sebe.
1. Telo treba da ima unutrašnji razmak za vrh i dno `50px`.
1. Telo treba da ima gornje i donje ivice `0`, kao i leve i desne ivice koje su centrirane.
1. Telo treba da ima širinu postavljenu prema širini vidokruga.
1. Telo treba da koristi funkciju `calc` za postavljanje atributa `min-height` na 100% visine vidokruga, isključujući sve unutrašnje razmake postavljene gore i dole tela.
1. Trebalo bi da imaš najmanje jedan element `hr` unutar svojih letaka.
1. Trebalo bi da postaviš `width` za sve elemente `hr` i `section` na procenat vrednost u odnosu na roditeljski element.

**Napomena:** Obavezno povežite vaš list modela sa HTML i koristite svoj CSS.

# --hints--

Trebalo bi da ima element `header` unutar elementa `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Trebalo bi da imaš element `img` kao prvo dete u `header`.

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

Trebalo bi da imaš element `h1` kao drugog deteta u `header`.

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

Trebalo bi da ima samo jedan element `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`h1` Trebalo bi da ima naslov vašeg događaja.

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

Trebalo bi da imaš element `main` unutar elementa `body`.

```js
assert.exists(document.querySelector('body > main'));
```

Trebalo bi da ima najmanje dva elementa `section` unutar tvog elementa `main`.

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

Elementi vašeg `section`, svaki mora da sadrži `h2` unutar sebe.

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

Vaše komponente za `h2` ne smeju biti prazne.

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

Element vašeg `body` trebalo bi da ima unutrašnji razmak iznad i ispod `50px`.

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

Tvoj element od `body` mora imati gornju i donju granicu od `0`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

Tvoj element `body` treba da ima levi i desni rub od `auto`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

Tvoj element za `body` treba da postavi svoju širinu koristeći `vw`.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

Tvoj `body` treba da koristi funkciju `calc` kako bi postavio `min-height` na `100vh - 100px`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

Trebalo bi da ima barem jedan element od `hr`.

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

Širina vaših elemenata za `hr` treba biti postavljena korišćenjem procentualne vrednosti.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

Širina vaših elemenata za `section` bi trebalo postaviti korišćenjem procentualne vrednosti.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
