---
id: 66a207974c806a19d6607073
title: Napravi program za biblioteku knjiga
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Treba da imaš element `h1` sa tekstom `Book Inventory`.
1. Treba da imaš element `table` sa kolonama `Title`, `Author`, `Category`, `Status` i `Rate`.
1. Tabela treba da ima najmanje četiri reda: prvi red za zaglavlja kolona, a ostali redovi treba da sadrže informacije o knjigama.
1. Svaki red tabele u telu tabele treba da ima klasu `read`, `to-read` ili `in-progress`.
1. `td` elementi u koloni `Status` treba da sadrže element `span` sa klasom `status` koji obuhvata tekst `Read`, `To Read` ili `In Progress`, u skladu sa klasom reda.
1. `td` elementi u koloni `Rate` treba da sadrže element `span` sa klasom `rate` koji obuhvata tri prazna elementa `span`.
1. Elementi `.rate` unutar redova `.read` treba da imaju dodatnu klasu sa vrednošću `one`, `two` ili `three`, u skladu sa pojedinačnom ocenom. Ta vrednost treba da bude navedena posle `rate`.
1. Treba da napraviš tri selektora atributa za ciljanje elemenata sa klasama `read`, `to-read` i `in-progress` i da im postaviš svojstvo `background-image` koristeći `linear-gradient` po svom izboru.
1. Treba da postaviš svojstvo `display` svih elemenata `span` na `inline-block`.
1. Treba da koristiš selektor atributa za ciljanje elemenata `span` sa klasom `status` koji su potomci elemenata `tr` sa klasom `to-read` i da im postaviš svojstva `border` i `background-image`.
1. Treba da koristiš selektor atributa za ciljanje elemenata `span` sa klasom `status` koji su potomci elemenata `tr` sa klasom `read` i da im postaviš svojstva `border` i `background-image`.
1. Treba da koristiš selektor atributa za ciljanje elemenata `span` sa klasom `status` koji su potomci elemenata `tr` sa klasom `in-progress` i da im postaviš svojstva `border` i `background-image`.
1. Treba da koristiš selektor atributa za ciljanje elemenata `span` sa klasom `status` i elemenata `span` čija vrednost klase počinje sa `rate`, i da im postaviš svojstva `height`, `width` i `padding`.
1. Treba da koristiš selektor atributa za ciljanje elemenata `span` koji su direktna deca elemenata `span` čija vrednost klase počinje sa `rate`, i da im postaviš svojstva `border`, `border-radius`, `margin`, `height`, `width` i `background-color`.
1. Treba da koristiš selektor atributa za ciljanje prvog roditelja `span` elemenata čija vrednost atributa `class` sadrži reč `one` i da njegovom svojstvu `background-image` dodeliš `linear-gradient`.
1. Treba da koristiš selektor atributa za ciljanje prva dva roditelja `span` elemenata čija vrednost atributa `class` sadrži reč `two` i da njihovom svojstvu `background-image` dodeliš `linear-gradient`.
1. Treba da koristiš selektor atributa za ciljanje tri `span` elementa koji su potomci `span` elemenata čija vrednost atributa `class` sadrži reč `three` i da njihovom svojstvu `background-image` dodeliš `linear-gradient`.

# --hints--

Trebalo bi da imaš element sa tekstom `h1` koji sadrži tekst `Book Inventory`.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

Trebalo bi da ima samo jedan element `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Trebalo bi da imaš element `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Trebalo bi da ima jedan element `thead` i jedan element `tbody` unutar `table`.

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

Unutar `thead` treba da ima jedan element od `tr` sa pet komponenti od `th`.

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

Tvoja prva stranica bi trebalo da ima tekst `Title` kao naslov sekcije.

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

Vaš drugi red mora da sadrži tekst `Author` kao naslov sekcije.

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

Vaš treći red bi trebalo da sadrži tekst `Category` kao zaglavlje sekcije.

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

Vaš četvrti red treba da ima tekst `Status` kao naslov sekcije.

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

Vaš peti red bi trebalo da sadrži tekst `Rate` kao naslov sekcije.

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

Tabela bi trebalo da ima najmanje četiri reda.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

Svaki red mora imati 5 kolona uvek.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

Svaki red tabele, osim reda zaglavlja sekcije, treba da ima klasu `read`, `to-read` ili `in-progress`.

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

Elementi od ``td`` u redu ``Status`` treba da imaju atribut/element ``span``.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

Svaki element redova `span` mora imati klasu `status` iz reda `Status`.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

Svaki element sa `.status` mora imati tekst `Read`, `To Read` ili `In Progress`, u zavisnosti od klase njegovog reda.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

Elementi za `td` iz reda `Rate` treba da imaju element/atribut `span`.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

Svaki element `span` koji je direktni potomak elementa `td` iz redova `Rate` mora imati klasu `rate` kao prvu klasu.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

Svaki element `.rate` mora imati tri prazna elementa `span`.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

Elementi ``.rate`` postavljeni unutar linija ``.read`` treba da imaju dodatnu klasu nakon klase ``rate`` sa vrednošću ``one``, ``two``, ili ``three``, zavisno od pojedinačnih zahteva.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

Trebalo bi da imate selektor atributa elementa koji cilja linije sa klasom `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

Treba da koristite selektor atributa elementa za ciljanje redova sa klasom `read` i postavite njihov atribut `background-image` na kombinaciju linija boje koje birate.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

Trebalo bi da imate selektor atributa elementa koji cilja linije sa klasom `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

Treba da koristiš selektor atributa elementa za ciljanje linija sa klasom `to-read` i postavljanje njihovog atributa `background-image` na kombinaciju boje linije koju biraš.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

Trebalo bi da imaš selektor atributa elementa koji cilja linije sa klasom `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

Treba da koristite selektor atributa elementa da ciljate na linije sa klasom `in-progress` i postavite njihov atribut `background-image` na kombinaciju boje linije koju birate.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

Trebalo bi da postavite atribut `display` svakom elementu `span` na `inline-block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

Trebalo bi da imate selektor atributa elementa za ciljanje elemenata `span` sa klasom `status` koji su potomci elemenata `tr` sa klasom `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

Treba koristiti selektor atributa elementa za ciljanje elemenata `span` koji imaju klasu `status`, koji su potomci elemenata `tr` koji imaju klasu `to-read` i postavljanje njihovog atributa `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

Trebalo bi da koristite selektor svojstva elementa za ciljanje elemenata `span` sa klasom `status`, koji su potomci elemenata `tr` sa klasom `to-read` i postavljanjem njihovog svojstva `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

Treba vam selektor atributa elementa koji cilja elemente ``span`` koji imaju klasu ``status``, a koji su potomci elemenata ``tr`` koji imaju klasu ``read``.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

Trebalo bi da se koristi selektor atributa za ciljanje elemenata `span` koji imaju klasu `status`, a koji su potomci elemenata `tr` koji imaju klasu `read` i postaviti njihov atribut `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

Treba koristiti selektor atributa elementa za ciljanje elemenata `span` sa klasom `status`, koji su potomci elemenata `tr` sa klasom `read` i postaviti njihov atribut `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

Treba vam selektor atributa elemenata koji cilja na elemente `span` sa klasom `status`, a koji su potomci elemenata `tr` sa klasom `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

Treba koristiti selector atributa za ciljanje elemenata `span` sa klasom `status` koji su potomci elemenata `tr` sa klasom `in-progress` i postaviti njihov atribut `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

Treba koristiti selektor atributa elementa da cilja na elemente `span` sa klasom `status` koji su potomci elemenata `tr` sa klasom `in-progress` i postaviti njihov atribut `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

Treba vam selektor atributa elementa da biste ciljali elemente `span` koji imaju klasu `status` i elemente `span` sa vrednošću klase koja počinje od `rate`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

Treba koristiti selektor atributa elementa za ciljanje elemenata `span` sa klasom `status` i elemenata `span` sa vrednošću klase koja počinje sa `rate`, te postavljanjem njihovog atributa `height`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

Trebalo bi da koristite selektor atributa elemenata za ciljanje elemenata `span` koji imaju klasu `status` i elemente `span` koji imaju vrednost klase koja počinje od `rate`, i postavljanjem njihovog atributa `width`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

Trebalo bi da koristite selektor atributa elementa za ciljanje elemenata `span` koji imaju klasu `status` i elemenata `span` koji imaju vrednost klase počevši od `rate`, i postaviti njihov atribut `padding`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

Trebalo bi da imaš selektor atributa elementa koji cilja elemente `span` koji su direktni potomci elemenata `span` sa vrednošću `class`, počevši od `rate`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

Treba da koristite selektor atributa elementa za ciljanje elemenata `span` koji su direktni potomci elemenata `span` sa vrednošću `class` počevši od `rate` i postavljanjem njihovog atributa `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

Trebalo bi da se koristi selektor atributa elementa da cilja na elemente `span` koji su direktni potomci elemenata `span` sa vrednošću od `class` počevši od `rate` i postavljanjem njihovog atributa `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

Treba koristiti selektor atributa elementa za ciljanje elemenata `span` koji su direktni potomci elemenata `span` sa vrednošću `class` počevši od `rate` i postavljanjem njihovog atributa `margin`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

Trebalo bi da koristite selektor atributa elementa za ciljanje elemenata `span` koji su direktni potomci elemenata `span` sa vrednošću `class` počevši od `rate` i postavljanjem njihovog atributa `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

Trebalo bi da koristiš selektor atributa elementa za ciljanje elemenata `span` koji su direktni potomci elemenata `span` sa vrednošću `class` počevši od `rate` i postave njihov atribut `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

Trebalo bi da koristite selektor atributa elementa za ciljanje elemenata `span` koji su direktni potomci elemenata `span` sa vrednošću `class`, počevši od `rate`, i postavite njihov atribut `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

Morate imati selektor svojstva elementa koji cilja prvog roditelja elemenata sa `span`, a koji sadrže reč `one` kao deo njihove vrednosti od `class`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Trebalo bi da koristite selektor atributa elementa za ciljanje prvog roditelja elemenata sa `span` koji sadrže reč `one` kao deo njihove vrednosti `class` i postavite njegov atribut `background-image` koristeći `linear-gradient`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

Trebalo bi da ima selektor svojstva elementa koji cilja na prvih dvoje roditelja elemenata `span`, koji sadrže reč `two` kao deo njihove vrednosti `class`.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Treba koristiti selektor atributa elementa da cilja na dva prvog roditelja elemenata sa `span` koji sadrže reč `two` kao deo njihove vrednosti `class`, i postaviti njihove atribute `background-image` koristeći `linear-gradient`.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

Treba da ima selektor atributa elementa koji cilja elemente sa `span`, koji su potomci elemenata sa `span`, koji sadrže reč `three` kao deo njihove vrednosti `class`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Trebalo bi da koristite selektor atributa elementa za ciljanje elemenata od `span` koji su potomci elemenata od `span` koji sadrže reč `three` kao deo njihove vrednosti od `class` i postaviti njihove atribute od `background-image` koristeći `linear-gradient`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
