---
id: 673b3d6b7ef7318eef926d5a
title: Emocionalna ploča
challengeType: 25
dashedName: build-a-mood-board
demoType: onClick
---

# --description--

Emocionalna tabla je kolekcija slika i teksta koja opisuje opštu ideju, cilj ili osećanje o određenoj temi.

U ovoj laboratoriji, izgradite senzorsku ploču koristeći komponente od React ponovno upotrebljive. CSS je već dat vam.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Unapaswa kuunda na kusafirisha sehemu ya `MoodBoardItem` inayokubali vigezo vitatu vinavyopitishwa: `color`, `image`, na `description`.
2. Sehemu yako ya `MoodBoardItem` inapaswa kurudisha `div` yenye darasa `mood-board-item` kama kipengele chake cha ngazi ya juu.
3. Unapaswa kuweka rangi ya nyuma ya kipengele cha `.mood-board-item` kuwa thamani ya kigezo cha `color` ukitumia mitindo ya ndani.
4. Unapaswa kuonyesha kipengele cha `img`, chenye darasa `mood-board-image` na sifa yake ya `src` kuwekwa kuwa thamani ya kigezo cha `image`, ndani ya kipengele cha `.mood-board-item`.
5. Unapaswa kuonyesha kipengele cha `h3`, chenye darasa `mood-board-text` na maandishi yake kuwa thamani ya kigezo cha `description`, ndani ya kipengele cha `.mood-board-item`.
6. Unapaswa kuunda na kusafirisha `MoodBoard`.
7. Sehemu yako ya `MoodBoard` inapaswa kurudisha `div` kama kipengele chake cha ngazi ya juu.
8. Sehemu yako ya `MoodBoard` inapaswa kuonyesha kipengele cha `h1` chenye darasa `mood-board-heading` na maandishi `Destination Mood Board`.
9. Sehemu yako ya `MoodBoard` inapaswa kuonyesha `div` yenye darasa `mood-board`.
10. Sehemu yako ya `MoodBoard` inapaswa kuonyesha angalau sehemu tatu za `MoodBoardItem` ndani ya kipengele cha `.mood-board`, kila moja ikipitisha vigezo `color`, `image`, na `description` vyenye thamani halali.

Možete koristiti sledeće slike na svojoj tabli raspoloženja ako želite:

- `https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/shore.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/grass.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/ship.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg`

# --hints--

Treba da prebaciš deo `MoodBoardItem`.

```js
assert.isFunction(window.index.MoodBoardItem, 1);
```

Tvoj deo ``MoodBoardItem`` bi trebalo da vrati ``div`` sa klasom ``mood-board-item`` kao njegovu komponentu višeg nivoa.

```js
const item = document.getElementsByClassName('mood-board-item')[0];
assert.exists(item);
assert.equal(item.tagName, 'DIV');
```

Pozadinska boja elementa `.mood-board-item` treba da se postavi na vrednost parametra `color` korišćenjem inline stilova.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { color: "red" });
  const moodBoardItem = container.querySelector(".mood-board-item");
  assert.equal(moodBoardItem.style.backgroundColor, "red");
```

Vaš deo `MoodBoardItem` treba da prikaže element `img` sa klasom `mood-board-image` i njegov atribut `src` postavljen na vrednost parametra `image`.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { image: "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg" });
  const img = container.querySelector(".mood-board-image");
  assert.exists(img);
  assert.equal(img.tagName, "IMG");
  assert.equal(img.src, "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg");
```

Tvoj deo `MoodBoardItem` treba da prikaže element `h3` sa klasom `mood-board-text`, a njegov tekst je postavljen na vrednost parametra/atributa `description`.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { description: "Caribbean" });
  const text = container.querySelector(".mood-board-text");
  assert.exists(text);
  assert.equal(text.tagName, "H3");
  assert.equal(text.textContent, "Caribbean");
```

Treba ti da prebaciš deo `MoodBoard`.

```js
assert.isFunction(window.index.MoodBoard, 1);
```

Vaš deo `MoodBoard` treba da vrati `div` kao svoj element na višoj razini.

```js
const root = document.getElementById('root');
assert.equal(root.children[0].tagName, 'DIV');
```

Tvoja sekcija od `MoodBoard` treba da prikaže element `h1` sa klasom `mood-board-heading` i tekstom `Destination Mood Board`.

```js
const heading = document.getElementsByClassName('mood-board-heading')[0];
assert.exists(heading);
assert.equal(heading.tagName, 'H1');
assert.equal(heading.textContent, 'Destination Mood Board');
```

Vaš deo za `MoodBoard` mora da prikaže najmanje tri dela od `MoodBoardItem`, svaki koji zadovoljava kriterijume definisane/prolazne `color`, `image`, i `description` sa validnim vrednostima.

```js
const moodboard = document.getElementsByClassName('mood-board')[0];
const items = moodboard.querySelectorAll(".mood-board-item");
assert.isAtLeast(items.length, 3);     

items.forEach((item) => {
  const color = item.style.backgroundColor;
  const img = item.querySelector("img");
  const text = item.querySelector("h3");
  
  assert.isAtLeast(color.length, 1);
  assert.isAtLeast(img.src.length, 1);
  assert.isAtLeast(text.textContent.length, 1);
});
```

Vaš deo `MoodBoard` mora biti prikazan u elementu `#root` na stranici.

```js
const root = document.getElementById('root');
const moodBoard = root.getElementsByClassName('mood-board')[0];
assert.exists(moodBoard);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx

```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx
export function MoodBoardItem(props) {
  return (
    <div className='mood-board-item' style={{ backgroundColor: props.color }}>
      <img src={props.image} alt='Mood' className='mood-board-image' />
      <h3 className='mood-board-text'>{props.description}</h3>
    </div>
  );
}

export function MoodBoard() {
  return (
    <div>
      <h1 className='mood-board-heading'>Destination Mood Board</h1>
      <div className='mood-board'>
        <MoodBoardItem
          color='#2da64f'
          image='https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg'
          description='Caribbean'
        />
        <MoodBoardItem
          color='#8e44ad'
          image='https://cdn.freecodecamp.org/curriculum/labs/shore.jpg'
          description='Gawadar Beach'
        />
        <MoodBoardItem
          color='#3498db'
          image='https://cdn.freecodecamp.org/curriculum/labs/grass.jpg'
          description='Cape Town'
        />
        <MoodBoardItem
          color='#bf3d7e'
          image='https://cdn.freecodecamp.org/curriculum/labs/ship.jpg'
          description='Suez Canal'
        />
        <MoodBoardItem
          color='#e74c3c'
          image='https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg'
          description='Santorini'
        />
        <MoodBoardItem
          color='#95a5a6'
          image='https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg'
          description='Istanbul'
        />
      </div>
    </div>
  );
}
```
