---
id: 66db57ad34c7089b9b41bfd6
title: Izgradi pregledač za lightbox
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

Lightbox se koristi na veb stranicama za prikaz multimedijalnog sadržaja na privlačan način putem korišćenja popap-up prozora ili modala preko glavnog sadržaja stranice.

U ovom laboratoriju, kreiraćeš lightbox biblioteku koja prikazuje slike u punoj veličini kada se klikne na minijaturu. Za svaku sliku, izdaju se dve verzije: minijatura i slika u punoj veličini. Slika u punoj veličini je isto što i minijatura, ali bez sufiksa `-thumbnail`.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratorijum.

Priče o korisniku:

1. Trebalo bi da imate `div` sa klasom `gallery` unutar vašeg `body`.
1. Unutar elementa `.gallery`, trebalo bi da imate tri minijature slika, svaka sa klasom `gallery-item`. Trebalo bi da koristite sledeće linkove za minijature slika:
   
    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`

1. Treba da imaš `div` sa klasom `lightbox` unutar svog `body` elementa.
1. Treba da imaš `button` sa atributom `id` vrednosti `close-btn` unutar elementa `.lightbox`. Ako želiš, možeš koristiti `&times;` kao njegov tekst.
1. Treba da imaš `img` sa atributom `id` vrednosti `lightbox-image` unutar elementa `.lightbox`.
1. Tvoj element `.lightbox` treba da ima fiksnu poziciju kako bi se prikazivao iznad postojećih slika.
1. Tvoj element `.lightbox` treba da prekriva celu vidljivu oblast tako što će imati visinu i širinu od 100% kontejnera. Takođe treba da obezbediš da element `.lightbox` počinje u gornjem levom uglu kontejnera.
1. Element `.lightbox` treba da ima boju pozadine. Njegovo svojstvo `display` u početku treba da bude postavljeno na `none` kako bi bio sakriven.
1. Kada klikneš na bilo koji od svojih elemenata `.gallery-item`, svojstvo `display` elementa `.lightbox` treba da bude postavljeno na `flex`, kako bi element `.lightbox` i njegova dva unutrašnja elementa postali vidljivi.
1. Kada klikneš na bilo koji od svojih elemenata `.gallery-item`, atribut `src` elementa `#lightbox-image` treba da bude postavljen na verziju slike u punoj veličini tako što će se ukloniti `-thumbnail` iz atributa `src` slike. Slike u punoj veličini nalaze se na sledećim adresama:

- `https://cdn.freecodecamp.org/curriculum/labs/stonehenge.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees.jpg`
  
1. Vaš element za `.lightbox` kada je vidljiv i kada kliknete na element `#close-btn` ili `.lightbox`, svojstvo `.lightbox` elementa `display` bi trebalo da se resetuje na `none`.

**Napomena:** Uverite se da ste povezali svoj stilski fajl sa fajlom JavaScript u vašem HTML.

# --hints--

Trebalo bi da imaš `div` sa klasom `gallery` unutar svog elementa `body`.

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

Unutar elementa `.gallery`, mora imati tri elementa od `img` koji imaju klasu `gallery-item`.

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

Unutar elementa `.gallery`, treba da ima element `img` sa svojstvom elementa `src` postavljeno na `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Unutar elementa za `.gallery`, treba da postoji element `img` koji ima atribut elementa `src` postavljen na `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Unutar elementa `.gallery`, mora imati element `img` sa atributom elementa `src` postavljenog na `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Trebalo bi da ima element sa `div` koji ima klasu `lightbox` unutar vašeg elementa `body`.

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

Unutar vašeg elementa od `.lightbox`, treba da ima komponentu od `button` sa atributom elementa `id` postavljen na `close-btn`.

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

Unutar elementa za `.lightbox`, mora imati element `img` sa atributom elementa `id` postavljen na `lightbox-image`.

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

Vaša komponenta `.lightbox` treba da ima trajno mesto.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

Vaš element od `.lightbox` treba da pokrije celokupno vidljivo područje.

```js
const style = new __helpers.CSSHelp(document).getStyle(".lightbox");

const width = style?.width;
const height = style?.height;

assert.oneOf(
  width,
  ["100%", "100vw"]
);

assert.oneOf(
  height,
  ["100%", "100vh"]
);
```

Tvoja komponenta za `.lightbox` treba da se poveže sa gornjim levim uglom kontejnera.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

Tvoj element za `.lightbox` bi trebalo da ima pozadinsku boju.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

Vaša funkcija za `.lightbox` treba da bude inicijalno sakrivena.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

Kada kliknete na jedan od vaših komponenti sa `.gallery-item`, svojstvo `.lightbox` elementa `display` mora biti postavljeno na `flex` kako bi se učinilo da su `.lightbox` i dva elementa unutar njega vidljiva.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

Kada kliknete na jedan od vaših elemenata za `.gallery-item`, svojstvo `#lightbox-image` elementa `src` treba postaviti na verziju slike u punoj veličini uklanjanjem `-thumbnail` iz svojstva slike elementa `src`.

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src; 
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

Vaš element `.lightbox` kada je vidljiv i kliknete na dugme `#close-btn`, svojstvo `.lightbox` elementa `display` treba da se vrati na `none`.

```js
const lightbox = document.querySelector(".lightbox");
const closeBtn = document.getElementById("close-btn"); 
const galleryItem = document.querySelector(".gallery-item");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

closeBtn.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

Vaš element `.lightbox` je vidljiv i kada kliknete na element `.lightbox`, svojstvo `.lightbox` elementa `display` mora biti vraćeno na `none`.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

lightbox.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
    <link rel="stylesheet" href="styles.css">
  </head>

  <body>
    <h1 class="title">Museum Gallery</h1>

    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
    </div>
    <div id="lightbox" class="lightbox">
      <button id="close-btn" class="close-btn">&times;</button>
      <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);  
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; 
}

function closeLightbox() {
  lightbox.style.display = "none"; 
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox(); 
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", "")); 
  });
});
```
