---
id: 69899123889dcf5b19be1847
title: Implementiraj algoritam dubinskog pretraživanja prvo
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

U ovoj laboratoriji, implementiraćete algoritam za pretraživanje grafa koji se zove <dfn>Pretraživanje u dubinu</dfn> (DFS).

Dok <dfn>široko-prvi pretraga</dfn> (BFS) istražuje susjedima korak po korak, <dfn>duboko-prvi pretraga</dfn> ulazi što je moguće dublje u jednu putanju konekcija prije povratka.

Kada algoritam stigne do kraja putanje (čvora bez neistraženih suseda), on se vraća nazad do najbližeg čvora koji još uvek ima neistražene veze i nastavlja pretragu odatle.

Da biste implementirali ovaj algoritam cikličnim putem, biće vam potrebno da koristite **stack**. U JavaScript, uobičajeno `Array` se dobro koristi kao stack koristeći metode `.push()` i `.pop()`, prateći princip **Last-In-First-Out (LIFO)**. Ovo osigurava da je sused koji je najnovije dodan na stack prvi obrađivan.

Jednostavan rezultat ovog algoritma je lista čvorova koji su dostupni od određenog čvora. Stoga, takođe ćete morati da pratite čvorove koje posećujete.

**Cilj**: Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Treba da definiš funkciju nazvanu `dfs`.
1. Funkcija `dfs` treba da primi dva argumenta:

- `graph`: Matrica susjednosti (matrica redova podataka).
    - `root`: Etiketa broja čvora (početni indeks) koja je vrijednost broja čvora između `0` i `n - 1`, gdje je `n` ukupan broj čvorova u grafu.

1. Funkcija mora vratiti listu svih oznaka dostižnih čvorova od početnog čvora.

# --hints--

Trebalo bi da imaš funkciju nazvanu `dfs` koja prima dva argumenta.

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` treba da vrati listu koja sadrži `1`, `2`, `3`, na `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` trebalo bi da vrati listu koja sadrži `1`, `2`, `3`, i na `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` treba da vrati `[3]`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` trebalo bi da vrati listu koja sadrži `3` i `2`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` trebalo bi da vrati listu koja sadrži `0` i `1`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
