---
id: 68e43f94365a964f26c88f78
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Sada je vreme da radimo sa novom metodom sekvenciranja karaktera.

Počnite kreiranjem novog parametra pod nazivom `exampleSentence` i postavite mu vrednost stringa/niza karaktera `"I love cats and cats are so much fun!"`.

Zatim ispod tog parametra, napišite u konzolu `"Original sentence:"`. Zatim ispod toga `console.log()`, dodajte još jedan sa `console.log()` koji ima parametar `exampleSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da imaš kriterijum `exampleSentence`.

```js
assert.exists(exampleSentence);
```

Vaš parametar `exampleSentence` mora biti sekvenca slova.

```js
assert.isString(exampleSentence);
```

Trebalo bi da postavite vrednost niza karaktera `"I love cats and cats are so much fun!"` za parametar `exampleSentence`.

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

Trebalo bi da napišete `"Original sentence:"` u konzolu.

```js
assert.equal(getLogs()[4], "Original sentence:");
```

Trebalo bi da upišete parametar `exampleSentence` u konzolu.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
