---
id: 68e4465fba0bd256868b7deb
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Sada je vreme za rad sa novom metodom sekvenciranja slova/karaktera.

Počnite kreiranjem parametra nazvanog `learningSentence` i postavite vrednost na niz znakova `"I love learning!"`.

Zatim ispod tog parametra, napiši u konzolu `"Original learning sentence:"`. Ispod `console.log()`, dodaj još jedan `console.log()` za parametar `learningSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da imate kriterijum `learningSentence`.

```js
assert.exists(learningSentence);
```

Vaš parametar `learningSentence` mora biti niz slova.

```js
assert.isString(learningSentence);
```

Trebalo bi da unesete vrednost niza karaktera `"I love learning!"` za parametar `learningSentence`.

```js
assert.equal(learningSentence, "I love learning!");
```

Treba da napišete `"Original learning sentence:"` u konzoli.

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

Treba da upišete parametar za `learningSentence` u konzolu.

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
