---
id: 68e448ec4ac4205968c0a805
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Biće dobro da se reč `"love"` ponovi tri puta u rečenici `"I love learning!"`.

Kao što vas pamtićete iz prethodne lekcije, možete ponoviti sekvencu slova određeni broj puta koristeći metod `repeat()`. Evo primera:

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

Pošto string znakova ne može biti modifikovan, ova metoda neće promeniti originalni niz znakova. Vraćaće novi niz znakova sa ponovljenim sadržajem.

Kreiraj varijablu nazvanu `repeatedLove` i postavi vrednost na `"love ".repeat(3)`. Zatim napišite u konzolu varijablu `repeatedLove` da biste videli rezultat.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da imate parametar koji se zove `repeatedLove`.

```js
assert.exists(repeatedLove);
```

Vaš parametar `repeatedLove` treba da bude niz slova.

```js
assert.isString(repeatedLove);
```

Trebalo bi da postaviš vrednost `"love ".repeat(3)` u parametar `repeatedLove`.

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

Trebalo bi da napišete parametar `repeatedLove` u konzoli.

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
