---
id: 68e43e884ef09d4d2ab05631
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada je vreme da vidimo rezultate korišćenjem metode `replace()`.

Počnite dodavanjem `console.log()` sa vrednošću `"After using the replace() method:"`.

Zatim, ispod `console.log()`, dodaj drugi `console.log()` koji ima varijablu `replacedString`.

Pogledajte u konzolu i trebalo bi da vidite da nova rečenica kaže `"I love dogs."` umesto `"I love cats."` prethodnu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da napišete `"After using the replace() method:"` u konzoli.

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

Morate napisati varijablu `replacedString` u konzoli.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
