---
id: 68e44abea678a65b877864dc
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Biće dobro da vidimo tu reč ponovljenu u rečenici. Kreiraj novi parametar nazvan `newSentence` i postavi mu vrednost rezultata spajanja reči `"I "` sa parametrom `repeatedLove`, a zatim spaja niz znakova `" learning."`. Zatim ispiši vrednost parametra `newSentence` u konzolu. Možete izabrati da koristite ili literalan `+` ili znakove naslova za spajanje niza znakova.

Sada bi trebalo da vidiš tekst `"I love love love  learning."`.

*MOLIM VAS*: Postoji dodatni prostor nakon zadnjeg validnog `"love"`. Na poslednjem koraku, uklonićete taj dodatni prostor.

# --hints--

Trebalo bi da imaš parametar koji se zove `newSentence`.

```js
assert.exists(newSentence);
```

Tvoj parametar za `newSentence` mora biti niz slova.

```js
assert.isString(newSentence);
```

Treba da spojiš reč `"I "` sa parametrom `repeatedLove`, a zatim i niz znakova `" learning."`. Ovi rezultati treba da se dodeli vrednosti parametra `newSentence`. Obrati pažnju na razmake.

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

Trebalo bi da postavite niz znakova `"I love love love  learning."` kao vrednost parametra `newSentence`. Molimo vas da ponovo pročitate uputstva na ispravan način.

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
