---
id: 68e4417b7317835208145323
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Sledeće, kreirajte novi parametar nazvan `dogsOnlySentence` i postavite vrednost `exampleSentence.replace("cats", "dogs")`.

Zatim ispod tog parametra, napišite u konzolu `"Replacing all occurrences of cats with dogs:"`. Ispod `console.log()`, dodajte još jedan `console.log()` za parametar `dogsOnlySentence`.

Trebalo bi da vidiš tekst `"I love dogs and cats are so much fun!"` napisan u konzoli. Međutim, to nije rezultat koji želiš. U sledećem koraku, naučićeš kako da popraviš da tekst izgleda kao `"I love dogs and dogs are so much fun!"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da imate kriterijum za `dogsOnlySentence`.

```js
assert.exists(dogsOnlySentence);
```

Vaš parametar za `dogsOnlySentence` mora biti sekvenca slova.

```js
assert.isString(dogsOnlySentence);
```

Trebalo bi da postavite vrednost `exampleSentence.replace("cats", "dogs")` za vaš parametar `dogsOnlySentence`.

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

Treba da napišete `"Replacing all occurrences of cats with dogs:"` u konzoli.

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

Trebalo bi da napišete parametar `dogsOnlySentence` u konzoli.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
