---
id: 68e42de6396f5c49d90d1951
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

U prethodnoj lekciji, naučili ste o metodi `replace()` koja se koristi za pronalaženje malog dela niza karaktera unutar niza karaktera i zamenjivanje ga drugom vrednošću.

Evo primera:

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

Imajte na umu da sekvenca karaktera ne može biti promenjena, što znači da originalna sekvenca karaktera `text` nije promenjena u tom primeru. Umesto toga, kreira se nova sekvenca karaktera.

Kreiraj novi parametar nazvan `replacedString` i postavi vrednost `originalString.replace("cats", "dogs")` unutar njega.

# --hints--

Trebalo bi da imaš parametar nazvan `replacedString`.

```js
assert.exists(replacedString);
```

Vaš parametar `replacedString` mora biti niz slova.

```js
assert.isString(replacedString);
```

Treba da postaviš vrednost `originalString.replace("cats", "dogs")` u parametar `replacedString`.

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
