---
id: 68e42b62d7093345c6975f22
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Ispod vaše varijable za `originalString`, dodaj `console.log()` sa nizom slova `"Original string:"`.

Zatim ispod tog `console.log()`, dodajte drugi `console.log()` koji ima varijablu od `originalString`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Trebalo bi da ukucate niz karaktera `"Original string:"` u konzolu.

```js
assert.equal(getLogs()[0], "Original string:");
```

Treba da napišete varijablu `originalString` u konzolu.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
