---
id: 68e44519c769bf545bb3f22a
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Da biste dobili rezultate koje želite za `"I love dogs and dogs are so much fun!"`, moraćete koristiti metodu `replaceAll()` umesto metode `replace()`. Ova metoda vraća novu sekvencu karaktera za svako poklapanje podstringa.

Ispravite `exampleSentence.replace("cats", "dogs")` ili koristite putanju od `replaceAll()` umesto toga. Sada bi trebalo da vidite ispravan tekst u konzoli.

# --hints--

Trebalo bi da ispravite liniju koda `exampleSentence.replace("cats", "dogs")` ili koristite putanju `replaceAll()`.

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
