---
id: 68e41084bb275539f72de9fe
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, navik ćete se da radite sa metodama `replace()`, `replaceAll()` i `repeat()` za izradu programa koji menja sekvencu karaktera.

Da biste počeli, definišite parametar nazvan `originalString` i postavite mu vrednost `"I love cats."`.

# --hints--

Trebalo bi da imate parametar nazvan `originalString`.

```js
assert.exists(originalString);
```

Vaš parametar za `originalString` mora biti niz slova.

```js
assert.isString(originalString);
```

Treba da uneseš niz karaktera `"I love cats."` u parametar `originalString`.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
