---
id: 587d7dac367417b2b2512b73
title: Kreiraj osnovni JavaScript objekat
challengeType: 1
forumTopicId: 301317
dashedName: create-a-basic-javascript-object
---

# --description--

Zamislite stvari koje ljudi vide svakog dana, kao što su automobili, prodavnice i ptice. Hivi ni <dfn>stvari</dfn>: stvari koje ljudi vide i sa kojima mogu da interaguju.

Koje su osobine ovih stvari? Auto ima točkove. Prodavnice prodaju robu. Ptice imaju krila.

Ove osobine, ili <dfn>sifa</dfn>, opisuju ono što čini objekat. Zapamtite da slični objekti imaju iste osobine, ali mogu imati različite vrednosti tih osobina. Na primer, svi automobili imaju točkove, ali ne i svi automobili imaju isti broj točkova.

Predmeti u JavaScript se koriste za simulaciju stvarnih stvari iz sveta, dajući im atribute i ponašanje kao stvarni predmeti. Evo primera koji koristi ove koncepte za kreiranje nečega kao što je `duck`:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
```

Ovaj predmet za `duck` ima dva para atribut/vrednost: `name` od `Aflac` i `numLegs` sa 2.

# --instructions--

Kreiraj stavku od `dog` sa svojstvima za `name` i `numLegs`, sa vrednošću alfanumerične sekvence, kontrolisan.

# --hints--

`dog` Trebalo bi da bude nešto.

```js
assert(typeof dog === 'object');
```

Objekat `dog` treba da ima svojstvo `name` postavljeno na string.

```js
assert(typeof dog.name === 'string');
```

`dog` trebalo bi imati svojstvo `numLegs` postavljeno kao broj.

```js
assert(typeof dog.numLegs === 'number');
```

# --seed--

## --seed-contents--

```js
let dog = {

};
```

# --solutions--

```js
let dog = {
  name: '',
  numLegs: 4
};
```
