---
id: 587d7dae367417b2b2512b7b
title: Razumite lične parametre
challengeType: 1
forumTopicId: 301326
dashedName: understand-own-properties
---

# --description--

U sledećem primeru, komponenta `Bird` definiše dva parametra: `name` i `numLegs`:

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let duck = new Bird("Donald");
let canary = new Bird("Tweety");
```

`name` i `numLegs` nazivaju se <dfn>privatni parametri</dfn>, jer su definisani direktno na modelu objekta. To znači da svaki od `duck` i `canary` ima sopstvenu kopiju ovih parametara. Zapravo, svaka instanca `Bird` će imati svoju kopiju ovih parametara. Sledeći kod dodaje sve privatne parametre od `duck` u red podataka `ownProps`:

```js
let ownProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  }
}

console.log(ownProps);
```

Konsola će prikazati vrednost `["name", "numLegs"]`.

# --instructions--

Dodajte lične parametre za `canary` u red podataka `ownProps`.

# --hints--

`ownProps` trebalo bi da uključuje vrednosti za `numLegs` i `name`.

```js
assert(ownProps.indexOf('name') !== -1 && ownProps.indexOf('numLegs') !== -1);
```

Trebalo bi da rešiš ovaj izazov bez korišćenja metode izgrađene `Object.keys()`.

```js
assert(!/Object(\.keys|\[(['"`])keys\2\])/.test(__helpers.removeJSComments(code)));
```

Ovaj izazov treba da rešiš bez direktnog upisivanja niza `ownProps`.

```js
assert(
  !/\[\s*(?:'|")(?:name|numLegs)|(?:push|concat)\(\s*(?:'|")(?:name|numLegs)/.test(
    __helpers.removeJSComments(code)
  )
);
```

# --seed--

## --seed-contents--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
let ownProps = [];
// Only change code below this line
```

# --solutions--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
function getOwnProps (obj) {
  const props = [];

  for (let prop in obj) {
    if (obj.hasOwnProperty(prop)) {
      props.push(prop);
    }
  }

  return props;
}

const ownProps = getOwnProps(canary);
```
