---
id: 587d7db1367417b2b2512b87
title: Dodaj metode nakon nasleđivanja
challengeType: 1
forumTopicId: 301315
dashedName: add-methods-after-inheritance
---

# --description--

Funkcija konstruktora koja nasleđuje `prototype` od nadtipa može da ima i sopstvene metode, pored onih koje je nasledila.

Na primer, `Bird` je konstruktor koji nasleđuje `prototype` od konstruktora `Animal`:

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
Bird.prototype.constructor = Bird;
```

Pored svega što nasleđuje od `Animal`, želiš da dodaš ponašanje koje je jedinstveno za objekte `Bird`. Ovde će `Bird` dobiti metodu `fly()`. Metoda se dodaje u `Bird.prototype` na isti način kao i kod bilo kog drugog konstruktora:

```js
Bird.prototype.fly = function() {
  console.log("I'm flying!");
};
```

Sasa matoleo ya `Bird` yatakuwa na njia za `eat()` na `fly()`:

```js
let duck = new Bird();
duck.eat();
duck.fly();
```

`duck.eat()` će ispisati string `nom nom nom` u konzoli, a `duck.fly()` će ispisati string `I'm flying!`.

# --instructions--

Dodaj sav potreban kod kako bi objekat `Dog` nasledio `Animal`, a konstruktor `Dog` imao svoj `prototype` pravilno postavljen. Zatim dodaj metodu `bark()` objektu `Dog` kako bi `beagle` mogao da koristi metode `eat()` i `bark()`. Metoda `bark()` treba da ispiše `Woof!` u konzoli.

# --hints--

`Animal` ne treba da ima metodu `bark()`.

```js
assert(typeof Animal.prototype.bark == 'undefined');
```

`Dog` treba da nasledi metodu `eat()` od objekta `Animal`.

```js
assert(typeof Dog.prototype.eat == 'function');
```

Prototip za `Dog` mora imati način za `bark()`.

```js
assert('bark' in Dog.prototype);
```

`beagle` treba da bude instanca (`instanceof`) objekta `Animal`.

```js
assert(beagle instanceof Animal);
```

Polje `beagle` treba da se postavi na `Dog`.

```js
assert(beagle.constructor === Dog);
```

`beagle.eat()` mora napisati sekvencu slova `nom nom nom`

```js
const spy = __helpers.spyOn(console, 'log');
beagle.eat();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'nom nom nom');
```

`beagle.bark()` treba da napiše niz slova `Woof!`

```js
const spy = __helpers.spyOn(console, 'log');
beagle.bark();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'Woof!');
```

# --seed--

## --seed-contents--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }

// Only change code below this line




// Only change code above this line

let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Dog.prototype.bark = function () {
  console.log('Woof!');
};
let beagle = new Dog();

beagle.eat();
beagle.bark();
```
