---
id: 587d7dad367417b2b2512b78
title: Koristi komponentu za kreiranje stvari
challengeType: 1
forumTopicId: 18233
dashedName: use-a-constructor-to-create-objects
---

# --description--

Evo komponente `Bird` iz prethodnog izazova:

```js
function Bird() {
  this.name = "Albert";
  this.color  = "blue";
  this.numLegs = 2;
}

let blueBird = new Bird();
```

**NAPOMENA:** `this` unutar komponente uvek se odnosi na stvar koja se kreira.

Smatrajte da operater `new` se koristi prilikom pozivanja komponente. Ovo govori JavaScript da kreira novu instancu `Bird` nazvanu `blueBird`. Bez operatera `new`, `this` unutar komponente ne bi se referisao na novokreirani objekat, što bi dovelo do nepredvidivih rezultata. Sada `blueBird` ima sve deklarisane parametre unutar komponente `Bird`:

```js
blueBird.name;
blueBird.color;
blueBird.numLegs;
```

Kao i svakom drugom objektu, njegovim svojstvima može se pristupiti i mogu se menjati:

```js
blueBird.name = 'Elvira';
blueBird.name;
```

# --instructions--

Koristi konstruktor `Dog` iz prethodne lekcije da napraviš novu instancu `Dog` i sačuvaš je u promenljivoj `hound`.

# --hints--

`hound` treba se kreirati koristeći komponentu od `Dog`.

```js
assert(hound instanceof Dog);
```

Vaš kod bi trebalo da koristi operater `new` za kreiranje instance od `Dog`.

```js
assert(__helpers.removeJSComments(code).match(/new/g));
```

# --seed--

## --seed-contents--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
// Only change code below this line
```

# --solutions--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
const hound = new Dog();
```
