---
id: 587d7db1367417b2b2512b86
title: Resetuj nasleđeno svojstvo konstruktora
challengeType: 1
forumTopicId: 301324
dashedName: reset-an-inherited-constructor-property
---

# --description--

Kada objekat nasledi `prototype` od drugog objekta, on takođe nasleđuje svojstvo `constructor` njegovog nadtipa.

Evo primera:

```js
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
let duck = new Bird();
duck.constructor
```

Međutim, `duck` i sve instance objekta `Bird` treba da pokazuju da su kreirane pomoću konstruktora `Bird`, a ne `Animal`. To možeš postići tako što ćeš ručno postaviti svojstvo `constructor` objekta `Bird` na sam konstruktor `Bird`:

```js
Bird.prototype.constructor = Bird;
duck.constructor
```

# --instructions--

Izmeni kod tako da `duck.constructor` i `beagle.constructor` vraćaju svoje odgovarajuće konstruktore.

# --hints--

`Bird.prototype` treba da bude instanca objekta `Animal`.

```js
assert(Animal.prototype.isPrototypeOf(Bird.prototype));
```

`duck.constructor` treba da vrati `Bird`.

```js
assert(duck.constructor === Bird);
```

`Dog.prototype` treba da bude instanca objekta `Animal`.

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

`beagle.constructor` inapaswa kurudisha `Dog`.

```js
assert(beagle.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Animal() { }
function Bird() { }
function Dog() { }

Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);

// Only change code below this line



let duck = new Bird();
let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
function Bird() { }
function Dog() { }
Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Bird.prototype.constructor = Bird;
let duck = new Bird();
let beagle = new Dog();
```
