---
id: 587d7daf367417b2b2512b7d
title: Pregledajte sve kriterijume
challengeType: 1
forumTopicId: 301320
dashedName: iterate-over-all-properties
---

# --description--

Sada ste videli dve vrste parametara: <dfn>sopstveni parametri</dfn> i parametre za `prototype`. Sopstveni parametri su definisani direktno u primeru tog objekta. I parametri prototipa su definisani na `prototype`.

```js
function Bird(name) {
  this.name = name;  //own property
}

Bird.prototype.numLegs = 2; // prototype property

let duck = new Bird("Donald");
```

Ovo je kako dodati sopstvenih parametara za `duck` u red podataka `ownProps` i parametre za `prototype` u red podataka `prototypeProps`:

```js
let ownProps = [];
let prototypeProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  } else {
    prototypeProps.push(property);
  }
}

console.log(ownProps);
console.log(prototypeProps);
```

`console.log(ownProps)` će prikazati `["name"]` na konzoli, i `console.log(prototypeProps)` će prikazati `["numLegs"]`.

# --instructions--

Dodaj sve sopstvene parametre za `beagle` u red podataka `ownProps`. Dodaj sve parametre za `prototype` od `Dog` u red podataka `prototypeProps`.

# --hints--

Set podataka `ownProps` treba da sadrži samo `name`.

```js
assert.deepEqual(ownProps, ['name']);
```

Podatni red `prototypeProps` treba da ima samo `numLegs`.

```js
assert.deepEqual(prototypeProps, ['numLegs']);
```

Trebalo bi da rešiš ovaj izazov bez korišćenja predgrađenog načina `Object.keys()`.

```js
assert(!/\Object.keys/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];
for (let prop in beagle) {
  if (beagle.hasOwnProperty(prop)) {
    ownProps.push(prop);
  } else {
    prototypeProps.push(prop);
  }
}
```
