---
id: 587d7db1367417b2b2512b88
title: Ghairi njia zilizopokelewa
challengeType: 1
forumTopicId: 301322
dashedName: override-inherited-methods
---

# --description--

U prethodnim lekcijama, naučili ste da nešto može nasleđivati svoje karakteristike (metodu) iz drugog objekta referenciranjem njegovog `prototype`:

```js
ChildObject.prototype = Object.create(ParentObject.prototype);
```

Zatim, `ChildObject` je primio sopstvene puteve za povezivanje sa svojim `prototype`:

```js
ChildObject.prototype.methodName = function() {...};
```

Moguće je zameniti nasleđenu metodu. To se radi na isti način – dodavanjem metode u `ChildObject.prototype` sa istim imenom kao metoda koju želiš da zameniš. Evo primera u kojem `Bird` zamenjuje metodu `eat()` nasleđenu od `Animal`:

```js
function Animal() { }
Animal.prototype.eat = function() {
  return "nom nom nom";
};
function Bird() { }

Bird.prototype = Object.create(Animal.prototype);

Bird.prototype.eat = function() {
  return "peck peck peck";
};
```

Ako imaš instancu `let duck = new Bird();` i pozoveš `duck.eat()`, JavaScript će ovako tražiti metodu u lancu `prototype` objekta `duck`:

1.  `duck` => Je, `eat()` imefafanuliwa hapa? Hapana.
2.  `Bird` => Je, `eat()` imefafanuliwa hapa? => Ndiyo. Iite na acha kutafuta.
3. `Animal` => `eat()` je takođe definisana, ali je JavaScript prekinuo pretragu pre nego što je stigao do ovog nivoa.
4. `Object` => JavaScript je prekinuo pretragu pre nego što je stigao do ovog nivoa.

# --instructions--

Zameni metodu `fly()` za `Penguin` tako da vraća string `Alas, this is a flightless bird.`

# --hints--

`penguin.fly()` treba da vrati niz slova `Alas, this is a flightless bird.`

```js
assert(penguin.fly() === 'Alas, this is a flightless bird.');
```

Metoda `bird.fly()` treba da vrati string `I am flying!`

```js
assert(new Bird().fly() === 'I am flying!');
```

# --seed--

## --seed-contents--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;

// Only change code below this line



// Only change code above this line

let penguin = new Penguin();
console.log(penguin.fly());
```

# --solutions--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;
Penguin.prototype.fly = () => 'Alas, this is a flightless bird.';
let penguin = new Penguin();
console.log(penguin.fly());
```
