---
id: 587d7daf367417b2b2512b7f
title: Promeni početni primer u nešto novo
challengeType: 1
forumTopicId: 301316
dashedName: change-the-prototype-to-a-new-object
---

# --description--

Sada ste dodavali karakteristike za `prototype` jednu po jednu:

```js
Bird.prototype.numLegs = 2;
```

Ovo postaje umorno nakon samo nekoliko pohvala.

```js
Bird.prototype.eat = function() {
  console.log("nom nom nom");
}

Bird.prototype.describe = function() {
  console.log("My name is " + this.name);
}
```

Najbolji način je da se `prototype` postavi kao nešto novo što već ima te karakteristike. Na ovaj način, karakteristike se dodaju zajedno odjednom:

```js
Bird.prototype = {
  numLegs: 2, 
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --instructions--

Dodaj funkciju `numLegs` na dva načina `eat()` i `describe()` u `prototype` od `Dog` postavljanjem `prototype` da bude nova stavka.

# --hints--

`Dog.prototype` Trebalo bi ga postaviti kao nešto novo.

```js
assert(/Dog\.prototype\s*?=\s*?{/.test(__helpers.removeJSComments(code)));
```

`Dog.prototype` trebalo bi da ima karakteristiku `numLegs`.

```js
assert(Dog.prototype.numLegs !== undefined);
```

`Dog.prototype` Trebalo bi da ima putanju `eat()`.

```js
assert(typeof Dog.prototype.eat === 'function');
```

`Dog.prototype` treba da ima putanju `describe()`.

```js
assert(typeof Dog.prototype.describe === 'function');
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype = {
  // Only change code below this line

};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
numLegs: 4,
  eat () {
    console.log('nom nom nom');
  },
  describe () {
    console.log('My name is ' + this.name);
  }
};
```
