---
id: 587d7dae367417b2b2512b7a
title: Proveri tip objekta koristeći instanceof
challengeType: 1
forumTopicId: 301337
dashedName: verify-an-objects-constructor-with-instanceof
---

# --description--

Svaki put kada komponentna funkcija kreira nešto novo, to se naziva <dfn>mfano</dfn> njegove komponente. JavaScript nudi jednostavan način da se ovo potvrdi korišćenjem izraza `instanceof`. `instanceof` omogućava vam da uporedite element sa komponentom, vraćajući `true` ili `false` zavisno od toga da li je taj element kreiran tom komponentom ili ne. Evo primera:

```js
let Bird = function(name, color) {
  this.name = name;
  this.color = color;
  this.numLegs = 2;
}

let crow = new Bird("Alexis", "black");

crow instanceof Bird;
```

Ova njia za `instanceof` će vratiti `true`.

Ako se nešto kreira bez korišćenja komponente, `instanceof` to će potvrditi da nije primer te komponente:

```js
let canary = {
  name: "Mildred",
  color: "Yellow",
  numLegs: 2
};

canary instanceof Bird;
```

Njia hii ya `instanceof` itarudisha `false`.

# --instructions--

Napravi novu instancu konstruktora `House` pod nazivom `myHouse` i prosledi broj spavaćih soba. Zatim koristi operator `instanceof` da potvrdiš da je to instanca objekta `House`.

# --hints--

Objekat `myHouse` treba da ima svojstvo `numBedrooms` postavljeno na broj.

```js
assert(typeof myHouse.numBedrooms === 'number');
```

Trebalo bi da potvrdite da je `myHouse` primer za `House` koristeći operator `instanceof`.

```js
assert(/myHouse\s*instanceof\s*House/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}

// Only change code below this line
```

# --solutions--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}
const myHouse = new House(4);
console.log(myHouse instanceof House);
```
