---
id: 587d7db1367417b2b2512b85
title: Postavi model roditelja kao model deteta
challengeType: 1
forumTopicId: 301325
dashedName: set-the-childs-prototype-to-an-instance-of-the-parent
---

# --description--

U prethodnom izazovu otkrio si prvi korak u nasleđivanju ponašanja od nadtipa (ili roditelja) `Animal`: kreiranje nove instance objekta `Animal`.

Ovaj izazov se tiče sledećeg koraka: postavi `prototype` podtipa (ili deteta)—u ovom slučaju, `Bird`—da bude primer za `Animal`.

```js
Bird.prototype = Object.create(Animal.prototype);
```

Zapamti da je `prototype` kao "recept" za kreiranje nečega. Na neki način, recept od `Bird` sada uključuje sve ključne sastojke iz `Animal`.

```js
let duck = new Bird("Donald");
duck.eat();
```

`duck` nasleđuje sva svojstva objekta `Animal`, uključujući metodu `eat`.

# --instructions--

Promenite kod ili primere iz `Dog` nasleđeni iz `Animal`.

# --hints--

`Dog.prototype` bi trebalo da bude primer od `Animal`.

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }

// Only change code below this line


let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);

let beagle = new Dog();
beagle.eat();
```
