---
id: 587d7daf367417b2b2512b80
title: Zapamtite da dodelite atribut komponenti kada menjate model
challengeType: 1
forumTopicId: 301323
dashedName: remember-to-set-the-constructor-property-when-changing-the-prototype
---

# --description--

Postoji jedan važan nuspojavak ručnog postavljanja modela na nešto novo. Briše svojstvo `constructor`! Ovo svojstvo se može koristiti za potvrđivanje koja je komponentna funkcija kreirala model, ali pošto je ovo svojstvo prepisano, sada daje netačne rezultate:

```js
duck.constructor === Bird;
duck.constructor === Object;
duck instanceof Bird;
```

Redosledom, ova fraza bi bila smatrana kao `false`, `true`, i `true`.

Kada ovo menjate, šablon će biti ručno postavljen kao novi element, zapamtite da definišete svojstvo za `constructor`:

```js
Bird.prototype = {
  constructor: Bird,
  numLegs: 2,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name); 
  }
};
```

# --instructions--

Definišite svojstvo `constructor` na `Dog` `prototype`.

# --hints--

`Dog.prototype` treba da postavi svojstvo od `constructor`.

```js
assert(Dog.prototype.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
Dog.prototype = {

  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
  constructor: Dog,
  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```
