---
id: 587d7db0367417b2b2512b82
title: Razumite lanac prototipa
challengeType: 1
forumTopicId: 301329
dashedName: understand-the-prototype-chain
---

# --description--

Sve stvari u JavaScript (uključujući neke retke slučajeve) imaju `prototype`. Takođe, `prototype` je stvar.

```js
function Bird(name) {
  this.name = name;
}

typeof Bird.prototype;
```

Zato što `prototype` je nešto, `prototype` može imati svoje `prototype`! U ovom slučaju, `prototype` od `Bird.prototype` je `Object.prototype`:

```js
Object.prototype.isPrototypeOf(Bird.prototype);
```

Koliko je ovo važno? Možda se sećaš načina za `hasOwnProperty` iz prethodnog izazova:

```js
let duck = new Bird("Donald");
duck.hasOwnProperty("name");
```

Putanja za `hasOwnProperty` je definisana u `Object.prototype`, koja se može postići sa `Bird.prototype`, što se zatim može postići sa `duck`. Ovo je primer lanca od `prototype`. U ovom lancu od `prototype`, `Bird` je `supertype` za `duck`, dok je `duck` `subtype`. `Object` je `supertype` za `Bird` i `duck`. `Object` je `supertype` za sve stavke u JavaScript. Stoga, svaki element može koristiti putanju `hasOwnProperty`.

# --instructions--

Promenite kod ili prikažite ispravan lanac prototipa.

# --hints--

Vaš kod bi trebalo da pokaže da `Object.prototype` je prototip za `Dog.prototype`

```js
assert(/Object\.prototype\.isPrototypeOf/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

Dog.prototype.isPrototypeOf(beagle);  // yields true

// Fix the code below so that it evaluates to true
???.isPrototypeOf(Dog.prototype);
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
Object.prototype.isPrototypeOf(Dog.prototype);
```
