---
id: 587d7db2367417b2b2512b8b
title: Razumite šta je anonimna funkcija koja se izvršava samo jednom (IIFE)
challengeType: 1
forumTopicId: 301328
dashedName: understand-the-immediately-invoked-function-expression-iife
---

# --description--

Standardno je u JavaScript izvršiti funkciju čim se objavi:

```js
(function () {
  console.log("Chirp, chirp!");
})();
```

Ovo je poziv anonimne funkcije koja se izvršava jednom, i vraća `Chirp, chirp!` jednom.

Zapamti da funkcija nema ime i nije sačuvana ni u jednoj promenljivoj. Dve obične zagrade `()` na kraju deklaracije funkcije uzrokuju da se ona odmah izvrši. Ovaj obrazac je poznat kao <dfn>odmah pozvani izraz funkcije</dfn> ili <dfn>IIFE</dfn>.

# --instructions--

Ponovo napišite funkciju `makeNest` i uklonite njen poziv, ili ga umesto toga učinite anonimnom funkcijom koja se odmah izvršava (IIFE).

# --hints--

Funkcija mora biti anonimna.

```js
assert(/\((function|\(\))(=>|\(\)){?/.test(__helpers.removeJSComments(code).replace(/\s/g, '')));
```

Funkcija bi trebalo da ima standardne zagrade na kraju izjave kako bi se pozvala odmah.

```js
assert(/\(.*(\)\(|\}\(\))\)/.test(__helpers.removeJSComments(code).replace(/[\s;]/g, '')));
```

# --seed--

## --seed-contents--

```js
function makeNest() {
  console.log("A cozy nest is ready");
}

makeNest();
```

# --solutions--

```js
(function () {
  console.log("A cozy nest is ready");
})();
```

---

```js
(function () {
  console.log("A cozy nest is ready");
}());
```

---

```js
(() => {
  console.log("A cozy nest is ready");
})();
```

---

```js
(() =>
  console.log("A cozy nest is ready")
)();
```
