---
id: 587d7dad367417b2b2512b75
title: Napravi put na nečemu
challengeType: 1
forumTopicId: 301318
dashedName: create-a-method-on-an-object
---

# --description--

Stvari mogu imati poseban tip karakteristika, nazvan <dfn>njia</dfn>.

Metoda je svojstvo koje je funkcionalno. Ovo dodaje različito ponašanje objektu. Evo primera za `duck` sa metodama:

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + duck.name + ".";}
};
duck.sayName();
```

Ovaj primer dodaje metodu `sayName`, koja vraća rečenicu sa imenom objekta `duck`. Obrati pažnju da metoda koristi `duck.name` u naredbi `return`. Sledeći izazov će pokazati još jedan način da se to uradi.

# --instructions--

Ako zavisiš od `dog`, daj put koji se zove `sayLegs`. Taj put bi trebalo da vrati rečenicu `This dog has 4 legs.`

# --hints--

`dog.sayLegs()` Trebalo bi da bude funkcija.

```js
assert(typeof dog.sayLegs === 'function');
```

`dog.sayLegs()` Trebalo bi da vrati sekvencu karaktera koja je data – zapamtite da su interpunkcijski znaci i razmaci važni.

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,

};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
