---
id: 587d7dae367417b2b2512b7c
title: Koristi atribute za prototip kako bi smanjio ponavljajući kod
challengeType: 1
forumTopicId: 301336
dashedName: use-prototype-properties-to-reduce-duplicate-code
---

# --description--

Pošto `numLegs` može imati istu vrednost za sve instance od `Bird`, stoga ima promenljivu `numLegs` koja se ponavlja u svakoj instanci od `Bird`.

Ovo verovatno nije problem gde ima samo dva događaja, ali zamislite da postoji miliona događaja. To će biti mnogo ponavljajućih varijabli.

Najbolji način je korišćenje `prototype` od `Bird`. Atributi u `prototype` se dele između SVIH događaja `Bird`. Evo kako dodati `numLegs` na `Bird prototype`:

```js
Bird.prototype.numLegs = 2;
```

Sve događaje od `Bird` sada poseduju karakteristiku `numLegs`.

```js
console.log(duck.numLegs);
console.log(canary.numLegs);
```

Pošto sve događaje imaju te karakteristike direktno u `prototype`, smatrajte `prototype` kao "recept" za kreiranje stvari. Zapamtite da je `prototype` za `duck` i `canary` deo komponente `Bird` kao `Bird.prototype`.

# --instructions--

Dodaj svojstvo `numLegs` u `prototype` objekta `Dog`.

# --hints--

Objekat `beagle` treba da ima svojstvo `numLegs`.

```js
assert(beagle.numLegs !== undefined);
```

`beagle.numLegs` Trebalo bi da bude broj.

```js
assert(typeof beagle.numLegs === 'number');
```

`numLegs` treba da bude svojstvo `prototype`, a nije sopstveno svojstvo.

```js
assert(beagle.hasOwnProperty('numLegs') === false);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}



// Only change code above this line
let beagle = new Dog("Snoopy");
```

# --solutions--

```js
function Dog (name) {
  this.name = name;
}
Dog.prototype.numLegs = 4;
let beagle = new Dog("Snoopy");
```
