---
id: 587d7dac367417b2b2512b74
title: Koristi tačke pisanja da dobiješ atribute objekta
challengeType: 1
forumTopicId: 301333
dashedName: use-dot-notation-to-access-the-properties-of-an-object
---

# --description--

Poslednji izazov stvorio je nešto sa različitim karakteristikama. Sada ćete videti kako da dobijete vrednosti tih karakteristika. Evo primera:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
console.log(duck.name);
```

Notacija sa tačkom se koristi za ime objekta, `duck`, praćeno imenom svojstva, `name`, da bi se dobila vrednost `Aflac`.

# --instructions--

Ispišite oba svojstva objekta `dog` u vašoj konzoli.

# --hints--

Vaš kod bi trebalo da koristi `console.log` za štampanje vrednosti svojstva `name` objekta `dog`.

```js
assert(/console.log\(.*dog\.name.*\)/g.test(__helpers.removeJSComments(code)));
```

Tvoj kod treba da koristi `console.log` za štampanje vrednosti svojstva `numLegs` objekta `dog`.

```js
assert(/console.log\(.*dog\.numLegs.*\)/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
// Only change code below this line
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
console.log(dog.name);
console.log(dog.numLegs);
```
