---
id: 587d7dad367417b2b2512b76
title: Neka kod bude ponovo upotrebljiv korišćenjem ključne reči this
challengeType: 1
forumTopicId: 301321
dashedName: make-code-more-reusable-with-the-this-keyword
---

# --description--

Poslednji izazov je predstavio način pristupa svojstvima objekta `duck`. Koristio je tačkastu notaciju `duck.name` da bi dobio vrednost svojstva `name` u naredbi `return`:

```js
sayName: function() {return "The name of this duck is " + duck.name + ".";}
```

Iako je ovo validan način za dobijanje svojstva objekta, postoji rizik ovde. Ako se ime svojstva promeni, bilo koji kod koji referencira originalno ime takođe će zahtevati korekciju. U kratkoj definiciji objekta, to nije problem, ali ako objekat ima mnogo referenci na svoja svojstva, velika je verovatnoća greške.

Način za izbegavanje ovih problema je korišćenjem glavne reči `this`:

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + this.name + ".";}
};
```

`this` je široka tema, a primer naveden gore je samo jedan način za korišćenje toga. U ovom kontekstu, `this` se odnosi na nešto sa čime je ovaj način povezan: `duck`. Ako se ime neke stvari promeni u `mallard`, nije potrebno tražiti sve reference za `duck` u kodu. Ovo čini kod ponovo upotreivim i lakšim za čitanje.

# --instructions--

Promeni putanju `dog.sayLegs` da ukloni sve reference na `dog`. Koristi primer `duck` kao vodič.

# --hints--

`dog.sayLegs()` Trebalo bi da vrati sekvencu slova koja je data.

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

Vaš kod bi trebalo da koristi glavni termin od `this` da dohvati atribut `numLegs` iz `dog`.

```js
assert(__helpers.removeJSComments(code).match(/this\.numLegs/g));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs: function() {return "This dog has " + dog.numLegs + " legs.";}
};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
