---
id: 587d7db0367417b2b2512b83
title: Tumia urithi ili usijirudie mwenyewe
challengeType: 1
forumTopicId: 301334
dashedName: use-inheritance-so-you-dont-repeat-yourself
---

# --description--

U programiranju postoji princip poznat kao <dfn>Ne ponavljaj se (DRY)</dfn>. Problem sa ponavljanjem koda je što svaka promena zahteva izmenu na više mesta. To obično znači više posla za programere i veću mogućnost za greške.

Pogledaj u sledećem primeru kako metodu `describe` dele i `Bird` i `Dog`:

```js
Bird.prototype = {
  constructor: Bird,
  describe: function() {
    console.log("My name is " + this.name);
  }
};

Dog.prototype = {
  constructor: Dog,
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

Putanja za `describe` se ponavlja na dva mesta. Kod se može izmeniti prateći princip DRY kreiranjem `supertype` (ili roditeljski) nazvanog `Animal`:

```js
function Animal() { };

Animal.prototype = {
  constructor: Animal, 
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

Pošto `Animal` uključuje metod `describe`, možete ga ukloniti iz `Bird` i `Dog`:

```js
Bird.prototype = {
  constructor: Bird
};

Dog.prototype = {
  constructor: Dog
};
```

# --instructions--

Putanja za `eat` se ponavlja u `Cat` i `Bear`. Refaktoriši kod duhom DRY premeštanjem putanje `eat` na `Animal` `supertype`.

# --hints--

`Animal.prototype` treba da ima svojstvo `eat`.

```js
assert(Animal.prototype.hasOwnProperty('eat'));
```

`Bear.prototype` ne bi trebalo da ima karakteristiku od `eat`.

```js
assert(!Bear.prototype.hasOwnProperty('eat'));
```

`Cat.prototype` ne smije imati karakteristiku od `eat`.

```js
assert(!Cat.prototype.hasOwnProperty('eat'));
```

# --seed--

## --seed-contents--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,

};
```

# --solutions--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
```
