---
id: 587d7daf367417b2b2512b7e
title: Razumite karakteristiku komponenta
challengeType: 1
forumTopicId: 301327
dashedName: understand-the-constructor-property
---

# --description--

Kuna posebna karakteristika za `constructor` koja se nalazi u događajima objekata `duck` i `beagle` koji su kreirani u prethodnim izazovima:

```js
let duck = new Bird();
let beagle = new Dog();

console.log(duck.constructor === Bird); 
console.log(beagle.constructor === Dog);
```

Oba poziva od `console.log` će prikazati `true` u konzoli.

Zapamtite da je svojstvo `constructor` referenca na funkciju komponente koja je kreirala taj događaj. Prednost svojstva `constructor` je što je moguće proveriti ovo svojstvo kako bi se znalo kakva vrsta objekta je reč o njemu. Evo primera kako se ovo može koristiti:

```js
function joinBirdFraternity(candidate) {
  if (candidate.constructor === Bird) {
    return true;
  } else {
    return false;
  }
}
```

**Napomena:** Pošto svojstvo `constructor` se može prepisati (što će biti obtrađeno u sledećih dva izazova), generalno je bolje koristiti metodu `instanceof` za proveru tipa objekta.

# --instructions--

Napiši funkciju za `joinDogFraternity` koja prima argument od `candidate` i, koristeći svojstvo `constructor`, vrati `true` ako je kandidat `Dog`, inače vrati `false`.

# --hints--

`joinDogFraternity` Trebalo bi da se definiše kao funkcija.

```js
assert(typeof joinDogFraternity === 'function');
```

`joinDogFraternity` treba da vrati `true` ako `candidate` je događaj od `Dog`.

```js
assert(joinDogFraternity(new Dog('')) === true);
```

`joinDogFraternity` mora da koristi svojstvo od `constructor`.

```js
assert(/\.constructor/.test(__helpers.removeJSComments(code)) && !/instanceof/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
function joinDogFraternity(candidate) {

}
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
function joinDogFraternity(candidate) {
  return candidate.constructor === Dog;
}
```
