---
id: 587d7db0367417b2b2512b81
title: Razumi/Shvati odakle `prototype` nečega potiče
challengeType: 1
forumTopicId: 301330
dashedName: understand-where-an-objects-prototype-comes-from
---

# --description--

Kao što ljudi dobijaju gener od svojih roditelja, objekat dobija svoj `prototype` direktno od komponente funkcije koja ga je kreirala. Na primer, ovde komponenta `Bird` stvara objekat od `duck`:

```js
function Bird(name) {
  this.name = name;
}

let duck = new Bird("Donald");
```

`duck` dobija svoj `prototype` iz funkcije komponente `Bird`. Možete prikazati ovu vezu putem `isPrototypeOf`:

```js
Bird.prototype.isPrototypeOf(duck);
```

Ovo će vratiti `true`.

# --instructions--

Koristi `isPrototypeOf` da vidiš `prototype` od `beagle`.

# --hints--

Trebalo bi da pokažeš da je `Dog.prototype` deo `prototype` od `beagle`

```js
assert(/Dog\.prototype\.isPrototypeOf\(beagle\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
```
