---
id: 587d7dad367417b2b2512b77
title: Objasni funkciju komponente
challengeType: 1
forumTopicId: 16804
dashedName: define-a-constructor-function
---

# --description--

<dfn>Konstruktor</dfn> je funkcija koja stvara nove stvari. Kreira svojstva i ponašanja koja će pripadati novom objektu. Zamislite ga kao nacrt za kreiranje novih stvari.

Evo primera komponente:

```js
function Bird() {
  this.name = "Albert";
  this.color = "blue";
  this.numLegs = 2;
}
```

Ovaj komponent kreira objekat ``Bird`` sa atributima ``name``, ``color`` i ``numLegs``, postavljeni na Albert, blue i 2, respektivno. Komponenta prati nekoliko ugovora:

<ul><li>Komponente se pišu imenom koje počinje velikim slovom kako bi se razlikovale od drugih funkcija koje nisu <code>vijenzi</code>.</li><li>Komponente koriste ključnu reč <code>this</code> da bi definisale atribute nečega što će kreirati. Unutar komponente, <code>this</code> se odnosi na nešto novo što će biti kreirano.</li><li>Komponente definišu atribute i ponašanje umesto da vraćaju vrednost kao što to mogu druge funkcije.</li></ul>

# --instructions--

Napravi konstruktor `Dog` sa svojstvima `name`, `color` i `numLegs`, postavljenim redom na string, string i broj.

# --hints--

`Dog` treba da ima svojstvo `name` postavljeno na string.

```js
assert(typeof new Dog().name === 'string');
```

`Dog` mora imati svojstvo `color` postavljeno kao niz karaktera.

```js
assert(typeof new Dog().color === 'string');
```

`Dog` treba da ima svojstvo `numLegs` postavljeno na broj.

```js
assert(typeof new Dog().numLegs === 'number');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function Dog (name, color, numLegs) {
  this.name = 'name';
  this.color = 'color';
  this.numLegs = 4;
}
```
