---
id: 587d7db2367417b2b2512b8c
title: Koristi IIFE za kreiranje modula
challengeType: 1
forumTopicId: 301332
dashedName: use-an-iife-to-create-a-module
---

# --description--

Izraz funkcije koji se odmah izvršava (IIFE) često se koristi za grupisovanje povezanih funkcionalnosti u jednom objektu ili <dfn>moduli</dfn>. Na primer, prvobitni problem identifikovao je dve mixin komponente:

```js
function glideMixin(obj) {
  obj.glide = function() {
    console.log("Gliding on the water");
  };
}
function flyMixin(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  };
}
```

Možemo spojiti ove mixine u modul kao što sledi:

```js
let motionModule = (function () {
  return {
    glideMixin: function(obj) {
      obj.glide = function() {
        console.log("Gliding on the water");
      };
    },
    flyMixin: function(obj) {
      obj.fly = function() {
        console.log("Flying, wooosh!");
      };
    }
  }
})();
```

Zapamtite da imate izraz funkcije koji se zove odmah pozvana izjava funkcije (IIFE) koji vraća objekat `motionModule`. Ovaj vraćeni objekat ima sve karakteristike miksina kao svojstva tog objekta. Prednost modularne strukture je što sve karakteristike pokreta mogu biti pakovane u jedan objekat koji se zatim može koristiti od strane drugih delova vašeg koda. Evo primera kako ga koristiti:

```js
motionModule.glideMixin(duck);
duck.glide();
```

# --instructions--

Kreiraj modul nazvan `funModule` koji pokriva dva mixina `isCuteMixin` i `singMixin`. `funModule` trebalo bi da vrati nešto.

# --hints--

`funModule` Trebalo bi da bude definisano i vrati nešto.

```js
assert(typeof funModule === 'object');
```

`funModule.isCuteMixin` Treba da ima funkcionalni pristup.

```js
assert(typeof funModule.isCuteMixin === 'function');
```

`funModule.singMixin` trebalo bi da ima funkcionalni pristup.

```js
assert(typeof funModule.singMixin === 'function');
```

# --seed--

## --seed-contents--

```js
let isCuteMixin = function(obj) {
  obj.isCute = function() {
    return true;
  };
};
let singMixin = function(obj) {
  obj.sing = function() {
    console.log("Singing to an awesome tune");
  };
};
```

# --solutions--

```js
const funModule = (function () {
  return {
    isCuteMixin: obj => {
      obj.isCute = () => true;
    },
    singMixin: obj => {
      obj.sing = () => console.log("Singing to an awesome tune");
    }
  };
})();
```
