---
id: 587d7db2367417b2b2512b89
title: Koristi mixin da dodaš zajednička svojstva između nepovezanih entiteta
challengeType: 1
forumTopicId: 301331
dashedName: use-a-mixin-to-add-common-behavior-between-unrelated-objects
---

# --description--

Kao što ste videli, karakteristike se nasleđuju putem nasleđivanja. Međutim, postoje slučajevi gde nasleđivanje nije najbolje rešenje. Nasleđivanje se ne uklapa dobro za stvari koje nemaju veze kao što su `Bird` i `Airplane`. Sve može da preskoči, ali `Bird` nije tip od `Airplane` i obrnuto.

Za nezavisne stvari, bolje je koristiti <dfn>mixins</dfn>. Mixin omogućava drugim objektima da koriste kolekciju funkcija.

```js
let flyMixin = function(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  }
};
```

`flyMixin` uzima bilo šta i pruža put do `fly`.

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let plane = {
  model: "777",
  numPassengers: 524
};

flyMixin(bird);
flyMixin(plane);
```

Ovde `bird` i `plane` se propušta unutar `flyMixin`, što zatim postavlja funkcionalnost od `fly` za sve. Sada `bird` i `plane` svi mogu da preskoče:

```js
bird.fly();
plane.fly();
```

Konsola prikazuje sekvencu karaktera `Flying, wooosh!` dva puta, jednom za svaki poziv `.fly()`.

Pogledajte kako mixin dozvoljava da se ista metoda za `fly` ponovo koristi i neusklađenim/nepovezanim `bird` i `plane`.

# --instructions--

Napravi mixin pod nazivom `glideMixin` koji deklarše metodu pod nazivom `glide`. Zatim koristi `glideMixin` da mu/njoj pruži `bird` i `boat` mogućnost pregledavanja.

# --hints--

Vaš kod mora da deklariše varijablu tipa `glideMixin`, koja je funkcija.

```js
assert(typeof glideMixin === 'function');
```

Vaš kod bi trebalo da koristi `glideMixin` na objektu od `bird` da pruži put ka `glide`.

```js
assert(typeof bird.glide === 'function');
```

Vaš kod treba da koristi `glideMixin` u objektu od `boat` da pruži put ka `glide`.

```js
assert(typeof boat.glide === 'function');
```

# --seed--

## --seed-contents--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};

// Only change code below this line
```

# --solutions--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};
function glideMixin (obj) {
  obj.glide = () => 'Gliding!';
}

glideMixin(bird);
glideMixin(boat);
```
