---
id: 69708d59833ceded7d00541f
title: Izvršite operacije za povezanu listu
challengeType: 26
dashedName: implement-linked-list-operations
---

# --description--

U ovoj laboratoriji, implementiraćete dodatne operacije za strukturu podataka spojenog popisa, gradeći na osnovnoj funkcionalnosti dodavanja i uklanjanja čvorova.

**Cilj**: Popuni priče korisnika ispod i prođi sve testove kako bi se završio laboratorijum.

**Priče korisnika:**

1. Mora da postoji metoda za `contains` koja prima linkovanu listu i element. Trebalo bi vratiti `true` ako je pomenuti element u linkovanoj listi, i `false` inače.
1. Mora da postoji metoda za `getAt` koja prima linkovanu listu i šablon. Trebalo bi vratiti element koji je na šablonu datom u linkovanoj listi. Ako je šablon van granica, treba da vrati `undefined`.
1. Mora da postoji metoda za `insertAt` koja prima linkovanu listu, šablon i element. Trebalo bi umetnuti element dat u poziciju navedenu u linkovanoj listi. Ako je šablon van granica, neće menjati listu.
1. Mora da postoji metoda za `removeAt` koja prima linkovanu listu i šablon. Trebalo bi ukloniti čvor koji je na šablonu datom u linkovanoj listi. Ako je šablon van granica, neće menjati listu.
1. Mora da postoji metoda za `clear` koja prima linkovanu listu. Trebalo bi ukloniti sve elemente iz linkovane liste, efikasno je vraćajući u prazan stanja.

**Napomena:** Neki budući testovi zavise od prethodnih metoda. Na primer, ako `getAt` nije implementiran ispravno, funkcionalni testovi kao što su `insertAt` i `removeAt` mogu da propadnu čak i ako te funkcije budu skoro ispravne.

# --hints--

`initList` treba da vrati stavku sa `head` postavljenu na `null` i `length` postavljenu na `0`.

```js
const list = initList();
assert.isNull(list.head);
assert.strictEqual(list.length, 0);
```

`add` Trebalo bi da poveća dužinu liste za jedan.

```js
const list = initList();
add(list, 'apple');
assert.strictEqual(list.length, 1);
```

`remove` Trebalo bi da smanji dužinu liste za jedan.

```js
const list = initList();  
add(list, 'apple');  
add(list, 'banana');  
remove(list, 'apple');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(list.head.element, 'banana');  
assert.isNull(list.head.next);
```

Trebalo bi da imaš funkciju `contains`.

```js
assert.isFunction(contains);
```

`contains` trebalo bi da vrati `true` ako je element u listi.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
add(list, 'cherry');
assert.isTrue(contains(list, 'banana'));
```

`contains` treba da vrati `false` ako element nije u listi.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
assert.isFalse(contains(list, 'cherry'));
```

`contains` treba da vrati `false` na prazan spisak.

```js
const list = initList();
assert.isFalse(contains(list, 'apple'));
```

Trebalo bi da ima funkciju `getAt`.

```js
assert.isFunction(getAt);
```

`getAt` Trebalo bi da vrati element koji je u priloženom modelu/šemi.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
```

`getAt` treba da vrati `undefined` za nevažeće parametre.

```js
const list = initList();  
add(list, 'apple');  
assert.isUndefined(getAt(list, -1));  
assert.isUndefined(getAt(list, 1));  
assert.isUndefined(getAt(list, 5));
```

`getAt` treba da vrati `undefined` na praznu listu.

```js
const list = initList();  
assert.isUndefined(getAt(list, 0));
```

Trebalo bi da imaš funkciju `insertAt`.

```js
assert.isFunction(insertAt);
```

`insertAt` Element treba da se ubaci na početak liste.

```js
const list = initList();  
add(list, 'second');  
add(list, 'third');  
insertAt(list, 0, 'first');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` Trebalo bi umetnuti element na sredini liste.

```js
const list = initList();  
add(list, 'first');  
add(list, 'third');  
insertAt(list, 1, 'second');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` Treba umetnuti element na kraju liste.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
insertAt(list, 2, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt` treba da se unese u prazan spisak na šablonu od `0`.

```js  
const list = initList();  
insertAt(list, 0, 'first');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.isNull(list.head.next); 
```

`insertAt` Ne bi trebalo da menja listu sa nevažećim podacima.

```js
const list = initList();  
add(list, 'apple');  
insertAt(list, -1, 'invalid');  
insertAt(list, 10, 'invalid');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
assert.isUndefined(getAt(list, 1));
```

Trebalo bi da imate funkciju `removeAt`.

```js
assert.isFunction(removeAt);
```

`removeAt` Trebalo bi da ukloni element iz pomenutog modela.

```js
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'third');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` Trebalo bi da ukloni prvi element tačno.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
removeAt(list, 0);
assert.strictEqual(getAt(list, 0), 'second');
assert.strictEqual(list.length, 1);
assert.isNull(list.head.next);
```

`removeAt` trebalo bi da ukloni poslednji element tačno.

```js  
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 2);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` Trebalo bi da obnovi listu prilikom uklanjanja samo jednog elementa.

```js
const list = initList();  
add(list, 'only');  
removeAt(list, 0);  
assert.strictEqual(list.length, 0);  
assert.isNull(list.head);  
assert.isTrue(isEmpty(list));
```

`removeAt` Ne smete da promenite listu sa nevažećim podacima.

```js
const list = initList();  
add(list, 'apple');  
removeAt(list, -1);  
removeAt(list, 10);  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
```

Trebalo bi da imate funkciju `clear`.

```js
assert.isFunction(clear);
```

`clear` Trebalo bi da ukloni sve elemente iz liste.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
clear(list);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

`clear` Trebalo bi da omogući ponovnu upotrebu liste.

```js  
const list = initList();  
add(list, 'first');  
clear(list);  
clear(list);  
add(list, 'second');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'second');  
assert.isNull(list.head.next); 
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  
}

function getAt(list, index) {
 
}

function insertAt(list, index, element) {
  
}

function removeAt(list, index) {
 
}

function clear(list) {

}
```

# --solutions--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  let current = list.head;
  while (current !== null) {
    if (current.element === element) {
      return true;
    }
    current = current.next;
  }
  return false;
}

function getAt(list, index) {
  if (index < 0 || index >= list.length) return undefined;
  let current = list.head;
  let i = 0;
  while (i < index) {
    current = current.next;
    i++;
  }
  return current.element;
}

function insertAt(list, index, element) {
  if (index < 0 || index > list.length) return;
  const node = { element, next: null };
  if (index === 0) {
    node.next = list.head;
    list.head = node;
    list.length++;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = node;
  node.next = current;
  list.length++;
}

function removeAt(list, index) {
  if (index < 0 || index >= list.length) return;
  if (index === 0) {
    list.head = list.head.next;
    list.length--;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = current.next;
  list.length--;
}

function clear(list) {
  list.head = null;
  list.length = 0;
}
```
