---
id: 66ea8239b9eeb76f12299030
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sledeći deo radionice je pregled kako generisati nasumičan broj između dve vrednosti. 

Počnite dodavanjem još jednog `console.log()` koji loguje poruku `"Now, generate a random number between two values."` na konzolu.

Zatim, kreirajte varijablu nazvanu `min` i dodelite joj vrednost `1`, pa kreirajte varijablu nazvanu `max` i dodelite joj vrednost `100`. 

U sledećem koraku, generisaćete nasumičan broj između ovih dve vrednosti.

# --hints--

You should have a `console.log()` that logs the message `"Now, generate a random number between two values."`.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

You should have a variable called `min`.

```js
assert.isNotNull(min);
```

Your `min` variable should be a number.

```js
assert.isNumber(min);
```

You should assign the number `1` to your `min` variable.

```js
assert.equal(min, 1);
```

You should have a variable called `max`.

```js
assert.isNotNull(max);
```

Your `max` variable should be a number.

```js
assert.isNumber(max);
```

You should assign the number `100` to your `max` variable.

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
