---
id: 66ea7fc048a6e26e33c598fe
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Sada je vreme da generišete nasumičan broj koristeći metod `Math.random()`.

Kreirajte varijablu nazvanu `randomNum` i dodelite joj vrednost rezultata pozivanja metode `Math.random()`.

Zatim, logujte varijablu `randomNum` na konzolu.

Pokušajte dodati razmak u kod da ponovo pokrenete bota i vidite različite nasumične brojeve koji se generišu.

# --hints--

You should have a variable called `randomNum`.

```js
assert.isNotNull(randomNum);
```

You should assign the result of calling the `Math.random()` method to the `randomNum` variable.

```js
assert.match(code, /randomNum\s*=\s*Math\.random\(\);?/);
```

You should not hardcode a random number for the `randomNum` variable. Make sure to use the `Math.random()` method to generate the random number.

```js
assert.notMatch(code, /randomNum\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum\s*=\s*\d+;?/);
```

You should log the `randomNum` variable to the console.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
