---
id: 66ea845f54195c7020cc521b
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Formula za generisanje nasumičnog broja između dve vrednosti je sledeća:

```js
Math.random() * (maximum - minimum) + minimum;
```

Ovo će proizvesti rezultat koji je broj sa decimalnim zarezom (floating point) između dve vrednosti.

Kreirajte varijablu nazvanu `randomNum2` i dodelite joj rezultat generisanja vrednosti između vrednosti `min` i `max`.

Zatim, logujte varijablu `randomNum2` da biste videli rezultat. Pokušajte osvežiti stranicu kako biste videli različite rezultate.

# --hints--

You should have a variable called `randomNum2`.

```js
assert.isNotNull(randomNum2);
```

Your `randomNum2` variable should be a number.

```js
assert.isNumber(randomNum2);
```

You should assign the result of generating a value between the `min` and `max` values. Make sure not to hardcode a value. You should use the formula provided above.

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

You should log the `randomNum2` variable to the console.

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
