---
id: 66eb488c0129e192542f0865
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

U prethodnim lekcijama, naučili ste kako da radite sa metodom `Math.ceil()` na ovaj način:

```js
const price = 10.01;
Math.ceil(price); // 11
```

Metoda `Math.ceil()` zaokružuje vrednost naviše do najbližeg celog broja.

Kreirajte varijablu nazvanu `numRoundedUp` i dodelite joj rezultat zaokruživanja broja sa plutajućom zarezom `3.2` naviše do najbližeg celog broja.

Zatim, logujte varijablu `numRoundedUp` u konzolu da biste videli rezultat.

# --hints--

You should have a `numRoundedUp` variable.

```js
assert.isNotNull(numRoundedUp);
```

The `numRoundedUp` variable should be a number.

```js
assert.isNumber(numRoundedUp);
```

Your `numRoundedUp` variable should equal the result of rounding the floating point number `3.2` down to the nearest whole integer.

```js
assert.equal(numRoundedUp, 4);
```

You should not hardcode the value of `4` for the `numRoundedUp` variable.

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

You should log the `numRoundedUp` variable to the console.

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
