---
id: 69218e46f4490b33f2481a1a
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Formula za generisanje nasumičnog celog broja između dve vrednosti je sledeća:

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

Ovo će proizvesti rezultat koji je ceo broj između dve vrednosti.

Kreirajte varijablu nazvanu `randomInt` i dodelite joj rezultat generisanja vrednosti između vrednosti `min` i `max`.

Zatim, logujte varijablu `randomInt` da biste videli rezultat. Pokušajte osvežiti stranicu kako biste videli različite rezultate.

# --hints--

You should have a variable called `randomInt`.

```js
assert.isNotNull(randomInt);
```

Your `randomInt` variable should be a number.

```js
assert.isNumber(randomInt);
```

You should assign the result of generating a value between the `min` and `max` values. Make sure not to hardcode a value. You should use the formula provided above.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

You should log the `randomInt` variable to the console.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
