---
id: 69218a15fe536ee1be38aefb
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Sledeći deo radionice je pregled kako generisati nasumičan ceo broj između dve vrednosti.

Počnite dodavanjem još jednog `console.log()` koji loguje poruku `"Now, generate a random integer between two values."` na konzolu.

U sledećem koraku, generisaćete nasumičan ceo broj između vrednosti promenljivih `min` i `max` koje ste već kreirali.

# --hints--

You should have a `console.log()` that logs the message `"Now, generate a random integer between two values."`.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+integer\s+between\s+two\s+values\.\s*\1\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

--fcc-editable-region--

--fcc-editable-region--
```
