---
id: 587d7fad367417b2b2512be1
title: Obradi klik događaje korišćenjem JavaScript-a koristeći atribut `onclick`
challengeType: 6
forumTopicId: 301503
dashedName: handle-click-events-with-javascript-using-the-onclick-property
---

# --description--

Želite da se vaš kod izvrši čim stranica završi učitavanje. Za tu svrhu, možete priložiti događaj (event) `JavaScript` na dokument nazvan ``DOMContentLoaded``. Ovo je kod koji to radi:

```js
document.addEventListener('DOMContentLoaded', function () {});
```

Možete implementirati rukovaoce događajima koji ulaze u funkciju `DOMContentLoaded`. Možete implementirati rukovaoc događaja za `onclick` koji se aktivira kada korisnik klikne na element `#getMessage`, dodavanjem sledećeg koda:

```js
document.getElementById('getMessage').onclick = function () {};
```

# --instructions--

Dodaj rukovaoca događaja klika unutar funkcije za `DOMContentLoaded` za element sa ID-om `getMessage`.

# --hints--

Tvoj kod bi trebalo da koristi metodu `document.getElementById` za odabir elementa sa ID-om `getMessage`.

```js
assert.match(code, /document\s*\.getElementById\(\s*('|")getMessage\1\s*\)/g);
```

Vaš kod bi trebalo da doda rukovaoca događaja za `onclick`.

```js
assert.isFunction(document.getElementById('getMessage').onclick);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    document.getElementById('getMessage').onclick = function () {};
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
