---
id: 587d7faf367417b2b2512be9
title: Pošalji podatke putem JavaScript XMLHttpRequest-a
challengeType: 6
forumTopicId: 301504
dashedName: post-data-with-the-javascript-xmlhttprequest-method
---

# --description--

U prethodnim primerima, primili ste podatke sa eksternog resursa. Takođe možete poslati podatke na eksternom resursu, pod uslovom da taj resurs podržava AJAX zahteve i znate URL.

Putanja za JavaScript na `XMLHttpRequest` takođe se koristi za slanje podataka na server. Evo primera:

```js
const xhr = new XMLHttpRequest();
xhr.open('POST', url, true);
xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
xhr.onreadystatechange = function () {
  if (xhr.readyState === 4 && xhr.status === 201) {
    const serverResponse = JSON.parse(xhr.response);
    document.getElementsByClassName('message')[0].textContent =
      serverResponse.userName + serverResponse.suffix;
  }
};
const body = JSON.stringify({ userName: userName, suffix: ' loves cats!' });
xhr.send(body);
```

Vi ste već videli neke od ovih metoda. Ovde metoda `open` inicijalizuje zahtev kao `POST` za dostupan eksterni resurs URL, i prosleđuje `true` kao treći parametar - što ukazuje da se ta operacija izvršava asinhrono.

Metoda `setRequestHeader` postavlja vrednost zaglavlja zahteva za HTTP, koje sadrži informacije o pošiljaocu i zahtevu. Mora se pozvati nakon metode `open`, ali pre metode `send`. Dva parametra su ime zaglavlja i vrednost za postavljanje kao telo tog zaglavlja.

Sledeći, handler događaja za `onreadystatechange` obrađuje promene statusa zahteva. `readyState` od `4` znači da je operacija završena, a `status` od `201` znači da je zahtev uspešan. Stoga, HTML dokumenta se može ažurirati.

Konačno, putanja `send` šalje zahtev sa vrednošću `body`. `body` uključuje `userName` i funkciju `suffix`.

# --instructions--

Ispravite kod da napravite zahtev za `POST` koristeći krajnju tačku (endpoint) API. Zatim unesite svoje ime u polje za unos i kliknite na `Send Message`. Vaša AJAX funkcija bi trebalo da promeni `Reply from Server will be here.` sa podacima sa servera. Formatirajte odgovor da prikažete vaše ime uz tekst od ` loves cats`.

# --hints--

Vaš kod treba da kreira `XMLHttpRequest` novi.

```js
assert.match(code, /new\s*XMLHttpRequest\(\s*\)/g);
```

Vaš kod bi trebalo da koristi metodu `open` za iniciranje zahteva `POST` na serveru.

```js
assert.match(code, /\.open\(\s*('|")POST\1\s*,\s*url\s*,\s*true\s*\)/g);
```

Tvoj kod bi trebalo da koristi metodu `setRequestHeader`.

```js
assert.match(
  code,
  /\.setRequestHeader\(\s*('|")Content-Type\1\s*,\s*('|")application\/json;\s*charset=UTF-8\2\s*\)/g
);
```

Vaš kod bi trebalo da ima hendlera događaja za `onreadystatechange`, postavljenog u metodu.

```js
assert.match(code, /\.onreadystatechange\s*=/g);
```

Tvoj kod bi trebalo da dohvati element sa klasom `message` i promeni njegovu `textContent` u `userName loves cats`.

```js
assert.match(
  code,
  /document\.getElementsByClassName\(\s*('|")message\1\s*\)\[0\]\.textContent\s*=\s*.+?\.userName\s*\+\s*.+?\.suffix/g
);
```

Vaš kod bi trebalo da koristi `send`.

```js
assert.match(code, /\.send\(\s*body\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line

      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message box">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line
      const xhr = new XMLHttpRequest();
      xhr.open('POST', url, true);
      xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
      xhr.onreadystatechange = function () {
        if (xhr.readyState === 4 && xhr.status === 201) {
          const serverResponse = JSON.parse(xhr.response);
          document.getElementsByClassName('message')[0].textContent =
            serverResponse.userName + serverResponse.suffix;
        }
      };
      const body = JSON.stringify({
        userName: userName,
        suffix: ' loves cats!'
      });
      xhr.send(body);
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```
