---
id: 587d7fae367417b2b2512be3
title: Dobijanje JSON-a pomoću JavaScript XMLHttpRequest
challengeType: 6
forumTopicId: 301502
dashedName: get-json-with-the-javascript-xmlhttprequest-method
---

# --description--

Takođe možete zatražiti podatke iz eksternog izvora. Ovde se API unos.

Imajte na umu da API - ili Interfejsi softver-za-softver - su alati koje računari koriste za komunikaciju među sobom. Naučićete kako da ažurirate HTML i podatke koje dobijamo od API korišćenjem tehnologije nazvane AJAX.

API Mreže prenose podatke u format koji se zove JSON. JSON predstavlja JavaScript Object Notacija.

Sintaksa JSON izgleda veoma slično sa sintaksom fraze objekta od JavaScript. JSON ima svojstva objekta i njihove trenutne vrednosti, postavljene između `{` i `}`.

Ove osobine i njihove vrednosti se često nazivaju "parovi ključ-vrednost".

Međutim, JSON koja se šalje od API se šalje kao `bytes`, i vaša aplikacija prima ga kao `string`. Ovo se može konvertovati u stvari tipa JavaScript, ali ne u stvari tipa JavaScript podrazumevano. Metoda `JSON.parse` analizira sekvencu karaktera i gradi nešto od JavaScript što je opisano tom sekvencom.

Možete zatražiti JSON iz API od Piktura mačaka sa freeCodeCamp-a. Ovo je kod koji možete postaviti u vaš događaj za klik da biste ovo uradili:

```js
const req = new XMLHttpRequest();
req.open('GET', '/json/cats.json', true);
req.send();
req.onload = function () {
  const json = JSON.parse(req.responseText);
  document.getElementsByClassName('message')[0].innerHTML =
    JSON.stringify(json);
};
```

Ovde je sažetak onoga šta svaki deo radi. Objekat JavaScript `XMLHttpRequest` ima nekoliko metoda koje se koriste za prenos podataka. Prvo, primer objekta `XMLHttpRequest` je kreiran i sačuvan u parametru `req`. Zatim, metoda `open` inicira zahtev – ovaj primer traži podatke od API, pa je to zahtev za `GET`. Drugi argument za `open` je URL od API koji traži podatke. Treći argument je vrednost tačno ili netačno, gde `true` čini zahtev asinhronim. Metoda `send` šalje zahtev. Na kraju, rukovatelji događaja `onload` analiziraju povratke podatke i koriste metodu `JSON.stringify` da konvertuju objekat JavaScript u niz karaktera. Ovaj niz se zatim unosi kao tekst poruke.

# --instructions--

Ispravite kod da kreirate i pošaljete zahtev za `GET` za API na freeCodeCamp Cat Pictures. Zatim kliknite na dugme `Get Message`. Vaša AJAX funkcija će promeniti tekst `The message will go here` i sirove rezultate JSON iz API.

# --hints--

Vaš kod treba da kreira `XMLHttpRequest` novi.

```js
assert.match(code, /new\s+?XMLHttpRequest\(\s*?\)/g);
```

Tvoj kod bi trebalo da koristi metodu `open` za iniciranje zahteva `GET` za API od Cat Pictures of freeCodeCamp.

```js
assert.match(
  code,
  /\.open\(\s*('|")GET\1\s*,\s*('|")\/json\/cats\.json\2\s*,\s*true\s*\)/g
);
```

Vaš kod bi trebalo da koristi metodu `send` za slanje zahteva.

```js
assert.match(code, /\.send\(\s*\)/g);
```

Vaš kod bi trebalo da ima *event handler* za `onload` postavljen u funkciju.

```js
assert.match(
  code,
  /\.onload\s*=\s*(function|\(\s*\))\s*(\(\s*\)|\=\>)\s*{/g
);
```

Vaš kod treba da koristi metodu `JSON.parse` za analizu `responseText`.

```js
assert.match(code, /JSON\s*\.parse\(\s*.*\.responseText\s*\)/g);
```

Vaš kod bi trebalo da dobije element sa klasom `message` i promeni njegov interni HTML u niz podataka za/od JSON.

```js
assert.match(
  code,
  /document\s*\.getElementsByClassName\(\s*?('|")message\1\s*?\)\[0\]\s*\.innerHTML\s*?=\s*?JSON\.stringify\(.+?\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = () => {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
