---
id: 587d7fae367417b2b2512be5
title: Prevedi podatke iz JSON-a u HTML
challengeType: 6
forumTopicId: 16807
dashedName: convert-json-data-to-html
---

# --description--

Kada počnete da dobijate podatke od API iz JSON, možete ga prikazati u HTML.

Možete koristiti metodu `forEach` za pregled podataka jer su stavke slika mačaka sačuvane u nizu podataka. Kada dođete do svakog elementa, možete promeniti atribute HTML.

Prvo, deklariši atribut HTML-a sa `let html = "";`.

Zatim, iteriraj JSON, dodajući HTML u parametar koji prolazi kroz imena ključeva u oznakama `strong`, praćen vrednošću. Kada završiš iteraciju, prikaži.

Ovo je kod koji radi ovo:

```js
let html = '';
json.forEach(function (val) {
  const keys = Object.keys(val);
  html += "<div class = 'cat'>";
  keys.forEach(function (key) {
    html += '<strong>' + key + '</strong>: ' + val[key] + '<br>';
  });
  html += '</div><br>';
});
```

**Napomena:** Za ovaj zadatak, morate dodati nove komponente za HTML na stranici, pa ne možete se oslonjati na `textContent`. Umesto toga, morate koristiti `innerHTML`, što može učiniti sajt ranjivim na napade cross-site scriptinga.

# --instructions--

Uspostujte način za `forEach` okruživanja podataka JSON i kreiranje komponenti HTML za prikazivanje.

Ovo je primer za JSON:

```json
[
  {
    "id": 0,
    "imageLink": "https://cdn.freecodecamp.org/curriculum/legacy-json-apis-ajax/funny-cat.jpg",
    "altText": "A white cat wearing a green helmet shaped melon on its head. ",
    "codeNames": ["Juggernaut", "Mrs. Wallace", "Buttercup"]
  }
]
```

# --hints--

Vaš kod treba da sačuva podatke u parametru `html`

```js
assert.match(__helpers.removeWhiteSpace(code), /html(\+=|=html\+)/g);
```

Vaš kod bi trebalo da koristi metodu ``forEach`` za obavotavanje podataka `JSON` iz `API`.

```js
assert.match(code, /json\.forEach/g);
```

Vaš kod treba da okruži imena polja u oznakama `strong`.

```js
assert.match(code, /<strong>.+<\/strong>/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        let html = '';
        // Add your code below this line

        // Add your code above this line
        document.getElementsByClassName('message')[0].innerHTML = html;
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        let html = '';
        // Add your code below this line
        json.forEach(function (val) {
          var keys = Object.keys(val);
          html += "<div class = 'cat'>";
          keys.forEach(function (key) {
            html += '<strong>' + key + '</strong>: ' + val[key] + '<br>';
          });
          html += '</div><br>';
        });
        // Add your code above this line
        document.getElementsByClassName('message')[0].innerHTML = html;
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
