---
id: 587d7faf367417b2b2512be8
title: Dobavljanje podataka o lokaciji korisnika ili dobijanje GPS koordinata korisnika
challengeType: 6
forumTopicId: 18188
dashedName: get-geolocation-data-to-find-a-users-gps-coordinates
---

# --description--

Još jedna zanimljiva stvar koju možete učiniti je dobiti pristup lokaciji vašeg trenutnog korisnika. Svaki pretraživač ima ugrađen pregledač koji vam može pružiti ove informacije.

Pretraživač će privremeno dobiti dimenzije korisnika.

Videćeš upozorenje koje traži dozvolu ili sprečava ovaj sajt da zna tvoju trenutnu lokaciju. Izazov se može završiti na bilo koji način, pod uslovom da je kod tačan.

Prilikom odabira dozvola, videćete tekst na ekranu rezultata koji će promeniti svoju širinu i visinu.

Ovde je kod koji radi ovo:

```js
if (navigator.geolocation) {
  navigator.geolocation.getCurrentPosition(function (position) {
    document.getElementById('data').innerHTML =
      'latitude: ' +
      position.coords.latitude +
      '<br>longitude: ' +
      position.coords.longitude;
  });
}
```

Prvo, proverava se da li postoji objekat `navigator.geolocation`. Ako postoji, poziva se metoda `getCurrentPosition` na tom objektu, koja pokreće asinhroni zahtev za lokaciju korisnika. Ako je zahtev uspešan, funkcija povratne informacije u toj metodi se izvršava. Ova funkcija dobija vrednosti objekta `position` po širini i dužini koristeći dot notation (pisanje tačaka) i ažurira HTML.

# --instructions--

Dodajte primer kod unutar tagova `script` ili proverite trenutni korisnički deo i unesite ga u HTML.

# --hints--

Vaš kod bi trebalo da koristi `navigator.geolocation` da dobije trenutnu lokaciju korisnika.

```js
assert.match(code, /navigator\.geolocation\.getCurrentPosition/g);
```

Vaš kod treba da koristi `position.coords.latitude` za prikaz širine korisnika.

```js
assert.match(code, /position\.coords\.latitude/g);
```

Vaš kod bi trebalo da koristi `position.coords.longitude` za prikaz punog imena korisnika.

```js
assert.match(code, /position\.coords\.longitude/g);
```

Treba da prikažete poziciju korisnika unutar elementa ``div`` koristeći ``id="data"``.

```js
assert.match(
  code,
  /document\.getElementById\(\s*('|")data\1\s*\)\.innerHTML/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  // Add your code below this line

  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data"></div>
```

# --solutions--

```html
<script>
  // Add your code below this line
  if (navigator.geolocation) {
    navigator.geolocation.getCurrentPosition(function(position) {
      document.getElementById('data').innerHTML = "latitude: " + position.coords.latitude + "<br>longitude: " + position.coords.longitude;
    });
  }
  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data">

</div>

</section>
```
