---
id: 587d7fae367417b2b2512be4
title: Preuzmi JSON podatke iz API-ja
challengeType: 6
forumTopicId: 301499
dashedName: access-the-json-data-from-an-api
---

# --description--

U prethodnom izazovu, naučiti kako dobiti podatke za JSON iz API slika mačaka sa freeCodeCamp-a.

Sada ćete detaljnije pregledati povratke podatke kako biste bolje razumeli strukturu JSON. Obratite pažnju na određene zapise u JavaScript:

<blockquote>[ ] -> Kvadratne zagrade predstavljaju red podataka.<br>{ } -> Okrugle zagrade predstavljaju objekat.<br>" " -> Navodnici predstavljaju niz znakova. Takođe se koriste za imena ključeva u JSON.</blockquote>

Razumeti strukturu podataka koja se vraća sa API je važno jer utiče na način na koji dobijate vrednosti koje vam trebaju.

Levo, klikni na dugme `Get Message` ili preuzmi/pakuj JSON od API za sliku mačke freeCodeCamp-a na HTML.

Prvi i poslednji atribut koji vidite u podacima za JSON su uglate zagrade `[ ]`. To znači da su vraćeni podaci niz podataka. Drugi atribut u podacima za JSON su vitičaste zagrade `{` koje nešto započinju. Ako pažljivo pogledate, videćete tri različite stvari. Podaci za JSON su niz od tri stavke, gde svaka stavka sadrži informacije o slici mačke.

Naučio si prvobitno da stvari imaju „parove ključ-vrednost“ razdvojene zarezima. Na primer, za slike mačaka, prvi element ima `"id":0` gde je `id` ključ, a `0` njegova odgovarajuća vrednost. Slično tome, postoje ključevi za `imageLink`, `altText` i `codeNames`. Svaki element slike mačka ima ove iste ključeve, ali sa različitim vrednostima.

Još par zanimljivih "ključ-vrednost" parova u prvom objektu je `"codeNames":["Juggernaut","Mrs. Wallace","ButterCup"]`. Ovde je `codeNames` ključ, a njegova vrednost je niz podataka sa sekvencijom od tri karaktera. Takođe je moguće imati niz objekata zajedno sa ključem čija je vrednost niz.

Zapamtite kako dobiti podatke u nizovima podataka i objektima. Nizovi podataka koriste notaciju uglastih zagrada da dobiju specifičnu instancu elementa. Objekti koriste notaciju uglastih zagrada ili notaciju tačke da dobiju vrednost određenog atributa/svojstva. Evo primera koji ispisuje svojstvo `altText` prve slike mačke - zapamtite da su podaci za JSON analizirani u editoru sačuvani u parametru nazvan `json`:

```js
console.log(json[0].altText);
```

Konsola će prikazati niz slova `A white cat wearing a green helmet shaped melon on its head.`.

# --instructions--

Za mačku sa `id` od 2, ispišite na konzolu drugu vrednost u redovima `codeNames`. Treba da koristite notaciju uglastih zagrada i tačku na objektu (koji je sačuvan u parametru `json`) da biste dobili tu vrednost.

# --hints--

Tvoj kod treba da koristi zagrade i tačke kako bi dobio ispravno ime koda, i odštampajte `Loki` u konzolu.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*json\s*\[2\]\s*(\.\s*codeNames|\[\s*('|`|")codeNames\2\s*\])\s*\[\s*1\s*\]\s*\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line

        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line
        console.log(json[2].codeNames[1]);
        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
