---
id: 5ccfad82bb2dc6c965a848e5
title: Dohvati JSON koristeći JavaScript fetch
challengeType: 6
forumTopicId: 301501
dashedName: get-json-with-the-javascript-fetch-method
---

# --description--

Drugi način za zahtev eksternih podataka je korišćenje metode JavaScript `fetch()`. Slično je `XMLHttpRequest`, ali se njegov sintaksa smatra lak za razumevanje.

Ovo je kod za slanje GET zahteva ka `/json/cats.json`

```js
fetch('/json/cats.json')
  .then(response => response.json())
  .then(data => {
    document.getElementById('message').innerHTML = JSON.stringify(data);
  });
```

Napomena: Putanja `fetch()` koristi `GET` kao putanju `HTTP` podrazumevano. To znači da ne morate eksplicitno navoditi za dohvaćanje osnovnih podataka.

Pogledajte svaki deo ovog koda.

Prvi red je onaj koji pravi zahtev. Stoga, `fetch(URL)` pravi zahtev za `GET` za URL koji je naveden. Ovaj način vraća Obećanje.

Nakon što je Obećanje restaurirano, ako je zahtev uspešan, putem `then` se izvršava put, koji uzima odgovor i konvertuje ga u format od JSON.

Putanja za `then` takođe vraća Obećanje, koje je obrađeno sledećom putanjom od `then`. Argument u `then` drugi je nešto od JSON što tražite!

Sada, bira element koji će primiti podatke koristeći `document.getElementById()`. Zatim menja kod elementa HTML unosom niza karaktera generisanog iz objekta JSON koji je povučen iz zahteva.

# --instructions--

Ispravite kod ili kreirajte i pošaljite zahtev za `GET` za API Cat Pictures of freeCodeCamp. Međutim, ovog puta koristite metodu `fetch` umesto `XMLHttpRequest`.

# --hints--

Vaš kod bi trebalo da koristi dostupne podatke za modifikovanje internog HTML

```js
const catData = 'dummy data';
const ref = fetch;
fetch = () => Promise.resolve({ json: () => catData });

  try {
    document.getElementById('getMessage').click();
    await new Promise((resolve, reject) => setTimeout(() => resolve(), 250));
  } catch (error) {
    console.log(error);
  } finally {
    fetch = ref;
    assert.equal(
      document.getElementById('message').textContent,
      JSON.stringify(catData)
    );
  }

```

Vaš kod bi trebalo da pošalje zahtev za `GET` ka `fetch`.

```js
assert.match(code, /fetch\s*\(\s*('|")\/json\/cats\.json\1\s*\)/g);
```

Vaš kod mora da koristi `then` da promeni odgovor u JSON.

```js
assert.match(
  code,
  /\.then\s*\(\s*\(?(?<var>\w+)\)?\s*=>\s*\k<var>\s*\.json\s*\(\s*\)\s*\)/g
);
```

Vaš kod bi trebalo da koristi `then` za obradu podataka konvertovanih u JSON i drugog tipa, `then`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.then\(\(?\w+\)?=>{[^}]*}\)/g);
```

Vaš kod mora da dobije element sa identifikatorom `message` i modifikuje svoj interni HTML u sekvencu podataka za JSON.

```js
assert.match(
  __helpers.removeWhiteSpace(code),
  /document\.getElementById\(('|")message\1\)\.innerHTML=JSON\.stringify\(?\w+\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      fetch('/json/cats.json')
        .then(response => response.json())
        .then(data => {
          document.getElementById('message').innerHTML = JSON.stringify(data);
        });
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
