---
id: 587d8248367417b2b2512c3b
title: Zuia IE otvaranje HTML neiskorištenog za helmet.ieNoOpen()
challengeType: 2
forumTopicId: 301584
dashedName: prevent-ie-from-opening-untrusted-html-with-helmet-ienoopen
---

# --description--

Kao podsjetnik, ovaj projekat je baziran na osnovnom projektu kopiranom iz <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Neki veb programi omogućavaju preuzimanje nepouzdanog HTML-a. Neke verzije Internet Explorer-a podrazumevano otvaraju te HTML datoteke u kontekstu vaše veb stranice. To znači da nepouzdan HTML sadržaj može izvršavati zlonameran kod u kontekstu vaše stranice. Ovaj posrednički program postavlja zaglavlje `X-Download-Options` na `noopen`. To sprečava korisnike Internet Explorer-a da izvršavaju preuzete datoteke u kontekstu pouzdane veb stranice.

# --instructions--

Koristi putanju `helmet.ieNoOpen()` na tvom serveru.

# --hints--

Centralni program helmet.ieNoOpen() Trebalo bi da se postavi pravilno

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'ienoopen');
assert.equal(data.headers['x-download-options'], 'noopen');
```
