---
id: 587d8249367417b2b2512c40
title: Sanidi Helmet koristeći centralni program ‘parent’ helmet()
challengeType: 2
forumTopicId: 301575
dashedName: configure-helmet-using-the-parent-helmet-middleware
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`app.use(helmet())` uključuje sve centralne programe navedene iznad, osim `noCache()` i `contentSecurityPolicy()`, ali ih je moguće omogućiti ako je potrebno. Takođe možete onemogućiti ili konfigurisati bilo koji drugi centralni program pojedinačno, koristeći element za konfiguraciju.

**Primer:**

```js
app.use(helmet({
  frameguard: {         // configure
    action: 'deny'
  },
  contentSecurityPolicy: {    // enable and configure
    directives: {
      defaultSrc: ["'self'"],
      styleSrc: ['style.com'],
    }
  },
  dnsPrefetchControl: false     // disable
}))
```

Predstavili smo svaki middleware na različite načine radi edukacije i lakoće verifikacije. Korišćenje 'parent' middleware-a `helmet()` je jednostavno za implementaciju u stvarnom projektu.

# --hints--

Nije ispit - to je izazov objašnjenja

```js
assert(true);
```
