---
id: 587d8247367417b2b2512c38
title: Smanite rizik od clickjackinga pomoću helmet.frameguard()
challengeType: 2
forumTopicId: 301582
dashedName: mitigate-the-risk-of-clickjacking-with-helmet-frameguard
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Vaša stranica može biti postavljena u `<frame>` ili `<iframe>` bez vaše dozvole. Ovo može izazvati napade clickjacking-a, između ostalog. Clickjacking je tehnika prevare korisnika da interaguje sa drugom stranicom od one koju misli. To se može postići pokretanjem vaše stranice u opasnom kontekstu, putem iframinga. U tom kontekstu, napadač može postaviti skriveni sloj preko vaše stranice. Skrivena dugmad mogu se koristiti za pokretanje zlonamernog koda. Ovaj middleware postavlja zaglavlje X-Frame-Options. On sprečava ko može da stavi vaš sajt u okvir (frame). Ima tri stanja: DENY, SAMEORIGIN, i ALLOW-FROM.

Ne trebaju nam da naša aplikacija bude postavljena u okvir.

# --instructions--

Koristi `helmet.frameguard()` kada prolaziš kroz nešto za konfiguraciju `{action: 'deny'}`.

# --hints--

Za centralni program, helmet.frameguard() treba biti pravilno postavljen.

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(
  data.appStack,
  'frameguard',
  'helmet.frameguard() middleware is not mounted correctly'
);
```

helmet.frameguard() 'action' treba ga postaviti na 'DENY'

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.property(data.headers, 'x-frame-options');
assert.equal(data.headers['x-frame-options'], 'DENY');
```
