---
id: 587d8248367417b2b2512c3c
title: Tražite da pretraživači koriste vaš sajt samo putem HTTPS-a kwa helmet.hsts()
challengeType: 2
forumTopicId: 301573
dashedName: ask-browsers-to-access-your-site-via-https-only-with-helmet-hsts
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

HTTP Strict Transport Security (HSTS) je protokolarna sigurnosna politika koja pomaže u zaštiti veb stranica od napada degradacije protokola i krađe kolačića (cookie hijacking). Ako je vaša veb stranica dostupna putem HTTPS, možete tražiti da korisnički pregledači izbegnu korišćenje nesigurnog HTTP. Postavljanjem zaglavlja Strict-Transport-Security, govorite pregledačima da koriste HTTPS za dolazne zahteve tokom definisanog vremenskog perioda. Ovo će važiti za sve naredne zahteve nakon prvobitnog zahteva.

# --instructions--

Konfigurišite `helmet.hsts()` koristeći HTTPS za narednih 90 dana. Proverite element konfiguracije `{maxAge: timeInSeconds, force: true}`. Možete kreirati parametar `ninetyDaysInSeconds = 90*24*60*60;` koristeći za `timeInSeconds`.

Napomena: Konfigurisanje HTTPS na specifičnoj veb stranici zahteva kupovinu domena, i sertifikat za SSL/TLS.

# --hints--

centralni program helmet.hsts() mora biti pravilno podešen

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hsts');
assert.property(data.headers, 'strict-transport-security');
```

maxAge treba biti jednak 7776000 s (dana 90)

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.match(
  data.headers['strict-transport-security'],
  /^max-age=7776000;?/
);
```
