---
id: 587d8247367417b2b2512c36
title: Instaliraj i Ita Helmet
challengeType: 2
forumTopicId: 301581
dashedName: install-and-require-helmet
---

# --description--

Rad na ovim izazovima će zahtevati pisanje vašeg koda korišćenjem jedne od sledećih metoda:

- Kopiraj <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">hii repozitorijum od GitHub</a> i završi ove projekte interno.
- Koristi omiljeni alat za izradu veb sajtova da završiš projekat. Proveri da uključuješ sve fajlove iz našeg repozitorijuma GitHub.

Kaska pomaže da zaštiti vaše programe za Express postavljanjem različitih zaglavlja od HTTP.

# --instructions--

Svi vaši kodovi za ove lekcije postavljeni su u fajl `myApp.js` između linija koda koje smo započeli za vas. Molimo Vas da ne menjate niti brišete kod koji smo dodali.

Verzija Helmet-a `3.21.3` je već postavljena, pa će biti kao `helmet` u `myApp.js`.

# --hints--

Verzija od `helmet` za `3.21.3` treba da bude u `package.json`

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
const packJson = JSON.parse(data);
const helmet = packJson.dependencies.helmet;
assert(helmet === '3.21.3' || helmet === '^3.21.3');
```
