---
id: 587d8247367417b2b2512c37
title: Informacije koje mogu predstavljati opasnost prilikom korišćenja helmet.hidePoweredBy()
challengeType: 2
forumTopicId: 301580
dashedName: hide-potentially-dangerous-information-using-helmet-hidepoweredby
---

# --description--

Kao podsetnik, ovaj projekat je izgrađen na osnovu osnovnog projekta uzet iz <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Hakeri mogu da koriste poznastu ranjivost u Express/Node ako vide da je vaš sajt pokren na Express. `X-Powered-By: Express` se šalje u svakom zahtevu koji dolazi iz Express po defaultu. Koristite middleware `helmet.hidePoweredBy()` da uklonite X-Powered-By zaglavlje.

# --hints--

Centralni program helmet.hidePoweredBy() treba ga postaviti na odgovarajuće mesto

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hidePoweredBy');
assert.notEqual(data.headers['x-powered-by'], 'Express');
```
