---
id: 587d8249367417b2b2512c3e
title: Onemogućavanje čuvanja kopije na strani klijenta za helmet.noCache()
challengeType: 2
forumTopicId: 301576
dashedName: disable-client-side-caching-with-helmet-nocache
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Ako objavljujete ažuriranje za vaš sajt i želite da korisnici uvek preuzimaju najnoviju verziju, možete (pokušati) da isključite keširanje u klijentskom pretraživaču. Ovo može biti korisno i tokom razvoja. Keširanje ima prednosti u performansama koje ćete izgubiti, pa koristite ovu opciju samo kada postoji stvarna potreba.

# --instructions--

Koristi putanju `helmet.noCache()` na tvom serveru.

# --hints--

Srednji program helmet.noCache() mora biti pravilno postavljen

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nocache');
assert.equal(
  data.headers['cache-control'],
  'no-store, no-cache, must-revalidate, proxy-revalidate'
);
```
