---
id: 587d8248367417b2b2512c3d
title: Početni DNS upit za helmet.dnsPrefetchControl()
challengeType: 2
forumTopicId: 301577
dashedName: disable-dns-prefetching-with-helmet-dnsprefetchcontrol
---

# --description--

Kao podsetnik, ovaj projekat je baziran na osnovnom projektu preuzetom iz <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Radi poboljšanja performansi, većina pregledača dohvaća DNS zapise za linkove koji se nalaze na stranici. Na taj način, adresa IP odredišta je već poznata kada korisnik klikne na link. Ovo može dovesti do prekomernog korišćenja DNS usluge (ako imate veliku sajt, posetioca u milionima...), problema sa privatnošću (jedan tracker može utvrditi da ste na određenoj stranici), ili promena statistike stranice (drugi linkovi mogu izgledati kao da su posećeni iako nisu). Ako vam je potrebna visoka bezbednost, možete onemogućiti pred-dobijanje DNS zapisa, ali na uštrinu smanjenja performansi.

# --instructions--

Koristi putanju `helmet.dnsPrefetchControl()` na tvom serveru.

# --hints--

Centralni program helmet.dnsPrefetchControl() mora biti pravilno podešen

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'dnsPrefetchControl');
assert.equal(data.headers['x-dns-prefetch-control'], 'off');
```
