---
id: 587d8248367417b2b2512c3a
title: Izbegavanje Vrste MIME odgovora za helmet.noSniff()
challengeType: 2
forumTopicId: 301574
dashedName: avoid-inferring-the-response-mime-type-with-helmet-nosniff
---

# --description--

Kao podsetnik, ovaj projekat je napravljen na osnovu osnovnog projekta preuzetog iz <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Pretraživači mogu koristiti manipulaciju sadržajem ili MIME tipovima da izmene zaglavlje odgovora ``Content-Type`` kako bi predvideli i obradili podatke koristeći nejasno definisan tip sadržaja. Iako ovo može biti korisno u nekim slučajevima, takođe može izazvati opasne napade. Ovaj posrednički program postavlja zaglavlje ``X-Content-Type-Options`` na ``nosniff``, naređujući pretraživaču da ne koristi (ili: *da se ne oslanja na*) navedeni ``Content-Type``.

# --instructions--

Koristi putanju `helmet.noSniff()` na tvom serveru.

# --hints--

Centralni program helmet.noSniff() mora biti pravilno postavljen

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nosniff');
assert.equal(data.headers['x-content-type-options'], 'nosniff');
```
