---
id: 58a25bcff9fc0f352b528e7d
title: Resetovanje lozinke i poređenje putem jednokratne metode
challengeType: 2
forumTopicId: 301578
dashedName: hash-and-compare-passwords-asynchronously
---

# --description--

Kao podsetnik, ovaj projekat je izgrađen na osnovu baznog projekta uzetog sa <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Pošto je hashovanje (hashing) namenjeno da ima visok računarski opterećenje, savetuje se da se izvrši asinhrono na vašem serveru kako bi se izbeglo blokiranje dolaznih konekcija tokom hashovanja. Sve što je potrebno da biste asinhrono hashovali lozinke jeste napraviti poziv

```js
bcrypt.hash(myPlaintextPassword, saltRounds, (err, hash) => {
  /*Store hash in your db*/
});
```

# --instructions--

Dodajte ovu funkciju za maskiranje na svoj server (već smo definisali parametre koji se koriste u funkciji da biste ih koristili) i napišite je u konzolu da vidite! Tamo bi obično sačuvali hash u svoju bazu podataka.

Kada želite da utvrdite da li je novi unos isti podatak kao heš, koristićete samo funkciju za poređenje.

```js
bcrypt.compare(myPlaintextPassword, hash, (err, res) => {
  /*res == true or false*/
});
```

Dodajte ovo unutar vaše postojeće funkcije (jer morate sačekati da se hash završi pre nego što pozovete funkciju za poređenje) nakon što upišete završen hash i ukucate 'res' u konzolu unutar poređenja. Trebalo bi da vidite hash u konzoli, a zatim će biti ispisano 'true'! Ako promenite 'myPlaintextPassword' u funkciji za poređenje na 'someOtherPlaintextPassword', onda će reći false.

```js
bcrypt.hash('passw0rd!', 13, (err, hash) => {
  console.log(hash);
  //$2a$12$Y.PHPE15wR25qrrtgGkiYe2sXo98cjuMCG1YwSI5rJW1DSJp0gEYS
  bcrypt.compare('passw0rd!', hash, (err, res) => {
    console.log(res); //true
  });
});

```

Dostavite svoju stranicu kada smatrate da je dobro urađena.

# --hints--

Hash mora biti tačno generisan i upoređen.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash.*myPlaintextPassword( |),( |)saltRounds( |),( |).*err( |),( |)hash[^]*END_ASYNC/gi,
  'You should call bcrypt.hash on myPlaintextPassword and saltRounds and handle err and hash as a result in the callback'
);
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash[^]*bcrypt.compare.*myPlaintextPassword( |),( |)hash( |),( |).*err( |),( |)res[^]*}[^]*}[^]*END_ASYNC/gi,
  'Nested within the hash function should be the compare function comparing myPlaintextPassword to hash'
);
```
