---
id: 58a25bcff9fc0f352b528e7e
title: Sakrij i uporedi lozinku pomoću koda pozicije
challengeType: 2
forumTopicId: 301579
dashedName: hash-and-compare-passwords-synchronously
---

# --description--

Kao podsetnik, ovaj projekat je izgrađen na osnovu baznog projekta uzetog sa <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Keširanje pomoću pozicionog koda je takođe lako za izvođenje, ali može uzrokovati kašnjenje ako se koristi na serverskoj strani sa velikim troškovima ili prečesto kešira. Keširanje na ovaj način je jednostavno kao funkциони poziv

```js
var hash = bcrypt.hashSync(myPlaintextPassword, saltRounds);
```

Dodaj ovaj način šifrovanja u svoj kod, a zatim napiši rezultate u konzolu. Štaviše, parametri koji su već korišćeni definisani su na serveru, pa nećete morati ništa da menjate. Možete primetiti da čak i ako šifrujete istu lozinku kao u funkciji async, rezultati na konzoli su drugačiji — ovo je zato što se sol generiše nasumično svaki put, kao što je vidljivo u prvih 22 karaktera treće sekvence šifrovanja. Sada za poređenje unosa lozinke i novog stanja, koristićete metodu compareSync:

```js
var result = bcrypt.compareSync(myPlaintextPassword, hash);
```

Rezultat će biti vrednost Tačno ili Netačno.

# --instructions--

Dodajte funkciju i ispišite rezultate u konzolu da biste videli kako funkcioniše.

Predaj svoju stranicu kada smatrate da je dobro urađena.

# --hints--

Statusni kod mora biti generisan i upoređen sa preciznošću.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_SYNC[^]*hash.*=.*bcrypt.hashSync.*myPlaintextPassword( |),( |)saltRounds[^]*END_SYNC/gi,
  'You should call bcrypt.hashSync on myPlaintextPassword with saltRounds'
);
assert.match(
  data,
  /START_SYNC[^]*result.*=.*bcrypt.compareSync.*myPlaintextPassword( |),( |)hash[^]*END_SYNC/gi,
  'You should call bcrypt.compareSync on myPlaintextPassword with the hash generated in the last line'
);
```
