---
id: 65ef1be2b958cabf712296cd
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Sada ćete ponovano smanjivati interval dobijanjem srednje vrednosti trenutnog intervala i poređenjem kvadrata srednje vrednosti sa ciljnom vrednošću.

Za to, unutar bloka `else`, kreiraj kolo `for` koje radi do puta `max_iterations`.

Za vašu petlju, koristite funkciju ``range``, koja generiše sekvencu brojeva preko kojih možete iterirati. Sintaksa je ``range(start, stop, step)``, gde je ``start`` početni broj (uključujući), ``stop`` krajnji broj (isključujući), a ``step`` razlika između brojeva u sekvenci.

Pia, koristi `_` kao parametar petlje. `_` funkcioniše kao zamena i važno je kada moraš koristiti parametar, ali ti ne treba njegova stvarna vrednost.

# --hints--

Treba da kreirate kolo `for` koje se ponavlja preko `range(max_iterations)`. Koristite `_` kao parametar kola.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
