---
id: 65ef190c6b51e9b5a5f7ed29
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Navedite funkciju za `square_root_bisection` sledeći parametri:

- `square_target`: Broj čiji želite pronaći kvadratni koren.
- `tolerance` (opciono): Dozvoljena razlika između kvadrata procenjene vrednosti korena i stvarne ciljne vrednosti (podrazumevano je `1e-7`). Tolerancija od `1e-7` znači da će rešenje biti tačno u roku od `0.0000001` stvarne vrednosti i predstavlja dobar podrazumevani izbor koji balansira preciznost i performanse.
- `max_iterations` (opciono): Maksimalan broj iteracija (podrazumevano je `100`). Ako metoda ne pronađe rešenje unutar ovog limita, smatraće se da rešenje nije dostupno.# --hints--

Vaša funkcija mora imati ove parametre: `square_target`, `tolerance = 1e-7`, i `max_iterations = 100`. Red je važan.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```
