---
id: 65ef19425d1b27b6c930bae6
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Ako je broj čiji kvadratni koren želiš da izračunaš negativan, kod bi trebalo da izbaci grešku jer kvadratni koren negativnog broja nije definisan u realnim brojevima.

Ukloni deo od `pass` i kreiraj deo od `if` ili proveri da li `square_target` je manji od `0`.

# --hints--

Trebalo bi da uklonite glavnu reč `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

Trebalo bi da kreirate deklaraciju za `if` ili je proverite kao `square_target < 0`.

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
