---
id: 65ef1d5e3d2927c5e0f4997b
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

U Python, glavna funkcija od `is` proverava identifikaciju objekta. Koristi se za određivanje da li dva parametra ukazuju na isti objekat u memoriji. Za razliku od `is`, operator jednakosti (`==`) proverava da li su vrednosti dva objekta jednake, bez obzira na to da li je u pitanju isti objekat u memoriji.

Izvan ciklusa for, kreiraj izjavu za proveru da li je `if` jednako kao `root is None`. Ako je istina, odštampaj poruku `'Failed to converge within {max_iterations} iterations.'`. Zapamti da kreiraš tu poruku koristeći f-string sekvencu.

# --hints--

Treba da napraviš deklaraciju `if` koja proverava da li je `root is None`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[0].is_equivalent("root is None")`));

    }
})
```

Trebalo bi da prođeš argument `f'Failed to converge within {max_iterations} iterations.'` kroz funkciju `print`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].is_equivalent("if root is None:\\n    print(f'Failed to converge within {max_iterations} iterations.')")`));
       
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--

--fcc-editable-region--

```
