---
id: 65ef198fde24dfb7ff675b42
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Sentenca za `raise` vam omogućava da prisilite određeni tip diskriminacije. Ona uključuje glavnu reč `raise` praćenu tipom diskriminacije i omogućava vam da kreirate prilagođeno poruku o grešci:

```py
raise ValueError("Invalid value")
```

Kada se kod gore izvrši, `ValueError` i poruka sa `"Invalid value"` prikazuju se korisniku.

Ako `square_target` je manje od `0`, kvadratni koren realne vrednosti ne može se izračunati. Stoga, upload-ujte `ValueError` priloženo poruci od `'Square root of negative number is not defined in real numbers'`. Ne zaboravite da uklonite glavnu reč iz `pass`.

# --hints--

Trebalo bi vam da uklonite glavnu reč `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

Morate priložiti `ValueError` kada ga prilažete poruci od `'Square root of negative number is not defined in real numbers'` unutar tela `if`.

```js
({
    test: () => {
        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
