---
id: 65ef1a50049cf9bada13266f
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Kreiraj deklaraciju za `elif` ili proveri da li je `square_target` jednako `0`. Ako su jednaki, postavi vrednost `0` za parametar `root`. Takođe, odštampi poruku `'The square root of {square_target} is 0'`. Zapamti da kreiraš tu poruku koristeći f-string sekvencu.

# --hints--

Trebalo bi da imate deklaraciju `elif` ili proverite uslov `square_target == 0`.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

Trebalo bi da postaviš vrednost `0` za parametar `root` i proslediš argument `f'The square root of {square_target} is 0'` pozivu `print`.


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    
--fcc-editable-region--
```
