---
id: 65ef181dad3d96b14a21a03e
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovom projektu, pronaći ćete približan kvadratni koren broja dobijen korišćenjem metode bisekcije.

Metoda biseksije je tehnika za pronalaženje korena funkcije sa realnim vrednostima. Radiše se na smanjenjem intervala koji sadrži koren, sve dok ne pristigne do vrednosti unutar definisanog tolerancnog opsega.

Počnite kreiranjem funkcije nazvane `square_root_bisection`. Ostavite parametre prazne za sada.

# --hints--

Trebalo bi da definišete funkciju `square_root_bisection` bez ikakvih parametara.

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
