---
id: 6606a4641ec48b3a9fe8c2fc
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Funkcija unutar ``enumerate`` uzima iterable kao svoj argument i vraća enumerate objekat preko kojeg se može iterirati.
Pruža brojač (koji podrazumevano počinje od nule) i vrednost iz iterable.

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

Petlja iz gore navedenog primera bi ispisała tuple `0, a`, `1, b`, i `2, c`.

Unutar puta `find_empty_cell`, zameni `pass` sa petljom `for` koja koristi funkciju `enumerate()` za ponavljanje svakog reda na Sudoku tabli. Koristi `row` kao indeks trenutnog reda i `contents` za elemente trenutnog reda.


# --hints--

Treba da zameniš `pass` sa ciklusom `for` koji se ponavlja preko `enumerate(self.board)`. Koristi `row` i `contents` kao parametre ciklusa.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
