---
id: 66069d65162e61357c793e0c
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Rešivi Sudoku zagonetka će biti lista lista, kao što sledi:

```py
[
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]
```

Zapamtite da su prazne ćelije popunjene sa nuli.

Objavite promene `puzzle` i postavite njegovu vrednost lista lista, kao što je prikazano u gornjem primeru.

# --hints--

Treba da objaviš promenu `puzzle` postavljanjem njegove vrednosti u listu koordinata (2D) koji je dat.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("puzzle").is_equivalent("puzzle = [[0, 0, 2, 0, 0, 8, 0, 0, 0], [0, 0, 0, 0, 0, 3, 7, 6, 2], [4, 3, 0, 0, 0, 0, 8, 0, 0], [0, 5, 0, 0, 3, 0, 0, 9, 0], [0, 4, 0, 0, 0, 0, 0, 2, 6], [0, 0, 0, 4, 6, 7, 0, 0, 0], [0, 8, 6, 7, 0, 4, 0, 0, 0], [0, 0, 0, 5, 1, 9, 0, 0, 8], [1, 7, 0, 0, 0, 6, 0, 0, 5]]")`)) })
```

# --seed--

## --seed-contents--

```py

class Board:
    def __init__(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
gameboard = Board()
```
