---
id: 660699119472f332798860ad
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Sada, zameni `pass` pozivom `print` i prođi niz znakova `'Spam!'`.

# --hints--

Trebalo bi da obrišeš `pass` i odštampa sekvencu slova `'Spam!'` unutar puta `spam`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass
    
gameboard = Board()
--fcc-editable-region--
```
