---
id: 6606cc754a8834509cd0afb6
title: Korak 36
challengeType: 20
dashedName: step-36
---

# --description--

Još nešto što treba proveriti je da li se broj može uneti u 3x3 kvadrat.

U klasi za `Board`, kreiraj metod nazvan `valid_in_square` sa četiri parametra: `self`, `row`, `col`, i `num`. Gde `row`, `col` i `num` predstavljaju primer reda, primer kolone i broj koji treba proveriti, odnosno.

# --hints--

Treba ti napraviti novi put nazvan `valid_in_square` unutar klase `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_square")`)) })
```

Tvoj put/način za `valid_in_square` mora imati četiri parametra: `self`, `row`, `col`, na `num`, u ovom redosledu.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_square").has_args("self, row, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
