---
id: 6606b224a69a293f98f8db8f
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Putanja `.index()` izaziva grešku `ValueError` kada vrednost nije dostupna. Da bi se sprečilo zaustavljanje izvršavanja programa, postavite ovu liniju koda unutar bloka `try`. Klauzula `try` koristi se za obuhvatanje koda koji može izazvati grešku. Klauzula `except`, s druge strane, pruža alternativni kod za izvršavanje ako dođe do greške:

```py
try:
    <code>
except:
    <code>
```

Postavi šablon `col` unutar bloka `try`. Zatim, kreiraj sekciju `except` i popuni njen sadržaj sa `pass`.

# --hints--

Treba da umetnete varijablu `col` unutar bloka `try` i kreirate sekciju za `except`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\nexcept:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            col = contents.index(0)
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
