---
id: 66069b0b36053733a2f012fe
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Sada, obriši svoj poziv za `spam`.

# --hints--

Ne bi trebalo da imaš `gameboard.spam()` u svom kodu.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
gameboard.spam()
--fcc-editable-region--
```
