---
id: 6606beade9200b49aaeecd94
title: Korak 26
challengeType: 20
dashedName: step-26
---

# --description--

Ako `num` nije u redu, matematički izraz potvrđuje da je `True` i to znači da je broj validan za unos.

Ako je `num` u redu, matematički izraz je potvrđen da je `False` i unos će kršiti pravilo.

Dodaj glavnu reč `return` pre matematičkog izraza unutar tela putanje `valid_in_row`, kako bi se proverila validnost broja.

# --hints--

Trebalo bi da vratiš `num not in self.board[row]`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").find_body().is_equivalent("return num not in self.board[row]")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    def valid_in_row(self, row, num):
        num not in self.board[row]
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
