---
id: 6606933d6813a8308c962dd1
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Da bi bila metoda, potrebni su joj specifični parametar, nazvan `self` za ugovor. Ovaj parametar je referenca na primer klase objekata i uvek mora biti prvi parametar.

Dodajte parametar za `self` na vašoj putanji od `spam`.

# --hints--

Tvoja putanja ``spam`` mora imati parametar ``self``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
