---
id: 66069f86f58f85371d47123e
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Atribut polja je parametar povezan sa nekim objektom, koji se koristi za skladištenje podataka kao standardnih parametara.

Unutar metode ``__init__``, postavi vrednost parametra ``board`` (koji se prosleđuje prilikom kreiranja instanci klase ``Board``) za svojstvo polja instance ``board`` koristeći ``self.board``.

`self.board` se odnosi na svojstvo atributa `board` instance klase. To je svojstvo koje poseduje objekat kreiran iz klase `Board`.

# --hints--

Trebalo bi da obrišeš `pass` i postavi vrednost parametra `board` za `self.board` unutar puta `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").find_body().is_equivalent("self.board = board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
