---
id: 66068fb0bfddba2b7977eb60
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovom projektu, naučićete o klasama objekata i objektima kreiranjem rešenja Sudoku zagonetke.

U Python, klasa je šema za kreiranje objekata. Objekti kreirani iz klase su instanc(e) te klase. Možete kreirati klasu koristeći ovu sintaksu:

```py
class ClassName:
    pass
```

Gde `class` je ključna reč potrebna za definisanje klase, i `ClassName` je ime klase, napisano prema konvenciji *PascalCase*.

Počnite kreiranjem klase za `Board`.

# --hints--

Trebalo bi da kreirate klasu sa imenom `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
