---
id: 66069b992c1c5e3451f3deb0
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Konstruktor kreira prazan objekat. Metoda ``__init__`` je posebna metoda koja vam omogućava da kreirate objekat u prilagođenom stanju. Kada klasa objekata izvrši metodu ``__init__``, ``__init__`` se automatski poziva tokom kreiranja.

U vašoj klasi za `Board`, obrišite metodu `spam` i zameni je metodom `__init__` koja uključuje parametar `self`.

# --hints--

Ne bi trebalo da imaš put do `spam` u tvom kursu od `Board`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

Treba da definišete metodu `__init__` u vašoj klasi `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

Tvoja putanja ``__init__`` mora imati parametar ``self``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
