---
id: 66069167b3307b2f4067b22b
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Nova klasa objekata kreira se korišćenjem poziva funkcije, što uključuje dodavanje para vitičastih zagrada na ime klase.

Izvan definicije klase, kreiraj instancu klase `Board` i postavi njenu vrednost na parametar nazvan `gameboard`.

# --hints--

Trebalo bi da objaviš atribut `gameboard` iz klase `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

Vaš parametar za `gameboard` treba da ima vrednost `Board()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
