---
id: 66069e5759b800364707988e
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Vraća se putem `__init__`, zahteva dodatni parametar koji predstavlja zagonetku za rešavanje.

Dodaj drugi parametar nazvan `board` u metodi `__init__` i modifikuj inicijalizaciju `gameboard` prosleđivanjem liste `puzzle`, kao što prosleđuješ argumente u pozivu funkcije.

# --hints--

Tvoja putanja ``__init__`` mora imati dva parametra u sledećem redosledu: ``self``, i ``board``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self, board")`)) })
```

Vaš parametar za `gameboard` mora imati vrednost od `Board(puzzle)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board(puzzle)")`)) })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board()
--fcc-editable-region--
```
