---
id: 6606927d010be4300a4e5330
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Kreiranje objekta stvara prazan entitet. Međutim, klase objekata mogu imati metode, koje su kao interne funkcije za svaki objekat. Unutar klase objekata, metode se deklarišu na sledeći način:

```py
class ClassName:
    def method_name():
        pass
```

U klasi objekata `Board`, zameni `pass` sa praznom metodom `spam`.

# --hints--

Trebalo bi da deklarišete putanju sa imenom `spam` unutar klase objekta `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
    
gameboard = Board()
--fcc-editable-region--
```
