---
id: 660699aabc59c532f2d556e5
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Da biste izbrisali primer, morate koristiti tačku za prekid (nokten):

```py
instance_name.method_name()
```

Gde je `instance_name` primer ili objekat, a `method_name` je putanju koju želite da nazovete.

Put za `spam` za predmet `gameboard`.

# --hints--

Treba pozvati metodu `spam` objekta `gameboard` za `gameboard.spam()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
