---
id: 6606bc4e5535c0484990ccd5
title: Korak 24
challengeType: 20
dashedName: step-24
---

# --description--

Sledeće, radićete na metodi koja proverava da li određeni broj može biti unet u specifičan red na Sudoku tablu.

Unutar klase `Board`, kreiraj metodu nazvanu `valid_in_row` i sa tri parametra: `self`, `row`, i `num`. Gde `self` predstavlja primer klase, a `row` i `num` su predstavnici reda i proverovanog broja, respektivno.


# --hints--

Treba da napraviš novi put nazvan `valid_in_row` unutar klase `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_row")`)) })
```

Tvoja metoda `valid_in_row` treba da ima tri parametra: `self`, `row` i `num`, tim redosledom.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").has_args("self, row, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
