---
id: 691f7bbf76229bb97c60827d
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, naviknićete se na korišćenje petlji ``for...of`` za izradu funkcije koja broji koliko puta se sekvenca slova pojavljuje u nizu podataka o sekvencama slova.

Pre nego što kreirate tu funkciju, započeti ćete sa jednostavnijom funkcijom koja ispisuje svako slovo niza znakova na konzolu.

Započnite sa definisanjem prazne funkcije koja se zove `printCharacters`, sa parametrom `str`.

# --hints--

Trebalo bi da kreirate funkciju za `printCharacters`.

```js
assert.isFunction(printCharacters);
```

Tvoja funkcija od `printCharacters` treba da ima jedan parametar pod nazivom `str`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
