---
id: 6920d189e36b6b60a6330fee
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Sada, unutar petlje, unesite vrednost `char` u konzolu.

# --hints--

Trebalo bi da imaš deklaraciju `console.log` sa `char` kao argument u ciklusu.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)[1]

assert.match(loopCodeBlock, /console\.log\(char\)/);
```

Kada funkcija `printCharacters` bude pozvana sa argumentom `'test'`, izjava `console.log()` unutar funkcije treba da se pozove četiri puta sa sledećim argumentima, po redosledu: `'t'`, `'e'`, `'s'`, i `'t'`. Umesto toga, korišćen je sa sledećim argumentima: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
