---
id: 6920c9cda1c7f056f75e4b1a
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Sada je vreme da napišete svoju prvu petlju za `for...of`. U prethodnim lekcijama, naučili ste da možete koristiti petlju `for...of` za iteraciju preko vrednosti iz iterabilnog objekta (kao što su niz znakova ili niza podataka).

Ovo je primer kola za `for...of`:

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

Dodajte petlju za `for...of` sa praznim delom koda unutar vaše funkcije. Trebalo bi da iterirate preko svakog karaktera u argumentu `str`.

# --hints--

Tvoja funkcija od `printCharacters` mora da ima petlju od `for...of` unutar svog kodnog dela.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

Tvoja petlja za `for...of` mora da iterira po argumentu tvoje funkcije `str`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
