---
id: 6921e10ae2681205b516fbe6
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Sada je vreme da dodaš logiku u svoj ciklus koji odgovarajuće dodaje `count`.

Unutar petlje, koristi uslovnu izjavu za dodavanje `count` na `1` ako parametar `word` odgovara parametru `match`. Inače, ostavi vrednost `count` nepromenjenom.

# --hints--

Pozivanje `getMatchedWordCount([], "bar")`, `getMatchedWordCount(["foo"], "bar")`, `getMatchedWordCount(["foo", "bar", "foo"], "bar")`, na `getMatchedWordCount(["foo", "bar", "foo"], "foo")` po tom redosledu trebalo je da vrati vrednost tipa `"number"`; umesto toga, vratio je vrednosti tipa `--fcc-actual--`.

```js
const obj = { method: (arg1, arg2) => typeof getMatchedWordCount(arg1, arg2) };
const spy = __helpers.spyOn(obj, 'method');
obj.method([], "bar");
obj.method(["foo"], "bar");
obj.method(["foo", "bar", "foo"], "bar");
obj.method(["foo", "bar", "foo"], "foo");
assert.deepEqual(spy.returns, ["number", "number", "number", "number"]);
spy.restore();
```

`getMatchedWordCount([], "bar")` trebalo bi da vrati `0`, ali vraća `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount([], "bar"));
assert.strictEqual(resultString, "0");
```

`getMatchedWordCount(["foo"], "bar")` bi trebalo da vrati `0`, ali vraća `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo"], "bar"));
assert.equal(resultString, "0");
```

`getMatchedWordCount(["foo", "bar", "foo"], "bar")` trebalo bi da vrati `1`, ali vraća `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "bar"));
assert.equal(resultString, "1");
```

`getMatchedWordCount(["foo", "bar", "foo"], "foo")` trebalo bi da vrati `2`, ali vraća `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "foo"));
assert.equal(resultString, "2");
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
