---
id: 6921de8d2a111d0341cff702
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Zdravo, tvoja funkcija vraća `0`, ali i dalje ništa ne računa.

Kako biste ovo popravili, unutar `getMatchedWordCount`, kreirajte petlju za `for...of` sa praznim kodnim segmentom koji se ponavlja preko svake reči u `sentence`.

# --hints--

Tvoja funkcija za `getMatchedWordCount` treba da sadrži ciklus od `for...of` unutar svog kodnog dela.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

Tvoj ciklus za `for...of` treba da ponavlja preko argumenta `sentence` tvoje funkcije.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
--fcc-editable-region--
  
--fcc-editable-region--
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
