---
id: 6921d688b8d58581d34459b9
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

U sledećim koracima, kreiraćeš funkciju koja broji koliko puta se niz slova pojavljuje u nizu podataka sa nizom slova.

Definiši praznu funkciju sa imenom `getMatchedWordCount` i parametrima `sentence` i `match` u tom redosledu.

# --hints--

Trebalo bi da kreirate funkciju za `getMatchedWordCount`.

```js
assert.isFunction(getMatchedWordCount);
```

Tvoja funkcija od `getMatchedWordCount` mora imati dva parametra: `sentence` i `match`, u tom redosledu.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
