---
id: 6921dae10700e300f26994e3
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Da biste videli rezultat vaše funkcije ``getMatchedWordCount`` u njenom trenutnom stanju, dodajte deklaraciju ``console.log`` ispod nje. Unutar ``console.log()``, pozovite tu funkciju sa argumentima ``["I", "really", "really", "really", "like", "to", "code"]`` za parametar ``sentence`` i ``"really"`` za parametar ``match``.

# --hints--

Treba da napravite poziv funkcije `getMatchedWordCount` unutar deklaracije `console.log`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

Treba da izvršite poziv funkcije `getMatchedWordCount` sa argumentom `["I", "really", "really", "really", "like", "to", "code"]` za parametar `sentence` i `"really"` za parametar `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
