---
id: 69220a03729c43082ad3e9f6
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Završi sa radom na svom modulu/funkciji `getMatchedWordCount`!

Sada ćete početi da ga testirate sa različitim podacima kako biste videli kako reaguje/odgovara.

Dodajte novu deklaraciju za `console.log` ispod postojećeg koji poziva `getMatchedWordCount` sa ovim parametrima:

- sentence: `["Do", "not", "fear", "the", "dandy", "lion"]`
- match: `"dandy"`

Čestitam! Završili smo ovaj radionica.

# --hints--

Trebalo bi da ima drugu deklaraciju `console.log` koja poziva funkciju `getMatchedWordCount` unutar sebe.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/g)[1];
assert.exists(normalizedConsoleLogCode);
```

U drugoj deklaraciji od `console.log`, izvršiće funkciju `getMatchedWordCount` koristeći `["Do", "not", "fear", "the", "dandy", "lion"]` kao argument za `sentence` i `"dandy"` kao argument za `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/g)[1];
assert.match(normalizedConsoleLogCode, /\[('|"|`)Do\1,('|"|`)not\2,('|"|`)fear\3,('|"|`)the\4,('|"|`)dandy\5,('|"|`)lion\6\],('|"|`)dandy\7/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
    if (word === match) {
      count++;
    }
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
    if (word === match) {
      count++;
    }
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);

console.log(
  getMatchedWordCount(
    ["Do", "not", "fear", "the", "dandy", "lion"],
    "dandy"
  )
);
```
