---
id: 5dc17dc8f86c76b9248c6eb5
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Ostavljanje komentara omogućava vam da ostavite poruku bez uticaja na način na koji pregledač prikazuje stranicu. Takođe vam omogućava da učinite kod neupotrebljivim. Komentari u HTML počinju sa `<!--`, sadrže bilo koje tekstualne linije i završavaju se sa `-->`.

Ovo je primer komentara sa `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Dodajte komentare na sekciju `p` za ovaj tekst:

`TODO: Add link to cat photos`

# --hints--

Vaši komentari bi trebalo da počnu sa `<!--`. Nedostaje vam jedno ili više slova koja označavaju početak komentara.

```js
assert.match(code, /<!--/);
```

Vaši komentari bi trebalo da se završe sa `-->`. Nedostaje vam jedno ili više slova koja označavaju kraj komentara.

```js
assert.match(code, /-->/);
```

Vaš kod ne bi trebalo da ima dodatne karaktere za otvaranje/zatvaranje komentara. Imate `<!--` ili `-->` dodatne karaktere vidljive u pretraživaču.

```js
const noSpaces = code.replace(/\s/g, '');
assert.isBelow(noSpaces.match(/<!--/g)?.length, 2)
assert.isBelow(noSpaces.match(/-->/g)?.length, 2);
```

Vaši komentari bi trebalo da budu iznad elementa `p`. Postavili ste ga u pogrešnom redosledu.

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

Vaš komentar treba da sadrži tekst `TODO: Add link to cat photos`.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
  </body>
</html>
```
