---
id: 5dc1798ff86c76b9248c6eb3
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Ispod elementa `h1`, dodaj element `h2` sa ovim tekstovima:

`Cat Photos`

# --hints--

Tvoj element od `h1` mora imati otvorenu oznaku. Otvarajuće oznake imaju ovu sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Tvoj element `h1` treba da ima zatvarajući tag. Zatvarajući tagovi imaju ovu sintaksu: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Trebalo bi da ima samo jedan element od `h1`. Uklonite višak.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Tekst vašeg elementa/polja `h1` treba da bude 'CatPhotoApp'. Ostavili ste tekst ili postoje greške u pravopisu.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```

Etikete za `h2` bi trebalo da budu u malim slovima. Prema ugovoru, sve etikete za HTML se pišu malim slovima.

```js
assert.notMatch(code, /<\/?H2>/);
```

Tvoj element od `h2` treba da ima otvarajuću oznaku. Otvarajuće oznake imaju ovu sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

Vaš element za `h2` treba da ima zatvarajući tag. Zatvarajući tagovi sadrže `/` nakon karaktera `<`.

```js
assert.match(code, /<\/h2\>/);
```

Sadržaj vašeg elementa za `h2` treba da bude `Cat Photos`. Umestite tekst `Cat Photos` samo između otvarajućih i zatvarajućih oznaka za `h2`.

```js
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat photos');
```

Vaš element za `h2` treba biti ispod elementa `h1`. Element za `h1` je važniji i treba biti iznad elementa `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
