---
id: 5dc17d3bf86c76b9248c6eb4
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Kreiraj element `p` ispod tvog elementa `h2` sa sledećim tekstom:

`Everyone loves cute cats online!`

# --hints--

Vaš element `p` mora imati oznaku za otvaranje. Oznake za otvaranje imaju sledeću sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

Vaše oznake za `p` treba da budu u malim slovima. Prema sporazumu, sve oznake za HTML su napisane malim slovima.

```js
assert.notMatch(code, /<\/?P>/);
```

Vaš element `p` mora imati zatvarajući tag. Zatvarajući tagovi su `/` odmah nakon karaktera `<`.

```js
assert.match(code, /<\/p\>/);
```

Tekst tvog `p` elementa treba da bude `Everyone loves cute cats online!`. Nedostaje tekst ili postoji greška u pisanju.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Tvoj element za `p` bi trebalo da bude ispod elementa za `h2`. Postavio si ga u pogrešan redosled.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
