---
id: 5dc174fcf86c76b9248c6eb2
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, nastavićete da radite sa osnovnim elementima HTML kao naslovima sekcija, pasusima i listama za izradu programa sa slikama mačke.

Počnite radionicu dodavanjem elementa `h1` koji sadrži tekst `CatPhotoApp`.

# --hints--

Reč `CatPhotoApp` mora biti u kodu. Možda želite proveriti pravopis.

```js
assert.match(code, /catphotoapp/i);
```

Vaša komponenta `h1` treba da ima etiketu za otvaranje. Etikete za otvaranje imaju ovu sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Etiketa za `h1` treba da bude u malim slovima. Prema sporazumu, sve etikete za HTML su napisane u malim slovima.

```js
assert.notMatch(code, /<\/?H1>/);
```

Vaš element `h1` treba da ima zatvarajući tag. Zatvarajući tagovi imaju ovu sintaksu: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Tekst elementa `h1` bi trebalo da bude `CatPhotoApp`. Možda ste ostavili te reči, postoje greške u pravopisu, ili se ne nalaze između otvarajućeg i zatvarajućeg taga elementa `h1`.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
