---
id: 5dc24165f86c76b9248c6ebc
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Svi elementi sa `img` moraju imati atribut `alt`. Tekst atributa `alt` koristi ga čitač ekrana za poboljšanje pristupačnosti i prikazuje se ako slika ne učita.

Ovo je primer elementa od `img` sa karakteristikom elementa od `alt`:

```html
<img src="cat.jpg" alt="A cat">
```

Unutar elementa `img`, dodaj svojstvo elementa `alt` sa ovim tekstom:

`A cute orange cat lying on its back`

# --hints--

Tvoj kod bi trebalo da ima element `img`. Ukloni element `img` iz prethodnog koraka.

```js
assert.exists(document.querySelector('img'));
```

Vaš element `img` nema svojstvo elemenata `alt`. Uverite se da postoji razmak nakon imena otvarajuće oznake i/ili da postoji razmak pre svih imena svojstava elemenata.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

Vrednost atributa elementa `img` u elementu `alt` je postavljena da bude drugačija od `A cute orange cat lying on its back`. Proverite da je vrednost atributa elementa `alt` okružena navodnicima.

```js
const altText = document
  .querySelector('img')
  ?.alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert.match(altText, /A cute orange cat lying on its back\.?$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
