---
id: 5dc2385ff86c76b9248c6eb7
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 ima elemente koji identifikuju različite sekcije sadržaja. Ovi elementi čine vaš HTML lakšim za čitanje i pomažu u Optimizaciji pretraživača (SEO) i pristupačnosti.

Element `main` se koristi za predstavljanje glavnog sadržaja tela dokumenta HTML. Sadržaj unutar elementa `main` mora biti jedinstven za taj dokument i ne sme se ponavljati u drugim delovima tog dokumenta.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Pronađite glavni deo ove stranice dodavanjem otvarajućeg taga `<main>` pre elementa `h1`, i zatvarajućeg taga `</main>` nakon elementa `p`.

# --hints--

Element vašeg `main` mora imati oznaku za otvaranje. Oznake za otvaranje imaju ovu sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('main'));
```

Vaše etikete za `main` treba da budu malim slovima. Po sporazumu, sve etikete za HTML su napisane malim slovima.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0]
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

Vaša komponenta za `main` mora imati zatvarajući tag. Zatvarajući tagovi su `/` odmah nakon karaktera `<`.

```js
assert.match(code, /<\/main\>/);
```

Otvarajući tag tvog elementa `main` treba da bude ispod otvarajućeg taga elementa `body`. Postavili ste ih u pogrešnom redosledu.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Otvor za vaš element `main` mora biti iznad elementa `h1`. Postavili ste ih u pogrešnom redosledu.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

Zatvarajući deo tvog elementa ``main`` treba da bude ispod elementa ``p``. Postavili ste ga u pogrešan redosled.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
