---
id: 5dc24073f86c76b9248c6ebb
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Atributi elementa za HTML su posebne reči koje se koriste unutar otvarajućeg taga elementa za kontrolisanje ponašanja tog elementa. Atribut elementa `src` u elementu `img` opisuje/određuje URL slike (mesto gde se slika nalazi).

Ovde je primer elementa za `img` sa svojstvom komponente `src` koje se odnosi na logo freeCodeCamp:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

Unutar postojećeg elementa `img`, dodaj atribut elementu `src` sa ovim URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Vaš kod bi trebalo da sadrži element `img`. Možda ste uklonili element `img` ili niste obavili vrednost atributa elemenata `src` sa navodnicima.

```js
assert.exists(document.querySelector('img'));
```

Vaš element sa `img` treba da ima svojstvo elementa `src`. Možda ste propustili to svojstvo ili postoje greške u kucanju. Proverite da postoji razmak između imena elementa i imena svojstva elementa.

```js
assert.exists(document.querySelector('img')?.src);
```

`src` atribut `img` elementa treba da bude postavljen na `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`. Možda si ostavio pogrešan URL ili postoji greška u kucanju. Velika i mala slova u URL-u su važna.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Iako ste definisali atribut elementa `img` elementa `src` za URL, savetuje se da vrednost atributa elementa uvek obavijate navodnicima.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

Tvoja funkcija za `img` mora biti zatvorena sa `>` ili `/>`.

```js
assert.match(code, /<img\s+src\s*=\s*("|')https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```
