---
id: 5dc23f9bf86c76b9248c6eba
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Možete dodati sliku na vašu veb stranicu koristeći element `img`. Elementi za `img` imaju otvarajući tag bez zatvarajućeg taga. Element koji nema zatvarajući tag poznat je kao <dfn>prazan element</dfn>.

Dodajte funkciju `img` ispod funkcije `p`. Trenutno, slika se neće prikazati u pretraživaču.

# --hints--

Tvoj `img` element treba da ima početnu oznaku. Početne oznake imaju sledeću sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

Vaš element za `img` ne sme imati zatvorni tag `</img>`.

```js
assert.notMatch(code, /<\/img\>/);
```

Trebalo bi da ima samo jedan element `img`. Uklonite sve suvišne stvari.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

Vaš element `img` treba da bude ispod elementa `p`. Postavljen je u pogrešan redosled.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
