---
id: 5dc23991f86c76b9248c6eb8
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

U prethodnom koraku, umetnuli ste elemente `h1`, `h2`, komentare i `p` unutar elementa `main`. Ovo se zove *ugnježđivanje*. Ugnježdene elemente treba postaviti za dva razmakta desnije od elementa u koji su ugnježđeni. Ovaj razmak se naziva prekid linije i koristi se kako bi HTML bio lak za čitanje.

Evo primera za formatiranje unutra i prekid linije:

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Element od `h1`, element od `h2` i komentari imaju uvlačenje za dva razmaka više nego element od `main` u kodu navedenom ispod. Koristite taster razmak (space bar) na vašoj tastaturi da dodate dva razmaka pre elementa od `p`, kako bi i ono bilo pravilno uvlačeno.

# --hints--

Vaš kod bi trebalo da ima element sa `h2` i tekstom `Cat Photos`. Možda ste ga slučajno obrisali, nedostaju otvoreni i zatvoreni tagovi, ili se tekst promenio.

```js
assert.exists(document.querySelector('h2'));
assert.match(code, /<\/h2\>/);
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(),'cat photos')
```

Ne treba da dodajete elemente `ul` ili `li` iz primera.

```js
assert.notExists(document.querySelector('ul'));
assert.notExists(document.querySelector('li'));
```

Ne smete menjati prekid linije na redu sa vašim elementom `h2`. Njegov otvor bi trebalo da počinje na poziciji 6 od početka linije. Možete ponoviti ovaj korak da biste vratili prvobitni prekid linije.

```js
assert.match(code.toLowerCase(), /<\/h1\>\s*\n\s{6}<h2>/);
```

Vaš kod bi trebalo da ima komentare. Izbrisali ste komentare iz prethodnog koraka.

```js
assert.match(code, /<!--.*-->/);
```

Tekst komentara treba da bude `TODO: Add link to cat photos`. Ne menjajte tekst ili lokacije komentara.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i);
```

Ne smete menjati prekid reda na redu sa vašim elementom komentara. Njegova etiketa za otvaranje mora početi na poziciji 6 od početka reda. Možete ponovo inicijalizovati ovaj korak da biste vratili prvobitni prekid reda.

```js
assert.match(
  code
    .toLowerCase(),
    /<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/
);
```

Tvoj kod treba da ima element `p`. Izbrisao si element `p` iz prethodnog koraka.

```js
assert.exists(document.querySelector('p'));
```

Tekst elementa `p` bi trebalo da bude `Everyone loves cute cats online!`. Ne menjajte tekst, razmak ili znakove pisanja elementa `p`.

```js
assert.match(
  document
    .querySelector('p')
    ?.innerText?.trim().toLowerCase(),
    /^everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

Otvarajući tag za `p` treba da ima prekid linije koji odgovara vašim elementima od `h2` i komentarima. Njegov otvarajući tag treba da počinje na poziciji 6 od početka linije.

```js
assert.match(code.toLowerCase(), /-->\s*\n\s{6}<p>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
    </main>
  </body>
</html>
```
