---
id: 5dc24614f86c76b9248c6ebd
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Možete se povezati sa drugom stranicom korišćenjem anker atributa (`a`).

Ovo je primer povezivanja sa `https://www.freecodecamp.org`:

```html
<a href="https://www.freecodecamp.org"></a>
```

Dodajte ankorni element nakon paragrafa koji se povezuje na `https://freecatphotoapp.com`. Do sada, link nije vidljiv u početnom prikazu.

# --hints--

Vaš element sidra (`a`) bi trebalo da ima etiketu za otvaranje. Etikete za otvaranje imaju ovu sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

Tvoj anchor (`a`) element treba da ima završnu oznaku. Završne oznake imaju `/` odmah nakon znaka `<`.

```js
assert.match(code, /<\/a\>/);
```

Tvoj anchor (`a`) element treba da bude ispod `p` elementa. Postavio si ih u pogrešnom redosledu.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

Tvoj ankorni element (`a`) nema atribut `href`. Osigurajte da postoji razmak nakon imena otvarajućeg taga i/ili da postoji razmak pre svih imena atributa.

```js
assert.isTrue(document.querySelector('a')?.hasAttribute('href'));
```

Tvoj anker element (`a`) mora se povezati sa `https://freecatphotoapp.com`. Ostavio ste URL ili ste napravili pravopisne greške.

```js
assert.equal(
  document.querySelector('a')?.getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

Iako ste postavili atribut `a` za atribut sidra (`href`) za ispravnu vezu, uvek se preporučuje obavijanje vrednosti atributa i navodnika.

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
