---
id: 5ddbd81294d8ddc1510a8e56
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Tekst u liniji mora biti između otvarajućeg i zatvarajućeg taga anker elementa (`a`).

Ovde se nalazi primer veze sa tekstom `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Dodajte tekst za sidro `cat photos` u element sidra. Ovo će biti tekst veze/linka.

# --hints--

Vaš element sidra (`a`) bi trebalo da ima etiketu za otvaranje. Etikete za otvaranje imaju ovu sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

Tvoj anker element (`a`) treba da ima zatvarajući tag. Zatvarajući tagovi imaju `/` odmah nakon slova `<`.

```js
assert.match(code, /<\/a\>/);
```

Sadržaj vašeg ankora elementa (`a`) treba da bude `cat photos`. Uverite se da umetnete tekst linka između otvarajućeg i zatvarajućeg taga ankora elementa (`a`).

```js
assert.equal(
  document.querySelector('a')?.innerText?.trim().toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
