---
id: 5dfa2407b521be39a3de7be1
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Za otvaranje linkova u novom tabu, možete koristiti atribut elementa `target` na elementu anchor taga (`a`).

Atribut polja `target` opisuje gde će se povezani dokument otvoriti. `target="_blank"` otvara povezani dokument u novoj karti ili prozoru.

Ovo je osnovna sintaksa za element `a` koji ima atribut elementa `target`:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Dodaj atribut elementu `target` sa vrednošću `_blank` na atribut za otvaranje elementa anker taga (`cat photos` `a`), kako bi se link otvorio u novoj karti.

# --hints--

Vaš element za `p` bi trebalo da ima ugrađen element ankornog taga (anchor tag) `a` sa tekstom `cat photos`. Možda ste ga izbrisali ili postoje pravopisne greške.

```js
const anchor = document.querySelectorAll('p > a')[1];
assert.equal(
    anchor?.innerText?.toLowerCase().trim().replace(/\s+/g, ' '), 'cat photos'
);
```

Tvoj element ankornog taga (`a`) nema atribut elementa `target`. Osigurajte da postoji razmak nakon imena otvarajuće oznake i/ili da postoji razmak pre svih imena atributa elemenata.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

Vrednost atributa elementa `target` bi trebalo da bude `_blank`. Možda ste zaboravili vrednost ili postoji greška u kucanju. Zapamtite da atributi elementa moraju biti obavijeni navodnicima.

```js
assert.equal(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
