---
id: 5dfa3589eacea3f48c6300ae
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

U okviru drugog elementa od `section`, dodajte novi element sa `h2` koji sadrži tekst `Cat Lists`.

# --hints--

Vaš element za `section` bi trebalo da ima otvorajući label. Otvarajući labeli imaju ovu sintaksu: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Tvoj element od `h2` treba da ima otvarajuću oznaku. Otvarajuće oznake imaju ovu sintaksu: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

Tvoj element za `h2` bi trebalo da ima zatvarajući tag. Zatvarajući tagovi imaju `/` odmah nakon karaktera `<`.

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

Vaš drugi element `h2` treba da bude odmah iznad zatvarajućeg oznake drugog elementa `section`. Nije na pravilnoj poziciji.

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection?.lastElementChild?.nodeName, 'H2');
```

Drugi deo od `h2` bi trebalo da ima tekst `Cat Lists`. Ostavili ste tekst prazan ili postoje pravopisne greške.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
