---
id: 5dfa37b9eacea3f48c6300b0
title: Korak 22
challengeType: 0
dashedName: step-22
---

# --description--

Da biste kreirali neuređenu listu stavki, možete koristiti funkciju `ul`.

Nakon elementa `h3` sa tekstom `Things cats love:`, dodajte stavku nepoređenog spiska (`ul`). Zapamtite da trenutno ništa neće biti prikazano.

# --hints--

Tvoj `ul` element treba da ima početnu oznaku. Početne oznake imaju sledeću sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('ul'));
```

Tvoj `ul` element treba da ima završnu oznaku. Završne oznake imaju `/` odmah nakon znaka `<`.

```js
assert.match(code, /<\/ul>/);
```

Element `ul` bi trebalo da bude iznad zatvarajuće oznake drugog elementa od `section`.

```js
const secondSectionLastElemNode =
  document.querySelectorAll('main > section')[1]?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName, 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
