---
id: 5dfb6a35eacea3f48c6300b4
title: Korak 26
challengeType: 0
dashedName: step-26
---

# --description--

Element za opis slika (`figcaption`) se koristi za dodavanje opisa slike koji opisuje sliku koja je unutar elementa `figure`.

Evo primera `figcaption` elementa sa opisom `A cute cat`:

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

Nakon slike ugnježdene u `figure` element, dodaj `figcaption` element sa postavljenim tekstom:

`Cats love lasagna.`

# --hints--

Vaša komponenta za `figcaption` mora imati oznaku za otvaranje. Oznake za otvaranje imaju sledeću sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('figcaption'));
```

Tvoj `figcaption` element treba da ima završnu oznaku. Završne oznake imaju `/` odmah nakon znaka `<`.

```js
assert.match(code, /<\/figcaption\>/);
```

Element `figcaption` treba da bude smešten unutar elementa `figure`.

```js
assert.exists(document.querySelector('figure > figcaption'));
```

Element `img` za lazanje treba da bude smešten unutar elementa `figure`.

```js
assert.exists(document.querySelector('figure > img'))
assert.equal(
    document.querySelector('figure > img')?.getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Element `figcaption` smešten unutar elementa `figure` treba da se nalazi ispod elementa `img`. Postavili ste ih pogrešnim redosledom.

```js
assert.equal(
  document.querySelector('figcaption')?.previousElementSibling.nodeName, 'IMG'
);
```

Tekst vašeg elementa `figcaption` treba da bude `Cats love lasagna.`. Nedostaje tekst ili postoji greška u pisanju.

```js
assert.match(
  document.querySelector('figcaption')?.innerText?.trim(), /Cats love lasagna.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          
--fcc-editable-region--          
        </figure>
      </section>
    </main>
  </body>
</html>
```
