---
id: 5dfb655eeacea3f48c6300b3
title: Korak 25
challengeType: 0
dashedName: step-25
---

# --description--

Element `figure` predstavlja samostalan sadržaj i omogućava povezivanje slike sa opisom slike.

Ugnjezdi sliku koju si nedavno dodao unutar elementa `figure`.

# --hints--

Tvoj `figure` element treba da ima početnu oznaku. Početne oznake imaju sledeću sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('figure'));
```

Tvoj `figure` element treba da ima završnu oznaku. Završne oznake imaju `/` odmah nakon znaka `<`.

```js
assert.match(code, /<\/figure\>/);
```

Element `figure` treba da se nalazi odmah iznad završne oznake drugog `section` elementa.

```js
assert.equal(document.querySelectorAll('section')[1]?.lastElementChild.nodeName, 'FIGURE');
```

`img` element za lazanje treba da bude ugnježden unutar elementa `figure`.

```js
assert.exists(document.querySelector('figure > img'));
assert.equal(
    document.querySelector('figure > img').getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        
        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
