---
id: 5ef9b03c81a63668521804ea
title: Korak 38
challengeType: 0
dashedName: step-38
---

# --description--

Napomenite da je sve što ste do sada postavili na stranicu unutar elementa `body`. Svi elementi sadržaja stranice koji treba da se prikažu na stranici nalaze unutar elementa `body`. Međutim, druge važne informacije nalaze se unutar elementa `head`.

Element ``head`` se koristi za skladištenje meta-podataka o dokumentu, kao što su njegov naslov, veze ka stilskim listovima i skriptama. Meta-podaci su informacije o stranici koje nisu direktno prikazane na toj stranici.

Dodajte funkciju `head` na funkciju `body`.

# --hints--

Obrisao ste element `body` ili mu nedostaje otvarajući ili zatvarajući tag.

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

Vaš element `head` bi trebalo da ima etiketu za otvaranje. Etikete za otvaranje imaju sledeću sintaksu: `<elementName>`.

```js
assert.match(code, /\<head\>/);
```

Vaša komponenta za `head` treba da ima zatvorni tag. Zatvorni tagovi imaju `/` odmah nakon slova `<`.

```js
assert.match(code, /\<\/head\>/);
```

Vaš element `head` treba da bude iznad labela za otvaranje komponente `body`. Postavili ste ga negde drugde.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
