---
id: 5ef9b03c81a63668521804e9
title: Korak 37
challengeType: 0
dashedName: step-37
---

# --description--

Promeni postojeći tekst `freeCodeCamp.org` u link koji kruži oko elementa sidra (`a`). Atribut `href` treba postaviti na `https://www.freecodecamp.org`.

# --hints--

Tvoj ankorni element (`a`) treba da bude postavljen unutar elementa `p`, unutar `footer`. Proverite da dodate otvarajući i zatvarajući tag za ankorni element (`a`).

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

Tvoj anker element (`a`) treba da ima zatvarajući tag. Zatvarajući tagovi imaju `/` odmah nakon karaktera `<`.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

Vaš ankorni element (`a`) mora imati atribut `href` sa vrednošću `https://www.freecodecamp.org`. Možda ste zaboravili atribut/vrednost, ili postoje pravopisne greške.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

Tekst linka bi trebalo da bude `freeCodeCamp.org`. Možda ste zaboravili tekst ili postoje greške u kucanju.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'freecodecamp.org'
);
```

Nakon ugradnje jangkarnog elementa (`a`), sadržaj elementa `p` koji je vidljiv u pretraživaču mora biti `No Copyright - freeCodeCamp.org`. Proverite da ponovo pregledate tekst, razmak ili znake za pisanje za `p` i ugrađeni jangkar element.

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
