---
id: 5ef9b03c81a63668521804eb
title: Korak 39
challengeType: 0
dashedName: step-39
---

# --description--

Element `title` određuje koje pretraživače se prikazuju u traci zaglavlja ili kartici stranice.

Dodaj element za `title` unutar elementa `head` koristeći tekst ispod:

`CatPhotoApp`

# --hints--

Izbrisali ste element za `head` ili mu nedostaje otvarajući ili zatvarajući tag.

```js
assert.match(code, /\<head\>/);
assert.match(code, /\<\/head\>/);
```

Vaš element `title` treba da bude postavljen unutar elementa `head`. Proverite da ste dodali tag za otvaranje i tag za zatvaranje za element `title`.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<head\>\<title\>.*\<\/title\>\<\/head\>/);
```

Vaš element za `title` mora imati zatvarajući label. Zatvarajuće labeli su `/` odmah nakon slova `<`.

```js
assert.match(code, /\<\/title\>/);
```

Tekst vašeg elementa za `title` bi trebalo da bude `CatPhotoApp`. Ostavili ste tekst ili postoje greške u pravopisu.

```js
assert.equal(document.title?.toLowerCase(), 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <head>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
