---
id: 5ef9b03c81a63668521804e7
title: Korak 35
challengeType: 0
dashedName: step-35
---

# --description--

Element ``footer`` se koristi za definisanje u podnožju dokumenta ili sekcije. Podnožje obično sadrži informacije o autoru dokumenta, podatke o autorskim pravima, linkove ka uslovima korišćenja, kontakt informacije i više.

Nakon elementa `main`, dodajte element `footer`.

# --hints--

Izbrisali ste element `main`, ili mu nedostaje otvarajuća ili zatvarajuća oznaka.

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

Vaš element ``footer`` mora imati labelu za otvaranje. Labeli za otvaranje imaju sledeću sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('footer'));
```

Tvoj element za `footer` treba da ima zatvarajući tag. Zatvarajući tagovi imaju `/` odmah nakon slova `<`.

```js
assert.match(code, /<\/footer\>/);
```

Vaša komponenta `footer` mora biti ispod zatvarajućeg taga komponente `main`. Stavio ste je negde drugde.

```js
assert.equal(document.querySelector('main')?.nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
