---
id: 5ef9b03c81a63668521804d0
title: Korak 27
challengeType: 0
dashedName: step-27
---

# --description--

Da biste naglasili određenu reč ili frazu, možete koristiti karakteristiku `em`.

Sisistiza neno `love` u delu `figcaption` pokrivanjem sa elementom naglaska `em`.

# --hints--

Element naglašavanja za `em` treba da ima otvorenu oznaku/tag. Otvarajući tagovi imaju ovu sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('em'));
```

Vaš element naglaska za `em` treba da ima zatvarajući tag. Zatvarajući tagovi su na `/` odmah nakon slova `<`.

```js
assert.match(code, /<\/em\>/);
```

Izbrisali ste element za `figcaption` ili mu nedostaje oznaka za otvaranje ili zatvaranje.

```js
assert.exists(document.querySelector('figcaption'));
assert.match(code, /<\/figcaption\>/);
```

Vaš element za naglašavanje ``em`` bi trebalo da okruži tekst ``love``. Ili ste izostavili taj tekst ili postoje greške u pravopisu.

```js
assert.equal(
  document.querySelector('figcaption > em')?.innerText?.trim().toLowerCase(), 'love'
);
```

Tekst za `figcaption` bi trebalo da bude `Cats love lasagna`. Proverite pravopisne greške i osigurajte da su potrebni razmaci blizu tagova otvaranja i zatvaranja elementa `em`.

```js
assert.match(
  document
    .querySelector('figcaption')
    ?.innerText?.trim().replace(/\s+/gi, ' '),
    /^cats love lasagna\.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
