---
id: 5f07be6ef7412fbad0c5626b
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Pre dodavanja bilo kakvog novog sadržaja, trebalo bi da koristite element ``section`` za razdvajanje sadržaja slika mačaka od budućeg sadržaja.

Element ``section`` se koristi za definisanje sekcija u dokumentu, kao što su poglavlja, naslovi, podnaslovi ili bilo koji drugi deo dokumenta. To je semantički element koji pomaže pri SEO-u i pristupačnosti.

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

Uzmi svoj element od `h2`, dva elementa od `p`, i element za anker tag (`a`) i komponente unutar elementa `section`.

# --hints--

Vaš element za `section` mora imati otvarajući tag. Otvarajući tagovi imaju sledeću sintaksu: `<elementName>`.

```js
assert.exists(document.querySelector('section'));
```

Tvoj element za `section` treba da ima zatvarajući tag. Zatvarni tagovi su `/` odmah nakon slova `<`.

```js
assert.match(code, /<\/section\s*>/i);
```

Onemogućen element od `section` treba biti između otvarajućih i zatvarajućih tagova za `main`.

```js
assert.equal(document.querySelector('section')?.parentNode.nodeName, 'MAIN');
```

Postojeći elementi `h2`, `p`, i element ankora taga (`a`) treba da budu između otvarajućeg i zatvarajućeg taga elementa `section`.

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4)
```

Element `h1` ne smije biti postavljen unutar elementa `section`.

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      
      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
