---
id: 5f07c98cdb9413cbd4b16750
title: Korak 19
challengeType: 0
dashedName: step-19
---

# --description--

Vreme je da se doda nova sekcija. Dodajte funkciju `section` drugi ispod postojećeg elementa/funkcije `section`.

# --hints--

Vaš element za `section` bi trebalo da ima otvorajući label. Otvarajući labeli imaju ovu sintaksu: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

Morate dodati samo jedan otvarajući tag `section`. Molimo vas da uklonite sve suvišnosti.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

Vaš element `section` treba da ima zatvarački tag. Zatvarački tagovi imaju `/` odmah nakon slova `<`.

```js
assert.isAtLeast(code.match(/<\/section>/g)?.length, 2);
```

Trebalo bi da dodaš samo jednu etiketu za zatvaranje `section`. Molimo te, ukloni sve viškove.

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Element drugi od `section` ne bi trebalo da se postavi unutar elementa prvi od `section`.

```js
const el = document.querySelector('main > section');
assert.isNotNull(el);
const childrenOf1stSection = [
  ...el.children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

Oba elementa od `section` bi trebalo da budu između otvarajućeg i zatvarajućeg taga elementa `main`.

```js
const el = document.querySelector('main');
assert.isNotNull(el);
const childrenOfMain = [...el.children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
