---
id: 671141feba228a35cefba82d
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Promeni postojeći tekst `cute cats` u sidro element koji se povezuje sa:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

Nova ankerna funkcionalnost mora da postoji u prvom elementu `p`.

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

Tekst sidra bi trebalo da bude `cute cats`.

```js
assert.strictEqual(document.querySelector('a')?.innerText?.trim(), "cute cats");
```

Atribut elemenata `href` novog ankornog elementa treba da bude `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`.

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

Trebalo bi da postoji samo jedna instanca fraze `cute cats` u vašem kodu.

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

Sadržaj elementa za `p` treba da ostane kao `Everyone loves cute cats online!`.

```js
assert.strictEqual(document.querySelector('p')?.innerText?.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
