---
id: 671b6e873249bb35c9debfcf
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Dodajte tekst `See more ` pre jangornog elementa i ` in our gallery` posle jangornog elementa.

# --hints--

Trebalo bi da koristiš postojeću jangkir funkciju, nemoj kreirati novu funkciju.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Tekst unutar anker elementa ne sme da se menja, mora biti `cat photos`.

```js
assert.strictEqual(document.querySelector('a')?.innerText?.replace(/\s+/g, ' ').trim(), "cat photos");
```

Trebalo bi da imaš `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` u svom kodu.

```js
assert.match(code, /See more\s+<a href="https:\/\/freecatphotoapp\.com">cat photos<\/a>\s+in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
