---
id: 646c5d7057c45f432fcdd46c
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Koristeći selektor klase elemenata, postavi svojstvo za `.cat-head` širinu od `205px` i visinu od `180px`. Takođe, postavi granicu za `1px solid #000` i `border-radius` od `46%`.

# --hints--

Trebalo bi da imaš birač za `.cat-head`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Vaš selektor za `.cat-head` mora imati `width` postavljen na `205px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width, '205px');
```

Vaš selektor za `.cat-head` mora imati `height` postavljen na `180px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height,'180px')
```

Vaš selektor za `.cat-head` mora imati `border`, postavljen na `1px solid #000`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border, '1px solid rgb(0, 0, 0)')
```

Vaš selektor za `.cat-head` mora imati `border-radius` postavljen na `46%`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius, '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
