---
id: 646ce9d790d2a44de5f99e04
title: Korak 19
challengeType: 0
dashedName: step-19
---

# --description--

Kreiraj dva elementa za `div`, prvi unutar elementa `.cat-left-ear` sa klasom `cat-left-inner-ear`, i drugi unutar elementa `.cat-right-ear` sa klasom `cat-right-inner-ear`.

# --hints--

Morate da promenite element `div` koji se nalazi u klasi objekata `cat-left-ear`.

```js
assert.lengthOf(document.querySelectorAll('div.cat-left-ear'), 1);
```

Ne smete da menjate polje `div` koje se nalazi u klasi objekata `cat-right-ear`.

```js
assert.lengthOf(document.querySelectorAll('div.cat-right-ear'), 1);
```

Trebalo bi da imaš jedan element `div` unutar tvog elementa `.cat-left-ear`.

```js
assert.lengthOf(document.querySelectorAll('.cat-left-ear div'), 1);
```

Trebalo bi da imaš jedan element `div` unutar tvog elementa `.cat-right-ear`.

```js
assert.lengthOf(document.querySelectorAll('.cat-right-ear div'), 1);
```

Nova karakteristika `div` unutar `.cat-left-ear` mora imati klasu objekata `cat-left-inner-ear`.

```js
assert.isTrue(document.querySelectorAll('.cat-left-ear div')[0]?.classList.contains('cat-left-inner-ear'));
```

Nova funkcija za `div` unutar `.cat-right-ear` treba da ima klasu objekata `cat-right-inner-ear`.

```js
assert.isTrue(document.querySelectorAll('.cat-right-ear div')[0]?.classList.contains('cat-right-inner-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          --fcc-editable-region--
          <div class="cat-left-ear">

          </div>
          <div class="cat-right-ear">
            
          </div>
          --fcc-editable-region--
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
