---
id: 646cfc2b8e6fe95c20a819d5
title: Korak 30
challengeType: 0
dashedName: step-30
---

# --description--

Uši treba postaviti iznad dela glave koji se preklapa. Ovo možete uraditi koristeći atribut `z-index`.

`z-index` je svojstvo koje možete koristiti za definisanje rasporeda presekujućih elemenata HTML. Svaki element sa gornjim `z-index` će biti postavljen iznad elementa sa donjim `z-index`.

Da biste videli `z-index` kako radi, postavite svojstvo `z-index` levog uha na `-1`.

# --hints--

Tvoj selektor za `.cat-left-ear` mora imati svojstvo `z-index` od `-1`. Ne zaboravi da dodaš tačku.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.zIndex, '-1')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  --fcc-editable-region--
  
  --fcc-editable-region--
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
