---
id: 646ceb843412c74edee27a79
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Koristeći selektor klase objekta, dodaj elementu `.cat-right-ear` svojstva iz `height` i `width` koja su jednaka `100px`. Postavi `background-color` na `white`. Postavi levu i desnu granicu na `35px solid blue`. Postavi gornju i donju granicu na `35px solid red`.

# --hints--

Trebalo bi da imate birač za `.cat-right-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Vaš selektor za `.cat-right-ear` treba da ima svojstvo `height` postavljeno na `100px`. Ne zaboravite da dodate decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '100px')
```

Vaš selektor za `.cat-right-ear` mora imati atribut `width` postavljen na `100px`. Ne zaboravite da dodate decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '100px')
```

Tvoj selektor za `.cat-right-ear` mora imati svojstvo `background-color` postavljeno na `white`. Ne zaboravi da dodaš decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor, 'white')
```

Tvoj selektor za `.cat-right-ear` mora imati svojstvo `border-left` postavljeno na `35px solid blue`. Ne zaboravi da dodaš decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid blue')
```

Tvoj selektor za `.cat-right-ear` mora imati svojstvo `border-right` postavljeno na `35px solid blue`. Ne zaboravi da dodaš decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid blue')
```

Vaš selektor za `.cat-right-ear` mora imati svojstvo `border-top` postavljeno na `35px solid red`. Ne zaboravite da dodate decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid red')
```

Tvoj selektor za `.cat-right-ear` mora imati svojstvo `border-bottom` postavljeno na `35px solid red`. Ne zaboravi da dodaš decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
