---
id: 646c59e942f35541923104bf
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Idite u vaš fajl HTML, kreirajte element `main`. Unutar tog elementa `main`, dodajte element `div` koji ima klasu za stvari `cat-head`.

# --hints--

Trebalo bi da imaš element `main`.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Trebalo bi da imaš element `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

Vaš element za `div` treba da ima klasu objekata `cat-head`.

```js
assert.equal(document.querySelector('div')?.getAttribute('class'), 'cat-head');
```

Element vašeg `div` treba da bude unutar vaše oznake `main`.

```js
assert.lengthOf(document.querySelectorAll('main div'), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
