---
id: 646c48df8674cf2b91020ecc
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Dostavljen vam je osnovni dizajn/struktura. Počnite dodavanjem elementa `link` unutar vaše komponente `head`. Za taj element `link`, postavite svojstvo (property) elementa `rel` na `stylesheet` i `href` na `./styles.css`.

# --hints--

Trebalo bi da imaš element `link`.

```js
assert.exists(document.querySelector("head>link"));
```

Vaš element za `link` treba da ima svojstva elemenata `rel` i `href`, postavljena na `stylesheet` i `./styles.css`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
