---
id: 6476f7a4827bcc61682f2347
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Kada se koristi vrednost `absolute` za vaš atribut `position`, element je uklonjen iz normalnog toka dokumenta, a zatim se njegova pozicija određuje atributima `top`, `right`, `bottom` i `left`.

Postavi atribut `position` tvog elementa `.cat-head` na `absolute`, zatim postavi atribute za `top` i `left` sa bilo kojom pozitivnom pixel vrednošću.

<!-- **Napomena**: Možete pokušati sa funkcijama `top`, `left`, `bottom`, i `right` ovde, ali test će proći samo za `top` od `300px`, i sa leve strane za `400px`. -->

# --hints--

Tvoj selektor za `.cat-head` mora imati svojstvo `position` postavljeno na `absolute`. Proveri da dodaš decimalnu tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

Vaš selektor za `.cat-head` treba da ima atribut `top` podešen na bilo koji broj u `px`. Uverite se da dodajete decimalnu tačku.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/);
```

Tvoj selektor za `.cat-head` mora imati atribut `left` postavljen na bilo koji broj u `px`. Uverite se da dodate decimalnu tačku.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
