---
id: 6476f5c17f99146071ee884c
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Moglo je videti da ništa se ne dešava. Element `.cat-head` nije pomeren uprkos postavljanju `top` i `left` za `100px`, svaki. Ali situacija nije takva za konfiguraciju `relative`.

Kada koristite vrednost `relative`, element je i dalje postavljen u skladu sa normalnim tokom dokumenta, ali vrednosti za `top`, `left`, `bottom`, i `right` postaju aktivne.

Umesto `static`, daj svoj `.cat-head` prostor za `relative`, i ostavi atribute `top` i `left` kako jesu.

# --hints--

Tvoj selektor za `.cat-head` mora imati svojstvo `position` postavljeno na `relative`. Proveri da dodaš decimalnu tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'relative');
```

Vaš selektor za `.cat-head` mora imati atribut `top` postavljen na `100px`. Proverite da dodate decimalnu tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

Tvoj selektor za `.cat-head` mora imati svojstvo `left` postavljeno na `100px`. Uverite se da dodate tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
