---
id: 6476fc5cf14b276e6d04e82a
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Koristi selektor klase objekata da mu dodeli svoj atribut od `.box` širine od `200px`, visine od `600px` i boje pozadine od `#000`. Takođe, dodeli mu `position` za `absolute`, `top` za `800px` i `left` za `650px`.

# --hints--

Trebalo bi da imaš birač za `.box`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.box'));
```

Tvoj selektor za `.box` treba da ima svojstvo `width` postavljeno na `200px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.width, '200px');
```

Tvoj selektor za `.box` mora imati svojstvo `height` postavljeno na `600px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.height, '600px');
```

Vaš selektor za `.box` mora imati svojstvo `background-color` postavljeno na `#000`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor, 'rgb(0, 0, 0)');
```

Vaš selektor za `.box` mora imati svojstvo `position` postavljeno na `absolute`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.position, 'absolute');
```

Tvoj selektor za `.box` mora imati svojstvo `left` postavljeno na `650px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.left, '650px');
```

Selektor za `.box` mora da ima atribut/svojstvo `top` postavljeno na `800px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.top, '800px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

