---
id: 6477062778c85972eb648030
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Sada treba da staviš glavu mačke u sredinu.

Postavi atribut `.cat-head` za element `position` i postavi ga na `absolute`. Postavi vrednost `0` za atribute `right`, `left`, `top`, `bottom`, a zatim postavi njegov atribut `margin` sa svih strana na `auto`. Ovo je jedan način za centriranje elementa vertikalno i horizontalno korišćenjem rasporeda CSS.

# --hints--

Vaš element sa `.cat-head` bi trebalo da ima svojstvo od `position` postavljeno na `absolute`. Proverite da dodate decimalnu tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

Tvoj selektor za `.cat-head` mora imati svojstvo `top` postavljeno na `0`. Uverite se da dodate decimalnu tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

Tvoj selektor za `.cat-head` treba da ima svojstvo `left` postavljeno na `0`. Proveri da dodaš decimalnu tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '0px');
```

Tvoj selektor za `.cat-head` mora imati svojstvo `right` postavljeno na `0`. Uverite se da ste dodali decimalnu tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right, '0px');
```

Tvoj selektor za `.cat-head` treba da ima svojstvo `bottom` postavljeno na `0`. Proveri da dodaš decimalnu tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom, '0px');
```

Tvoj selektor za `.cat-head` mora imati svojstvo `margin` postavljeno na `auto`. Proveri da dodaš decimalnu tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
