---
id: 64770351e8586671ec0911f0
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Kod `sticky` je kombinacija koda `relative` i `fixed`. Omogućava **skupljanje** elementa na određenoj lokaciji unutar kontejnera ili vidljivog prostora, u zavisnosti od pozicije skrolovanja.

Promeni vrednost atributa `position` sa `.cat-head` u `sticky`, promeni `top` u `0`, zatim obriši `left` i njegovu vrednost.

**Napomena**: Da biste videli kako `sticky` radi, morate staviti malo teksta pre i posle vašeg elementa za `.cat-head`. Ako skrolujete dole nakon toga, videćete da `.cat-head` ostaje pri vrhu i zadržava se tamo.

# --hints--

Tvoj selektor za `.cat-head` treba da ima svojstvo od `position` postavljeno na `sticky`. Proverite da dodate decimalnu tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'sticky');
```

Tvoj selektor za `.cat-head` mora imati atribut `top` postavljen na `0`. Proveri da dodaš decimalnu tačku.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

Ne bi trebalo da imaš svojstvo `left` i njegovu vrednost u tvom kodu.

```js
assert.notMatch(code, /left:\s*100px;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
