---
id: 6676a8a8d9861319ee901a90
title: Korak 22
challengeType: 0
dashedName: step-22
---

# --description--

Sada možeš početi da radiš na ušima svog mačka.

Očisti svoj kod za reviziju uklanjanjem selektora `.cat-right-ear` i svih njegovih atributa.

Kada koristite selektor klase objekata, postavite elementu `.cat-left-ear` levne i desne granice od `35px solid transparent` za oba. Takođe, postavite donju granicu na `70px solid #5e5e5e`.

# --hints--


Tvoja selekcija od `.cat-right-ear` ne bi trebalo da ima svojstvo `height`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
}
```

Vaš selektor za `.cat-right-ear` ne bi trebalo da ima svojstvo `width`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
}
```

Vaš selektor za `.cat-right-ear` ne bi trebalo da ima karakteristiku/svojstvo od `border-left`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
}
```

Vaš [objekat] sa `.cat-right-ear` ne bi trebalo da ima karakteristiku od `border-right`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
}
```

Tvoj selektor za `.cat-right-ear` ne bi trebalo da ima svojstvo `border-top`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
}
```

Tvoj selektor za `.cat-right-ear` ne bi trebalo da ima svojstvo `border-bottom`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
}
```

Trebalo bi da imate selektor za `.cat-right-ear`.

```js
assert.notExists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Trebalo bi da imaš birač za `.cat-left-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-left-ear'))
```

Tvoj selektor za `.cat-left-ear` mora imati svojstvo `border-left` postavljeno na `35px solid transparent`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft, '35px solid transparent')
```

Tvoj selektor za `.cat-left-ear` mora imati svojstvo `border-right` postavljeno na `35px solid transparent`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight, '35px solid transparent')
```

Vaš selektor za `.cat-left-ear` mora imati svojstvo `border-bottom` postavljeno na `70px solid #5e5e5e`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom, '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
