---
id: 6676a8b01e56ec1a1e07c202
title: Korak 21
challengeType: 0
dashedName: step-21
---

# --description--

Molimo vas da pogledate da imate beli kvadrat sa debelim okvirima. Ti okviri imaju oštri uglovi koji se mogu koristiti za izradu trouglova. Biće vam potrebni oni za pravljenje uši vašeg mačka.

Unutar selektora klase jednakih elemenata, promenite `height` i `width` u `0`. Postavite levu, desnu i gornju granicu na `transparent`.

Uklonite atribut `background-color` sa ugrađenog trougla.

# --hints--


Trebalo bi da imate birač za `.cat-right-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Tvoj selektor za `.cat-right-ear` mora imati svojstvo `height` postavljeno na `0`. Ne zaboravi da dodaš decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '0px')
```

Selektor vašeg `.cat-right-ear` mora imati svojstvo `width`, postavljeno na `0`. Ne zaboravite da dodate decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '0px')
```

Tvoj vaš račun za `.cat-right-ear` treba da ima atribut `border-left` postavljen na `35px solid transparent`. Ne zaboravi da dodaš tačku i zarez.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid transparent')
```

Tvoj selektor za `.cat-right-ear` mora imati svojstvo od `border-right` postavljeno na `35px solid transparent`. Ne zaboravi da dodaš decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid transparent')
```

Tvoj selektor za `.cat-right-ear` mora imati svojstvo `border-top` postavljeno na `35px solid transparent`. Ne zaboravite da dodate decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid transparent')
```

Tvoj selektor za `.cat-right-ear` treba da ima svojstvo od `border-bottom` postavljeno na `35px solid red`. Ne zaboravi da dodaš decimalnu tačku.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

Vaš izbor od `.cat-right-ear` ne bi trebalo da ima karakteristiku `background-color`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
