---
id: 67ef8f695db52583424bd118
title: Napravi stranicu za pregled filma
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da imate element `main`.
2. Unutar elementa `main`, trebalo bi da postoji element `h1` za naslov filma.
3. Ispod elementa `h1`, treba da postoji element `img` koji prikazuje poster filma. Vaš element `img` bi trebalo da sadrži tekst `alt` koji opisuje sliku. Slobodno možete koristiti sledeću sliku ako želite: `https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`.
4. Trebalo bi da imate element `p` sa kratkim opisom filma.
5. Treba vam još jedan element `p` za prikaz ocene filma. Unutar njega, treba da imate sledeće stavke u navedenom redosledu:
   - Element `strong` sa tekstom `Movie Rating:`.
   - Element `span` sa atributom elementa `aria-hidden` postavljenim na `true` koji prikazuje ocenu u zvezdama `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`.
   - Numerička vrednost koja predstavlja ocenu, unutar običnih zagrada (npr. `9.2/10`) nakon elementa span.
6. Trebalo bi da imate element `h2` sa tekstom `Cast Members`.
7. Trebalo bi da imate element `ul`.
8. Unutar elementa `ul`, trebalo bi da postoji više elemenata `li`, svaki sa elementom `strong` za ime glumca praćeno imenom lika koje prethodi tekst `as`. (npr. `James Holloway as Ethan Carter`).

# --hints--

Trebalo bi da imaš element `main`.

```js
assert.isNotNull(document.querySelector("main"));
```

Unutar elementa `main`, bi trebalo da ima element `h1` sa naslovom filma.

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

Ispod elementa `h1`, trebalo bi da ima element `img` sa svojstvom elementa `alt` koje opisuje.

```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

Trebalo bi da imaš element `p` nakon slike sa filmskim opisom.

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

Trebalo bi da imaš još jedan atribut od tipa `p` sa nivoom.

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

Unutar drugog elementa od `p`, treba da ima element `strong` sa tekstom `Movie Rating: `.

```js
const strongEl = document.querySelector("main p:nth-of-type(2) strong");
assert.match(strongEl?.innerText.trim(), /^Movie Rating:?$/);
```

Element za `span` mora biti prisutan unutar standardnog paragrafa.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

Element `span` unutar standardnog paragrafa mora imati deset zvezdica, popunjene (⭐) ili prazne (☆), praćeno numeričkom vrednošću unutar običnih zagrada nakon `span`.

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

Element `span` unutar standardnog odlomka mora imati atribut elementa `aria-hidden` postavljen na `true`.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

Trebalo bi da imaš element `h2`.

```js
assert.exists(document.querySelector("h2"));
```

Tvoj element `h2` mora da sledi nakon tvog drugog elementa `p`.

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

Polje vašeg `h2` treba da ima tekst `Cast Members`.

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

Trebalo bi da imaš element `ul` nakon `h2`.

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

Funkcija od `ul` mora da ima mnogo funkcija od `li`.

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

Svaki element `li` mora imati element `strong` sa imenom glumca, praćeno imenom lika koje prethodi tekstualnom sadržaju `as`. (primer, `James Holloway as Ethan Carter`).

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const strong = li.querySelector("strong");
  assert.isNotNull(strong);
  const fullText = li.textContent.trim();
  const actorName = strong.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <strong>Movie Rating:</strong>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><strong>James Holloway</strong> as Ethan Carter</li>
        <li><strong>Olivia Sterling</strong> as Lena Mitchell</li>
        <li><strong>William Lancaster</strong> as Professor Adams</li>
        <li>
          <strong>Maria Collins</strong> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
