---
id: 68af273960e009597dd861e2
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Želite čitljiv prikaz za debagovanje i ispisivanje. Kreirajte metod `__str__` koji vraća string sa formatom `title (year) - duration min, director` (gde se `title`, `year`, `duration`, i `director` treba zameniti vrednostima odgovarajućih atributa) kako bi ispisivanje filma pružilo prijatan sažetak.

# --hints--

Your `Movie` class should have a `__str__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

Vaša `__str__` metoda treba da vrati string u formatu `title (year) - duration min, director`, gde `title`, `year`, `duration` i `director` treba zameniti vrednostima odgovarajućih atributa.

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    
--fcc-editable-region--
```
