---
id: 68b804f3be853a355f055fef
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Funkcija `super` se koristi za referenciranje roditeljske klase bez eksplicitnog navođenja njenog imena. Na primer, `super().__init__(arg1, arg2, ...)` poziva metod `__init__` definisan u roditeljskoj klasi.

Trenutno ste efektivno prebrisali metod `__init__`, nasleđen od `Movie`. To znači da će objekat `TVSeries` imati samo atribut `seasons` i atribut `total_episodes`. Ovo nije ono što želite.

Da zadržite sav kod definisan u klasi `Movie` ' `__init__`, koristite `super` za pozivanje roditeljskog `__init__`. Napomena da ćete morati da prosledite argumente koje metod zahteva. U ovom slučaju, `title`, `year`, `director` i `duration`.

# --hints--

You should call `super().__init__(title, year, director, duration)` inside your `__init__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("TVSeries").find_function("__init__").has_stmt("super().__init__(title, year, director, duration)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
class TVSeries(Movie):
    def __init__(self, title, year, director, duration, seasons, total_episodes):
        
        self.seasons = seasons
        self.total_episodes = total_episodes
--fcc-editable-region--
class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'
        
        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

catalogue = MediaCatalogue()

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    catalogue.add(movie1)
    movie2 = Movie('Inception', 2010, 'Christopher Nolan', 148)
    catalogue.add(movie2)
    print(catalogue)
except ValueError as e:
    print(f'Validation Error: {e}')
```
