---
id: 63ee3f71381756f9716727ef
title: Osnovni CSS selektor A
challengeType: 25
dashedName: css-foundations-exercise-a
---

# --description--

**Cilj:**
U ovom vežbanju, ćete vežbati dodavanje CSS u fajl HTML koristeći sve tri metode: CSS eksterni, CSS interni i CSS inline. Trebalo bi da koristite samo selektore tipa za ovo vežbanje prilikom dodavanja stilova putem eksternih i internih metoda. Takođe, trebalo bi da koristite ključne reči za boje (npr. "blue") umesto korišćenja vrednosti `RGB` ili `HEX`.

## Priče za korisnika

- Trebalo bi da vidite element `div` sa tekstom.
  - Mora imati pozadinu od `red`, tekst od `white`, veličinu fonta od `32px`, sa tekstom centriranim pomoću `bold`.
  - CSS za element `div` treba dodati eksternalno, koristeći selector tipa.
- Trebalo bi da vidite element `p` sa tekstom.
  - Mora imati pozadinu od `green`, tekst od `white` i veličinu fonta od `18px`.
  - CSS za element `p` treba dodati interno, koristeći selector tipa.
- Trebalo bi da vidite element `button` sa tekstom.
  - Element `button` mora imati pozadinu od `orange` i veličinu fonta od `18px`.
  - CSS za element `button` treba dodati korišćenjem inline stilova.

# --hints--

Trebalo bi da imaš jedan element sa tekstom `div`.

```js
const divElement = document.querySelector('div');

assert.isNotNull(divElement);
assert.isAtLeast(divElement?.innerText.length, 1);
```

Trebalo bi da imate listu eksternih stilova sa stilom elementa `div`.

```js
const styleSheet = new __helpers.CSSHelp(document).getStyleSheet();
const isExternal = styleSheet?.ownerNode.classList.contains('fcc-injected-styles');
const divStyle = new __helpers.CSSHelp(document).getStyle('div');

assert.isTrue(isExternal);
assert.isNotNull(divStyle);
```

Vaš element `div` ne bi trebalo da ima CSS, koji je dodan korišćenjem internih ili inline stilova.

```js
const styleElement = document.querySelector('style:not([class])');

assert.isNotTrue(styleElement?.innerText.includes('div'));
assert.isNotTrue(document.querySelector('div')?.hasAttribute('style'));
```

Vaša komponenta za `div` treba da ima `background-color` od `red` i `color` od `white`.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const divBGColor = divStyle?.getPropertyValue('background-color');
const divColor = divStyle?.getPropertyValue('color');

assert.equal(divBGColor, 'red');
assert.equal(divColor, 'white');
```

Vaše polje ``div`` treba da ima ``font-weight`` postavljeno na ``bold``, ``font-size`` postavljeno na ``32px``, i ``text-align`` postavljeno na ``center``.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const textAlign = divStyle?.getPropertyValue('text-align');
const fontSize = divStyle?.getPropertyValue('font-size');
const fontWeight = divStyle?.getPropertyValue('font-weight');

assert.equal(textAlign, 'center');
assert.equal(fontSize, '32px');
assert.equal(fontWeight,'bold');
```

Trebalo bi da imaš jedan element sa `p` i taj treba da sadrži tekst.

```js
const pElement = document.querySelector('p');

assert.isNotNull(pElement);
assert.isAtLeast(pElement?.innerText.length, 1)
```

Element vašeg `p` treba da ima stilove dodate unutra koristeći element `style`.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let isStyled = false;

if (rules && rules.selectorText === 'p') {
  isStyled = true;
}

assert.isTrue(isStyled);
```

Vaš element `p` bi trebalo da ima `font-size` od `18px`, a to je postavljeno na `color` i postavljen je na `white`.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let fontSize, color;

if (rules && rules.selectorText === 'p') {
  fontSize = rules.style.fontSize;
  color = rules.style.color;
}

assert.equal(fontSize, "18px");
assert.equal(color, 'white');
```

Trebalo bi da imaš jedan element sa tekstom `button`.

```js
const btnElement = document.querySelector('button');

assert.isNotNull(btnElement);
assert.isAtLeast(btnElement?.innerText.length, 1);
```

Tvoj element za `button` treba da ima inline stil unutar linije.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('style'));
```Vaš element od `button` treba da ima svoj `background-color` podešen na `orange`.

```js
assert.equal(document.querySelector('button')?.style.backgroundColor, 'orange')
```

Vaš element za `button` mora imati `font-size`, postavljen na `18px`.

```js
assert.equal(document.querySelector('button')?.style.fontSize, '18px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Exercise A</title>
  </head>
  <body>

  </body>
</html>
```

```css
/* styles.css */
```

# --solutions--

```html
<!DOCTYPE html>
<html>
<head>
  <title>My Styling Example</title>
  <link rel="stylesheet" type="text/css" href="styles.css">
  <style>
    p {
      background-color: green;
      color: white;
      font-size: 18px;
    }
  </style>
</head>
<body>
  <div>Hello World!</div>
  <p>This is a paragraph.</p>
  <button style="background-color: orange; font-size: 18px;">Click Me</button>
</body>
</html>
```

```css
div {
  background-color: red;
  color: white;
  font-size: 32px;
  text-align: center;
  font-weight: bold;
}
```
