---
id: 63ee3fe4381756f9716727f0
title: Osnovna veĹľba CSS B
challengeType: 25
dashedName: css-foundations-exercise-b
---

# --description--

**Cilj:** Postoji nekoliko komponenti u datoteci HTML koja je dostavljena, koje Ä‡ete morati dodati svojstvo objekta klase ili ID-ja. Zatim Ä‡ete morati dodati pravilo u datoteku CSS koja je dostavljena koristeÄ‡i ispravnu sintaksu selektora.

## PriÄŤe korisnika

1. Trebalo bi da vidite pozadinsku boju za `yellow` za sve stavke sa jedinstvenim brojevima u spisku.

1. Treba da postoji selektor za `class` koji se koristi za sve elemente sa jedinstvenim brojevima u listi.

Trebalo bi da vidite da drugi element u listi ima tekst `blue` i `font-size` od `36px`.

1. `font-size` i boja teksta na drugom elementu treba da se postavi koristeÄ‡i svojstvo polja `id`.

1. Trebalo bi da vidite da treÄ‡i element u spisku ima `font-size` od `24px`.

1. `font-size` u treÄ‡em elementu treba ga postaviti koristeÄ‡i atribut elementa ``class``.

1. Trebalo bi da vidite da ÄŤetvrti element u listi ima pozadinsku boju od `red`, `font-size` na `24px`, i `font-weight` na `bold`.

1. `font-size` ÄŤetvrti elementi mora biti postavljen na svojstvo elementa `class`. `font-weight` i boja mora biti postavljena na svojstvo elementa `id`.

# --hints--

Svaki element sa brojem dvadeset i jedan mora imati svojstvo elementa `class`.

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasClass = p?.every((paragraph) => paragraph.classList.length > 0);

assert(everyPHasClass);
```

VaĹˇi elementi sa jedinstvenim brojevima moraju imati atribut elementa `background-color` od `yellow`.

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasBackgroundColor = p?.every((paragraph) => {
  const style = getComputedStyle(paragraph);
  
  return style?.backgroundColor === 'rgb(255, 255, 0)';
})

assert.equal(everyPHasBackgroundColor, true);
```

Tvoj drugi element treba da ima tekst od `blue` i `font-size` iz `36px`.

```js
const secondElementId = document.querySelectorAll('div')?.[0]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${secondElementId}`);

assert.equal(style?.color, 'blue');
assert.equal(style?.fontSize, '36px');
```

VaĹˇ treÄ‡i element treba da sadrĹľi tekst i `font-size` iz/od `24px`.

```js
const thirdElement = document.querySelectorAll('p')?.[1];

assert(thirdElement?.innerText?.length > 0);

const thirdElementClasses = Array.from(thirdElement?.classList?.values());

assert(thirdElementClasses.some(thirdElementClass => {

  const style = new __helpers.CSSHelp(document).getStyle(`.${thirdElementClass}`);

  return style?.fontSize === '24px';

}))

``` 

ÄŚetvrti element bi trebalo da ima `font-size` od `24px`.

```js
const fourthElementClass = document.querySelectorAll('div')?.[1]?.classList[0];

const style = new __helpers.CSSHelp(document).getStyle(`.${fourthElementClass}`);

assert(style?.fontSize === '24px');
```

ÄŚetvrti element treba da ima `red` `background-color`.

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.backgroundColor === 'red');
```

ÄŚetvrti element treba da ima `font-weight` od `bold`.

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.fontWeight === 'bold');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p>Number 1 - I'm a class!</p>
    <div>Number 2 - I'm one ID.</div>
    <p>Number 3 - I'm a class, but cooler!</p>
    <div>Number 4 - I'm another ID.</div>
    <p>Number 5 - I'm a class!</p>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p class="odd">Number 1 - I'm a class!</p>
    <div id="two">Number 2 - I'm one ID.</div>
    <p class="odd adjust-font-size">Number 3 - I'm a class, but cooler!</p>
    <div id="four" class="adjust-font-size">Number 4 - I'm another ID.</div>
    <p class="odd">Number 5 - I'm a class!</p>
  </body>
</html>
```

```css
.odd {
  background-color: yellow;
  font-family: Verdana, "DejaVu Sans", sans-serif;
}

.adjust-font-size {
  font-size: 24px;
}

#two {
  color: blue;
  font-size: 36px;
}

#four {
  background-color: red;
  font-weight: bold;
}
```
