---
id: 694bf6eb530e19a9c48c59f1
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Trenutno, vaša adresa izgleda nepotpuna. Takođe želite dodati broj etaže gde živi zaposleni, pa stoga trebate izvršiti izmene.

Ako želite dodati sadržaj na kraju postojećeg niza znakova, možete koristiti izraz **mgawo wa thamani ulioboreshwa** ,`+=`. Ovo je kraće od pisanja `var = var + 'new text'`. Primer:

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

Zapamtite da se sekvenca slova ne može promeniti, stoga ova operacija ne menja originalnu sekvencu slova. Umesto toga, ona kreira novu sekvencu slova i postavlja je u izmene.

Koristi atribut `+=` da dodaš niz znakova `, Apartment 4B` u tvoje izmene za `address`.

# --hints--

Trebalo bi da koristiš atribut `+=` za dodavanje sekvence slova `, Apartment 4B` u izmene `address`.

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
