---
id: 694bf6eb530e19a9c48c59f7
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Takođe možete koristiti negativan broj da biste uzeli deo sekvence karaktera od kraja. Na primer, `-1` predstavlja poslednji karakter, `-2` predstavlja drugi karakter od kraja, i tako dalje.

Da biste dobili tri poslednja karaktera sekvence, možete koristiti `string[-3:]`. Proverite kako parametar isključenja ostaje neuključen nakon vertikalne linije (pipe). To znači da bi izdvajanje dela sekvence trebalo da se nastavi do kraja sekvence.

Kreiraj parametar nazvan `last_three`. Koristi negativan obrazac da uzme poslednja tri karaktera iz `employee_code` (koji predstavljaju identifikacioni broj). Na kraju, ispiši `last_three` na terminalu.

Dakle, radionica za zaposlenu karticu je završena.

# --hints--

Trebalo bi da imaš parametar koji se zove `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

Trebalo bi da uzmeš poslednja tri slova od `employee_code` koristeći negativni obrazac i staviš ih u svoj parametar `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

Treba da zalepiš `last_three` u terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
