---
id: 694bf6ea530e19a9c48c59f0
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Ako kombinujete dve sekvence karaktera kao `'John' + 'Doe'`, rezultat je `'JohnDoe'` bez razmaka. Da biste ovo popravili, morate kombinovati sekvencu karaktera koja ima razmak (`' '`) između njih.

Ispravite izmene promenljive `full_name` ili spojite `first_name`, razmak i `last_name`.

# --hints--

Treba da dodaš razmak (`' '`) između `first_name` i `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
