---
id: 694bf6eb530e19a9c48c59f5
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Kada radiš sa nizom znakova, često ćeš morati da izvučeš određeni deo tog niza. To se zove **seckanje (slicing)**.

Sintaksa je `string[start:stop]`, gde:

* `start` je obrazac koji označava početak sečenja (**uključujući**).
* `stop` je obrazac koji označava kraj sečenja (**isključujući**).

Na primer, ako `text = 'Python'`, onda `text[0:2]` izbacuje `'Py'`.

Definiši `employee_code` kao `'DEV-2026-JD-001'`. Nakon toga, kreiraj parametar `department` i postavi vrednost sekcije/delova `employee_code` od primera `0` do `3`. Zatim ispisi `department` na terminal.

# --hints--
Trebalo bi da imaš parametar nazvan `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

Trebalo bi da unesete niz znakova `DEV-2026-JD-001` u parametar vašeg `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

Trebalo bi da ima kriterijum nazvan `department`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

Moraš da izvučeš prvih tri karaktera iz `employee_code` i staviš ih u svoj parametar/polje `department`.

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

Trebalo bi da odštampate `department` u terminalu.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
