---
id: 694bf6eb530e19a9c48c59f4
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Spajanje sekvenca mnogo slova korišćenjem `+` i menjanje brojeva korišćenjem `str()` može biti teško i izazvati zbunjivanje prilikom čitanja.

Python 3.6 uvodi **f-strings** za rešavanje ovoga. Dodavanjem slova `f` pre otvarajuće navodnike, možete postaviti varijable i matematičke izraze unutar delova za ubacivanje predstavljenih vitičnim zagradama `{}`. Primer:

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

Kreiraj parametar `employee_card` i postavi mu vrednost f-stringa koja prikazuje `Employee:` praćen razmakom i vrednošću parametra `full_name`.

# --hints--

Trebalo bi da imate parametar nazvan `employee_card`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

Vaš parametar za `employee_card` treba imati vrednost `f'Employee: {full_name}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
