---
id: 694bf6eb530e19a9c48c59f3
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Sada ćete koristiti funkciju `str()` još jednom. Kao i kod uzrasta, morate konvertovati bilo koju numeričku vrednost u niz karaktera pre spajanja sa drugim tekstom.

Kreirajte parametar nazvan `experience_years` i postavite vrednost punog broja `5`.

Zatim, kreiraj parametar `experience_info`. Postavi vrednost stringa dobijenog konkateniranjem `'Experience: '`, parametra `experience_years` (koji je konvertovan u string), i `' years'`. Ispiši rezultate na terminal.

# --hints--

Trebalo bi da imaš parametar nazvan `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

Trebalo bi da unesete celobrojnu vrednost `5` u vaš parametar `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

Trebalo bi da imate parametar nazvan `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

Trebalo bi da postavite vrednost niza znakova dobijenu spajanjem ``'Experience: '``, ``str(experience_years)``, i ``' years'`` u vaš parametar ``experience_info``.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

Trebalo bi da odštampate `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
