---
id: 694bf6eb530e19a9c48c59f2
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Sada kompletiraj rečenicu ispravljanjem dodeljivanja vrednosti za `employee_info` i takođe spoji niz znakova ` years old` na kraju. Zapamti da uključi razmak na početku svog niza znakova.

Konačno, objavi `employee_info`.

# --hints--

Trebalo bi da povežete `' years old'` na kraj parametra `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age) + ' years old'")`)) })
```

Trebalo bi da odštampate `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_info)")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + str(employee_age)

--fcc-editable-region-- 
```
