---
id: 694bf6eb530e19a9c48c59f6
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Možete izrezati bilo koji deo sekvence karaktera, ne samo početak. I ovo je korisno u mnogim situacijama.

Kreiraj parametar `year_code` i postavi mu vrednost kao deo `employee_code` počevši od modela `4` do `8`. Ovo će dati `2026`.

Napravite parametar `initials` i postavi mu vrednost kao deo `employee_code`, počevši od modela `9` do `11`. Ovo će vratiti `JD`.

Konačno, ispišite oba parametra na terminal.

# --hints--

Trebalo bi da imaš parametar nazvan `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

Trebalo bi vam je da izaberete `employee_code` počevši od elementa `4` do `8` i unesete rezultate u vaš parametar/polje `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

Trebalo bi da imaš parametar nazvan `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

Trebalo bi ti da izvučeš `employee_code` počevši od elementa `9` do `11` i stavi rezultate u svoj parametar za `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

Moraš da zalepiš `year_code` u terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

Treba da odštampate `initials` u terminalu.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
