---
id: 695c9f21a8b1d53cf1b97841
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Sada kreiraj parametar nazvan `employee_age` i postavi mu vrednost na pun broj `28`.


# --hints--

Trebalo bi da imaš parametar koji se zove `employee_age`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

Potrebno je da unesete kompletnu numeričku vrednost `28` za vaš parametar `employee_age`. Nemojte okruživati tekst navodnicima.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
