---
id: 695c9c7c61d446f1451e0ae7
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Sledeće, napravite izmene `address` za čuvanje adrese zaposlenog. Postavite mu vrednost na niz karaktera `123 Main Street`, i na kraju ispišite `address`.

# --hints--

Trebalo bi da imaš promenu sa imenom `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

Trebalo bi da postavite vrednost sekvence slova `123 Main Street` za vaše izmene u `address`. Zapamtite da tekst stavite između pojedinačnih ili dvostrukih navodnika.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

Trebalo bi da objaviš promene za `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```
