---
id: 695c96f82af4fb78e59fa9ed
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Niz znakova je sekvenca znakova koja se koristi za skladištenje tekstualnih podataka. Kao što ste možda zapamtili iz prethodnih lekcija, možete kreirati niz znakova postavljanjem teksta unutar pojedinačnih navodnika (`'`) ili dvostrukih navodnika (`"`). Primer:

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

Kreirajte promenljivu `first_name` koja sadrži nisku `John` i promenljivu `last_name` koja sadrži nisku `Doe`. Zatim ispišite `first_name` i `last_name`.

# --hints--

Trebalo bi da imaš parametar nazvan `first_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

Trebalo bi da postaviš vrednost sekvence slova `John` u svoj parametar `first_name`. Zapamti da tekst mora biti između pojedinačnih ili dvostrukih navodnika.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

Trebalo bi da imaš parametar nazvan `last_name`.

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

Trebalo bi da postavite vrednost sekvence slova `Doe` u vaš parametar `last_name`. Zapamtite da tekst stavite između pojedinačnih ili dvostrukih navodnika.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

Trebaće ti da odštampaš svoj parametar `first_name` u terminalu.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

Treba da odštampate vaš parametar `last_name` u terminalu.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
