---
id: 695df6480de23758af5a29cd
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Sada, želiš da kreiraš niz znakova koji pokazuje godine zaposlenog.

Počnite kreiranjem parametra `employee_info` i postavite mu vrednost rezultata spajanja:

- Atribut `full_name`.
- Niz znakova koji sadrži slovo `is`, prethodno i sledeće ga prati razmak.

# --hints--

Trebalo bi da imaš parametar koji se zove `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

Treba da postavite vrednost `full_name + ' is '` za vaš parametar `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
