---
id: 695ea1e55c0e619e597fc5c9
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Kako vidite, Python je proizvelo `TypeError: can only concatenate str (not "int") to str`. Ovo se dešava zato što Python nije dozvolilo kombinovanje teksta (sekvenca slova) i brojeva (pun broj) direktno.

Da biste ovo ispravili, morate prvo promeniti broj u niz slova koristeći funkciju ``str()``, koja vraća verziju niza slova objekta:

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

Ispraviti mapiranje za `employee_info` promenom `employee_age` u niz znakova koristeći `str(employee_age)`.

# --hints--

Trebalo bi da promenite dodelu vrednosti vašeg parametra `employee_info` tako što ćete promeniti `employee_age` u niz karaktera koristeći `str(employee_age)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
