---
id: 695fcfbba648b3a96a9495a4
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Izlaz u terminalu postaje sve veći. Pre nastavka, potrebno je da ga očistimo.

Uklonite sve pozive `print()` iz vašeg koda.

# --hints--

Ne smete imati nijedan poziv za `print` u vašem kodu.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_calls("print")) == 0`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'  
last_name = 'Doe'  
--fcc-editable-region--  
print(first_name)  
print(last_name)  
full_name = first_name + ' ' + last_name  
print(full_name)  
address = '123 Main Street'  
address += ', Apartment 4B'  
print(address)  
--fcc-editable-region--
```
