---
id: 695fd3dde03273875e5d6709
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Ispravite polje vrednosti za `employee_info` ili takođe povežite `employee_age` na kraju.

Kada ovo uradite, videćete `TypeError` u terminalu. U sledećem koraku, radićete na popravci toga.

# --hints--

Trebalo bi da ispraviš raspon vrednosti za `employee_info` ili dodaš `employee_age` na kraj.

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;  
assert.match(__helpers.python.removeComments(code), regex);  
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
