---
id: 69727ea5c9606bedfe36d3d2
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Sada je vreme za dodavanje završnih detalja na karticu.

Kreiraj varijablu nazvanu `position` sa vrednošću niza slova `Data Analyst` i varijablu nazvanu `salary` sa vrednošću celog broja `75000`.

Zatim, ažuriraj svoju f-string za `employee_card` ili spoji poziciju i platu. Mora da prati ovaj tačan format: `Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`. Promeni zamenske elemente (placeholders) i odgovarajuće parametre.

Konačno, pokrenite `employee_card` da biste videli rezultate.

# --hints--

Trebalo bi da imate predstavnika koji se zove `position`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

Trebalo bi da unesete vrednost niza slova `Data Analyst` u vaš atribut `position`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

Trebalo bi da imate predstavnika koji se zove `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("salary")`)) })
```

Treba da unesete vrednost kompletnog broja `75000` u vaš placeholder za `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

Trebalo bi da ažuriraš `employee_card` koristeći f-string `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

Treba da zalepiš `employee_card` u terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
