---
id: 6980d3ecfe800d8a54130f30
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Trenutno, `employee_card` prikazuje samo ime zaposlenog. Sada ćete dodati više informacija unutar njega.

Ispravite raspored vrednosti `employee_card` tako da uključuje uzrast zaposlenog. Konačni niz treba da izgleda ovako: `Employee: [name] | Age: [age]`, gde je `[name]` zamenjen imenom zaposlenog, a `[age]` njegovim uzrastom.

# --hints--

Trebalo bi da ispraviš `employee_card` ili koristi f-string `f'Employee: {full_name} | Age: {employee_age}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
