---
id: 698367b9d1c6914a22095aec
title: Algoritam za Bubble Sort
challengeType: 26
dashedName: implement-the-bubble-sort-algorithm
---

# --description--

Za ovu vežbu, implementiraćete algoritam bubble sort. Počinje na početku neuglavljenog niza podataka i „diže“ neuglavljene vrednosti ka kraju, ponavljajući kroz niz podataka dok nije potpuno uglavljen. Radi ovo poređenjem dve susedne elemente i zamena njihovih pozicija ako nisu u redosledu. Ovaj metod nastavlja da se ponavlja kroz niz podataka sve dok ne dođe do situacije kada nema promena, a tada je niz podataka sortiran.

Ova metoda zahteva iteriranje kroz skup podataka, a za prosečne i najgore slučajeve ima kvadratnu vremensku složenost. Iako je jednostavna, obično nije pogodno u mnogim slučajevima.

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili vežbu.

Priče o korisniku:

1. Napiši funkciju `bubbleSort` koja prima niz kompletnih brojeva kao ulaz i vraća niz ovih brojeva sortiranih od najmanjeg do najvećeg.

# --hints--

Trebalo bi da imaš funkciju koja se zove `bubbleSort`.

```js
assert.isFunction(bubbleSort);
```

`bubbleSort` treba da vrati sortirani skup podataka (od najmanjeg do najvećeg).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    bubbleSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`bubbleSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` Trebalo bi da vrati nepromenjenu rednu podatke osim za sortiranje.

```js
assert.sameMembers(
  bubbleSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`bubbleSort` ne bi trebalo da koristi način koji je ugrađen u `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    bubbleSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function bubbleSort(array) {
  for (let i = 0; i < array.length; i++) {
    let swapped = false;
    for (let j = 1; j < array.length; j++) {
      if (array[j - 1] > array[j]) {
        let temp = array[j-1];
        array[j-1] =  array[j];
        array[j] = temp;
        swapped = true;
      }
    }
    if (swapped === false) {
      break;
    }
  }
  return array;
}
```
