---
id: 6852db755c9c1274d75aae9a
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Na ovom radincu ćete kreirati ekstrakter PIN brojeva; cifre PIN-a su skrivene u svakom stihu pesme.

Počnite kreiranjem funkcije nazvane `pin_extractor` sa parametrom `poem`. Prazna funkcija daje grešku, pa privremeno dodajte `pass` unutar funkcije da je učinite važećom.

# --hints--

You should have a function named `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

The `pin_extractor` function should have a parameter named `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
