---
id: 68530231c05b93324e071104
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Dok radiš na funkciji, bilo bi korisno imati poziv funkcije za namenu debagovanja, pa kreiraj varijablu `poem`, van funkcije, koja sadrži ovaj poem:

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

Možeš koristiti višelinijski niz za to.

I onda, pozovi funkciju sa varijablom `poem` kao argumentom.

# --hints--

You should have a global variable named `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

The `poem` variable should have a value equal to `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`.

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

You should call the function `pin_extractor` with `poem` as argument.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
