---
id: 68531fbb934d735415053274
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Sada kada ste dodali poziv funkcije, možete videti bilo koju grešku kreiranu unutar funkcije u terminalu, i uvek možete koristiti `print` da biste videli vrednost stvari.

U svakoj liniji je skriven brojčani znak PIN-a, pa unutar funkcije koristite metod `split` da biste podelili niz u listu linija. Podelite linije po karakteru za novi red (`\n`) i dodeljujte rezultuju listi promenljivoj nazvanoj `lines`.

# --hints--

You should have a variable named `lines` within your `pin_extractor` function.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

The variable `lines` should have a value of `poem.split('\n')`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
