---
id: 6855918e87c33c6979d21d3f
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Sada funkcija dobro radi u izvlačenju koda iz pesme, ali postoji grančan slučaj koji treba uzeti u obzir: stih pesme bi mogao biti kraći nego što se očekuje, a to bi uzrokovalo grešku u funkciji (možete pokušati da ovo uradite uklanjanjem reči `bright` iz pesme).

Stavite liniju `secret_code += str(len(words[line_index]))` u izjavu `if` koja proverava da li ima dovoljno reči u listi `words`.

# --hints--

Inside the loop you should have an `if` statement.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

The if statement condition should be `len(words) > line_index`.


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

Inside the `if` statement you should have `secret_code += str(len(words[line_index]))`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
