---
id: 685578aae65e4106ad2b4300
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Kao što ste naučili u prethodnoj lekciji, funkcija `enumerate` omogućava vam da pratite indeks dok iterirate preko nečega što je iterabilno:

```py
fruits = ['apple', 'plum', 'bananas']

for index, item in enumerate(fruits):
    print(index, item)

# 0 apple
# 1 plum
# 2 bananas
```

n-ta cifra tajnog koda dolazi iz n-te reči na n-tom redu, pa morate znati koji je broj svakog reda.

Izmenite petlju tako da iterira na `enumerate(lines)`, i dodajte još jednu promenljivu petlje pre `line` pod nazivom `line_index`.

Takođe, ažurirajte poziv `print` na `print(line_index, line)`.

# --hints--

The loop should iterate over `enumerate(lines)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

The loop should have `line_index, line` as loop variable.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

You should have `print(line_index, line)` inside the loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    for line in lines:
        print(line)
--fcc-editable-region--
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
