---
id: 685535bf527b2c77236b3a9a
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

n-ta cifra PIN-a je skrivena kao dužina n-te reči u n-tom redu.

Da biste pronašli dužinu n-te reči, sledeći korak je razdvojiti red pesme u listu reči: unutar petlje, kreirajte varijablu `words` i dodelite joj vrednost `line` podeljenu na reči korišćenjem metode `split`.

Zatim, dodajte poziv `print` sa `words` kao njegovim argumentom.

# --hints--

You should have a `words` variable inside the loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

The `words` variable should have a value of `line.split()`.

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

You should have `print(words)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
