---
id: 685929ee6d500c6f2bce6751
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Pronađite funkciju ``print`` i prosledite joj vrednost iz ``pin_extractor(poem)`` kako biste videli šta ta funkcija vraća.

# --hints--

You should have the function call `print(pin_extractor(poem))`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('print(pin_extractor(poem))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        secret_code += str(len(words[line_index]))
    return secret_code


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

--fcc-editable-region--
pin_extractor(poem)
--fcc-editable-region--

```
