---
id: 68592730aaf34b61ae279168
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Dužina vraćena iz `len` je ceo broj, potrebno je da ga konvertujete u string korišćenjem funkcije `str` kako biste mogli da ga dodate u string `secret_code`.

Ažurirajte poziv `print` da ispiše dužinu kao string.

# --hints--

You should have `print(str(len(words[line_index])))` at the end of the loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(len(words[line_index]))
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
