---
id: 683ec95f2f8781355556ff7e
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Sada si izračunao model `mid`, sada moraš dobiti vrednost iz `search_list` u tom modelu. Ova vrednost će biti upoređena sa tvojim parametrom `value`, koji je cilj koji tražiš.

Kreiraj varijaciju od `value_at_middle` i postavi njegovu vrednost na element iz `search_list` koji se nalazi u modelu `mid`.

# --hints--

Trebalo bi da napraviš razliku od `value_at_middle` koja je postavljena na vrednost `search_list[mid]`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
--fcc-editable-region--
        
--fcc-editable-region--
```
