---
id: 683ec95f2f8781355556ff84
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Ako je uslov u `elif` tačan, onda ukloni glavnu reč iz `pass` i ažuriraj vrednost promene za `low` dodavanjem `1` za promenu `mid`.

Ovo će proširiti pretragu u desnu polovinu trenutnih zona pretrage na listi, jer ako je `value` veće od `value_at_middle`, to znači da `value` mora biti u desnoj polovini trenutne zone pretrage.

# --hints--

Trebalo bi da uklonite glavnu reč `pass` i ažurirate promene od `low` u `mid + 1`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_bodies()[1].is_equivalent("low = mid + 1")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
--fcc-editable-region--
            pass
--fcc-editable-region--    
        break
    
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4)) 
```
