---
id: 683ec95f2f8781355556ff83
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Možete da vidite da prvi poziv vraća `[3]`. Ovo je zato što `3` je srednja vrednost u `[1, 2, 3, 4, 5]`.

S druge strane, drugi poziv vraća prazan spisak. To se događa zato što, nakon proveravanja srednje vrednosti, petlja trenutno nema logiku za smanjenje opsega pretrage korišćenjem parametara `low` ili `high`. On samo izvršava jednu proveru srednje vrednosti, a zatim završava bez daljeg pretraživanja.

Da bi se omogućilo da binarno pretraživanje nastavi smanjujući svoju pretragu, dodajte element `elif` koji proverava da li je `value` veće od `value_at_middle`. Dodajte glavnu reč `pass` unutar `elif` za sada.

# --hints--

Trebalo bi da dodaš element `elif` kada je `value` veći od `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        
--fcc-editable-region--
        break
    
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
