---
id: 683ec95f2f8781355556ff7c
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Sada, nakon definisanja granica vašeg pretraživanja, vreme je da kreirate petlju koja će implementirati algoritam binarne pretrage.

Binarna pretraga radi tako što postupno smanjuje prostor pretrage. Ovaj proces se nastavlja sve dok postoji važeća lista elemenata za proveru.

Možete prikazati ovo za strugu `while`, koja će nastaviti projekat samo ako je vaš indikator `low` ispod ili jednak indikatoru `high`.

Ako ``low`` je veće od ``high``, to znači da je prostor pretrage istekao, i da ``value`` nije na spisku.

Unutar funkcije `binary_search`, kreiraj petlju `while` sa uslovom da se proveri da li `low` je nisko ili jednako visoko. Unutar te petlje, dodaj glavnu reč `pass` kao privremeni placeholder.

# --hints--

Treba da kreirate krug za `while` koji nastavlja projekat `low` nalazi se ispod ili jednako sa `high`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
