---
id: 683ec95f2f8781355556ff81
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Moraš da isprobaš stvari kako bi razumeo tok algoritma u ovoj početnoj fazi.

Da biste to uradili, prvo morate izaći iz ciklusa. To je zato što trenutno izvršavanje će dozvoliti samo jedan ciklus, pa ako se uslov ne ispuni, biće beskonačna petlja.

Nakon što se izvrši `if`izjava, koristi glavnu reč `break` iz petlje `while`. Zatim, nakon petlje `while`, vrati prazan spisak da bi ukazao na to da vrednost nije pronađena.

# --hints--

Trebalo bi da koristiš ključnu reč `break` iz konteksta `while` nakon pojma `if`.

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

Trebalo bi da vratite praznu listu nakon vašeg kruga `while`.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        
--fcc-editable-region--
```
