---
id: 683ec95f2f8781355556ff7a
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Algoritam binarnog pretraživanja koji gradite neće samo vratiti ``True`` ili ``False`` da pokaže da li je vrednost pronađena ili ne. Takođe će vratiti putanju koja je korišćena za dobijanje te vrednosti. I za to, imaćete razliku od ``path_to_target`` da pratite tu putanju.

Unutar vaše funkcije `binary_search`, promenite reč `pass` u vrednost `path_to_target` i počnite sa praznom listom.

# --hints--

Trebalo bi da kreirate razliku od `path_to_target` koja je postavljena kao prazna lista unutar vaše funkcije `binary_search`. Ne zaboravite da uklonite reč `pass`.

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
