---
id: 683ec95f2f8781355556ff80
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Sledeći korak je kreiranje stanja koje će proveriti da li `value` cilj je u sredini.

Unutar petlje `while`, proverava se da li `if` proveri da li je `value` jednak `value_at_middle`. Ako jeste, vratite vrednost promene za `path_to_target`.

# --hints--

Treba da napraviš `if` izraz koji proverava da li je `value` jednako `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("value == value_at_middle") or _cond.is_equivalent("value_at_middle == value")
`) })
```

Trebalo bi da vratiš izmene `path_to_target` unutar izraza `if`.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
--fcc-editable-region--
        
--fcc-editable-region--
```
