---
id: 683ec95f2f8781355556ff7b
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

U binarnoj pretrazi, postoji minimalni i maksimalni mogući indeks.

Najniži mogući indeks postavlja apsolutnu levu granicu u trenutnom opsegu pretraživanih vrednosti, u ovom slučaju, `search_list`. Najviši mogući indeks, s druge strane, postavlja apsolutnu desnu granicu u pretraživanim vrednostima.

Da biste uzeli u obzir ta dva, morate uzeti u obzir celu traženu listu. Stoga, unutar funkcije `binary_search`, definišite parametar `low` sa vrednošću `0`, i parametar `high` sa vrednošću koja uzima u obzir poslednji element u traženoj listi.

# --hints--

Trebalo bi ti da imaš parametar koji se zove `low` u tvojoj funkciji `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

Vaš parametar za `low` treba da bude postavljen na `0`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

Trebalo bi da imaš parametar koji se zove `high` u svojoj funkciji `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

Vaš parametar za `high` treba da se postavi na `len(search_list) - 1`.

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
--fcc-editable-region--
    
--fcc-editable-region--
```
