---
id: 683ec95f2f8781355556ff7d
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Da bi se započeo ciklus/petlja, potrebno je dobiti središnju vrednost trenutne pozicije pretrage.

Možete dobiti medijanu izračunavanjem proseka parametara za `low` i `high`. Morate koristiti podjelu sa zjenom da biste dobili prosek nakon sabiranja te dve vrednosti, tako da će odgovor uvek biti ceo broj zaokružen nadole, jer parametri moraju biti ceo brojevi.

Ukloni reč iz `pass` iz ciklusa `while` i kreirajte parametar `mid` postavljen na prosečnu vrednost od `low` i `high`.

# --hints--

Treba da uklonite reč `pass` iz opsega while i kreirate parametar `mid` definisan za `(low + high) // 2`.

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
--fcc-editable-region--
        pass
--fcc-editable-region--
```
