---
id: 683ec95f2f8781355556ff7f
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Sada, uključite `value_at_middle` u spisak `path_to_target` za praćenje preduzetih koraka tokom pretraživanja, bez obzira da li je to cilj.

# --hints--

Trebalo bi da koristiš putanju `append()` za dodavanje `value_at_middle` u listu `path_to_target`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid] \\npath_to_target.append(value_at_middle)")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
--fcc-editable-region--
        
--fcc-editable-region--
```
