---
id: 683ec95f2f8781355556ff82
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Zapamti da funkcija prima listu pod nazivom `search_list` i `value` koju tražiš.

Sada, izvrši funkzioni poziv za `binary_search([1, 2, 3, 4, 5], 3)` i `binary_search([1, 2, 3, 4, 5, 9], 4)` i ispisi te pozive odmah.

# --hints--

Trebalo bi da napravite funkcionalski poziv za `binary_search([1, 2, 3, 4, 5], 3)` i odštampate ga odmah.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5], 3))")
    `))
})
```

Trebalo bi da napravite funkcionalski poziv za `binary_search([1, 2, 3, 4, 5, 9], 4)` i odštapišete ga odmah.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5, 9], 4))")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
        break
    
    return []

--fcc-editable-region--

--fcc-editable-region--
```
