---
id: 683ec39565dbc7e3375263b1
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U prošloj lekciji, naučili ste o binarnoj pretrazi i kako se ona razlikuje od linearnog pretraživanja.

Na ovoj radionici, primenićete to praksom kreiranja binarnog algoritma pretrage.

Počni, kreiraj funkciju za `binary_search` sa parametrima `search_list` i `value`.

`search_list` je lista elemenata koje funkcija proverava unutra, i `value` je cilj koji funkcija pokušava da pronađe unutar `search_list`.

Dodajte glavnu reč za `pass` unutar funkcije trenutno da bi test prošao.

# --hints--

Trebalo bi da napravite funkciju za `binary_search`.

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

Tvoja funkcija `binary_search` mora imati parametre od `search_list` i `value`.

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
