---
id: 684169b00a073c9b9f0a703c
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Da li si primetio da drugi poziv funkcije sada ima vrednost `3, 5, 4` u spisku? To pokazuje da binarna pretraga radi kako je predviđeno.

Tako je bilo: Algoritam je prvo proverio `3` kao sredinu početne liste. Pošto je `4` veće od `3`, pretraga se okrenula udesno. Zatim je proverio `5` kao novu sredinu. Pošto je `4` manje od `5`, pretraga se okrenula ulevo, na kraju identifikujući `4` kao sredinu poslednje sekvence.

Da biste ponovo pokušali funkciju, pozovite je sa `[1, 3, 5, 9, 14, 22], 10` i ispišite poziv odmah. Ovo je stanje u kojem vrednost nije dostupna.

# --hints--

Moraš da pozoveš funkciju za `binary_search([1, 3, 5, 9, 14, 22], 10)` i odštampaš je jednom.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1
    
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
