---
id: 69757cc0faae0152c1418aad
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Na ovoj radionici, veštaćete se korišćenjem brojeva i matematičkih operacija kako biste napravili alat za podelu računa. Ovaj alat će izračunati iznos koji svako mora da plati nakon dodavanja troškova hrane i napojnice.

Na početku, potreban vam je način da pratite ukupan iznos kada se dodaju troškovi.

Kreiraj parametar pod nazivom `running_total` i postavi vrednost `0`.

# --hints--

Trebalo bi da imate parametar nazvan `running_total`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

Trebalo bi da unesete celobrojnu vrednost `0` u svoj parametar `running_total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
