---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Svaki kozi ima gharama. Potrebno je sačuvati ovaj iznos u parametru da ga kasnije koristite. Pošto ovaj iznos uključuje cent, koristićete tip float, koji se koristi za predstavljanje decimalnih brojeva. Evo primera parametra sa vrednošću float:

```py
change = 2.35
```

Postavi četiri parametra: `appetizers` postavi vrednost `37.89`, `main_courses` postavi vrednost `57.34`, `desserts` postavi vrednost `39.39`, na `drinks` postavi vrednost `64.21`.

# --hints--

Treba da definišete parametar nazvan `appetizers` i dodelite vrednost `37.89`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

Morate definisati parametar nazvan `main_courses` i dodeliti mu vrednost `57.34`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

Trebalo bi da definišete parametar nazvan `desserts` i postavite mu vrednost `39.39`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

Trebalo bi da definišete parametar koji se zove `drinks` i postavite vrednost `64.21`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
