---
id: 6976cc0b0c135686a4620b94
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Sledeće, treba da izračunate broj osoba koje dele račun. Sačuvajte ovu vrednost u varijablu, kao što ste uradili u prethodnom koraku.

Kreiraj varijablu nazvanu `num_of_friends` i postavi joj vrednost `4`. Ovo će se koristiti kasnije u radu za izračunavanje konačne podjele.

# --hints--

Trebalo bi da imate varijablu nazvanu `num_of_friends`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

Treba da staviš puni broj `4` u svoju varijablu `num_of_friends`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
