---
id: 69779dcb44016eccc05c15a4
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Usluga je bila greške dobra, pa je grupa odlučila da ostavi 25% napojnica. Da biste izračunali procenat u Python, možete ukupan iznos pomnožiti sa decimalnim ekvivalentom tog procenta.

Primer, da biste dobili 10 procenata vrednosti, pomnožili biste sa `0.10` koristeći operator `*`:

```py
tax = total * 0.10
```

Definiši parametar nazvan `tip` i postavi vrednost na rezultat množenja `running_total` sa `0.25`.

Konačno, koristi `print()` da prikaže niz slova `Tip amount:` praćen razmakom i vrednošću tvog parametra `tip`.

# --hints--

Trebalo bi da imaš parametar nazvan `tip`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

Vaš parametar za `tip` treba da bude rezultat `running_total * 0.25`.

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

Trebalo bi da ispišete niz znakova `Tip amount:` praćen razmakom i kriterijumom `tip`.

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Tip amount: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - ((37.89 + 57.34 + 39.39 + 64.21) * 0.25)) < 1e-6
`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
