---
id: 6982684f3a25f379e195a5fc
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Sada, nakon snimanja troškova za svaki, možete izračunati ukupno.

Imajte na umu da operater `+=` dodaje vrednost postojećem parametru i poboljšava ga odmah. Primer:

```py
total = 10
total += 2 + 2 + 1
print(total)  # total is now 15
```

Koristite operater `+=` jednom da dodate `appetizers`, `main_courses`, `desserts`, na `drinks` u `running_total`.

Na kraju, koristite `print()` da prikažete sekvencu slova `Total bill so far:` praćenu razmakom i vrednošću `running_total`.

**Napomena:** Možda ćete videti rezultate sa više decimalnih mesta nego što ste očekivali. Kao što ste naučili u prethodnoj lekciji, ovo se dešava jer su brojevi sačuvani u binarnom sistemu, a mnoge decimalne vrednosti se ne mogu tačno predstaviti u ovom sistemu, što uzrokuje zaokružene greške.

# --hints--

Trebalo bi da ažuriraš `running_total` koristeći kupon `+=` jednom da poveća sve četiri troška kursa.

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total += {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

Trebalo bi da odštampate niz karaktera `Total bill so far:` praćen razmakom i vrednošću `running_total`.

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Total bill so far: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - (37.89 + 57.34 + 39.39 + 64.21)) < 1e-6
`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
