---
id: 66ebd4ae2812430bb883c787
title: Izgraditi centar događaja
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

U ovom laboratoriju ćete koristiti semantičke elemente `HTML` za kreiranje strukture web stranice. Dodaćete sadržaj i slike kako bi izgledalo kao stvarni centar događaja.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da ima element `header`.
1. Unutar elementa `header`, treba da ima element `h1` sa tekstom `Event Hub` i element `nav`.
1. Unutar elementa `nav`, treba da ima neuređenu listu sa dva stavke koje sadrže linkove ka različitim delovima stranice. Prvi stavka treba da ima tekst `Upcoming Events`, a druga stavka treba da ima tekst `Past Events`.
1. Svaki link treba da bude predstavljen elementom `a` sa atributom elementa `href` koji povezuje ka odgovarajućim delovima stranice, `#upcoming-events` i `#past-events` respektivno.
1. Trebalo bi da ima element `main` sa različitim delovima stranice.
1. Unutar elementa `main`, treba da ima dva elementa/karakteristika `section`.
1. Prvi element `section` treba da ima atribut elementa `id` sa vrednošću `upcoming-events`.
1. Unutar sekcije/dela `#upcoming-events`, treba da ima:
   
   - Element `h2` sa tekstom `Upcoming Events`.
   - Dva elementa/karakteristika `article`. Svaki članak treba da predstavlja događaj i da ima:
      - Element `h3` za naslov događaja.
      - Element `p` za opis događaja. Možete dodati datum ispod ako želite.

1. Drugi element od `section` mora imati svojstvo polja `id` sa vrednošću `past-events`.
1. Unutar sekcije `#past-events`, treba da imate:

- Element sa tekstom `Past Events` za element `h2`.
   - Dva elementa `article`. Svaki element članka treba da predstavlja prošli događaj i mora imati:
     - Element `h3` za naslov događaja,
     - Element `p` za opis događaja. Možete dodati datum ispod ako želite.
     - Element slike sa atributom elementa `src` koji postavlja sliku i atribut elementa `alt` sa opisom slike.

**Napomena:** Možete koristiti bilo koji tekst za opis događaja i datum. Možete koristiti URL sledeće slike ako želite:

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`.
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

# --hints--

Trebalo bi da imaš element `header`.

```js
assert.isNotNull(document.querySelector("header"));
```

Tvoj element za `header` mora da dođe nakon otvarajuće label-e od `body`.

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

Unutar elementa `header`, treba da ima element `h1` sa tekstom `Event Hub`.

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

Unutar elementa `header`, nakon elementa `h1`, trebalo bi da postoji element `nav`.

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

Vaš element `nav` treba da ima neredovni spisak sa dva elementa.

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

Prvi element u neuređenoj listi mora biti link.

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

Drugi element u neuređenom spisku mora biti link.

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

Sadržaj prvog elementa u neuređenoj listi treba da bude `Upcoming Events`.

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

Prvi element u neuređenoj listi treba da ima `href` postavljen na `#upcoming-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

Sadržaj drugog elementa u neuređenom spisku treba da bude `Past Events`.

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

Drugi element u neuređenom spisku bi trebalo da ima `href`, postavljen na `#past-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

Trebalo bi da imaš element `main` nakon oznake zatvaranja elementa `header`.

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

Unutar elementa za `main`, treba da ima dva elementa sa `section`.

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

Prvi element od `section` mora imati svojstvo elementa `id` sa vrednošću `upcoming-events`.

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

Vaš drugi element koji pripada ``section`` mora imati atribut polja ``id`` sa vrednošću ``past-events``.

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

U sekciji `#upcoming-events`, trebalo bi da postoji element `h2` sa tekstom `Upcoming Events`.

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

U sekciji `#upcoming-events`, bi trebalo da ima dva elementa od `article` ispod elementa za `h2`.

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Svi elementi `article` unutar sekcije `#upcoming-events` moraju imati element/karakteristiku `h3` za zaglavlje događaja.

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Svi elementi unutar sekcije `article` u delu `#upcoming-events` moraju imati element paragrafa za opis događaja.

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

U delu `#past-events`, trebalo bi da postoji element/funkcija `h2` sa tekstom `Past Events`.

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

U sekciji `#past-events`, trebalo bi da ima dva elementa od `article` pod elementom `h2`.

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Sve komponente od `article` unutar sekcije `#past-events` moraju imati karakteristiku `h3` za zaglavlje događaja.

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Svi elementi od `article` unutar sekcije `#past-events` treba da imaju paragrafski element za opis događaja.

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Svi elementi unutar sekcije `article` u delu `#past-events` moraju imati sliku elementa.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

Svi elementi slike unutar sekcije `#past-events` moraju imati atribut elementa `src` koji postavlja fajl slike.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

Svi elementi slike unutar sekcije `#past-events` treba da imaju atribut elementa `alt` sa opisom slike.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

Svaki element sa `h3` mora imati naslov događaja.

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

Svaki element za `p` mora imati opis događaja.

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```
