---
id: a6e40f1041b06c996f7b2406
title: Napravi Finder Prvog Elementa
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

U ovoj laboratoriji, kreiraćeš funkciju koja prolazi kroz tok podataka i vraća prvi element koji prođe funkcionalni test (takođe poznat kao "test istine").

Funkcija prolazi kroz niz podataka i testira svaki element koristeći obuku za testiranje koja je data. Na kraju, vraća prvi element koji funkcija za testiranje vraća `true`.

Primer:

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**Cilj:** Popunite priče korisnika ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da postoji funkcija nazvana `findElement` koja prima niz podataka i funkciju kao argument.
2. Funkcija bi trebalo da vrati prvi element iz niza podataka koji prolazi booleanski test. To znači da, pozivanjem funkcije `func` i prosleđivanjem elementa `x`, booleanski test prolazi ako je `func(x)` jednak `true`.
3. Ako nijedan element ne prođe test, funkcija bi trebalo da vrati `undefined`.

# --hints--

Trebalo bi da imaš funkciju `findElement`.

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` treba da vrati `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` treba vratiti `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })` trebalo bi da vrati `3`.

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })` treba da vrati `"javascript"`.

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })` treba da vrati `undefined`.

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })` treba da vrati `2`.

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })` treba vratiti `undefined`.

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
