---
id: 69461e4f98594012939bc1c6
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Na ovom radionici, ćete vežbati tipove podataka za Python kreiranjem jednostavnog izveštaja rezultata.

Kao što ste naučili u prethodnim lekcijama, parametri se dodeljuju vrednost na ovaj način:

```py
greeting = 'Hello World'
```

U gornjem primeru, parametar `greeting` mu je dodeljena vrednost niza karaktera, koji je sekvenca slova okružena navodnicima.

Kreiraj parametar nazvan `name` i postavi mu vrednost stringa `Alice`. Zapamti da moraš obavijezovati tu vrednost sa jednim ili dvostrukim navodnicima, kao što je prikazano u primeru.

**Važno:** Python koristi razmake za uvlačenje (razmaci na početku linije) za formatiranje koda. Za ovaj radionica, osigurajte da nema dodatnih razmaka na početku svake kodne linije. Dodavanje dodatnih razmaka će izazvati `IndentationError` i sprečiti vaš kod da se izvrši.

# --hints--

Trebalo bi da imaš kriterijum `name`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

Parametar `name` mora da sačuva vrednost `'Alice'`.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
