---
id: 694648acde178bb8202d9517
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Rezultat je `True`, što potvrđuje da `score` jeste `float`.

Popunite izveštajni list tako što ćete ispisati vrednost `score` zajedno sa tipom podataka koristeći jednu deklaraciju od `print()`.

# --hints--

Trebalo bi da odštampate `score` i `type(score)` koristeći zarez kao separator.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
