---
id: 694648acde178bb8202d9513
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Ime studenta mora pratiti isti format kao i ostale specifikacije.

Ukloni privremene rezultate promenljive `name`. Zatim ispiši `name` i `type(name)`, razdvojene zarezom, u istom redu kao i prethodno.

# --hints--

Ne bi trebalo da imaš `print(name)` u svom kodu.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

Ne bi trebalo da imaš `print(type(name))` u svom kodu.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

Trebalo bi da odštampate `name` i `type(name)` u jednom pozivu `print()` koristeći separator zarez.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
