---
id: 694648acde178bb8202d9516
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Rezultat je `False`, što ukazuje na to da `score` nije `int`.

Drugi uobičajeni tip broja u Python je `float`, koji predstavlja decimalni broj. Promenite `int` u `float` u postojećem pozivu `isinstance()` da biste ovo potvrdili.

# --hints--

Treba da zameniš `int` sa `float` u pozivu `isinstance(score, int)` koji postoji.

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
