---
id: 694648acde178bb8202d9514
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Sada moraš da dodaš starost studenta na izveštajni dokument. Zbog toga ćeš koristiti puni broj, koji je jedna od vrsta numeričkih podataka u Python.

Deklarišete parametar pod nazivom `age` i dodeljuje vrednost celog broja `20`.

Zatim, ispiši vrednost i tip podataka za `age` zajedno, razdvojene zarezom. Proveri rezultate u terminalu koji prikazuju vrednost `age`, i njegov tip kao `int` (puni broj).

# --hints--

Trebalo bi da imaš kriterijum `age`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

Polje `age` treba da skladišti vrednost `20`. Ne obavijajte tu vrednost navodnicima.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

Trebalo bi da odštampate `age` i `type(age)` koristeći zarez kao razdvojnik.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
