---
id: 694648acde178bb8202d9515
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Sada, dodaj ocenu učenika.

Deklariši parametar nazvan `score` i postavi vrednost `80.5`.

Iako su i `age` i `score` brojevi, možda nisu iste vrste. Python pruža funkciju nazvanu `isinstance()` ili proveriti ovo.

```py
x = 10
print(isinstance(x, int)) # Output: True
```

Koristi `isinstance()` da proveriš da li je `score` jednako `int`, pa ispiši rezultat u terminal kao što je prikazano u primeru iznad.

# --hints--

Trebalo bi da imaš kriterijum `score`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

Parametar `score` treba da sačuva vrijednost `80.5`. Ne obavijte tu vrijednost navodnicima.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

Trebalo bi da ispišete rezultate `isinstance(score, int)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
