---
id: 694648acde178bb8202d9512
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Izveštaj takođe treba da pokaže da li je student trenutno registrovan. Ovo se može predstaviti korišćenjem vrednosti tačno ili netačno.

Booleanske vrednosti (istina ili laž) predstavljaju stanje da/ne, i često se koriste za donošenje odluka u kodu. Postoje samo dve booleanske vrednosti: `True` na `False`.

Objavi kriterijum nazvan `is_student` sa vrednošću `True`.

# --hints--

Trebalo bi da imaš kriterijum `is_student`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

Parametar `is_student` treba da sačuva vrednost `True`. Nemoj obavijati tu vrijednost navodnicima.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
