---
id: 694648abde178bb8202d9511
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Sada bi trebalo da vidiš ime studenta kako se prikazuje na terminalu.

Python daje funkciju nazvanu `type()` koju možete koristiti za proveru tipa vrednosti.

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

U gornjem primeru, rezultati `<class 'str'>` znače da parametar koji je prosleđen funkciji `type()` je niz znakova.

Koristi funkciju `type()` sa `name` kao njenim argumentom i ispiši rezultat kao u primeru. Proveri rezultate na terminalu koji prikazuje `name` da je tipa `str` (niz slova).

# --hints--

Trebalo bi da ispišete rezultate poziva `type(name)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
