---
id: 6947a0c2b6de035dd8d3eea1
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Možete ispisati vrednost parametra koristeći funkciju `print()`.

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

Naučićete više o funkcijama u narednim lekcijama. Trenutno, znajte da je funkcija ponovno upotrebljiv deo koda koji se može pozvati ili aktivirati kako bi izvršio svoj kod, a mu se mogu prosleđivati argumenti.

U navedenom primeru, `print(greeting)` je poziv funkcije, i `greeting` je argument te funkcije.

Povežite se na primer i ispišite parametar za `name`. Proverite rezultate u terminalu.

# --hints--

Trebalo bi da pozoveš `print()` sa `name` kao argumentom.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
