---
id: 6947a3e516bb8a4a681274a7
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Funkcija `print()` može prikazati više od jedne vrednosti istovremeno. Razdvojite vrednosti zarezom (`,`) kako biste ih ispisali na jednoj liniji.

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

Uneseti ``is_student`` i ``type(is_student)`` sve na jednoj liniji koristeći zarez ``,``, kao što je prikazano u primeru. Zatim, proverite rezultate u terminalu koji prikazuju vrednost od ``is_student``, i njen tip kao ``bool`` (tačna ili netačna).

# --hints--

Trebalo bi da ispišete `is_student` i `type(is_student)` koristeći zarez kao separator.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
