---
id: 657bdc8ba322aae1eac38390
title: Napravite konverter za rimska brojeva
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 16044
dashedName: build-a-roman-numeral-converter
---

# --description--

Rimski brojevi se izvode iz sedam simbola i mogu se zapisati korišćenjem različitih kombinacija za predstavljanje arapskih brojeva. Primer:

| Rimski brojevi | Arapski brojevi |
|-------------------|--------------------|
| M                 | 1000               |
| CM                | 900                |
| D                 | 500                |
| CD                | 400                |
| C                 | 100                |
| XC                | 90                 |
| L                 | 50                 |
| XL                | 40                 |
| X                 | 10                 |
| IX                | 9                  |
| V                 | 5                  |
| IV                | 4                  |
| I                 | 1                  |

**Cilj:** Napravite program koji radi kao <a href="https://roman-numeral-converter.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://roman-numeral-converter.freecodecamp.rocks</a>.

**Zahtevani funkcionalnosti (User Stories):**

1.  Morate imati `<input>` element sa `id`-om `"number"`.
2.  Morate imati `<button>` element sa `id`-om `"convert-btn"`.
3.  Morate imati `<div>`, `<span>` ili `<p>` element sa `id`-om `output`.
4.  Kada kliknete na element `#convert-btn` bez unosa vrednosti u element `#number`, element `#output` treba da sadrži tekst `"Please enter a valid number"`.
5.  Kada element `#number` sadrži broj `-1` i klikne se na element `#convert-btn`, element `#output` treba da sadrži tekst `"Please enter a number greater than or equal to 1"`.
6.  Kada element `#number` sadrži broj `4000` ili više i klikne se na element `#convert-btn`, element `#output` treba da sadrži tekst `"Please enter a number less than or equal to 3999"`.
7.  Kada element `#number` sadrži broj `9` i klikne se na element `#convert-btn`, element `#output` treba da sadrži tekst `"IX"`.
8.  Kada element `#number` sadrži broj `16` i klikne se na element `#convert-btn`, element `#output` treba da sadrži tekst `"XVI"`.
9.  Kada element `#number` sadrži broj `649` i klikne se na element `#convert-btn`, element `#output` treba da sadrži tekst `"DCXLIX"`.
10. Kada element `#number` sadrži broj `1023` i klikne se na element `#convert-btn`, element `#output` treba da sadrži tekst `"MXXIII"`.
11. Kada element `#number` sadrži broj `3999` i klikne se na element `#convert-btn`, element `#output` treba da sadrži tekst `"MMMCMXCIX"`.

Implementirajte sve zahtevane funkcionalnosti i prođite kroz sve testove ispod kako biste završili ovaj projekat. Dodajte svoj stil. Srećno kod programiranja!

# --hints--

Treba da imate `<input>` element sa `id`-om `"number"`.

```js
const el = document.getElementById('number');
assert.strictEqual(el?.nodeName?.toLowerCase(), 'input');
```

Treba da imate `<button>` element sa `id`-om `"convert-btn"`.

```js
const el = document.getElementById('convert-btn');
assert.strictEqual(el?.nodeName?.toLowerCase(), 'button');
```

Treba da imate `<div>`, `<span>` ili `<p>` element sa `id`-om `"output"`.

```js
const el = document.getElementById('output');
assert(['div', 'span', 'p'].includes(el?.nodeName?.toLowerCase()));
```

Kada kliknete na element `#convert-btn` bez unosa vrednosti u element `#number`, element `#output` treba da sadrži tekst `"Please enter a valid number"`.

```js
const numberInputEl = document.getElementById('number');
const convertBtnEl = document.getElementById('convert-btn');
const outputEl = document.getElementById('output');

numberInputEl.value = '';
convertBtnEl.click();
assert.strictEqual(outputEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'please enter a valid number');
```

Kada element `#number` sadrži broj `-1` i klikne se na element `#convert-btn`, element `#output` treba da sadrži tekst `"Please enter a number greater than or equal to 1"`.

```js
const numberInputEl = document.getElementById('number');
const convertBtnEl = document.getElementById('convert-btn');
const outputEl = document.getElementById('output');

numberInputEl.value = '-1';
numberInputEl.dispatchEvent(new Event('change'));
convertBtnEl.click();
assert.strictEqual(outputEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'please enter a number greater than or equal to 1');
```

Kada element `#number` sadrži broj `4000` ili više i klikne se na element `#convert-btn`, element `#output` treba da sadrži tekst `"Please enter a number less than or equal to 3999"`.

```js
const numberInputEl = document.getElementById('number');
const convertBtnEl = document.getElementById('convert-btn');
const outputEl = document.getElementById('output');

numberInputEl.value = '4000';
numberInputEl.dispatchEvent(new Event('change'));
convertBtnEl.click();
assert.strictEqual