---
id: 61329d52e5010e08d9b9d66b
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Još jedna važna karakteristika za `meta` u pogledu pristupačnosti i SEO je opis `description`. Vrednost atributa karakteristike `content` koristi se od strane pretraživača kako bi pružila detalje o vašoj stranici.

Dodaj element `meta` sa svojstvom elementa `name` postavljeno na `description`, i sa svojstvom elementa `content` korisno.

# --hints--

Moraš dodati novu funkciju za `meta` unutar `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

Treba da dodaš atribut `name` elementu `meta` sa vrednošću `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

Trebalo bi mu dati atribut elementa `meta` sa svojstvom `content`.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

Vrednost atributa polja `content` ne sme da premaši 165 karaktera. _Ovo je maksimalna dužina za opise na Google-u._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
