---
id: 613297a923965e0703b64796
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Možda ste već upoznati sa funkcionalnošću `meta`; koristi se za definisanje informacija o stranici, kao što su naslov, opis, ključne reči i autor.

Podesite svoju stranicu polje ``meta`` sa odgovarajućom vrednošću od ``charset``.

Atribut elementa `charset` opisuje kodiranje karaktera stranice, a danas je `UTF-8` jedino kodiranje koje podržava većina pregledača.

# --hints--

Trebalo bi da kreirate element `meta` unutar elementa `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Trebalo bi mu dati oznaku `meta` svojstvo elementa `charset` sa vrednošću `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
