---
id: 6133acc353338c0bba9cb553
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Na kraju `head`, element `title` je važan za čitač ekrana da razume sadržaj stranice. Štaviše, predstavlja važan deo _SEO_.

Dajte svoju stranicu `title` koja objašnjava i kratka je.

# --hints--

Trebalo bi dodaš komponentu za `title` u `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Ne bi trebalo da učiniš `title` dužim od 60 karaktera.

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

Pokušajte sa više detalja i funkcijom `title`. Savet: Najmanje 15 slova

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
