---
id: 613e275749ebd008e74bb62e
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Trenutno, `img` zauzima svoju standardnu veličinu, koja je veoma velika. CSS ima funkciju `max` koja vraća najveću vrednost među skupom vrednosti razdvojenih zarezima. Na primer:

```css
img {
  width: max(250px, 25vw);
}
```

U gornjem primeru, širina slike će biti 250px ako je širina vidokruga manja od 1000 piksela. Ako je širina vidokruga veća od 1000 piksela, širina slike će biti 25vw. To je zato što je 25vw jednako 25% širine vidokruga.

Smanite sliku koristeći `id` kao selektor, i postavite `width` na veće od `10rem` ili `18vw`.

# --hints--

Trebalo bi da koristiš selektor `#logo` za ciljanje elementa `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

Trebalo bi mu dati `img` `width` od `max(10rem, 18vw)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
