---
id: 613e2546d0594208229ada50
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Unutar `header`, dostavite kontekst za stranicu tako što ćete ga organizovati u jednom elementu od `img`, `h1`, i `nav`.

`img` mora prikazati `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`, imati `id` od `logo`, i sa tekstom od `alt` od `freeCodeCamp`.

`h1` treba da ima tekst od `HTML/CSS Quiz`.

# --hints--

Trebaš dodati element `img` na element `header`.

```js
assert.exists(document.querySelector('header > img'));
```

Treba da dodaš element `h1` u element `header`.

```js
assert.exists(document.querySelector('header > h1'));
```

Treba da dodaš element `nav` u element `header`.

```js
assert.exists(document.querySelector('header > nav'));
```

Trebalo bi da postavite elemente `img`, `h1` i `nav` u tom redosledu.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

Trebalo bi da dodaš atribut elementu sa ID-om `img`, koji je atribut elementa sa ID-om `src` i čija je vrednost `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

Trebalo bi mu dodeliti atribut elementu `img` sa atributom elementa `id` sa vrednošću `logo`.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

Trebalo bi mu dodeliti elementu `img` atribut elementa `alt` sa vrednošću `freeCodeCamp`.

```js
assert.equal(document.querySelector('img')?.alt, 'freeCodeCamp');
```

Treba da mu/jajš element `h1` sa tekstom `HTML/CSS Quiz`.

```js
assert.include(document.querySelector('h1')?.innerText.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
