---
id: 614202874ca576084fca625f
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Svako polje za `region` zahteva oznaku (label), koja pomaže korisnicima čitača ekrana da razume svrhu tog polja. Jedan način za dodavanje oznake je dodavanje elementa zaglavlja sekcije unutar tog polja, a zatim referenciranje atributom `aria-labelledby`.

Dodajte sledeće atribute za `aria-labelledby` za elemente od `section`:

- `student-info`
- `html-questions`
- `css-questions`

Zatim, unutar svakog elementa sa `section`, razmjestite jedan element sa `h2` koji ima svojstvo koje odgovara povezanom svojstvu `aria-labelledby` i `id`. Dostavite odgovarajući tekstualni sadržaj za svaki element sa `h2`.

# --hints--

Trebalo bi mu dati prvu instancu `section` atribut `aria-labelledby` sa vrednošću `student-info`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

Treba da mu date drugi element od `section` atribut sa `aria-labelledby` koji ima vrednost `html-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

Trebalo bi mu dati treći atribut za `section` svojstvo od `aria-labelledby` sa vrednošću `css-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

Treba da ga smestite unutar jednog elementa od `h2`, unutar prvog elementa od `section`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

Trebalo bi da ga postavite u prvi element od `h2`, unutar drugog elementa od `section`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

Treba da ga smestiš unutar jednog elementa `h2` i unutar trećeg elementa `section`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

Trebalo bi mu dati prvu instancu `h2` svojstvo `id` sa vrednošću `student-info`.

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

Trebalo bi mu dati drugi atribut `h2` svojstvo `id` sa vrednošću `html-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

Trebalo bi mu dati treći atribut `h2` svojstvo `id` sa vrednošću `css-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

Treba da mu/jajš prvi deo sadržaja teksta za `h2`. _Savet: Izabrao bih `Student Info`_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[0]?.textContent.trim());
```

Drugom elementu `h2` treba da dodaš odgovarajući tekstualni sadržaj. _Savet: Izabrao bih `HTML`._

```js
assert.isNotEmpty(document.querySelectorAll('h2')[1]?.textContent.trim());
```

Trećem elementu `h2` treba da dodaš odgovarajući tekstualni sadržaj. _Savet: Izabrao bih `CSS`._

```js
assert.isNotEmpty(document.querySelectorAll('h2')[2]?.textContent.trim());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
