---
id: 6143639d5eddc7094161648c
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Važno je povezati svaki `input` sa odgovarajućim elementom `label`. Ovo pruža korisnicima tehnologije vizuelni referentni pomoćnik za unos.

Ovo se dešava dodeljivanjem `label` atributa svojstvu `for`, koje ima `id` od `input`.

Ovaj deo će uzeti ime studenta, e-mail adresu i datum rođenja. Pružite odgovarajuće parametre za `label` za elemente `for`, zajedno sa tekstualnim sadržajem. Zatim, povežite elemente `input` sa odgovarajućim elementima `label`.

# --hints--

Trebalo bi mu dati prvi atribut od `label` odgovarajući atribut od `for`.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[0]?.htmlFor?.length, 1);
```

Trebalo bi mu dati drugi atribut od `label` odgovarajući atribut od `for`.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[1]?.htmlFor?.length, 1);
```

Trebalo bi mu dati treći atribut za `label` odgovarajući atribut od `for`.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[2]?.htmlFor?.length, 1);
```

Trebalo bi mu dati prvi deo sadržaja teksta za `label`.

```js
assert.isNotEmpty(document.querySelectorAll('label')[0]?.textContent.trim());
```

Trebalo bi mu dati drugi deo sadržaja teksta za `label`.

```js
assert.isNotEmpty(document.querySelectorAll('label')[1]?.textContent.trim());
```

Trebalo bi mu dati treći deo sadržaja teksta za `label`.

```js
assert.isNotEmpty(document.querySelectorAll('label')[2]?.textContent.trim());
```

Trebalo bi mu dati prvi atribut `input` svojstvo `id` koje odgovara svojstvu `for` prvog `label`.

```js
assert.equal(document.querySelectorAll('input')?.[0]?.id, document.querySelectorAll('label')?.[0]?.htmlFor);
```

Drugom elementu `input` treba da dodaš atribut `id` koji odgovara vrednosti atributa `for` drugog elementa `label`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.id, document.querySelectorAll('label')?.[1]?.htmlFor);
```

Treba mu dati treći element od `input` atribut/svojstvo `id` koje odgovara atributu/svojstvu `for` za treći od `label`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.id, document.querySelectorAll('label')?.[2]?.htmlFor);
```

Ne smete koristiti isti atribut od `id` za više od jednog elementa od `input`.

```js
const id = (n) => document.querySelectorAll('input')?.[n]?.id;
assert.notEqual(id(0), id(1));
assert.notEqual(id(0), id(2));
assert.notEqual(id(1), id(2));
```

Ne smete koristiti istu osobinu od `for` za više od jednog elementa iz `label`.

```js
const htmlFor = (n) => document.querySelectorAll('label')?.[n]?.htmlFor;
assert.notEqual(htmlFor(0), htmlFor(1));
assert.notEqual(htmlFor(0), htmlFor(2));
assert.notEqual(htmlFor(1), htmlFor(2));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
--fcc-editable-region--
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

```
