---
id: 614ccc21ea91ef1736b9b578
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Za ovaj seminar/radionicu, dobili ste osnovnu strukturu.

Započni ovaj put ka pristupačnosti tako što ćeš elementu `html` dodati atribut `lang`. Ovo će pomoći čitačima ekrana da prepoznaju jezik stranice.

Pia, zapamti da priložiš svoj list uzorka na stranicu.

# --hints--

Trebalo bi mu dati atribut elementa `html` svojstvo elementa `lang`. _Napomena: Možete koristiti vrednost `en` na engleskom._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

Trebalo bi da imaš element `link` unutar svog elementa `head`.

```js
assert.exists(document.querySelector("head>link"));
```

Vaš element za `link` mora imati `rel="stylesheet"` i `href="styles.css"`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
