---
id: 60a3e3396c7b40068ad69973
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Promeni `background-color` u dijagramski alat (drawing toolbar) u `#4d0f00`.

# --hints--

Trebalo bi da postaviš atribut `background-color` na `#4d0f00`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'rgb(77, 15, 0)');
assert(hasBackground);
```

Element za `.canvas` trebao bi imati `background-color` od `#4d0f00`.

```js
const canvasBackground = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('background-color');
assert(canvasBackground === 'rgb(77, 15, 0)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
