---
id: 60a3e3396c7b40068ad69970
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Pre nego što počnete sa dekorisanjem `div` koji ste dodali, morate povezati vaš CSS sa vašim HTML.

Dodaj funkciju `link` za povezivanje tvoje datoteke `styles.css`. Postavi `href` na `styles.css`, i zapamti da postaviš atribut funkcije `rel` na `stylesheet`.

# --hints--

Tvoj kod mora imati komponentu od `link`.

```js
assert(/<link/.test(code))
```

Vaš element od `link` mora imati svojstvo elementa `rel` sa vrednošću `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

Element vašeg `link` treba da ima svojstvo polja `href` sa vrednošću `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
