---
id: 60a3e3396c7b40068ad6996c
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Tokom prethodne obuke o modelu kutije za CSS, saznali ste da svaki element od HTML se tretira kao kutija sa četiri polja.

Zamislite da dobijate kutiju od vašeg omiljenog online prodavca – čija je sadržaj proizvodi unutar kutije, ili u vašem slučaju, naslov, pasus ili element slike.

Promeni atribut `src` u `<img>` od `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` do `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

Funkcija od `img` treba da ima svojstvo `src` za `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```
