---
id: 60a3e3396c7b40068ad6996e
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Rub je područje van kutije, i može se koristiti za kontrolu prostora između drugih kutija ili elemenata.

Ovaj donji element ima veliki gornji rub, što ga gura još niže na stranicu.

Nakon što sam brzo pregledao primer kutije za CSS, vreme je da počnemo sa slikom Rothka.

Uklonite element `img`.

# --hints--

Treba da imaš element `img` u svom kodu.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
