---
id: 60a3e3396c7b40068ad69975
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Napišite novi zakon koristeći alat za klasifikaciju objekata `.frame`.

Koristite kratki deo `border` da obuhvatite element `.frame` do čvrste, crne širine od `50px`.

# --hints--

Vaš kod bi trebalo da ima selektor za `.frame`.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

Treba da postaviš atribut `border` na `50px solid black`.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

Element vašeg `.frame` treba da ima svojstvo `border` sa vrednošću `50px solid black`.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
