---
id: 60a3e3396c7b40068ad69978
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Koristite ivicu da popravite prostor izvan elementa.

Ako se oslanjate na svojstvo `margin`, dajte elementu `.frame` vertikalnu marginu od `20px` i horizontalnu marginu od `auto`. Ovo će pomeriti okvir nadole za 20 piksela i centrirati ga horizontalno na stranici.

# --hints--

Trebalo bi postaviti atribut za `margin` na `20px auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

Vaše polje za `.frame` trebalo bi imati vrednost `margin` od `20px auto`.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
