---
id: 60a3e3396c7b40068ad69971
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

U vreme CSS.

Iako tvoj `<div>` nema tekst, i dalje se tretira kao kontejner sa sadržajem.
Napiši pravilo za CSS koje koristi selektor klase `.canvas` i postavi njegov atribut `width` na 500 piksela.
Evo pravila za CSS koje postavlja širinu klase `card` na 300 piksela:

```css
.card {
  width: 300px;
}
```

# --hints--

Vaš kod mora imati selektor za `.canvas`.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

Trebalo bi da postaviš vrednost `width` na `500px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

Vaš element od `.canvas` mora imati svojstvo `width` za `500px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
