---
id: 60a3e3396c7b40068ad6996f
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Dodaj komponentu `div` unutar `body`.

Postavi svojstvo elementa `class` jednako `canvas`.

Ovo će se koristiti kao platno za vaš slikovni rad.

# --hints--

Tvoj kod treba da sadrži element `div`.

```js
assert(document.querySelectorAll('div').length === 1)
```

Polje `div` mora imati svojstvo polja `class` sa vrednošću `canvas`.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
