---
id: 60a3e3396c7b40068ad69980
title: Korak 22
challengeType: 0
dashedName: step-22
---

# --description--

Dodajte `div` drugi sa vrednošću `class` od `two` na dnu vašeg elementa/polja `one`. Ovo će biti vaš drugi pravougaonik.

# --hints--

Ne smete menjati postojeći element od `.one`.

```js
assert(document.querySelectorAll('.one').length === 1);
```

Trebalo bi da imaš element `div` kao drugi unutar svog elementa `.canvas`.

```js
assert(document.querySelector('.canvas').children[1].tagName === 'DIV');
```

Polje vašeg `div` pilule mora imati vrednost `class` od `two`.

```js
assert(document.querySelector('.canvas').children[1].className.split(' ').includes('two'));
```

Vaša komponenta `.two` mora da dođe nakon vaše komponente `.one`.

```js
assert(document.querySelector('.two').previousElementSibling.className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
