---
id: 60a3e3396c7b40068ad69974
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Svaki crtež zahteva okvir.

Postavi ga unutar elementa `.canvas`, u drugi element `div`. Dodeli `div` klasu `frame`.# --hints--

Treba da dodate novu funkciju za `div`.

```js
assert(document.querySelectorAll('div').length === 2)
```

Tvoj element od `.canvas` treba da bude smešten unutar novog elementa od `div`.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

`div` novi treba da ima `class` sa vrednošću od `frame`.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

`div` novi treba da bude unutar vašeg elementa za `body`.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
    <div class="canvas">
    </div>
    
--fcc-editable-region--
  </body>
  ```
