---
id: 60a3e3396c7b40068ad6996b
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Pokrenite svoj projekat dodavanjem elementa `img` koji ima `src` od `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` unutar elementa `body`.

# --hints--

Vaš kod mora imati element `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

Element za `img` treba da ima `src` od `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
