---
id: 60a3e3396c7b40068ad6996d
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sadržaj je okružen prostorom nazvan unutrašnji prostor, slično zaštitnim omotačima od mjehurića koji razdvajaju element od okolne kutije.

Zamisli granicu kao kartonsku kutiju u koju je tvoj element poslat.

Promeni atribut elementa `src` u `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

# --hints--

Kipengele sa `img` treba da ima svojstvo kipengele `src` od `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
