---
id: 68c18b1d0fce081f85312d97
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

U prethodnoj lekciji, naučili ste da svojstvo elementa sa jednim prefiksom donje linije se koristi za interno korišćenje u ugovoru.

Promeni atribute od `name` na `level` u `_name` na `_level`, jer ne bi trebalo da se menjaju izvan njihove klase.

Zapamtite da ovo ne sprečava svojstvo elementa da se pristupi ili modifikuje izvan klase. Takođe, u Python uvek postoji način da se pristupi privatnim svojstvima elementa (koji imaju prefiks od dve podvučnice).

# --hints--

Tvoja putanja za `__init__` treba da postavi `self._name` na `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

Tvoja putanja ``__init__`` treba da postavi ``self._level`` na ``level``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
