---
id: 68c15a2a1f0e70ca7154aba5
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

Na ovoj radionici, pokreći ćete sistem za praćenje plate zaposlenih.

Počnite kreiranjem klase objekata nazvane `Employee`. Unutar nje creirajte metodu `__init__` sa parametrima `self`, `name` i `level`. Unutar metode `__init__`, postavite vrednosti `name` i `level` za svojstva primernog elementa sa istim imenima.

# --hints--

Trebalo bi da imaš klasu objekata nazvanu `Employee`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

Vaša klasa za `Employee` mora imati metodu `__init__` sa tri parametra `self`, `name`, i `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

Tvoj put za `__init__` bi trebalo da postavi `self.name` na `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Tvoj put za `__init__` bi trebalo da postavi `self.level` na `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
