---
id: 68c28d659d345270d4a61084
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Putanja `__repr__` je specijalna putanja koja se očekuje da vrati niz karaktera koji predstavlja nešto što se može koristiti za inicijalizovanje tog objekta.

Na primer, metoda `__repr__` od `Employee('Charlie', 'developer')` treba da vrati sekvencu karaktera `Employee('Charlie', 'developer')`, koja je isti niz korišćen za kreiranje tog objekta.

Pronađi klasu `Employee` sa metodom `__repr__` i parametrom `self`, i nateraj je da vrati niz znakova koji se može koristiti za inicijalizaciju tog objekta.

# --hints--

Vaša klasa za `Employee` treba imati metodu koja `__repr__` sa parametrom `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

Tvoja putanja ``__repr__`` treba da vrati sekvencu karaktera koja se koristi za inicijalizaciju toga.

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
