---
id: 68c28cfab52fef6913d18166
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Sada ukloni poziv poslednje dve `print`.

# --hints--

Ne bi trebalo da imaš `print(charlie_brown.name)` u svom kodu.

```js
({ test: () => assert(runPython(`not _Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

Ne smete imati `print(charlie_brown.level)` u svom kodu.

```js
({ test: () => assert(runPython(`not _Node(_code).has_stmt("print(charlie_brown.level)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
print(charlie_brown.name)
print(charlie_brown.level)
--fcc-editable-region--
```
