---
id: 68c27fa1697ceb7c578b2c9f
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Mkipaji `@property` se koristi u Python za konverziju putanje u atribut. Obično se koristi za definisanje ulaznih puteva, koji su putanje korišćene za dobijanje vrednosti atributa elementa:

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

Kreiraj svojstvo `name` sa parametrom `self` koristeći dekorator `@property`. Unutar metode vrati `self._name`. # --hints--

Tvoja klasa `Employee` treba da ima metodu `name` sa parametrom `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

Tvoja putanja ``name`` treba da bude dekorisana sa ``@property``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

Tvoja putanja ``name`` bi trebalo da vrati ``self._name``.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
