---
id: 68c28c6517e7de5e85f8ebe5
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Sada, pošto ima potrošače za svojstva funkcije `_name` i `_level`, ažurirajte vraćeni niz znakova `__str__` koristeći `self.name` i `self.level`. Ovo će pozvati potrošače umesto da se svojstva funkcije dobiju direktno.

# --hints--

Tvoja putanja ``__str__`` treba da vrati ``f'{self.name}: {self.level}'``

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_return("f'{self.name}: {self.level}'")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)
print(charlie_brown.level)
```
