---
id: 68c281cecda94ba2c6838e17
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Sada kada si definisao getter za `name`, možeš pristupiti atributu `_name` preko svojstva `name`. Zato ispiši `charlie_brown.name` u konzolu.

# --hints--

Trebalo bi da odštampate `charlie_brown.name` u konzolu.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
