---
id: 68c27eb292bfec6b9402d58f
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Dodajte metod za `__str__` u klasu objekata `Employee`. Napravite da vraća f-string sekvencu sa strukturom `name: level`, menjajući `name` i `level` za slična svojstva elemenata.

Nakon toga, zalepi `charlie_brown` na konzolu.

# --hints--

Vaša klasa za `Employee` mora imati metodu `__str__` sa parametrom `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

Tvoj metod `__str__` treba da vrati f-string u formatu `name: level`, zamenjujući `name` i `level` vrednostima `self._name` i `self._level`.

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

Treba da odštampate `charlie_brown` u konzolu.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--
```
