---
id: 68c7f7012a700243eff1cbc0
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Nivo koji se koristi za angažovanje zaposlenog treba da bude izabran između specifičnih nivoa. Koristićete `_base_salaries` da potvrdite nivo i postavite odgovarajuću platu za zaposlenog.

Nakon vaše deklaracije za `if` koja postoji, kreiraj još jednu deklaraciju za `if` koja proverava da li `level` nije u `Employee._base_salaries`.

U novoj deklaraciji od `if`, uklonite `ValueError` koji ima poruku `Invalid value '{level}' for 'level' attribute.`, gde se `{level}` treba zameniti vrednošću argumenta `level`.

# --hints--

Kada utvrdite da `level` nije u `Employee._base_salaries`, treba emitovati `ValueError` sa porukom `Invalid value '{level}' for 'level' attribute.`, gde bi `{level}` trebalo da bude zamenjen vrednošću argumenta od `level`.

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

Nije moguće izbaciti nijednu grešku kada `level` nalazi se u `Employee._base_salaries`.

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
