---
id: 68c9cab4b1118da59eecfc56
title: Korak 31
challengeType: 20
dashedName: step-31
---

# --description--

Nakon postojećeg deklarisanja od `if`, kreiraj drugi ili ukloni `ValueError` kada je `new_level` već izabrana vrednost.

Imajte na umu da `_level` još nije prisutan tokom inicijalizacije, pa koristite `hasattr(self, '_level')` da proverite da li postoji pre poređenja. Ovo sprečava `AttributeError` kada se nešto kreira prvi put.

U poruci, koristite `'{level}' is already the selected level.`, gde `{level}` treba zameniti sa trenutnim standardom.

# --hints--

Mora da ima treće deklaracije za `if` unutar vašeg seta (setter) od `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2]`)) })
```

Kada je `new_level` jednako `self.level`, treba da podigne `ValueError` i poruku `'{level}' is already the selected level.`, gde `{level}` treba da se promeni na trenutnu vrednost.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "'trainee' is already the selected level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

Tvoja izjava za `if` bi trebalo da koristi `hasattr(self, '_level')` da proveri da li postoji `_level` pre poređenja.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2].find_conditions()[0].is_equivalent("hasattr(self, '_level') and new_level == self.level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
