---
id: 69b014adbc624861e6ff3dd1
title: Korak 21
challengeType: 20
dashedName: step-21
---

# --description--

Nakon kreiranja settera za `name`, izmeni `__init__` tako da koristi `self.name = name` umesto `self._name = name`. Ovo omogućava dodeljivanje vrednosti preko settera, a potvrda će biti dodata u sledećem koraku.

Pia, popravi uslov za `if` ili samo proveri `level`, i ažuriraj poruku `TypeError` odgovarajuće:

# --hints--

Tvoj metod `__init__` treba da koristi `self.name = name` umesto `self._name = name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Tvoja putanja ``__init__`` bi trebalo samo da proveri ``isinstance(level, str)`` i ažurira poruku ``TypeError`` odgovarajuće.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_if("not isinstance(level, str)").find_body().has_stmt("raise TypeError(\\"'level' attribute must be of type 'str'.\\")")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
    def __init__(self, name, level):
--fcc-editable-region--
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]
--fcc-editable-region--

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        self._name = new_name
        
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary


charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
