---
id: 6995d0325fd7da40854fac40
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U prošloj lekciji, naučili ste o binarnoj pretrazi i kako se ona razlikuje od linearnog pretraživanja.

Na ovoj radionici, primenićete to praksom kreiranja binarnog algoritma pretrage.

Počnite, kreirajte funkciju za `binarySearch` sa parametrima `searchList` i `value`.

`searchList` je red podataka komponenti koje će funkcija pregledati unutar njega, i `value` je cilj koji će funkcija pokušati da pronađe unutar `searchList`.

# --hints--

Treba da napraviš funkciju za `binarySearch`.

```js
assert.isFunction(binarySearch);
```

Tvoja funkcija za `binarySearch` treba imati parametre od `searchList` i `value`.

```js
const binSearchParams = __helpers.getFunctionParams(binarySearch.toString());

const binarySearchParamsNames = binSearchParams.map(param => param.name)
assert.include(binarySearchParamsNames, "searchList");
assert.include(binarySearchParamsNames, "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
