---
id: 6996cb5869acf58de751ae81
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Zapamti da funkcija prima niz podataka nazvan `searchList` i `value` koji tražiš.

Sada, pozovi funkciju za `binarySearch([1, 2, 3, 4, 5], 3)` i `binarySearch([1, 2, 3, 4, 5, 9], 4)` i napiši oba izjave u konzolu.

# --hints--

Trebaće ti da pozoveš funkciju za `binarySearch([1, 2, 3, 4, 5], 3)` i napišeš je u konzolu.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*\]\s*,\s*3\s*\)\s*\)/
);
```

Treba da pozovete funkciju za `binarySearch([1, 2, 3, 4, 5, 9], 4)` i napišete to u konzolu.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*,\s*9\s*\]\s*,\s*4\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--  

--fcc-editable-region--  
```
