---
id: 6996c7363a4a437839f1f15d
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada, nakon definisanja granica vašeg pretraživanja, vreme je da kreirate petlju koja će implementirati algoritam binarne pretrage.

Binarna pretraga funkcioniše ponavljajućim smanjenjem opsega pretraživanja. Ovaj proces se nastavlja sve dok postoji validni opseg elemenata za proveru.

Možete ovo prikazati za strugu `while` koja nastavlja projekat samo ako je vaš savet od `low` ispod ili jednak savetu od `high`.

Ako `low` je veće od `high`, to znači da je pretraživačka oblast istekla, i `value` nije u redu podataka.

Unutar funkcije `binarySearch`, kreiraj petlju za `while` sa uslovom da se proveri da li `low` je nisko ili jednako visoko.

# --hints--

Trebalo bi vam kreirati ciklus od `while` koji se nastavlja kada `low` je manji ili jednak `high`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
