---
id: 6996c8f94fe6b250f11a2a1e
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Sada, uključite `valueAtMiddle` u red podatke `pathToTarget` za praćenje preduzetih koraka tokom pretrage, bez obzira da li je cilj. Možete to učiniti korišćenjem metode `push()`.

# --hints--

Trebalo bi da koristiš metod `push()` za dodavanje `valueAtMiddle` u red podataka `pathToTarget`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
