---
id: 6996c60281feca6718cc15c1
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Unutar vaše funkcije za `binarySearch`, kreirajte parametar `pathToTarget` i počnite postavljanjem na prazan niz podataka.

# --hints--

Trebalo bi da imaš parametar `pathToTarget` postavljen na prazan niz podataka `[]` unutar svoje funkcije `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
