---
id: 6996c8375f4301004459ee85
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Sada, nakon izračunavanja modela `mid`, morate dobiti vrednost od `searchList` za taj model. Ova vrednost će biti upoređena sa vašim parametrom `value`, koji je cilj koji tražite.

Definiši parametar za `valueAtMiddle` i postavi vrednost elementa iz `searchList` koji se nalazi u modelu/primeru `mid`.

# --hints--

Trebalo bi da napraviš parametar `valueAtMiddle` koji je jednak `searchList[mid]`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
