---
id: 6996d0a2c2ac55041145a13f
title: Korak 14
challengeType: 1
dashedName: step-14
---

# --description--

Da li si primetio da drugi poziv funkcije sada ima vrednost `3, 5, 4` u spisku? To pokazuje da binarna pretraga radi kako je predviđeno.

Tako je to bilo: Algoritam je prvo proverio `3` kao sredinu početne liste. Pošto je `4` veći od `3`, pretraga je prešla na desnu polovinu. Zatim je proverio `5` kao novu sredinu. Pošto je `4` manji od `5`, pretraga je prešla na levu polovinu, konačno identifikujući `4` kao sredinu poslednjeg opsega.

Da biste ponovo testirali funkciju, pozovite ga sa `[1, 3, 5, 9, 14, 22], 10` i ispišite rezultate u konzolu. Ovo je stanje gde vrednost nije dostupna.

# --hints--

Treba da pozoveš funkciju za `binarySearch([1, 3, 5, 9, 14, 22], 10)` i ispišeš rezultate u konzolu.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*3\s*,\s*5\s*,\s*9\s*,\s*14\s*,\s*22\s*\]\s*,\s*10\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
