---
id: 6996ca21a74104721f2c7e5d
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Moraš da isprobaš stvari kako bi razumeo tok algoritma u ovoj početnoj fazi.

Da biste to uradili, prvo morate izaći iz ciklusa/petlje. To je zato što trenutno izvršavanje će dozvoliti samo jednu iteraciju, pa ako uslov nije zadovoljen, biće beskonačna petlja.

Odmah nakon izjave `if`, koristi ključnu reč `break` iz petlje `while`. Zatim, nakon petlje `while`, vrati prazan red podataka da ukazuje da vrednost nije pronađena.

# --hints--

Trebalo bi da koristiš glavnu reč iz `break` nakon izreke `if` u ciklusu `while`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

Treba da resetuješ prazan set podataka nakon svog ciklusa od `while`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
