---
id: 6996c796385b6ea4e222d5e8
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Da bi se započeo ciklus/petlja, potrebno je dobiti središnju vrednost trenutne pozicije pretrage.

Možete dobiti srednju vrednost izračunavanjem proseka parametara ``low`` i ``high``. Biće vam potrebno da koristite celobrojno deljenje (floor division) kako biste dobili prosek nakon sabiranja ove dve vrednosti, tako da će rezultat uvek biti ceo broj zaokružen nadole, jer parametri moraju biti celi brojevi.

Unutar vašeg ciklusa od `while`, kreiraj parametar za `mid` koji je postavljen na prosek od `low` i `high`.

# --hints--

Unutar ciklusa while, treba kreirati parametar za `mid` postavljen za `Math.floor((low + high) / 2)`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
