---
id: 6996ccadf9048c79b857bb15
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Možete da vidite da prvi poziv vraća `[3]`. Ovo je zato što `3` je srednja vrednost u `[1, 2, 3, 4, 5]`.

S druge strane, drugi poziv vraća prazan niz podataka. To se dešava zato što, nakon proveravanja srednje vrednosti, petlja trenutno nema logiku da smanji opseg pretraživanja korigovanjem promena ``low`` ili ``high``. Izvršava samo jednu proveru srednje vrednosti, nakon čega završava bez daljeg pretraživanja.

Da bi se dopustilo da binarno pretraživanje nastavi smanjujući svoj opseg pretrage, dodajte izjavu `else if` koja proverava da li `value` je veće od `valueAtMiddle`.

# --hints--

Treba da dodaš deklaraciju `else if` sa uslovom `value > valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
--fcc-editable-region--
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
