---
id: 6996c95e5bede03bdf2937a2
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Sledeći korak je kreiranje stanja koje će proveriti da li `value` cilj je u sredini.

Unutar `while` petlje napravi `if` izraz koji proverava da li je ciljana vrednost `value` jednaka vrednosti `valueAtMiddle`. Ako jeste, vrati promenljivu `pathToTarget`.

# --hints--

Treba da kreiraš izraz sa stanjem `if` i `value === valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

Moraš vratiti parametar `pathToTarget` unutar izraza `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
