---
id: 6996c675b599baa7b103192d
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

U binarnoj pretrazi, postoji minimalni i maksimalni mogući indeks.

Najniža oznaka predstavlja apsolutnu levu granicu u trenutnom opsegu pretraženih vrednosti, u ovom slučaju, `searchList`. Najviša oznaka, s druge strane, predstavlja apsolutnu desnu granicu u pretraženim vrednostima.

Da bi se uzelo u obzir ta dva, morate uzeti u obzir ceo red pretraćenih podataka. Stoga, unutar funkcije ``binarySearch``, definišite razliku za ``low`` sa vrednošću ``0``, i razliku za ``high`` sa vrednošću koja uzima u obzir poslednji element u redu pretraćenih podataka.

# --hints--

Trebalo bi da imaš razliku nazvanu `low` postavljenu na `0` u tvojoj funkciji `binarySearch`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

Trebalo bi da imaš razliku nazvanu `high` postavljenu na `searchList.length - 1` u svojoj funkciji `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
