---
id: 6996cfb6c01e9a802add38be
title: Korak 13
challengeType: 1
dashedName: step-13
---

# --description--

Na kraju, ako `value` nije dostupan u sredini i takođe nije veći od `valueAtMiddle`, onda `value` mora biti manji od `valueAtMiddle` i mora biti na levoj strani. Poslednji deo od `else` će se baviti ovim.

Dakle, završi ciklus uklanjanjem glavne reči `break` i dodavanjem komada `else`. Unutar njega, ažuriraj parametar `high` uklanjanjem `1` iz `mid`.

# --hints--

Treba da promeniš glavnu reč od `break` sa delom `else` koji ispravlja parametar `high` u `mid - 1`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
--fcc-editable-region--      
    } 
    break;
--fcc-editable-region--
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
