---
id: 56533eb9ac21ba0edf2244b3
title: Prevedi temperaturu iz Celzijusa u Farenhajta
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

U ovoj laboratoriji, treba da napišete funkciju koja konvertuje temperaturu iz Celisiusa u Farenhajta. Formula za konverziju iz Celisiusa u Farenhajta je:

```js
fahrenheit = celsius * (9/5) + 32
```

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Trebalo bi da kreiraš funkciju nazvanu `convertCtoF`.
2. `convertCtoF` treba da primi jedan numerički argument, koji je temperatura u Celzijusima.
3. `convertCtoF` treba da vrati broj.

# --hints--

Trebalo bi da kreirate funkciju nazvanu `convertCtoF`.

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` Treba mu/Trebalo bi da prihvati jedan parametar.

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` Trebalo bi da vrati broj.

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` treba da vrati vrednost od `-22`.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` treba da vrati vrednost za `14`.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` treba da vrati vrednost `32`.

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` treba da vrati vrednost od `68`.

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` treba da vrati vrednost od `86`.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
