---
id: 67116c175a77cb64c5c10d49
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U prethodnim časovima, naučili ste kako da radite sa funkcijama višeg reda kao što su `map`, `filter` i `reduce`.

U ovom radionici, kreiraćete aplikaciju za upravljanje bibliotekom koja će vam pružiti priliku da vežbate rad sa različitim funkcijama višeg reda.

Da biste počeli, kreirajte varijablu nazvanu `library` i dodelite joj prazan niz.

# --hints--

You should create a `library` array.

```js
assert.isArray(library)
```

Your `library` array should be empty.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
