---
id: 6822694e3d34c807feeaba50
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

Upravo kao u prethodnom koraku, željeni rezultat bi izgledao bolje kao string umesto niza.

Povežite metod `join()` sa vašim postojećim metodom `map()` i prosledite `"\n"` kao separator u `join()`.

Zatim, ispod vaše funkcije `getBookSummaries`, dodajte `console.log()`. Izjava `console` bi trebalo da pozove funkciju `getBookSummaries` sa `library` kao argumentom.

Sada, u konzoli bi trebalo da vidite rezultat kao string.

# --hints--

Your `getBookSummaries` function should return a string.

```js
assert.isString(getBookSummaries(library));
```

Your `getBookSummaries` function should return a single string representing the final result.

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = `About A
About B
About C`;

assert.strictEqual(getBookSummaries(testLibrary), expected);
```

You should have a `console.log(getBookSummaries(library));` in your code.

```js
assert.match(code, /console\.log\(\s*getBookSummaries\(\s*library\s*\)\);?/gi);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`).join("\n");
}

console.log(getBookInformation(library));

console.log("\nList of book summaries:\n");

function getBookSummaries(catalog) {
--fcc-editable-region--
  return catalog.map((book) => book.about)
}

--fcc-editable-region--
```
