---
id: 61696ef7ac756c829f9e4048
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

U ovom projektu ćete koristiti eksternu datoteku CSS da stilizujete stranicu. Datoteka `styles.css` je već pripremljena za vas. Ali pre nego što je možete koristiti, moraćete ga povezati sa stranicom.

Umestite ga unutar elementa `link`, unutar elementa `head`. Atribut elementa `rel` sa vrednošću `"stylesheet"` i atribut elementa `href` sa vrednošću `"styles.css"`.

# --hints--

Trebalo bi da ima jedan element `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Vaš element za `link` treba da bude prazan element, ne sme imati zatvarajući tag `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Element vašeg `link` mora imati svojstvo elementa `rel` sa vrednošću `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Vaše polje `link` trebalo bi da ima svojstvo polja `href` sa vrednošću `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Tvoj element za `link` bi trebalo da bude unutar elementa za `head`.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
