---
id: 616968c2c94c8071b349c146
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U okviru `body`, postavite unutar elementa `h1` sa tekstom `CSS Color Markers`.

# --hints--

Vaš kod mora imati element `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

Trebalo bi da imate etiketu pakovanja za `h1`.

```js
assert.match(code, /<h1\s*>/i);
```

Trebalo bi da imaš etiketu za zatvaranje `h1`.

```js
assert.match(code, /<\/h1\s*>/i);
```

Tvoj element `h1` treba da bude unutar elementa `body`.

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

Vaš element od `h1` trebao bi da ima tekst od `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Zapamtite, velika i mala slova zajedno sa pravopisom su važne za pisanje `h1`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
