---
id: 616971b3cd990186b66c99a1
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sada nakon što je vaš eksterni fajl CSS pripremljen, možete početi da dekorujete stranicu.

Kao podsetnik, ovako označavate element pasusa i poravnajete ga desno:

```css
p {
  text-align: right;
}
```

Kreiraj novo pravilo za CSS koje cilja na karakteristiku `h1`, i postavi mu svojstvo `text-align` da bude `center`.

# --hints--

Trebalo bi da koristiš alat `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Vaše pravilo za CSS od `h1` trebalo bi da ima svojstvo `text-align` postavljeno na `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
