---
id: 616d55bd160a95e22453a081
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Sada si stavio jednu oznaku/tačku u sredini sa bojom, vreme je da dodaš druge oznake.

U elementu `.container`, dodaj dva elementa od `div` i daj svakom klasu `marker`.

# --hints--

Vaša prva nova funkcija za `div` mora imati oznaku za otvaranje.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][2]);
```

Tvoja prva nova karakteristika za `div` treba da ima zatvarajući tag.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][2]);
```

Tvoj novi drugi element `div` treba da ima labelu za otvaranje.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][3]);
```

Vaša druga nova komponenta za `div` mora imati zatvarajući tag.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][3]);
```

Tvoji novi elementi `div` bi trebalo da budu unutar `div` sa klasom `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
containerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(containerChildren, 3);
```

Svi tvoji novi elementi sa `div` moraju imati klasu `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert.isNotEmpty(containerChildren);
containerChildren.forEach(child => assert.isTrue(child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
