---
id: 616d595270d902f0e7086e18
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Ako imate tri različita elementa oznaka `div`, oni izgledaju kao jedan veliki pravougaonik. Trebalo bi da dodate razmak između njih kako biste lakše videli svaki element.

Kada atribut `margin` sa kratkim putem ima dve vrednosti, postavlja `margin-top` i `margin-bottom` za prvu vrednost, a `margin-left` i `margin-right` za drugu vrednost.

U vašem pravilu za `.marker` od CSS, postavite atribut `margin` na `10px auto`.

# --hints--

Vaše pravilo za `.marker` za CSS trebalo bi da ima svojstvo `margin` postavljeno na `10px auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

```
