---
id: 616d50b93ba424d6282c99cf
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Pozadinsku boju je primenjeno, ali pošto element oznake `div` nema sadržaj unutra, ne poseduje dužinu podrazumevano.

U vašem pravilu od `.marker` za CSS, postavite svojstvo `height` na `25px` i svojstvo `width` na `200px`

# --hints--

Tvoje pravilo za `.marker` od CSS trebalo bi da ima svojstvo `width` postavljeno na `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.width, '200px');
```

Vaše pravilo za `.marker` od CSS mora imati svojstvo `height` postavljeno na `25px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
