---
id: 616d525007b8c5d8b3308b1c
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Vaš znak će biti bolji ako je u centru stranice. Jednostavan način da to postignete je korišćenjem `margin` <dfn>funkcija skraćivanja</dfn>.

Možete postaviti oblast ivica elemenata pojedinačno sa atributima kao `margin-top` i `margin-left`, atribut skraćivanja `margin` omogućava lako postavljanje više oblasti ivice odjednom.

Da biste postavili svoj znak u sredinu stranice, postavite mu atribut od `margin` na `auto`. Ovo će postaviti `margin-top`, `margin-right`, `margin-bottom`, i `margin-left` sve na `auto`.

# --hints--

Vaše pravilo za `.marker` CSS mora imati svojstvo `margin` postavljeno na `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
