---
id: 616d3a67ccf800ad94ec89ae
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Sada ćete dodati neke komponente koje ćete na kraju pretvoriti u kolor markere.

Prvo, unutar elementa `body`, dodaj element `div` sa atributom elementa `class` postavljen na `container`. Osiguraj da je element `div` postavljen ispod elementa `h1`.

# --hints--

Trebalo bi da imate etiketu za otvaranje `div`.

```js
assert.match(code, /<div.*>/i);
```

Trebalo bi da imaš oznaku za zatvaranje `div`.

```js
assert.match(code, /<\/div\s*>/i);
```

Tvoj element `div` treba da bude unutar tvog elementa `body`.

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

Tvoj element `div` bi trebalo da bude posle elementa `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Vaš element od `div` bi trebalo da ima svojstvo elementa `class` postavljeno na `container`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}
```
