---
id: 61764f23ea21477b76f3b80f
title: Korak 12
challengeType: 0
dashedName: step-12
---

# --description--

Sledeće, uklonite atribut `background-color` i njegovu vrednost iz pravila `.marker` za CSS.

# --hints--

Tvoje pravilo za `.marker` od CSS ne bi trebalo da ima svojstvo i vrednost od `background-color`.

```js
assert.notInclude(code, 'background-color');
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
--fcc-editable-region--
  background-color: red;
--fcc-editable-region--
  margin: 10px auto;
}

```
