---
id: 6176519636a76b810ab1219a
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Kreiraj pravilo za CSS koje cilja na klasu objekata `two` i postavi mu svojstvo `background-color` da bude `green`.

Takođe, kreiraj različito pravilo za CSS koje cilja klasu objekata `three` i postavi `background-color` njegovo na `blue`.

# --hints--

Treba koristiti selektor klase objekata za ciljanje klase objekata `two`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.two'));
```

Vaše pravilo za CSS za `.two` trebalo bi da ima svojstvo `background-color` postavljeno na `green`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

Trebalo bi da koristiš klasifikator objekata za ciljanje klase objekata `three`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.three'));
```

Tvoja pravila za CSS koja su u `.three` treba da imaju karakteristiku `background-color` postavljenu na `blue`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
