---
id: 61764f7e5240eb7ccc7f6a0a
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Zatim, kreiraj novo pravilo za CSS koje cilja klasu objekata `one` i postavi mu svojstvo `background-color` na `red`.

# --hints--

Trebalo bi da koristiš klasifikator objekata za ciljanje klase objekta `one`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.one'));
```

Vaše pravilo za CSS od `.one` trebalo bi da ima svojstvo `background-color` postavljeno na `red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--

--fcc-editable-region--

```
