---
id: 61764c602bee6974e7790f35
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Da biste dodelili različite boje elementima, moraćete dodati jedinstvenu klasu za svaki. Možete dodati više klasama elementu navodeći ih u atribut ``class`` i razdvajajući ih razmakom. Na primer, sledeći kod dodaje klase ``animal`` i ``dog`` elementu ``div``:

```html
<div class="animal dog">
```

Ako dodate mnogo klasa na element HTML, stilovi prvih klasa koje navedete u stilskom listu mogu biti nadjačani kasnijim klasama.

Na početku, dodaj klasu za `one` za element prve oznake `div`.


# --hints--

Trebalo bi dodati klasu `one` za prvu oznaku `div`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

Tvoja prva `div` mora imati klase `marker` i `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('marker'));
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
--fcc-editable-region--
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
