---
id: 617b8b38f32bf2080a140675
title: Korak 17
challengeType: 0
dashedName: step-17
---

# --description--

Kao što vam je predano u prethodnim časovima, funkcija je deo koda koji može uzeti ulaz i izvršiti određenu akciju. Funkcija CSS `rgb` prima vrednosti, ili <dfn>hoja</dfn>, za crveno, zeleno i plavo, i kreira boje:

```css
rgb(red, green, blue);
```

Svaka značenje crvene, zelene i plave boje je broj od `0` do `255`. `0` znači da ta boja ima 0% procenta, i crna je. `255` znači da ta boja ima 100% procenta.

U pravilu `.one` od CSS, promeni osnovni naziv boje `red` u funkciju `rgb`. Za funkciju `rgb`, postavi vrednost crvene na `255`, vrednost zelene na `0`, i vrednost plave na `0`.

# --hints--

Tvoje pravilo za `.one` od CSS ne sme koristiti osnovnu boju `red` za postavljanje atributa `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'red');
```

Pravilo vaše `.one` od CSS bi trebalo da ima svojstvo `background-color` postavljeno na `rgb(255, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: red;
--fcc-editable-region--
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
