---
id: 617b65579ce424bf5f02ca73
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Kao što ste naučili u lekcijama, postoje dva glavna modela boja: aditivni model RGB (crvena, zelena, plava), koji se koristi u elektronskim uređajima, i subtraktywni model CMYK (cijan, magenta, žuta, crna), koji se koristi u štampariji.

U ovom projektu, koristićeš RGB model. To znači da boja počinje kao crna i menja se kako se uvode različiti nivoi crvene, zelene i plave. Jednostavan način za ovo je korišćenje funkcije CSS `rgb`.

Kreiraj novo pravilo za CSS koje cilja na klasu objekata `container` i postavi mu svojstvo `background-color` da bude crno za `rgb(0, 0, 0)`.

# --hints--

Trebalo bi da koristite selektor klase za ciljanje klase `container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.container'));
```

Tvoje pravilo za CSS od `.container` mora imati svojstvo `background-color` postavljeno na `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
