---
id: 617bd6ec666b1da2587e4e37
title: Korak 29
challengeType: 0
dashedName: step-29
---

# --description--

U RGB modelu aditivnih boja, <dfn>primarne boje</dfn> su boje koje, kada se pomešaju, daju čistu belu boju. Da bi se to postiglo, svaka boja mora biti na svom maksimalnom intenzitetu.

Pre nego što pomešate boje, vratite zelenu na čistu zelenu. U funkciji `rgb` CSS pravila `.two` postavite zelenu nazad na maksimalnu vrednost `255`.

Za Azure, ažuriraj funkciju `rgb` u pravilu CSS od `.two` da je zelena u `127` i da plava ima najvišu vrednost.

Na kwa waridi, ambao wakati mwingine hujulikana kama pinki angavu, sasisha kitendakazi cha `rgb` katika sheria ya CSS ya `.three` ili bluu iwe katika `127` na nyekundu iwe na thamani ya juu kabisa.

# --hints--

Vaše pravilo za CSS od `.one` treba da ima svojstvo `background-color` postavljeno na `rgb(127, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(127, 255, 0)');
```

Vaše pravilo za CSS od `.two` trebalo bi da ima svojstvo `background-color` postavljeno na `rgb(0, 127, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 255)');
```

Vaše CSS pravilo `.three` treba da ima svojstvo `background-color` postavljeno na `rgb(255, 0, 127)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
