---
id: 617b954d9f4f6217a749380e
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Sada povećajte vertikalni razmak malo više između vaših oznaka i ivica elementa `container` koji se nalaze unutra.

Da biste napravili chartreuse zelenu boju, ažurirajte funkciju `rgb` u CSS pravilu `.one` tako da crvena bude `127`, a zelena na maksimalnoj vrednosti.

# --hints--

Za ružičastu boju, koja se ponekad naziva i jarko roze, ažurirajte funkciju `rgb` u CSS pravilu `.three` tako da plava bude `127`, a crvena na maksimalnoj vrednosti.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

Sheria yako ya `.container` ya CSS inapaswa kuwa na sifa ya `padding` iliyowekwa kuwa `10px 0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.padding, '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
--fcc-editable-region--

--fcc-editable-region--
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
