---
id: 617b8e0d93a8d10d9a90e720
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Molimo proverite da je `background-color` vaša oznaka još uvek crvena. To je zato što ste postavili crvenu vrednost parametra `rgb` na vrednost veću od `255`, ili crveni procenat 100, i postavili sve zelene i plave vrednosti na `0`.

Sada koristi funkciju za `rgb` da postavi druge boje.

U pravilu CSS od `.two`, koristi funkciju `rgb` da postavi `background-color` na najveću zelenu vrednost, i `0` za druge vrednosti. I u pravilu CSS od `.three`, koristi funkciju `rgb` da postavi `background-color` na najveću plavu vrednost, i `0` za druge vrednosti.

# --hints--

Vaše pravilo za CSS od `.two` ne bi trebalo da koristi osnovni naziv boje `green` za postavljanje atributa `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

Tvoje pravilo za CSS od `.two` trebalo bi da ima svojstvo `background-color` postavljeno na `rgb(0, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

Tvoje pravilo za CSS od `.three` ne bi trebalo da koristi osnovni naziv boje iz `blue` za postavljanje atributa `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

Vaše pravilo za CSS, vezano za `.three`, mora imati atribut `background-color` postavljen na `rgb(0, 0, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
