---
id: 618a0927005553b74bfa05ae
title: Korak 31
challengeType: 0
dashedName: step-31
---

# --description--

Kolorni krug je krug gde su slične boje, ili <dfn>mionekano ya rangi</dfn>, blizu jedna drugoj, a različite su dalje. Na primer, čista crvena boja nalazi se između nijansi roze i narandžaste.

Dve boje koje se direktno susreću na točkići boja nazivaju se <dfn>rangi za nyongeza</dfn>. Ako se dve komplementarne boje pomešaju, biće sive. Ali kada se stave jedna pored druge, ove boje pružaju veliku razliku u vidljivosti i izgledaju življe.

U funkciji `rgb` za pravilo CSS od `.one`, postavite vrednost crvene boje na maksimum od `255` da biste dobili čistu crvenu boju. U funkciji `rgb` za pravilo CSS od `.two`, postavite vrednosti zelene i plave boje na maksimum od `255` da biste dobili ciano.

# --hints--

Vaše pravilo za CSS od `.one` trebalo bi da ima atribut `background-color` postavljen na `rgb(255, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

Vaše pravilo za CSS od `.two` trebalo bi imati atribut `background-color` postavljen na `rgb(0, 255, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
