---
id: 619b7424f43ec9215e538afe
title: Korak 39
challengeType: 0
dashedName: step-39
---

# --description--

Sledeće, izmeni klasu objekata `two` u `green` na drugom mestu `div` i izmeni klasu objekata `three` u `blue` na trećem mestu `div`.

# --hints--

Tvoja druga oznaka `div` ne sme da ima klasu objekata `two`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isFalse(containerSecondChild?.classList?.contains('two'));
```

Vaš drugi element `div` treba da ima klase objekata `marker` i `green`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('marker'));
assert.isTrue(containerSecondChild?.classList?.contains('green'));
```

Tvoja treća oznaka `div` ne sme imati klasu objekata `three`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isFalse(containerThirdChild?.classList?.contains('three'));
```

Tvoj treći element ``div`` treba da ima klase objekata ``marker`` i ``blue``.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('marker'));
assert.isTrue(containerThirdChild?.classList?.contains('blue'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      --fcc-editable-region--
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}

```
