---
id: 619b74fa777a2b2473c94f82
title: Korak 40
challengeType: 0
dashedName: step-40
---

# --description--

Ispravite konstruktor klase za CSS `.two` ili cilja na novu klasu `green`. Takođe, ispravite konstruktor klase za `.three` ili cilja na novu klasu `blue`.

# --hints--

Vaš kod više ne sme imati selektor klase za elemente `.two`.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.two'));
```

Treba da koristiš selektor klase objekata za ciljanje klase `green`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.green'));
```

Vaše pravilo za CSS od `.green` trebalo bi imati svojstvo `background-color` postavljeno na `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 0, 0)');
```

Tvoj kod ne bi trebalo da ima više selektora klase objekata `.three`.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.three'));
```

Trebalo bi da koristiš klasifikator objekata za ciljanje klase `blue`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.blue'));
```

Vaše pravilo za CSS od `.blue` mora imati svojstvo `background-color` postavljeno na `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
