---
id: 619b7c3c83de403126b69c1e
title: Korak 42
challengeType: 0
dashedName: step-42
---

# --description--

Verovatno već znate decimalne brojeve, ili bazu 10, koji počinju od 0 do 9. Šesnaestični brojevi, ili baza 16, počinju od 0 do 9, a zatim od A do F:

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

Za heksadecimalne boje, `00` je 0 procenata te boje, a `FF` je 100 procenata. Stoga `#00FF00` se tumači kao 0 procenata crveno, 100 procenata zeleno i 0 procenata plavo, i jednako je `rgb(0, 255, 0)`.

Smanite zelenu snagu postavljanjem zelene vrednosti boje na heksadecimalnu vrednost `7F`.

# --hints--

Vaše pravilo za `.green` od CSS trebalo bi da ima svojstvo `background-color` postavljeno na `#007F00`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
--fcc-editable-region--
  background-color: #00FF00;
--fcc-editable-region--
}

.blue {
  background-color: rgb(0, 0, 0);
}

```
