---
id: 61a49d15bdbb5e57cc6fd280
title: Korak 49
challengeType: 0
dashedName: step-49
---

# --description--

Tačke boja omogućavaju vam da ispravite gde su boje postavljene duž linije gradijenta. One se koriste u funkciji `linear-gradient` da bi specificirale gde se jedna boja završava i gde počinje sledeća promena boje. Bilo da je u pitanju procenat ili vrednost dužine, može se koristiti za definisanje pozicije tačke boje.

Na primer, u ovom crveno-crnom gradijentu boja, prelazak od crvene ka crnoj dešava se u području od 90% duž linije gradijenta boja, tako da crvena zauzima veliki deo dostupnog prostora:

```css
linear-gradient(90deg, red 90%, black);
```

U funkciji `linear-gradient`, dodaj tačku boje `75%` nakon prvog crvenog dela boje. Ne dodaj tačke boja za druge delove boje.

# --hints--

Vaše pravilo za `.red` CSS trebalo bi imati svojstvo `background` postavljeno na `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
