---
id: 61a489b8579e87364b2d2cdb
title: Korak 44
challengeType: 0
dashedName: step-44
---

# --description--

Navikovan ste na nekoliko načina postavljanja vertikalnih boja u CSS, ali možete koristiti i promenu boja, ili <dfn>gradient</dfn>, na elementu.

Gradijent je mesto gde jedna boja prelazi u drugu. Funkcija CSS `linear-gradient` vam je omogućila da kontrolišete pravac tih promena duž linije i koje se boje koriste.

Jedna stvar za pamćenje je da funkcija `linear-gradient` kreira element `image`, a obično je praćena svojstvom `background` koje može primiti sliku kao vrednost.

U pravilima CSS od `.red`, promeni atribut `background-color` u `background`.

# --hints--

Vaše pravilo za CSS od `.red` trebalo bi imati svojstvo `background` sa vrednošću `rgb(255, 0, 0)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
