---
id: 61ade49b2dad60c076cbd32d
title: Korak 78
challengeType: 0
dashedName: step-78
---

# --description--

Vaš logo izgleda dobro. Ali da bi izgledalo realnije, možete promeniti stil ivice u dve čvrste ivice.

Zbog skraćivanja `border-left`, promenite vrednost stila ivice od `solid` do `double`.

# --hints--

Vaše pravilo za `.sleeve` CSS trebalo bi imati svojstvo skraćivanja od `border-left` i vrednost `10px double black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.borderLeft, '10px double black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
--fcc-editable-region--
  border-left: 10px solid black;
--fcc-editable-region--
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
