---
id: 61b093429e7fc020b43b1bb6
title: Korak 71
challengeType: 0
dashedName: step-71
---

# --description--

Tvoje ruke izgledaju dobro, ali bi izgledale bolje da se postave na desnu stranu oznake. Jedan način da to uradiš je dodavanjem drugog elementa pre nego što ruka potisne udesno.

Dodaj `div` novi sa klasom `cap` pre elementa ručke `div`.

# --hints--

Vaša nova funkcija za `div` mora biti unutar crvenog označavanja od `div`.

```js
const redMarkerChildren = [...document.querySelector('div.red')?.children];
redMarkerChildren.forEach(child => {
  assert.strictEqual(child?.localName, 'div')
})
assert.lengthOf(redMarkerChildren, 2);
```

Vaš novi element za `div` treba da ima klasu `cap`.

```js
const redMarkerChildren = [...document.querySelector('div.red')?.children];
assert.isTrue(redMarkerChildren.some(child => child?.classList?.contains('cap')));
```

Tvoja kofija od `div` treba da bude pre ruke `div`.

```js
const redMarkerChildren = [...document.querySelector('div.red')?.children];
const cap = document.querySelector('div.cap');
const sleeve = document.querySelector('div.sleeve');
assert.isBelow(redMarkerChildren.indexOf(cap), redMarkerChildren.indexOf(sleeve));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      --fcc-editable-region--
        
      --fcc-editable-region--
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
