---
id: 61b093179e7fc020b43b1bb3
title: Korak 68
challengeType: 0
dashedName: step-68
---

# --description--

<dfn>Opacity</dfn> opisuje koliko je nešto neprozirno, ili ne-transparentno. Na primer, čvrsta zidna stena je neprozirna i kroz nju ne može da prođe svetlost. Ali čaša za piće je mnogo transparentnija, i možete videti kroz staklo na drugu stranu.

Sa svojstvom CSS `opacity`, možete kontrolisati kako je element skriven ili vidljiv. Za vrednost `0`, odnosno 0%, element će biti potpuno vidljiv, a za `1.0`, odnosno 100%, element će biti potpuno skriven kao po podrazumevanju.

Prema zakonu CSS `.sleeve`, postavi svojstvo `opacity` da bude `0.5`.

# --hints--

Tvoje pravilo za CSS `.sleeve` mora imati svojstvo `opacity` postavljeno na `0.5`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity, '0.5');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
--fcc-editable-region--
  
--fcc-editable-region--
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
