---
id: 612e89562043183c86df287c
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Pretraživači često primenjuju podrazumevane vrednosti margina i paddinga na određene elemente. Da biste bili sigurni da se vaš pianinski raspored prikazuje ispravno, morate resetovati box model.

Dodajte `html` selektor pravila u vaš CSS fajl i postavite svojstvo `box-sizing` na `border-box`.

# --hints--

You should have an `html` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

Your `html` selector should have the `box-sizing` property set to `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
