---
id: 612e8279827a28352ce83a72
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Sada, kopirajte skup od sedam `.key` elemenata i zalepite dva dodatna skupa u `.keys` div.

# --hints--

You should have 21 `.key` elements.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 21);
```

You should have 15 `.black--key` elements.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 15);
```

All `.key` elements should be within your `.keys` element.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.key'), 21);
```

All `.black--key` elements should be within your `.keys` element.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.black--key'), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
