---
id: 612ea4c4993aba52ab4aa32e
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Sada je vreme da koristite pseudo-selektore koje ste ranije postavili. Da biste kreirali crne tastere, dodajte novi `.key.black--key::after` selektor. On cilja elemente sa klasama i `key` i `black--key` i bira pseudo-element kreiran nakon ovih elemenata u HTML-u.

U novom selektoru, postavite `background-color` na `#1d1e22`. Takođe postavite svojstvo `content` na `""`. Ovo će učiniti pseudo-elemente praznim.

Svojstvo `content` koristi se za postavljanje ili prepisivanje sadržaja elementa. Podrazumevano, pseudo-elementi kreirani selektorima `::before` i `::after` su prazni, što znači da nisu prikazani na stranici. Postavljanjem svojstva `content` na prazan string `""`, osiguravate da se pseudo-elementi renderuju, a da pritom i dalje izgledaju prazno.

Ako želite da eksperimentišete, pokušajte da uklonite svojstvo `background-color` i postavite različite vrednosti za svojstvo `content`, kao što su `"♥"`. Zapamtite da ove promene poništite kada završite kako bi testovi prošli.

# --hints--

You should have a `.key.black--key::after` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Your `.key.black--key::after` selector should have a `background-color` property set to `#1d1e22`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor, 'rgb(29, 30, 34)');
```

Your `.key.black--key::after` selector should have a `content` property set to `""`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content, '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
