---
id: 612e89d254fe5d3df7d6693d
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Sada kada ste resetovali HTML box model, morate ovo primeniti i na elemente unutar njega. Da biste to uradili, moraćete da postavite svojstvo `box-sizing` svih ostalih elemenata na `inherited`, što osigurava da ciljani elementi usvoje istu vrednost kao i njihov roditeljski element.

Takođe ćete morati da ciljate pseudo-elemente, koji su specijalni ključni reči koje prate selektor. Dva pseudo-elementa koja ćete koristiti su `::before` i `::after` pseudo-elementi.

Selektor `::before` kreira pseudo-element koji je prvi potom izabranog elementa, dok selektor `::after` kreira pseudo-element koji je poslednji potom izabranog elementa. Ovi pseudo-elementi se često koriste za kreiranje kozmetičkog sadržaja, što ćete videti kasnije u ovom projektu.

Za sada, kreirajte CSS selektor za ciljanje svih elemenata koristeći `*`, i uključite pseudo-elemente `::before` i `::after`. Postavite svojstvo `box-sizing` na `inherit`.

# --hints--

You should have an `html` selector`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

Your `html` selector should have the `box-sizing` property set to `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

You should have a `*, ::before, ::after` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Your `*, ::before, ::after` selector should have the `box-sizing` property set to `inherit`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing, 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
