---
id: 6852ef95ccc27660275692cd
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Vaša klasa za `Email` treba da sačuva informacije o svakoj e-pošti kada se kreira.

U okviru vaše klase ``Email``, uklonite reč koja je ``pass`` i nalazi se u metodi ``__init__`` koja prima parametre ``self``, ``sender``, ``receiver``, ``subject`` i ``body``.

# --hints--

Morate definisati metodu `__init__` unutar klase `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

Tvoja putanja ``__init__`` mora imati parametre: ``self``, ``sender``, ``receiver``, ``subject``, i ``body``.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
