---
id: 6852ea0c3cadb54e897acdac
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovoj radionici, kreiraćete *Simulators e-pošte* koji simulira slanje, primanje i upravljanje e-porukama između različitih korisnika. Naučićete o klasama objekata, objektima i kako strukturirati kod koristeći pristup usmeren na objekte.

Počnite kreiranjem klase pod nazivom `Email`, koristeći glavni ključni izraz `class`.

# --hints--

Trebalo bi da kreirate klasu objekata nazvanu `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
