---
id: 6852f0c3fcedb964ab959921
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Sada pokušajte da vidite ako je atribut funkcije `read` ispravno dodato u klasu objekata `Email`. Pošto već imate e-mail sadržaj iz prethodnih koraka, objavite atribut funkcije `read` da biste videli da je sada `False` podrazumevano.

# --hints--

Trebalo bi da ispišete `email_obj.read` ili pokušate sa funkcijom čitanja.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
