---
id: 6852f11d412c29665985a915
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Sada korisnici mogu poslati e-poruke, a potrebna im je metoda za skladištenje primljenih e-poruka. Svaki korisnik bi trebalo da ima inbox za prikupljanje svojih e-poruka.

Dodajte svojstvo elementa `inbox` u metodu `User` klase objekata `__init__` i počnite ga sačuvati kao prazan niz `[]`. Ovo će sačuvati sve e-poruke koje korisnik prima.

# --hints--

Morate dodati `self.inbox = []` u metodu `User` klase objekata `__init__`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
