---
id: 6852f06fdde1b4634a263f03
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Sada probajte klasu `Email` kreiranjem objekta za e-poštu i proveravanjem da li pravilno čuva podatke. Ispišaćete dve osobine elementa kako biste potvrdili da sve radi.

Kreiraj stav e-pošte nazvan `email_obj` koristeći `alice@example.com` kao pošiljaoca, `bob@example.com` kao primaoca, `Hello` kao naslov, i `Hi Bob!` kao telo e-pošte. Zatim ispiši svojstva stavova `sender` i `subject` kao instrukcije za prikaz razlika kako bi se potvrdilo da su pravilno sačuvani.

# --hints--

Trebalo bi da kreirate e-mail sadržaj nazvan `email_obj` koristeći navedene parametre.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

Treba da ispišete `email_obj.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

Trebalo bi da odštampate `email_obj.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--

--fcc-editable-region--
```
