---
id: 6852f01116912861c7d1b0e6
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Unutar metode `__init__`, morate postaviti vrednosti parametara na svojstva objektnog polja kako bi svaki e-mail sačuvao svoje podatke.

U okviru vašeg procesa `__init__`, uklonite glavni termin iz `pass` i unesite vrednosti parametara u `self.sender`, `self.receiver`, `self.subject`, i `self.body`.

# --hints--

Trebalo bi da postaviš vrednost parametra `sender` na `self.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

Treba da postaviš vrednost parametra `receiver` na `self.receiver`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

Trebalo bi da postaviš vrednost parametra `subject` na `self.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

Trebalo bi da postaviš vrednost parametra `body` na `self.body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
