---
id: 6853f4584f508e95cc7a9f4d
title: Korak 30
challengeType: 20
dashedName: step-30
---

# --description--

Sada ćete kreirati metodu za listanje svih e-poštama u poštanskom sandučić́u. Ova metoda treba da obradi slučaj kada je poštanski sandučić prazan i prikaže listu adresa e-pošte tamo gde postoje poruke.

Dodaj metodu nazvanu `list_emails` u klasu `Inbox` koja uzima `self` kao parametar. Prvo, kreiraj metodu `if` da proveri da li je poštansko sanduče prazno proveravanjem liste `emails`. Ako je prazno, odštapi poruku `Your inbox is empty.\n`. Dodaj metodu `return` za izlazak iz te metode.

# --hints--

Treba da definišete metodu pod nazivom `list_emails` u klasu objekata `Inbox` koja uzima `self` kao parametar.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

Treba da proveriš da li je `self.emails` prazno korišćenjem `if not self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

Trebalo bi da odštampate poruku `Your inbox is empty.\n` ako nema e-pošte.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

Treba da imaš izraz `return` na kraju izraza `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    
--fcc-editable-region--

```
