---
id: 6853f2da9b8ed891e9fdc5fc
title: Korak 23
challengeType: 20
dashedName: step-23
---

# --description--

Sada ćemo prikazati informacije o pošiljaocu i primaocu. Da bismo prikazali ko je poslao i primio e-poštu, dodajte dve naredbe za štampanje na putanji `display_full_email` sa ovim formatom:

- `From: sender` gde je `sender` zamenjeno imenom pošiljaoca
- `To: receiver` gde je `receiver` zamenjeno imenom primaoca

# --hints--

Trebalo bi da koristiš f-stringe za štampanje imena pošiljaoca koristeći `From: {self.sender.name}`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'From: {self.sender.name}')")`));
    }
})
```

Trebalo bi da koristiš f-stringe za štampanje imena primaoca koristeći `To: {self.receiver.name}`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'To: {self.receiver.name}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
