---
id: 6853f4ded9f84896ff182dd8
title: Korak 31
challengeType: 20
dashedName: step-31
---

# --description--

Nakon pregleda praznog sandučića za poštu, objavi poruku `\nYour Emails:`. Kasnije, ponovi kroz sve e-poruke u inboxu koristeći petlju `for` koja broji. Ovo je sintaksa za brojanje od određenog broja: `enumerate(iterable, start=start_number)`.

Koristite brojač da dodelite e-mail adrese počevši od `1`. Koristite `i` (primer) i `email` (primer e-mail adrese) kao parametre petlje.

Unutar petlje, ispišite f-string sa varijablom petlje praćenom `.`, zatim razmak i reprezentaciju niza karaktera za varijablu e-pošte ovim formatom:

```md
i. string_representation
```

Gde je `i` šablon, a `string_representation` predstavlja šablon e-mail adrese.

# --hints--

Trebalo bi da odštampate naslov `\nYour Emails:` izvod iz `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

Trebalo bi da imaš ciklus od `for` koji se ponavlja preko `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

Vaš ciklus od `for` treba da koristi `i` i `email` za ponavljanje preko `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

Unutar petlje, treba da odštapi sažetak svake e-pošte koja ima prefiks broja praćen nizom slova generisanim iz metode `__str__`. Ne zaboravite da dodate `.` nakon broja.

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        
--fcc-editable-region--
```
