---
id: 6859400f8f7f872e50e1238c
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Darasa `User` zahteva način za slanje e-pošte drugim korisnicima. Kada pošaljete e-mail, kreiraćete novi objekat `Email` i dodati ga u kutiju primaoca.

Kreiraj metodu nazvanu `send_email` u svojoj klasi `User`. Ova metoda treba da primi parametre za `receiver`, `subject`, i `body`.

U okviru tog puta, kreiraj novi objekat za `Email` sa sledećim parametrima i postavi njegovu vrednost u parametar nazvan `email`:

- `sender`: `self` (Pošiljalac)
- `receiver`: `receiver` (Primaocu e-pošte)
- `subject`: `subject` (Predmet e-pošte)
- `body`: `body` (Sadržaj e-pošte)

# --hints--

Treba da definišete metodu za `send_email` u klasi `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

Putanja za `send_email` treba da ima parametre: `self`, `receiver`, `subject`, na `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

Ova putanja treba da kreira objekat `Email`, nazvanog `email` sa tačnim parametrima.

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
