---
id: 68594df7779a164e6a61e8ae
title: Korak 32
challengeType: 20
dashedName: step-32
---

# --description--

Sanduk za pošte zahtevaju način za čitanje specifičnih e-pošta. Kada korisnik želi da pročita e-poruku, on će navesti broj e-poruke (počevši od indeksa 0) koju želi da vidi, i taj način će prikazati kompletan sadržaj te e-poruke.

Dodaj metodu nazvanu `read_email` u vašu klasu `Inbox` koja prima parametar `index`. Prvo, proveri da li je poštansko sanduče prazno i odštapi poruku `Inbox is empty.\n` ako jeste. Dodaj izjavu `return` nakon toga da izađeš iz te metode.

# --hints--

Trebalo bi da definišete metodu nazvanu `read_email` u klasi `Inbox` koja prima parametre `self` i `index`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").has_args("self, index")`));
    }
})
```

Treba da proveriš da li je `self.emails` prazno korišćenjem `if not self.emails`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`));
    }
})
```

Trebalo bi da odštampate `Inbox is empty.\n` ako nema e-pošte i povratka.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[0].is_equivalent("print('Inbox is empty.\\\\n')")`));
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[1].is_equivalent("return")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
    
--fcc-editable-region--
```
