---
id: 68593ea22eac7e228ed585ce
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

E-mail poruke bi trebalo da prate da li su pročitane ili ne. Dodajte atribut elementa `read` na putanju `__init__` i postavite ga na `False` podrazumevano, jer nove e-poruke počinju kao nepročitane.

# --hints--

Trebalo bi da postaviš vrednost `False` za `self.read` u putanju `__init__`.

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
