---
id: 685c1e32254e5993a5c8823e
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Sada si kreirao klasu za `Email`, je vreme da kreiraš korisnike koji mogu slati i primati e-poštu. Svaki korisnik će imati ime i moći će da obavlja e-mail operacije.

Kreiraj klasu `User` sa metodom `__init__` koja prima parametre `self` i `name`. Unutar metode dodeli vrednost parametra `name` atributu `self.name`.

# --hints--

Treba da kreiraš klasu `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

Klasa za `User` trebalo bi da ima metodu `__init__` koja prima parametar `self` i `name`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

Morate postaviti vrednost parametra `name` za `self.name`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
