---
id: 685c1eba966c6396850736ea
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Kada korisnici mogu slati e-poruke i imati poštu, potrebna je posebna klasa za efikasno upravljanje operacijama poštanske sandu.

Klasa `Inbox` će čuvati listu e-poruka i obezbediti metode za dodavanje novih e-poruka, prikaz svih e-poruka i upravljanje njima.

Kreiraj novu klasu nazvanu `Inbox`, koja ima metodu `__init__` i koja inicijalizuje prazan spisak nazvan `emails`.

# --hints--

Trebalo bi da napraviš razred za `Inbox`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

Klasa `Inbox` treba da ima metodu `__init__` koja prima `self` kao parametar.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

Treba da inicijalizujete `self.emails = []` u metodi `Inbox` klase `__init__`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
