---
id: 685d2fd1f309b9191f257c55
title: Korak 40
challengeType: 20
dashedName: step-40
---

# --description--

Pre nego što unesete vremenske oznake u sistem e-pošte, vežbajte korišćenje formatiranja datuma i vremena. Funkcija `datetime.datetime.now()` daje nam trenutni datum i vreme, a možete koristiti metodu `strftime()` da ih organizujete na različite načine.

Ovo je kako `strftime()` radi sa kodom strukture:

```python
now = datetime.datetime.now()
print(now.strftime("%Y-%m-%d"))  # Output: 2024-03-15 (year-month-day with - separator)
```

Formatni kodovi poput `%Y` (godina), `%m` (mesec) i `%d` (dan) govore funkciji `strftime()` šta da uključi, a između njih možeš dodati razdvajače poput `-`.

Na samom dnu svog koda, kreiraj varijablu nazvanu `current_time` i postavi joj vrednost `datetime.datetime.now()`. Zatim koristi `strftime()` da odštampa vreme u formatu sat:minut:sekund koristeći `:` kao separator.

Koristite ove kodove strukture: `%H` za sate (format 24 sata), `%M` za minute, i `%S` za sekunde.

# --hints--

Treba da kreirate parametar nazvan `current_time` i postavite mu vrednost `datetime.datetime.now()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_stmt("current_time = datetime.datetime.now()")`));
    }
})
```

Trebalo bi da koristiš `strftime()` za štampanje vremena u formatu `%H:%M:%S`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(current_time.strftime('%H:%M:%S'))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

--fcc-editable-region--

--fcc-editable-region--
```
