---
id: 688754d6d59578df1313cb3c
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Sada je vreme za povezivanje klasa User i Inbox kako bi se e-pošta mogla isporučiti!

Korisnici treba da imaju stvarne stvari u Inbox-u umesto jednostavnog spiska, a način `send_email` bi trebalo da dostavi e-poštu u inbox primaoca.

Ispraviti klasu `User` `__init__` na način da svaki korisnik dobije stvarni sadržaj Inboxa umesto unapred generisanog praznog spiska.

# --hints--

Treba da ispraviš klasu korisnika (User) kako bi se inicijalizovao `self.inbox = Inbox()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
