---
id: 688754f3494b17e0b352d67d
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Sada je vreme da testirate kompletan sistem e-pošte! Kreirajte dva korisnika i vidite kako funkcionalnost e-pošte radi.

Kreiraj dva objekta korisnika: `alice` sa imenom `"Alice"` i `bob` sa imenom `"Bob"`. Ovo će pokazati kako korisnici mogu biti kreirani i međusobno interagovati.

# --hints--

Treba da napraviš nešto za Korisnika sa imenom `alice` i sa imenom `"Alice"`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("alice").is_equivalent('alice = User("Alice")')`));
    }
})
```

Treba da napraviš nešto za Korisnika sa imenom `bob` i sa imenom `"Bob"`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("bob").is_equivalent('bob = User("Bob")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--

--fcc-editable-region--
```
