---
id: 68875554c6c32de788afc64e
title: Korak 21
challengeType: 20
dashedName: step-21
---

# --description--

Korisnici treba da imaju način za pravilno čitanje svojih e-poštama.

Dodaj metodu `display_full_email` u klasu `Email` koja uzima `self` samo kao parametar. Unutar ove metode, pozovi metodu `mark_as_read()` da označi e-mail kao pročitano kada ga neko pregleda.

Ovaj način će se koristiti za prikaz kompletnog sadržaja e-pošte u lepšem formatu.

# --hints--

Trebalo bi da definiraš metodu `display_full_email` u klasi `Email`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

Putanja za `display_full_email` treba da uzme `self` kao jedini parametar.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

Ova metoda treba da pozove `self.mark_as_read()` da označi e-mail kao pročitan.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
