---
id: 6887550fe36f64e27074475a
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Neka Alice pošalje e-mail Bobu da vidi da li je slanje e-mailova u redu.

Iskoristi Alisin metod `send_email` da pošalješ Bobu e-poruku sa naslovom `"Hello"` i sadržajem `"Hi Bob, how are you?"`.

# --hints--

Treba da pozoveš `alice.send_email()` sa ispravnim parametrima.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call('alice.send_email(bob, "Hello", "Hi Bob, how are you?")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
alice = User("Alice")
bob = User("Bob")

--fcc-editable-region--
```
