---
id: 68b9350183ff05bd106c79e4
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Sada, pokušaj metodom `mark_as_read` koju si kreirao u klasi objekata `Email`. Koristi tu metodu na elementu/polju `Email` koji se nalazi u `email_obj` da promeniš stanje. Nakon postavljanja e-pošte kao pročitane, prikaži atribut elementa `read` iz `email_obj` da potvrdiš promene.

# --hints--

Trebalo bi da pozoveš rutu `mark_as_read` na `email_obj` kako bi označio e-mail kao pročitan.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

Trebalo bi da odštampate svojstvo atributa `read` za `email_obj` nakon pozivanja metode `mark_as_read`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)

--fcc-editable-region--
```
