---
id: 68b930ee0e7b8fb81895e59a
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Sada ćete dodati mogućnost označavanja e-pošte kao pročitane. Kada neko otvori e-poruku, želite da promenite njen status sa nepročitano na pročitano. Ovaj način će biti jednostavan – samo je potrebno postaviti atribut elementa `read` na `True`.

Dodaj metod nazvan `mark_as_read` u klasu vaših objekata `Email`. Ova metoda bi trebalo da uzima `self` kao jedini parametar i postavi svojstvo elementa `read` na `True`.

# --hints--

Trebalo bi definisati metodu nazvanu `mark_as_read` u klasi `Email`. koja mora da uzme `self` kao jedini parametar.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

Ta putanja bi trebalo da postavi `self.read` na `True`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
