---
id: 68b96d6a290d52d5e57a7485
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Uklonite `email_obj` sa stabla za štampanje sledećih.

# --hints--

Ne bi trebalo da imate liniju za kreiranje `email_obj` u vašem kodu.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`));
    }
})
```

Ne bi trebalo da imaš `print(email_obj.sender)` u svom kodu.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.sender)")`));
    }
})
```

Ne smete imati `print(email_obj.subject)` u vašem kodu.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.subject)")`));
    }
})
```

Ne bi trebalo da imaš `print(email_obj.read)` u svom kodu.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

Ne bi trebalo da imaš `email_obj.mark_as_read()` u svom kodu.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
email_obj.mark_as_read()
print(email_obj.read)
--fcc-editable-region--
```
