---
id: 68b9366fae8b13c1278761b1
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Sandući za poštu zahtevaju način za primanje novih poruka. Kada neko pošalje e-mail korisniku, on bi trebalo da se doda u njegov inbox.

Dodajte metodu nazvanu `receive_email` u klasu `Inbox` koja uzima `self` i objekat e-pošte `email` kao parametre. Unutar tela te metode, dodajte `email` u listu `emails` koristeći metodu `append`.

# --hints--

Trebalo bi da definišete metodu nazvanu `receive_email` u klasi objekata `Inbox` koja uzima `self` i `email` kao parametre.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

Treba da dodaš `email` u `self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
