---
id: 68be6a4447e61c4395ef2359
title: Korak 36
challengeType: 20
dashedName: step-36
---

# --description--

Vaš sandučić za poštu takođe zahteva način za brisanje e-pošte.

Neno kuu la Python `del` može se koristiti za brisanje stavki iz liste korišćenjem njihovog obrasca.

Dodajte putanju nazvanu `delete_email` u vašoj klasi za `Inbox`. Kao što je `read_email`, treba da:

- Uzeti parametar `index`, proveriti da li je poštansko sanduče prazno, i poslati poruku `Inbox is empty.\n` ako jeste. Koristiti `return` iz tog puta.

# --hints--

Trebalo bi da definišete metod nazvan `delete_email` u klasi `Inbox` koji prima `self` i `index` kao parametre.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").has_args("self,index")`))
    }
})

```

Treba vam prvo da proverite da li je poštanska kutija prazna koristeći `if not self.emails:`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})

```

Ako je poštanska kutija prazna, treba da odštampate poruku sa `Inbox is empty.\n` i koristeći `return` iz tog načina.

```js
({
    test: () => 
    {
    assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_body().is_equivalent("print('Inbox is empty.\\\\n')\\nreturn")`))
}
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()
--fcc-editable-region--

--fcc-editable-region--
```
