---
id: 64dc916138dbe03bdfe6cec6
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Cilj Tore na Hanoiju je da se prebace svi diskovi na poslednji šipac. Da biste to uradili, morate pratiti tri jednostavna pravila:

1. Moguće je da pomerate samo najgornje diskove.
2. Možete prebacivati samo jedan disk odjednom.
3. Ne možete staviti veliki disk na mali disk.

Ispod vašeg postojećeg koda, deklarište praznu funkciju nazvanu `move`. Kasnije, koristićeš tu funkciju za prebacivanje diska između šipaka. Trenutno, kako bi izbegao greške, koristi glavnu reč `pass` unutar tela funkcije.

# --hints--

Trebalo bi da objaviš praznu funkciju koja se zove `move`. Seti se da koristiš ključnu reč `pass` unutar tela funkcije u ispravnom smeru.

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
