---
id: 64dc922df2919e3c38cead61
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Na samom vrhu vašeg koda, deklariši parametar nazvan `NUMBER_OF_DISKS` za skladištenje broja diskova i sa podrazumevanom vrednošću od `3`. Zatim, promenite prvi argument koji se prosleđuje funkciji/metodi `range()` sa vašim novim parametrom.

# --hints--

Trebalo bi da objaviš parametar nazvan `NUMBER_OF_DISKS`.

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

Parametar za `NUMBER_OF_DISKS` mora imati vrednost `3`.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

Trebalo bi da promeniš vrednost prvog argumenta funkcije `range()` u parametar `NUMBER_OF_DISKS`.

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
