---
id: 64dc90949ff85e3b37be40d0
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Funkcija `range()` vraća neprekidnu sekvencu brojeva. Kao što vidite, tip podataka za `rods['A']` je `range`, ali želite da bude lista.

Pregledajte vaš poziv za `range()` za funkciju `list()` ili to uradite.

# --hints--

Morate proslediti svoj poziv `range()` funkciji `list()`.

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
