---
id: 64dc97005bc7943e2978df0a
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Sada možete ukloniti svoj poziv od `print()`. Zatim, unutar funkcije od `move()`, uklonite glavnu reč od `pass` i ispišite sadržaj svog reчника za `rods`.

# --hints--

Ne bi trebalo da imaš `print(number_of_moves)` u svom kodu.

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

Trebalo bi da uklonite glavnu reč `pass` i da ispišete rječnik `rods`.

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
