---
id: 64dc92a9718fb53ca3d1bad7
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Zadatak Torusa Hanoija može se rešiti u koracima 2<sup>n</sup> - 1, gde je n broj diskova. Deklariši parametar nazvan `number_of_moves` i postavi vrednost ukupnog broja koraka u ovaj parametar.

Funkcija snage u Python je `**`.

# --hints--

Trebalo bi da objavite parametar koji se zove `number_of_moves`.

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

Vrednost `number_of_moves` bi trebalo da bude matematički izraz koji računa broj koraka.

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
