---
id: 64dc90e837670a3ba26a1e61
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Pošto je tip `list` kako je zahtevano, možete ukloniti poziv `print()`.

# --hints--

Trebalo bi da uklonite `print(type(rods['A']))` iz vašeg koda.

```js
({ test: () => assert.match(code, /\}(?!\s+print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")A\1\s*\]\s*\)\s*\))/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
