---
id: 64dc8ea01436383a88256d7f
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Izazov počinje sa diskovima složenim na prvoj šipci, u opadajućoj veličini, gde je najmanji disk na vrhu, a najveći disk na dnu. Morate da popunite svoj prvi popis brojevima koji predstavljaju različite veličine diskova.

Umesto ručnog dodavanja stavki u prvu listu, generišite opadajući niz brojeva od `3` do `1` koristeći funkciju `range()` i postavite njegovu vrednost u `rods['A']`. Ovde, `3` predstavlja najveći disk ispod montaže, a `1` predstavlja najmanji disk iznad montaže.

Sintaksa je `range(x, y, h)`, gde `x` je početni broj (uključujući), `y` je krajnji broj (isključujući), i `h` je razlika između jednog broja i sledećeg u sekvenci.

# --hints--

Trebalo bi da koristite funkciju `range()` za postavljanje sekvencijalne brojne vrednosti na `rods['A']`. Sintaksa za pozivanje funkcije `range()` je `range(x, y, h)`.

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

Trebalo bi da koristiš funkciju `range()` za postavljanje serijske vrednosti brojeva od `3` do `1` u `rods['A']`.

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
