---
id: 64dc985e6720a23edac63b51
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Dozvoljene diskovne pozicije između stubova pokazuju obrazac koji se ponavlja svake tri pozicije. Na primer, pozicije između stubova `A` i stubova `C` su dozvoljene u prvoj, četvrtoj i sedmoj poziciji, gde je ostatak deljenja broja pozicije sa 3 jednak 1.

Unutar petlje ``for`` koja je prethodno kreirana, promenite postojeći poziv ``print()`` u deklaraciju ``if`` koja se pokreće kada je ``(i + 1) % 3 == 1``. Unutar deklaracije ``if``, ispišite ``f'Move {i + 1} allowed between {source} and {target}'`` koristeći f-string. Molimo, zapamtite da je ``i + 1`` indeks iteracije od trenutka kada je ``i`` bio nula tokom prve petlje.

# --hints--

Treba da imaš deklaraciju `if` koja bi trebalo da se izvrši tokom `(i + 1) % 3 == 1`.

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

Trebalo bi da ispišete sekvencu karaktera `f'Move {i + 1} allowed between {source} and {target}'` koristeći f-string.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
