---
id: 64dc8b1faf9c60304ca8a9b6
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovom projektu, rešićete matematičku zagonetku poznatu kao Kužica Hanoi. Ova zagonetka ima tri stuba i nekoliko diskova različitih veličina.

Cilj ove zagonetke je pomeranje diskova sa prve kolone na treću kolonu, uz poštovanje specifičnih pravila koja zabranjuju stavljanje većeg diska na manji disk.

Počnite kreiranjem praznog rečnika koji se zove `rods` ili predstavljate kolone.

# --hints--

Trebalo bi da imaš parametar nazvan `rods`.

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

Vaš parametar za `rods` treba da bude prazan rečnik.

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
