---
id: 64dc9004f658183afa85f4f1
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Sada proveri tip podataka vaših ključa ``'A'`` prosleđivanjem ga u funkciju ``type()`` i isprikaži rezultate na terminalu.

# --hints--

Morate proslediti vaš ključ za `'A'` u funkciju `type()`.

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

Trebalo bi da odštampate tip `rods['A']` koristeći funkciju `print()`.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
