---
id: 64de1469fccce61940f498c4
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Rečnik za `rods` će predstavljati tri šipke sa svojim diskovima. Daj niz karaktera `'A'`, `'B'`, i `'C'` kao ključeve i postavi svaki od njih na prazan spisak.

# --hints--

Vaš rečnik za `rods` mora imati ključ od `'A'`.

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']` Trebalo bi da bude prazna lista.

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

Vaš rječnik za `rods` treba da ima ključ od `'B'`.

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']` trebalo bi da bude prazna lista.

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

Rečnik za `rods` mora imati ključ od `'C'`.

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']` Trebalo bi da bude prazna lista.

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
