---
id: 64df496c6a8ddf1b38db1ed6
title: Korak 50
challengeType: 20
dashedName: step-50
---

# --description--

U prethodnom koraku, napisali ste kod za pomeranje najvećeg diska iz manjeg izazova ka ciljnoj mašti.

Sada, ono što treba da uradiš je dodati još jedan rekurzivni poziv za pomeranje diskova `n - 1` koje si već preneo. Kopiraj prvi rekurzivni poziv i zalepi ga na kraj elementa/sekcije `if`.

Zapamtite da operativni argumenti nisu u ispravnom redosledu. Pokušajte da utvrdite ispravan redosled.

# --hints--

Trebaće ti da pozoveš `move(n - 1,  auxiliary, source, target)` na kraju svoje funkcije `move`.

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
