---
id: 64df3f1011888113fbd3d81b
title: Korak 44
challengeType: 20
dashedName: step-44
---

# --description--

Da bismo rešili problem korišćenjem rekurzije (recursion), prva stvar je da podelimo početni problem na sve manje probleme.

Poslednji raspored diskova `n` koji su prikupljeni na treći šip u opadajućem poretku može se pronaći pomeranjem:

- diskove `n - 1` od izvora do podpore
- najveće diskove od izvora do destinacije
- i zatim diskove `n - 1` od podpore do destinacije.

Dakle, prva stvar koju funkcija ``move`` treba da uradi je da pozove sebe sa ``n - 1`` kao prvi argument. Ali ako pokušate to bez definisanja baznog slučaja, dobićete ``RecursionError``. To se dešava zato što funkcija nastavlja da poziva sebe neograničeno.

Pre vašeg komentara i poziva vašeg `print()`, dodajte rekurzivni poziv funkcije sa `n - 1` kao prvi argument i osigurajte da se telo funkcije izvršava samo kada `n` je veće od nule. Za sada, ostavite ostale argumente u istom redosledu.

# --hints--

Telo funkcije `move` trebalo bi početi sa dubinom od `if` koja je aktivirana kada `n` je veće od nule.

```js
({ test: () => assert.match(code, /def\s+move\(\s*n\s*,\s*source\s*,\s*auxiliary\s*,\s*target\s*\)\s*:\s+if\s+n\s*>\s*0\s*:/) })
```

Trebalo bi da pozoveš `move(n - 1, source, auxiliary, target)` unutar deklaracije `if`. Zapamti da ubaciš svoj poziv `print()`.

```js
const transformedCode = "\n" + code.replace(/\r/g, '');
const move = __helpers.python.getDef(transformedCode, "move");
const { function_parameters, function_body } = move;

assert.match(function_parameters, /n\s*,\s*source\s*,\s*auxiliary\s*,\s*target/);
const functionBodyIndent = function_body.match(/ +/)[0];
const re = new RegExp(`${functionBodyIndent}if\\s+n\\s*>\\s*0\\s*:\\s+^${functionBodyIndent}( +)move\\(\\s*n\\s*-\\s*1\\s*,\\s*source\\s*,\\s*auxiliary\\s*,\\s*target\\s*\\).+?^${functionBodyIndent}\\1print\\s*\\(\\s*rods\\s*,\\s*("|')\\\\n\\2\\s*\\)`, "ms");
assert.match(function_body, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
    
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
