---
id: 64df47b32b92301a815d5ef8
title: Korak 49
challengeType: 20
dashedName: step-49
---

# --description--

U početku, dodani rekurzivni poziv (recursive call) rešava mali izazov prebacivanja `n - 1` diska na drugi šapi.

Zbog toga, argument `target` se odnosi na vašu drugu poziciju, dok je argument `auxiliary` treća pozicija. Zapamtite da će menjati poziciju kako ponavljanje nastavlja.

Ispravite redosled argumenata menjanjem `target` i `auxiliary` u vašem ponovnom pozivu.

# --hints--

Trebalo bi da promeniš redosled argumenata u tvom pozivu za `move(n - 1, source, auxiliary, target)`.

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
