---
id: 650c6082e5586f9e3acfcd3b
title: Korak 52
challengeType: 20
dashedName: step-52
---

# --description--

Sada, obrišite rečnik `rods` i promenite njegove ključeve u parametre `A`, `B`, i `C`, za taj tok, čuvajući njihove vrednosti. Izmenite svoj kod da prikaže ove promene. Ako vidite rezultate u terminalu, uspešno ste izvršili ovaj korak.

# --hints--

Treba da promeniš `rods['A']` u parametar nazvan `A`, čuvajući njegovu vrednost.

```js
assert.match(code, /A\s*=\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/);
```

Trebalo bi da promeniš `rods['B']` u parametar nazvan `B`, čuvajući njegovu vrednost.

```js
assert.match(code, /B\s*=\s*\[\s*\]/);
```

Trebalo bi da promeniš `rods['C']` u parametar nazvan `C`, čuvajući njegovu vrednost.

```js
assert.match(code, /C\s*=\s*\[\s*\]/);
```

Trebalo bi da ispravite poziv `print()` da štampa `A`, `B`, `C`, i vrednost `rods`. Sačuvajte ponašanje za novi red u pozivu `print()`.

```js
assert.match(code, /print\s*\(\s*A\s*,\s*B\s*,\s*C\s*,\s*('|")\\n\1\s*\)/);
```

Trebalo bi ti da imaš `target.append(source.pop())` u svom kodu.

```js
assert.match(code, /target\.append\s*\(\s*source\.pop\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```
