---
id: 65782342850feb3b8d62f936
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Pošto ćete često koristiti matematički izraz `(i + 1) % 3`, lako je ga sačuvati u varijablu.

Nešto iznad tvog deklarisanja za `if`, deklariraj parametar `remainder` i postavi vrednost `(i + 1) % 3` za ovaj parametar.

# --hints--

Trebali biste da deklarišete parametar `remainder` malo iznad sekcije `if`.

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=/m) })
```

Vrednost vašeg parametra `remainder` bi trebalo da bude numerički izraz `(i + 1) % 3`.

```js
({ test: () => assert.match(code, /^\s{8}remainder\s*=\s*\(\s*i\s*\+\s*1\s*\)\s*%\s*3/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        
        if (i + 1) % 3 == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
