---
id: 657b667a772ed53e82962c81
title: Korak 54
challengeType: 20
dashedName: step-54
---

# --description--

Još postoji nešto što možete da uradite da poboljšate čitljivost svog koda.

Izmeni `if` ili ga učini da radi kada `n` je ispod ili jednako nuli i dodaj deklaraciju za `return` za zaustavljanje izvršavanja funkcije.

# --hints--

Trebalo bi da promenite postojeći `if` ili proverite da li je `n` ispod ili jednako `0`.

```js
({ test: () =>
  {
    const hanoi = __helpers.python.getDef(code, "move");
    const {function_body} = hanoi;    
    assert(function_body.match(/^\s{4}if\s+n\s*<=\s*0\s*:/m));
  }
})
```

Treba da dodaš deklaraciju `return` na prvu liniju tvog dela `if`.

```js
({ test: () =>
  {
    const hanoi = __helpers.python.getDef(code, "move");
    const {function_body} = hanoi;    
    assert(function_body.match(/^(\s{4})if\s+n\s*<=\s*0\s*:\s*\1\1return/m));
  }
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 5
A = list(range(NUMBER_OF_DISKS, 0, -1))
B = []
C = []
--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:

--fcc-editable-region--        
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        target.append(source.pop())
        
        # display our progress
        print(A, B, C, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, A, B, C)

```
