---
id: 66c051d13a6a20255a963695
title: Kreiraj stilizovanu listu zadataka
challengeType: 25
dashedName: build-a-stylized-to-do-list
demoType: onClick
---

# --description--

U ovoj laboratoriji ćete vežbati različite stilove koje je moguće primeniti na linkove kada se kliknu, pređe mišem preko njih, ponovo klikne i poseti.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da imate jednu neuređenu listu sa klasom `todo-list`.
2. Unutar neuređene liste, treba da imate četiri stavke liste.
3. Unutar svake stavke liste, treba da postoji:

- Element `input` tipa `checkbox` sa jedinstvenim `id`.
- Element `label` sa atributom `for` čija vrednost odgovara `id` atributu povezanog `input` elementa.
- Neuređena lista sa klasom `sub-item`.
- Stavka liste koja sadrži element sidra (`a`). Element sidra treba da ima klasu `sub-item-link`, ispravnu vrednost atributa `href` i atribut `target` koji otvara vezu u novoj kartici.

4. Vaši elementi za `a` ne bi trebalo da imaju nikakve tekstualne dekoracije.
5. Treba da postavite boju teksta za neposetavene linkove na željenu boju.
6. Kada su vaši linkovi posetaveni, boja bi trebalo da promeni na drugu željenu boju.
7. Kada se vaši linkovi kliknu, boja bi trebalo da promeni na drugu željenu boju.
8. Kada se pređe mišem preko vaših linkova (hover), treba da postoji koloristična linija koja okružuje link.
9. Kada se vaši linkovi kliknu, boja bi trebalo da promeni na drugu željenu boju.

**Napomena:** Uverite se da ste povezali svoj stilski list (stylesheet) u vaš HTML i koristili vaš CSS.

# --hints--

Trebalo bi da imate neuredan spisak sa klasom `todo-list`.

```js
assert.exists(document.querySelector('ul.todo-list'));
assert.lengthOf(document.querySelectorAll('ul.todo-list'), 1);
```

Trebalo bi da ima četiri elementa liste unutar neuređene liste.

```js
const li = document.querySelectorAll('ul.todo-list > li');
assert.lengthOf(li, 4);
```

`li` unutar `ul` sa klasom `todo-list` treba da ima karakteristiku od `input` koja ima `type` od `checkbox`.

```js
const checkboxes = document.querySelectorAll(
  'ul.todo-list li input[type="checkbox"]'
);
assert.lengthOf(checkboxes, 4);
```

`li` unutar `ul` sa klasom `todo-list` treba da ima karakteristiku `label`.

```js
const labels = document.querySelectorAll(
  'ul.todo-list li label'
);
assert.lengthOf(labels, 4);
```

Svi elementi od `input` moraju imati `id`.

```js
const inputs = document.querySelectorAll('input');
assert.isAbove(inputs.length, 0);
for (let input of inputs) {
  assert.isAbove(input?.id.length, 0);
}
```

Svi elementi `label` moraju imati atribut `for`.

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.htmlFor.length, 0);
}
```

Svi elementi `label` moraju imati tekst.

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.innerText.length, 0);
}
```

Atributi za `id` na `for` za elemente od `input` i `label` moraju imati odgovarajuće vrednosti.

```js
const labels = document.querySelectorAll("label");
assert.isNotEmpty(labels);

document.querySelectorAll("input").forEach((input, index) => {
  assert.equal(input.id, labels[index].htmlFor);
});
```

Nakon elemenata za `label`, treba da postoji neuređeni spisak sa klasom `sub-item`.

```js
const labelSbItems = document.querySelectorAll('ul.todo-list li label + ul.sub-item');
assert.lengthOf(labelSbItems, 4);
```

`li` unutar `ul` sa klasom `sub-item` treba da ima anker element sa klasom `sub-item-link`.

```js
const allItems = document.querySelectorAll(
  'ul.todo-list li label + ul.sub-item li a.sub-item-link'
);
assert.lengthOf(allItems, 4);
```

Svi delovi sidra moraju imati važeći atribut `href`.

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.href.length, 0);
}
```

Sve sidrenje mora imati natpis.

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.innerText.length, 0);
}
```

Svaki element `a` mora imati atribut `target` sa vrednošću `_blank`.

```js
const anchors = document.querySelectorAll('a');
assert.isAbove(anchors.length, 0);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

Vaši elementi za `a` ne bi trebalo da imaju nikakvo tekstualno formatiranje.

```js
assert.isTrue(
  new __helpers.CSSHelp(document).getStyle('a.sub-item-link')
    ?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('a')?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('.sub-item-link')
      ?.textDecoration === 'none'
);
```

Nepromenjeni elementi treba da imaju boju teksta postavljenu na željenu boju.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:link')?.color
);
```

Elementi bi trebalo da menjaju boju kada se klikne na njih.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:hover')?.color
);
```

Elementi bi trebalo da menjaju boju kada se kliknu.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:active')?.color
);
```

Zglobovi bi trebali imati bojni linijski znak kada se preglede.

```js
const css = new __helpers.CSSHelp(document);
assert.exists(
  css.getStyle('.sub-item-link:focus')?.outline ||
    css.getStyle('a:focus')?.outline ||
    css.getStyle('a.sub-item-link:focus')?.outline ||
    css.getStyle('.sub-item-link:focus')?.border ||
    css.getStyle('a:focus')?.border ||
    css.getStyle('a.sub-item-link:focus')?.border
);
```

Linkovi treba da menjaju boju kada su posetani.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:visited')?.color
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Styled To-Do List</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Styled To-Do List</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div class="container">
      <h1>My To-Do List</h1>
      <ul class="todo-list">
        <li>
          <input type="checkbox" id="task1" />
          <label for="task1">Explore gaming keyboards</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link">Store link</a>
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task2" />
          <label for="task2">Finish the report</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >Request Access</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task3" />
          <label for="task3">View Phone's Warranty</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Receipts</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task4" />
          <label for="task4">Check Processor Specs</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Model Number</a
              >
            </li>
          </ul>
        </li>
      </ul>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: #f0f0f0;
  margin: 0;
  padding: 20px;
}

.container {
  max-width: 400px;
  margin: 0 auto;
  background-color: #fff;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
  color: #333;
}

.todo-list {
  list-style-type: none;
  padding: 0;
}

.todo-list li {
  margin-bottom: 15px;
  line-height: 1.6;
  padding: 10px;
  background-color: #e5f1e5;
  border-radius: 10px;
  position: relative;
}

a.sub-item-link {
  margin-left: 10px;
  text-decoration: none;
}

.sub-item-link:link {
  color: blue;
}

.sub-item-link:visited {
  color: purple;
}

.sub-item-link:hover {
  color: red;
}

.sub-item-link:focus {
  outline: 2px solid blue;
  outline-offset: 2px;
}

.sub-item-link:active {
  color: green;
}

.sub-item {
  list-style-type: square;
}
```
