---
id: 61537c5f81f0cf325b4a854c
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Dodajte element `header` unutar elementa `body` i dodelite mu klasu `header`.

Unutar `header`, kreirajte `h1` sa tekstom `css flexbox photo gallery`.

# --hints--

Trebalo bi da imate element `header` unutar vašeg elementa `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Vaš `header` element bi trebalo da ima `class` sa `header` as vrednost.

```js
assert.isTrue(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Vaš `header` element trebao bi imati `h1` element unutar sebe.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

Vaš element `h1` bi trebalo da sadrži tekst `css flexbox photo gallery` unutar sebe.

```js
assert.equal(document?.querySelector('.header')?.querySelector('h1')?.innerText, 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
