---
id: 61537d86bdc3dd343688fceb
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Zatim, svakom `img` dodelite atribut `src` u skladu sa njegovim redosledom u dokumentu. Prvi `img` bi trebalo da ima `src` od `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`. Ostatak treba da bude isti, osim što zamenite `1` brojem na kojem se `img` nalazi u dokumentu.

# --hints--

Svi dev od vaših `img` elemenata bi trebalo da imaju atribut `src`.

```js
const images = [...document.querySelectorAll('img')];
assert.isTrue(images.every(image => image.getAttribute('src')));
```

Vaš prvi `img` element treba da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` set as the `src` atributnu vrednost.

```js
assert.equal(document.querySelectorAll('img')?.[0]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Vaš drugi `img` element treba da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` set as vrednost atributa `src`.

```js
assert.equal(document.querySelectorAll('img')?.[1]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

Vaš treći `img` element bi trebalo da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` set as vrednost atributa `src`.

```js
assert.equal(document.querySelectorAll('img')?.[2]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

Vaš četvrti `img` element bi trebalo da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` set as vrednost atributa `src`.

```js
assert.equal(document.querySelectorAll('img')?.[3]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

Vaš peti `img` element treba da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` set as vrednost atributa `src`.

```js
assert.equal(document.querySelectorAll('img')?.[4]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

Vaš šesti `img` element bi trebalo da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` set as vrednost atributa `src`.

```js
assert.equal(document.querySelectorAll('img')?.[5]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

Vaš sedmi `img` element bi trebalo da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` set as the `src` atributnu vrednost.

```js
assert.equal(document.querySelectorAll('img')?.[6]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

Vaš osmi `img` element bi trebalo da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` set as vrednost atributa `src`.

```js
assert.equal(document.querySelectorAll('img')?.[7]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

Deveti Vaš `img` element bi trebalo da ima `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` set as vrednost atributa `src`.

```js
assert.equal(document.querySelectorAll('img')?.[8]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
--fcc-editable-region--
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
