---
id: 6153893900438b4643590367
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Poravnajte vaÅĄ tekst `.header` po sredini. Napravite tekst velikim slovima koristeÄi svojstvo `text-transform` sa vrednoÅĄÄu `uppercase`.

Dajte mu padding od `32px` sa svih strana. Postavite pozadinu na `#0a0a23` i tekst na `#fff` kao vrednosti boja.

Dodajte `border-bottom` sa vrednoÅĄÄu `4px solid #fdb347`.

# --hints--

Trebalo bi da imate selektor `.header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

VaÅĄ selektor ``.header`` bi trebalo da ima svojstvo ``text-align`` `set` na ``center`` `as` vrednost.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

VaÅĄ `.header` selektor bi trebalo da ima `text-transform` svojstvo set do `uppercase` as vrednost.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

VaÅĄ `.header` selektor bi trebalo da ima svojstvo `padding` set za dodelu vrednosti do `32px` as.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

VaÅĄ `.header` selektor bi trebalo da ima `background-color` svojstvo set do `#0a0a23` as vrednost.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

VaÅĄ `.header` selektor bi trebalo da ima `color` svojstvo set za `#fff` as vrednost.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

VaÅĄ `.header` selektor bi trebalo da ima `border-bottom` svojstvo set na `4px solid #fdb347` as vrednost.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
