---
id: 6153908a366afb4d57185c8d
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Flexbox ima glavnu i poprečnu osu. Glavna osa je definisana svojstvom `flex-direction`, koje ima četiri moguće vrednosti:

- `row` (podrazumevano): horizontalna osa sa fleks elementima od leve ka desnoj strani
- `row-reverse`: horizontalna osa sa fleks elementima od desne ka leve strani
- `column`: vertikalna osa sa fleks elementima o tvore od vrha do dna
- `column-reverse`: vertikalna osa sa fleks elementima od dna ka vrhu

**Napomena**: Ose i smerovi će se razlikovati u zavisnosti od pravca teksta. Prikazane vrednosti su za levu ka desnoj (L-R) pravac teksta.

Isprobajte različite vrednosti da biste videli kako utiču na raspored.

Kada završite, postavite eksplicitno `flex-direction` od `row` na element `.gallery`.

# --hints--

Vaš `.gallery` selektor treba da ima `flex-direction` atribut set na `row` as vrednost.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection, 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}


.gallery {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
