---
id: 6494da0daf5df5197963671d
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Primećujete kako plavi okvir slike prelazi preko crvenog okvira galerije. To je zbog načina na koji pretraživači izračunavaju veličinu kontejnerskih elemenata.

Svojstvo `box-sizing` se koristi za postavljanje ovog ponašanja. Podrazumevano, koristi se model `content-box`. Sa ovim modelom, kada element ima određenu širinu, ta širina se izračunava samo na osnovu sadržaja elementa. Vrednosti padding-a i ivice dodaju se ukupnoj širini, tako da element raste kako bi smestio ove vrednosti.

Pokušajte eksplicitno postaviti `box-sizing` na `content-box`, koristeći globalni selektor `*`. U ovom trenutku, nećete videti nikakve promene, jer koristite podrazumevanu vrednost.

# --hints--

Trebalo bi da imate selektor `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Vaš `*` selektor treba da ima `box-sizing` svojstvo.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

Vaš `*` selektor bi trebalo da ima svoj `box-sizing` svojstvo set na `content-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
