---
id: 64dd076a1b33c86d84b4232f
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Da biste bolje vizualizovali kako su vaÅĄe elemente dimenzionirani, dodavanje ivice moÅūe biti korisno.

Dajte vaÅĄem elementu `.gallery` `width` od `50%` i `border` podeÅĄen na `5px solid red`.

Zatim dajte vaÅĄim elementima `img` `width` od `100%`, `padding` podeÅĄen na `5px`, i `border` podeÅĄen na `5px solid blue`.

# --hints--

Trebalo bi da imate selektor `.gallery`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

VaÅĄ `.gallery` selektor bi trebalo da ima svojstvo `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

VaÅĄ `.gallery` selektor bi trebalo da ima `width` svojstvo set do `50%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

Tvoj `.gallery` izabirnik bi trebalo da ima `border` svojstvo.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

Your `.gallery` selector should have a `border` property set to `5px solid red`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

Trebalo bi da imate selektor `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

VaÅĄ `img` selektor treba da ima `width` svojstvo.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

VaÅĄ `img` selektor bi trebalo da ima `width` svojstvo set do `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

VaÅĄ `img` selektor bi trebalo da ima `padding` svojstvo.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

VaÅĄ `img` selektor bi trebalo da ima `padding` svojstvo set do `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

VaÅĄ `img` selektor bi trebalo da ima `border` svojstvo.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

Tvoj `img` selektor bi trebalo da ima `border` svojstvo set do `5px solid blue`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
