---
id: 66c645b345a39bf04864dc50
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Da biste videli ažurirani spisak za kupovinu, moraćete da unesete trenutni spisak za kupovinu zajedno sa kratkom porukom u konzolu.

Pošto će ova poruka biti korišćena često na radionicama, bolje je kreirati funkciju koja se može ponovo koristiti.

Kreiraj funkciju nazvanu `getShoppingListMsg` koja prima niz podataka kao parametar i vraća sekvencu znakova `"Current Shopping List: "` praćenu sadržajem u dato nizu podataka.

Možete koristiti doslovne karaktere ili spojiti niz znakova koristeći parametar `+` da spojite niz znakova i niz podataka prosleđen funkciji.

# --hints--

Trebalo bi da imaš funkciju nazvanu `getShoppingListMsg` u svom kodu.

```js
assert.isFunction(getShoppingListMsg);
```

Tvoja funkcija `getShoppingListMsg` treba da vrati niz karaktera.

```js
assert.isString(getShoppingListMsg(shoppingList));
```

Tvoja funkcija ``getShoppingListMsg`` bi trebalo da primi red podataka i vrati poruku ``"Current Shopping List: "`` praćenu sadržajem tog reda. Proverite da nema grešaka u kucanju/pravopisnih grešaka.

```js
assert.oneOf(getShoppingListMsg(["Anchovies", "Olives"]), ["Current Shopping List: Anchovies,Olives", "Current Shopping List: Anchovies, Olives"]);
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
