---
id: 66c6491fe4c8e0f16845425f
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Sada je vreme da dodamo još jedno voće na listu.

Ako se oslanjate na isti način za red podataka kao i ranije, dodajte niz karaktera `"Grapes"` na kraj reda podataka `shoppingList`.

Zatim, dodajte `console.log` i pozovite funkciju `getShoppingListMsg` sa redom podataka `shoppingList` kao argument unutar `console.log` kako biste videli ažurirani spisak ispisani na konzolu.

# --hints--

Vaš red podataka `shoppingList` sada bi trebalo da ima dve stavke.

```js
assert.lengthOf(shoppingList, 2);
```

Vaš red podataka `shoppingList` mora imati sekvencu karaktera `"Grapes"`.

```js
assert.equal(shoppingList[1], "Grapes");
```

Treba da pozoveš funkciju `getShoppingListMsg(shoppingList)` unutar `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
