---
id: 66c64210413532ee9d3bd342
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Za ovu sledeću sekciju radionice, ćete vežbati dodavanje namirnica na listu za kupovinu.

Počnite koristeći `console.log` da napišete poruku `"It will be nice to have some fruit to eat."` u konzoli.

# --hints--

Trebalo bi da napišete poruku `"It will be nice to have some fruit to eat."` u konzolu.

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
