---
id: 66c64095f61166ede6e39a84
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Za ovu listu za kupovinu, koristićeš red podataka da predstavlja stvari koje želiš da kupiš.

U prethodnim lekcijama, naučili ste kako da kreirate redove podataka kao ovakve:

```js
const fruits = ['apple', 'banana', 'orange'];
```

U ovom koraku, kreiraj parametar nazvan `shoppingList` i postavi mu vrednost praznog niza podataka.

# --hints--

Trebalo bi da imaš parametar nazvan `shoppingList`.

```js
assert.isNotNull(shoppingList);
```

Trebalo bi da postavite vrednost kolone podataka na parametar `shoppingList`.

```js
assert.isArray(shoppingList);
```

Podatni red mora biti prazan.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
