---
id: 66c63bfa4685e4e3843fa0dc
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, nastavićete da učite o skupovima podataka kreiranjem liste za kupovinu namirnica.

Počnite dodavanjem `console.log` koja štampa sekvencu karaktera `"Grocery shopping list"` na konzoli.

# --hints--

Trebalo bi da imaš `console.log` u svom kodu.

```js
assert.match(code, /console\.log(.*)/);
```

`console.log` treba da izbaci poruku `"Grocery shopping list"`. Proverite pravopisne greške i velika slova.

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
