---
id: 66c644c08ebcd2ef90c750bd
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

U prethodnim lekcijama, naučili ste kako da dodate elemente na kraj reda podataka koristeći metodu `push`, kao što sledi:

```js
fruits.push('pear');
```

Koristeći metodu `push`, dodaj niz znakova `"Apples"` u red podataka `shoppingList`.

# --hints--

Safu yako ya data ya `shoppingList` trebala bi imati po jedan element unutar sebe.

```js
assert.lengthOf(shoppingList, 1);
```

Vaš red podataka za `shoppingList` treba da sadrži niz slova `"Apples"` unutar sebe. Proverite da li ste proverili pravopisne greške i velika slova (kapitalizaciju).

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
