---
id: 66c72a55418cc9247b710827
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Sledeće, dodajte `console.log` i pozovite funkciju `getShoppingListMsg` koristeći red podataka `shoppingList` kao argument unutar `console.log` kako biste videli ažurirani spisak snimljen na konzolu.

# --hints--

Trebalo bi da pozoveš `getShoppingListMsg(shoppingList)` unutar `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
