---
id: 66c730ee6ae076281721d0b9
title: Korak 14
challengeType: 1
dashedName: step-14
---

# --description--

U prethodnoj lekciji, naučili ste kako da uklonite stavke sa kraja niza podataka koristeći metodu `pop`.

Ovo je podsetnik kako da se koristi `pop`:

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

Koristi putanju `pop` da ukloni poslednji element iz niza podataka `shoppingList`.

# --hints--

Vaš red podataka za `shoppingList` sada bi trebalo da sadrži stavke `5` unutar sebe.

```js
assert.lengthOf(shoppingList, 5);
```

Vaš red podataka za `shoppingList` ne bi trebalo više da sadrži sekvencu karaktera `"Potato Chips"`.

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
