---
id: 66c72b0ffbc5522525768558
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

U prethodnim koracima, pregledali ste kako dodati element na kraju redova podataka koristeći metodu `push`.

Međutim, metoda `push` prihvata mnogo argumenata, pa možete dodati više elemenata na kraju reda podataka kao što sledi:

```js
array.push(item1, item2, item3);
```

U ovom koraku, koristite metodu `push` da dodate sekvencu karaktera `"Popcorn", "Beef Jerky", "Potato Chips"` do `shoppingList` u red podataka.

Redosled je važan, stoga se pobrinite da dodajete elemente u navedenom redosledu.

# --hints--

Vaš red podataka za `shoppingList` treba da ima ukupan broj znakova `6`.

```js
assert.lengthOf(shoppingList, 6);
```

`shoppingList` treba imati sekvencu slova `"Popcorn"`.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

`shoppingList` mora imati sekvencu slova `"Beef Jerky"`.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

`shoppingList` mora imati sekvencu slova `"Potato Chips"`.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
