---
id: 66c73a7798f6f62b2ae58f22
title: Korak 16
challengeType: 1
dashedName: step-16
---

# --description--

Sada je vreme da dodate još stavki na vrh liste za kupovinu.

Počnite dodavanjem deklaracije za `console.log` koja piše poruku `"It might be nice to get a dessert."`

Ispod te deklaracije od `console`, koristite odgovarajuću metodu niza podataka za dodavanje sekvence karaktera `"Chocolate Cake"` na početku niza podataka `shoppingList`.

Konačno, dodajte `console.log` i pozovite funkciju `getShoppingListMsg` sa redom podataka `shoppingList` kao argument unutar `console.log` da biste videli ažurirani spisak prikazan u konzoli.

# --hints--

Trebalo bi da imate `console.log` koji piše poruku `"It might be nice to get a dessert."`.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

Red vaših podataka za `shoppingList` bi trebalo da ima ukupan broj stavki `6`.

```js
assert.lengthOf(shoppingList, 6);
```

Trebalo bi da koristiš način `unshift` za dodavanje niza karaktera `"Chocolate Cake"` na početak kolone podataka `shoppingList`.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

Treba ti da pozoveš `getShoppingListMsg(shoppingList)` unutar `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
