---
id: 66c726c34ecf1e238aa9d7d5
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

U prethodnim lekcijama, naučili ste kako da dodate elemente na početak niza podataka koristeći metodu `unshift()`.

Ovo je podsetnik kako da se koristi metoda `unshift()`:

```js
array.unshift(item1, item2, ..., itemX);
```

Koristi putanju `unshift()` da dodaš niz karaktera `"Vegetable Oil"` na početak kolone podataka `shoppingList`.

# --hints--

Trebalo bi da imaš tri stavke u vašem nizu podataka `shoppingList`.

```js
assert.lengthOf(shoppingList, 3);
```

Trebalo bi da koristiš putanju `unshift()` za dodavanje `"Vegetable Oil"` na početak kolone podataka `shoppingList`.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
