---
id: 66c742d045c9fc2e09fa64b1
title: Korak 19
challengeType: 1
dashedName: step-19
---

# --description--

Poslednja promena koju ste izvršili na listi artikala iz prodavnice je ispravljanje prvog elementa na listi.

U prethodnoj lekciji, naučili ste kako da formatirate element koristeći zagradni zapis i specifikator elementa koji želite da formatirate.

Ovo je podsetnik kako da se ispravi element u redu podataka:

```js
const array = [1, 2, 3, 4, 5];

array[0] = 10;
// Result: [10, 2, 3, 4, 5]
console.log(array); 
```

Promenite prvi element u redu podataka `shoppingList` na `"Canola Oil"`.

# --hints--

Trebalo bi da ažuriraš prvi element u nizu podataka `shoppingList` u `"Canola Oil"`.

```js
assert.equal(shoppingList[0], "Canola Oil");
assert.notEqual(shoppingList[1], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

shoppingList.shift();

--fcc-editable-region--

--fcc-editable-region--
```
