---
id: 66cbe2319d3845545a293a0b
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Sada je vreme da se vidi poruka napisana na konzoli.

Dodajte `console.log` i pozovite funkciju `getShoppingListMsg` sa redom podataka `shoppingList` kao argument unutar `console.log` da biste videli poruku ispisanu na konzolu.

# --hints--

Treba da pozoveš funkciju `getShoppingListMsg(shoppingList)` unutar `console.log`.

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
