---
id: 68599af28ead38c22590a89d
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Potrebno je da postoji način povezivanja dugmadi sa sadržajem koji predstavljaju. To se postiže korišćenjem atributa `aria-controls` na svakom dugmetu. Kasnije ćete dodati sadržaj i odgovarajuće atribute `id` elementima sa sadržajem.

Počnite dodavanjem atributa `aria-controls` svakom od elemenata `button`. Vrednost za atribut `aria-controls` prvog dugmeta treba da bude postavljena na `"panel-earth"`. Atribut `aria-controls` drugog dugmeta treba da bude postavljen na `"panel-saturn"`. Atribut `aria-controls` trećeg dugmeta treba da bude postavljen na `"panel-mars"`.

# --hints--

Each of your `button` elements should have an `aria-controls` attribute.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[aria-controls]");
assert.lengthOf(btns, 3);
```

The first `button` element should have an `aria-controls` attribute set to `"panel-earth"`.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-controls"), "panel-earth");
``` 

The second `button` element should have an `aria-controls` attribute set to `"panel-saturn"`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-controls"), "panel-saturn");
```

The third `button` element should have an `aria-controls` attribute set to `"panel-mars"`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-controls"), "panel-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab">🌍 Earth</button>
        <button role="tab">🪐 Saturn</button>
        <button role="tab">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
