---
id: 685976a92ffab8b2dc744b72
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Unutar vašeg `.tabs` elementa, dodajte `div` element sa atributom `role` postavljenim na `"tablist"` i atributom `aria-labelledby` postavljenim na `"tabs-title"`.  

Atribut `aria-labelledby` je potreban ovde da poveže listu kartica sa vidljivom oznakom, pomažući asistivnim tehnologijama kao što su čitači ekrana da objave svrhu liste kartica korisnicima.

# --hints--

You should have a `div` element inside of your `.tabs` element.

```js
assert.exists(document.querySelector(".tabs div"));
```

Your `div` element should have a `role` attribute set to `"tablist"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

Your `div` element should have an `aria-labelledby` attribute set to `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
