---
id: 587d7b88367417b2b2512b44
title: Napiši funkciju strele sa parametrima
challengeType: 1
forumTopicId: 301223
dashedName: write-arrow-functions-with-parameters
---

# --description--

Kao običnu funkciju, možete proslediti argument unutar strelične funkcije.

```js
const doubler = (item) => item * 2;
doubler(4);
```

`doubler(4)` vraća vrednost od `8`.

Ako strelovna funkcija ima samo jedan parametar, standardne zagrade koje okružuju parametar mogu se izostaviti.

```js
const doubler = item => item * 2;
```

Moguće je da se proslede više od jedne tvrdnje unutar streličnog funkcionera.

```js
const multiplier = (item, multi) => item * multi;
multiplier(4, 2);
```

`multiplier(4, 2)` vraća vrednost od `8`.

# --instructions--

Ponovo napiši funkciju `myConcat` koja dodaje sadržaj `arr2` u `arr1`, tako da funkcija koristi sinaksu strelične funkcije.

# --hints--

Trebalo bi da promeniš reč `var`.

```js
assert.notMatch(code, /var/g);
```

`myConcat` treba da bude fiksni parametar (koristeći `const`).

```js
assert.match(code, /const\s+myConcat/g);
```

`myConcat` trebalo bi da bude strelična funkcija sa dva parametra

```js
assert(
  /myConcat=\(\w+,\w+\)=>/.test(__helpers.removeJSComments(code).replace(/\s/g, '')) &&
    typeof myConcat === 'function'
);
```

`myConcat()` treba da vrati `[1, 2, 3, 4, 5]`.

```js
assert.deepEqual(myConcat([1, 2], [3, 4, 5]), [1, 2, 3, 4, 5]);
```

Ova fraza `function` ne bi trebalo koristiti.

```js
assert.notMatch(code, /function/g);
```

# --seed--

## --seed-contents--

```js
var myConcat = function(arr1, arr2) {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```

# --solutions--

```js
const myConcat = (arr1, arr2) =>  {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```
