---
id: 587d7b8c367417b2b2512b58
title: Kreiraj alternativni način za eksportovanje korišćenjem export default
challengeType: 1
forumTopicId: 301199
dashedName: create-an-export-fallback-with-export-default
---

# --description--

U kursu za `export`, naučili ste se o sintaksi poznatoj kao <dfn>export sa imenom </dfn>. Ovo vam je omogućilo da funkcije i mnogi parametri možete koristiti u drugim fajlovima.

Postoji druga sintaksa za `export` koju treba da saznate, poznata kao <dfn>export default</dfn>. Obično koristite ovu sintaksu ako se iz fajla vraća samo jedna vrednost. Takođe se koristi za kreiranje alternativne vrednosti za fajl ili modul.

Ispod su primeri koji koriste `export default`:

```js
export default function add(x, y) {
  return x + y;
}

export default function(x, y) {
  return x + y;
}
```

Prvi je glagol sa imenom, a drugi je glagol bez imena.

Pošto `export default` se koristi za deklarisanje alternativne vrednosti za modul ili fajl, može imati samo jednu vrednost kao export default u svakom modulu ili fajlu. Pored toga, ne možete koristiti `export default` uz `var`, `let`, ili `const`

# --instructions--

Funkcija koja je navedena ispod treba da bude podrazumevana vrednost za modul. Molimo vas da dodate neophodni kod radi toga.

# --hints--

Vaš kod bi trebalo da koristi alternativnu metodu za `export`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /export\s+default\s+function(\s+subtract\s*|\s*)\(\s*x,\s*y\s*\)\s*{/g
  )
);
```

# --seed--

## --seed-contents--

```js
function subtract(x, y) {
  return x - y;
}
```

# --solutions--

```js
export default function subtract(x, y) {
  return x - y;
}
```
