---
id: 587d7b8c367417b2b2512b54
title: Koristi getter-e i setter-e za kontrolu pristupa objekta
challengeType: 1
forumTopicId: 301220
dashedName: use-getters-and-setters-to-control-access-to-an-object
---

# --description--

Možete dobiti vrednost iz objekta i postaviti vrednost svojstva unutar tog objekta.

Ove se obično zovu <dfn>getters</dfn> i <dfn>setters</dfn>.

Getter metoda je namenjena da vrati (dobije) vrednost privatnog atributa objekta korisniku, bez toga što korisnik direktno pristupa tom privatnom atributu.

Setter funkcija je namenjena za menjanje (postavljanje) vrednosti privatnog atributa objekta na osnovu vrednosti koja se prosleđuje setter funkciji. Ove promene mogu uključivati proračune, ili čak potpuno prepisanje originalne vrednosti.

```js
class Book {
  constructor(author) {
    this._author = author;
  }
  // getter
  get writer() {
    return this._author;
  }
  // setter
  set writer(updatedAuthor) {
    this._author = updatedAuthor;
  }
}
const novel = new Book('anonymous');
console.log(novel.writer);
novel.writer = 'newAuthor';
console.log(novel.writer);
```

Konsola će prikazati niz znakova `anonymous` i `newAuthor`.

Posmatrajte sintaksu koja se koristi za pozivanje getter-a i setter-a. Čak i ne izgledaju kao metode. Getter-i i setter-i su važni jer skrivaju detalje unutrašnje implementacije.

**Napomena:** To je mehanizam za definisanje imena privatnog parametra na donjoj liniji (`_`). Međutim, sama ta karakteristika ne čini parametar privatnim.

# --instructions--

Koristi glavnu reč `class` da kreira klasu `Thermostat`. `constructor` prihvata temperaturu u Farenhajtu.

U učionici, i `getter` dobijanje temperature u Celzijusima i `setter` koja prihvata temperaturu u Celzijusima.

Zapamtite da `C = 5/9 * (F - 32)` i `F = C * 9.0 / 5 + 32`, gde `F` je vrednost temperature u Fahrenhajtu, i `C` je ekvivalentna vrednost temperature u Celzijusu.

**Napomena:** Kada ovo implementirate, temperatura će biti sačuvana u prostoru koristeći jednu jedinicu mere, bilo da je to Fahrenheit ili Celsius.

Ovo je funkcija getera i seta. Napravljaš API za drugu aplikaciju, koja može dobiti tačne rezultate bez obzira na šta koristiš za skladištenje.

Drugim rečima, sakrivaš detalje implementacije od korisnika.

# --hints--

`Thermostat` treba da bude klasa od `class` sa metodom `constructor` definisanog.

```js
assert.isFunction(Thermostat);
assert.isFunction(Thermostat?.constructor);
```

Glavna reč za `class` mora se koristiti.

```js
assert.match(code, /class/);
```

`Thermostat` Trebalo bi da se inicijalizuje.

```js
const _t = new Thermostat(122);
assert.isObject(_t);
```

Kada se inicijalizuje sa vrednošću Farenhajta, `Thermostat` mora da postavi ispravan `temperature`.

```js
const _t = new Thermostat(122);
assert.strictEqual(_t?.temperature, 50);
```

`getter` mora biti specificirano.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.get);
```

`setter` Trebalo bi da bude definisano.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.set);
```

Postavljanje `setter` sa vrednošću Celzija mora da postavi `temperature`.

```js
const _t = new Thermostat(32);
_t.temperature = 26;
const _u = new Thermostat(32);
_u.temperature = 50;
assert.approximately(_t.temperature, 26, 0.1);
assert.approximately(_u.temperature, 50, 0.1);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```

# --solutions--

```js
class Thermostat {
  constructor(fahrenheit) {
    this._tempInCelsius = 5/9 * (fahrenheit - 32);
  }
  get temperature(){
    return this._tempInCelsius;
  }
  set temperature(newTemp){
    this._tempInCelsius = newTemp;
  }
}

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```
