---
id: 587d7b8a367417b2b2512b4d
title: Koristi model rešenja strukture za prenos stvari kao parametara funkcije
challengeType: 1
forumTopicId: 301217
dashedName: use-destructuring-assignment-to-pass-an-object-as-a-functions-parameters
---

# --description--

U nekim slučajevima, možete dekonstruisati strukturu objekta unutar samog argumenta funkcije.

Pogledajte kod dole:

```js
const profileUpdate = (profileData) => {
  const { name, age, nationality, location } = profileData;

}
```

Ovo otkriva strukturu objekta koji se efikasno šalje unutar uređaja. Ovo se takođe može uraditi na mestu:

```js
const profileUpdate = ({ name, age, nationality, location }) => {

}
```

Kada `profileData` prođe kroz funkciju iznad, vrednosti se uzimaju iz parametra funkcije za upotrebu unutar te funkcije.

# --instructions--

Koristite strukturalni razmak rešenja unutar funkcije `half` samo da prenesete `max` i `min` unutar funkcije.

# --hints--

`stats` treba biti `object`.

```js
assert(typeof stats === 'object');
```

`half(stats)` trebalo bi da bude `28.015`

```js
assert(half(stats) === 28.015);
```

Koristi rešenje dizajna.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/half=\({\w+,\w+}\)/));
```

Koristite dekomprimovani parametar.

```js
assert(!__helpers.removeJSComments(code).match(/stats\.max|stats\.min/));
```

# --seed--

## --seed-contents--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

// Only change code below this line
const half = (stats) => (stats.max + stats.min) / 2.0; 
// Only change code above this line
```

# --solutions--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

const half = ( {max, min} ) => (max + min) / 2.0;
```
