---
id: 587d7b88367417b2b2512b47
title: Koristi parametar `rest` zajedno sa parametrima funkcije
challengeType: 1
forumTopicId: 301221
dashedName: use-the-rest-parameter-with-function-parameters
---

# --description--

Da bi nam pomoglo da kreiramo fleksibilnije funkcije, ES6 uvodi `<dfn>rest` parametar `</dfn>` za parametre funkcija. Korišćenjem rest parametra, možete kreirati funkcije koje primaju promenljiv broj argumenata. Ovi argumenti se skladište u nizu koji se može pristupiti kasnije iz unutar funkcije.

Pogledajte ovaj kod:

```js
function howMany(...args) {
  return "You have passed " + args.length + " arguments.";
}
console.log(howMany(0, 1, 2));
console.log(howMany("string", null, [1, 2, 3], { }));
```

Konsole će prikazati sekvencu slova `You have passed 3 arguments.` i `You have passed 4 arguments.`.

Rest parametar eliminiše potrebu za korišćenjem `arguments` i omogućava korišćenje metoda nizova podataka na nizu parametara prosleđenih funkciji `howMany`.

# --instructions--

Izmenite funkciju `sum` koristeći REST parametar na način da funkcija `sum` može primiti bilo koji broj argumenata i vratiti njihov zbir.

# --hints--

Rezultat `sum(0,1,2)` treba da bude 3

```js
assert(sum(0, 1, 2) === 3);
```

Rezultat od `sum(1,2,3,4)` treba biti 10

```js
assert(sum(1, 2, 3, 4) === 10);
```

Rezultat za `sum(5)` treba da bude 5

```js
assert(sum(5) === 5);
```

Rezultat od `sum()` bi trebalo da bude 0

```js
assert(sum() === 0);
```

`sum` treba da bude strelna funkcija koja koristi sintaksu rest parametara (`...`) na parametru `args`.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/sum=\(\.\.\.args\)=>/));
```

# --seed--

## --seed-contents--

```js
const sum = (x, y, z) => {
  const args = [x, y, z];
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```

# --solutions--

```js
const sum = (...args) => {
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```
