---
id: 587d7b87367417b2b2512b42
title: Promeni red podataka deklarisan sa `const`
challengeType: 1
forumTopicId: 301206
dashedName: mutate-an-array-declared-with-const
---

# --description--

Ako ne znate o `const`, molimo pogledajte <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/declare-a-read-only-variable-with-the-const-keyword" target="_blank" rel="noopener noreferrer nofollow">ovo izazovno pitanje vezano za ključnu reč <code>const</code></a>.

Izjava o `const` ima široku primenu u JavaScript modernih.

Drugi programeri preferiraju da postave vrednosti za sve parametre korišćenjem `const` podrazumevano, osim ako ne znaju da će morati ponovo da postave vrednosti. Samo u tom slučaju koriste `let`.

Međutim, važno je razumeti da stavke (uključujući nizove podataka i funkcije) koje su dodeljene parametru koristeći `const` još uvek mogu biti izmenjene. Upotreba deklaracije `const` samo sprečava postavljanje nove vrednosti za identifikator parametra.

```js
const s = [5, 6, 7];
s = [1, 2, 3];
s[2] = 45;
console.log(s);
```

`s = [1, 2, 3]` to će uzrokovati grešku. Nakon unosa komentara na tu liniju, `console.log` će prikazati vrednost `[5, 6, 45]`.

Kao što vidite, možete promeniti nešto `[5, 6, 7]` same i parametar `s` će i dalje ukazivati na izmenjeni red podataka `[5, 6, 45]`. Ako sve redove podataka, elemente reda u `s` mogu se menjati, ali zato što je korišćeno `const`, ne možete koristiti identifikator parametra `s` da ukazujete na drugi red koristeći operator dodeljivanja.

# --instructions--

Redos podataka je deklarisan kao `const s = [5, 7, 2]`. Promeni taj redos u `[2, 5, 7]` koristeći dodelu vrednosti različitih elemenata.

# --hints--

Ne možete promeniti glavnu reč `const`.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`s` treba da bude trajni parametar (koristeći `const`).

```js
assert(__helpers.removeJSComments(code).match(/const\s+s/g));
```

Ne smete menjati početnu deklaraciju reda podataka.

```js
assert(__helpers.removeJSComments(code).match(
/const\s+s\s*=\s*\[\s*5\s*,\s*7\s*,\s*2\s*\]\s*;?/g
));
```

`s` treba biti jednako `[2, 5, 7]`.

```js
assert.deepEqual(s, [2, 5, 7]);
```

# --seed--

## --seed-contents--

```js
const s = [5, 7, 2];
function editInPlace() {
  // Only change code below this line

  // Using s = [2, 5, 7] would be invalid

  // Only change code above this line
}
editInPlace();
```

# --solutions--

```js
const s = [5, 7, 2];
function editInPlace() {
  s[0] = 2;
  s[1] = 5;
  s[2] = 7;
}
editInPlace();
```
