---
id: 587d7b8a367417b2b2512b4e
title: Napravi niz slova koristeći slova iz ćibice
challengeType: 1
forumTopicId: 301200
dashedName: create-strings-using-template-literals
---

# --description--

Novi element ES6 je <dfn>šablonski literal</dfn>. Ovo je poseban tip sekvence znakova koji pojednostavljuje kreiranje složenih sekvenci znakova.

Karakteri omogućavaju vam da kreirate višelinijske nizove karaktera i koristite komponente za unos vrednosti kako biste sastavili niz karaktera.

Razmotri kod koji je dole:

```js
const person = {
  name: "Zodiac Hasbro",
  age: 56
};

const greeting = `Hello, my name is ${person.name}!
I am ${person.age} years old.`;

console.log(greeting);
```

Konsole prikazuju sekvencu karaktera `Hello, my name is Zodiac Hasbro!` i `I am 56 years old.`.

Desilo se mnogo stvari ovde. Prvo, ovaj primer koristi oznake za šablonske stringove (`` ` ``), a ne navodnike (`'` ili `"`), za obuhvatanje stringa. Drugo, primetite da je string višelinijski, i u kodu i u izlazu. Ovo štedi potrebu za ubacivanjem `\n` unutar stringova. Sintaksa `${variable}` koja je korišćena iznad predstavlja mesto za ubacivanje vrednosti. Više nećete morati da koristite konkatenaciju pomoću operatora `+`. Da biste dodali promenljive u stringove, samo ubacite promenljivu u template string i obuhvatite je sa `${` i `}`. Na sličan način možete uključiti i druge izraze u string literal, na primer `${a + b}`. Ovaj novi način pravljenja stringova daje vam veću fleksibilnost za kreiranje složenijih stringova.

# --instructions--

Koristite sintaksu literala sa šablonima i markere za interpolaciju kako biste kreirali niz podataka u obliku inline koda/stringa (`li`). Svaki stav u listi mora imati jedan od elemenata niza podataka iz svojstva `failure` na objektu `result` i mora posedovati svojstvo elementa `class` sa vrednošću `text-warning`. Funkcija `makeList` treba da vrati niz podataka u obliku inline koda/stringa.

Koristite metode ponavljanja (bilo koji tip kruga/ciklus) da biste dobili očekivane rezultate (prikazani ispod).

```js
[
  '<li class="text-warning">no-var</li>',
  '<li class="text-warning">var-on-top</li>',
  '<li class="text-warning">linebreak</li>'
]
```

# --hints--

`failuresList` treba da bude podatni red koji sadrži poruku `result failure`.

```js
assert(
  typeof makeList(result.failure) === 'object' && failuresList.length === 3
);
```

`failuresList` Trebalo bi da bude jednako navedenim rezultatima.

```js
assert(
  makeList(result.failure).every(
    (v, i) =>
      v === `<li class="text-warning">${result.failure[i]}</li>` ||
      v === `<li class='text-warning'>${result.failure[i]}</li>`
  )
);
```

Koristite niz karakterima i unos vrednosti.

```js
assert.match(code, /(`.*\${.*}.*`)/);
```

Koristite ponovnu rutu.

```js
assert(__helpers.removeJSComments(code).match(/for|map|reduce|forEach|while/));
```

# --seed--

## --seed-contents--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  // Only change code below this line
  const failureItems = [];
  // Only change code above this line

  return failureItems;
}

const failuresList = makeList(result.failure);
```

# --solutions--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  return arr.map(val => `<li class="text-warning">${val}</li>`);
}

const failuresList = makeList(result.failure);
```
