---
id: 587d7b8c367417b2b2512b56
title: Koristite izvoz da podelite deo koda
challengeType: 1
forumTopicId: 301219
dashedName: use-export-to-share-a-code-block
---

# --description--

Zamislite fajl pod nazivom `math_functions.js` koji ima nekoliko funkcija vezanih za matematičke operacije. Jedna od njih je sačuvana u parametar, `add`, koji prima dva broja i vraća njihov zbir. Želite da koristite ovu funkciju u različitim fajlovima od JavaScript. Da biste podelili sa ovim drugim fajlovima, prvo vam je potrebno `export`.

```js
export const add = (x, y) => {
  return x + y;
}
```

Ovo je uobičajen način da se export jedna funkcija, ali možete postići isto na ovaj način:

```js
const add = (x, y) => {
  return x + y;
}

export { add };
```

Kada kreirate export parametar ili funkciju, možete ga uvesti u drugi fajl i koristiti bez ponovnog pisanja koda. Možete da export mnogo stvari ponavljanjem prvog primera za sve što želite da export, ili stavljanjem svega u deklaraciju export drugog primera, kao što je:

```js
export { add, subtract };
```

# --instructions--

Postoje dve funkcije koje su vezane za niz znakova u editoru. Izvezi obe koristeći svoj omiljeni metod.

# --hints--

Trebalo bi da -export `uppercaseString` prikladno.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+uppercaseString|export\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)})/g
  )
);
```

Unapaswa ku-export `lowercaseString` ipasavyo.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+lowercaseString|export\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)})/g
  )
);
```

# --seed--

## --seed-contents--

```js
const uppercaseString = (string) => {
  return string.toUpperCase();
}

const lowercaseString = (string) => {
  return string.toLowerCase()
}
```

# --solutions--

```js
export const uppercaseString = (string) => {
  return string.toUpperCase();
}

export const lowercaseString = (string) => {
  return string.toLowerCase()
}
```
