---
id: 587d7b89367417b2b2512b4a
title: Koristite model distribucije rešenja za postavljanje parametara vrednosti iz organizovanih stvari unutar njega
challengeType: 1
forumTopicId: 301214
dashedName: use-destructuring-assignment-to-assign-variables-from-nested-objects
---

# --description--

Možete koristiti iste principe iz dve prethodne lekcije da otkrijete vrednost u interno organizovanih elemenata.

Ako koristite nešto slično prethodnim primerima:

```js
const user = {
  johnDoe: { 
    age: 34,
    email: 'johnDoe@freeCodeCamp.com'
  }
};
```

Ovo je kako možete izvući vrednosti atributa objekta i postaviti ih u parametre sa istim imenima:

```js
const { johnDoe: { age, email }} = user;
```

Ovo je način na koji možete postaviti vrednosti svojstava objekta u svojstva sa različitim imenima:

```js
const { johnDoe: { age: userAge, email: userEmail }} = user;
```

# --instructions--

Promijenite postojeće mapiranje vrijednosti ES5 i odgovarajuće mapiranje rezolucije strukture. Još morate postaviti parametre `lowToday` i `highToday` vrijednosti za `today.low` na `today.high` iz objekta `LOCAL_FORECAST`.

# --hints--

Trebalo bi da uklonite sintaksu dodeljivanja vrednosti ES5.

```js
assert(
  !__helpers.removeJSComments(code).match(/lowToday = LOCAL_FORECAST\.today\.low/g) &&
    !__helpers.removeJSComments(code).match(/highToday = LOCAL_FORECAST\.today.high/g)
);
```

Trebalo bi da koristite strukturalni rešenje za kreiranje parametra `lowToday`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(low\s*:\s*lowToday[^}]*|[^,]*,\s*low\s*:\s*lowToday\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

Trebalo bi da koristite strukturalni rešenje za kreiranje parametra `highToday`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(high\s*:\s*highToday[^}]*|[^,]*,\s*high\s*:\s*highToday,?\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

`lowToday` mora biti jednako sa `64` i `highToday` mora biti jednako sa `77`.

```js
assert(lowToday === 64 && highToday === 77);
```

# --seed--

## --seed-contents--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};

// Only change code below this line
  
const lowToday = LOCAL_FORECAST.today.low;
const highToday = LOCAL_FORECAST.today.high;

// Only change code above this line
```

# --solutions--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};
 
const { today: { low: lowToday, high: highToday }} = LOCAL_FORECAST;
```
