---
id: 587d7b89367417b2b2512b4b
title: Koristite model distribuiranog rešenja da dodelite vrednosti parametrima iz kolona podataka
challengeType: 1
forumTopicId: 301213
dashedName: use-destructuring-assignment-to-assign-variables-from-arrays
---

# --description--

ES6 čini modelovanje nizova podataka tako jednostavnim kao modelovanje objekata.

Glavna razlika između spread operacije i rešenja za strukturu podataka u redu je što spread operacija otvara sav sadržaj reda podataka u listu razdvojenu zarezima. Stoga, ne možete odabrati elemente kojima želite dodeliti vrednosti za parametre.

Rešenje strukture redova podataka omogućava nam da uradimo upravo ovo:

```js
const [a, b] = [1, 2, 3, 4, 5, 6];
console.log(a, b);
```

Konsola će prikazati vrednosti za `a` i `b` kao `1, 2`.

Parametar `a` dobio je prvu vrednost niza podataka, a `b` dobio je drugu vrednost niza podataka. Takođe možemo dobiti vrednost na bilo kom indeksu niza podataka koristeći podešavanje obrasca sa zarezima kako bismo došli do željenog indeksa:

```js
const [a, b,,, c] = [1, 2, 3, 4, 5, 6];
console.log(a, b, c);
```

Konsola će prikazati vrednosti za `a`, `b`, i na `c` kao `1, 2, 5`.

# --instructions--

Koristite model rešenja strukture za promenljivanje vrednosti `a` i `b` ili `a` tako da dobiju vrednosti sačuvane u `b`, i da `b` dobije vrednost sačuvanu u `a`.

# --hints--

Vrednost `a` bi trebalo da bude `6`, nakon promene.

```js
assert(a === 6);
```

Vrednost `b` bi trebalo da bude `8`, nakon promene.

```js
assert(b === 8);
```

Trebalo bi da koristiš strukturalno rešenje za red podataka kako bi zamenu `a` na `b`.

```js
assert(/\[\s*(\w)\s*,\s*(\w)\s*\]\s*=\s*\[\s*\2\s*,\s*\1\s*\]/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let a = 8, b = 6;
// Only change code below this line
```

# --solutions--

```js
let a = 8, b = 6;
[a, b] = [b, a];
```
