---
id: 587d7b8c367417b2b2512b57
title: Koristite * da izvučete sve iz fajla
challengeType: 1
forumTopicId: 301210
dashedName: use--to-import-everything-from-a-file
---

# --description--

Zamislite da imate fajl i želite da prenesete sav njegov sadržaj u trenutni fajl. Ovo se može uraditi korišćenjem sintakse `import * as`. Evo primera gde su sadržaj fajla pod nazivom `math_functions.js` preneseni u jedan fajl u istom direktorijumu:

```js
import * as myMathModule from "./math_functions.js";
```

Rečenica ``import`` iznad će kreirati nešto nazvano ``myMathModule``. Ovo je samo ime parametra; možete dodeliti bilo koje ime. To će imati svaki `export` iz ``math_functions.js``, tako da možete doći do pristupa funkciji kao što radite sa bilo kojim drugim svojstvom objekta. Evo kako možete koristiti funkciju ``add`` i ``subtract`` koja je dostavljena:

```js
myMathModule.add(2,3);
myMathModule.subtract(5,3);
```

# --instructions--

Kod u ovoj datoteci zahteva sadržaj fajla: `string_functions.js`, koji se nalazi u istom direktorijumu kao i trenutna datoteka. Koristite sintaksu `import * as` da prenesete sve iz fajla u objekat nazvan `stringFunctions`.

# --hints--

Vaš kod bi trebalo da koristi sintaksu `import * as` pravilno.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*\*\s*as\s+stringFunctions\s+from\s*('|")\.\/string_functions\.js\1/g
  )
);
```

# --seed--

## --seed-contents--

```js

// Only change code above this line

stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```

# --solutions--

```js
import * as stringFunctions from "./string_functions.js";

// add code above this line
stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```
