---
id: 587d7b88367417b2b2512b46
title: Postavi podrazumevana kriterijum (parametre) za svoje funkcije
challengeType: 1
forumTopicId: 301209
dashedName: set-default-parameters-for-your-functions
---

# --description--

Da bi nam pomoglo da kreiramo fleksibilnije funkcije, ES6 uvodi <dfn>podrazumevane parametre</dfn> za funkcije.

Pogledajte ovaj kod:

```js
const greeting = (name = "Anonymous") => "Hello " + name;

console.log(greeting("John"));
console.log(greeting());
```

Konsola će prikazati sekvencu znakova `Hello John` i `Hello Anonymous`.

Podrazumevani parametar počinje da radi kada nije naveden argument (to je undefined). Kao što vidite u gornjem primeru, parametar `name` će primiti svoju podrazumevanu vrednost `Anonymous` kada ne pružite vrednost za taj parametar. Možete dodati podrazumevane vrednosti za više parametara koliko god želite.

# --instructions--

Izmijenite funkciju `increment` dodavanjem zadatih parametara da može da poveća za 1 `number` ako `value` nije navedeno.

# --hints--

Rezultati za `increment(5, 2)` bi trebalo da budu `7`.

```js
assert(increment(5, 2) === 7);
```

Rezultat od `increment(5)` treba da bude `6`.

```js
assert(increment(5) === 6);
```

Vrednost podrazumevanog parametra za `1` treba da se koristi za `value`.

```js
assert(__helpers.removeJSComments(code).match(/value\s*=\s*1/g));
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const increment = (number, value) => number + value;
// Only change code above this line
```

# --solutions--

```js
const increment = (number, value = 1) => number + value;
```
