---
id: 587d7b8a367417b2b2512b4f
title: Napišite tamo fupi za vidike literala koristeći skraćenice atributa stvari
challengeType: 1
forumTopicId: 301225
dashedName: write-concise-object-literal-declarations-using-object-property-shorthand
---

# --description--

ES6 dodaje odličnu podršku za jednostavno kreiranje literalnih vrednosti.

Razmotrite sledeći kod:

```js
const getMousePosition = (x, y) => ({
  x: x,
  y: y
});
```

`getMousePosition` je jednostavna funkcija koja vraća nešto sa dve karakteristike. ES6 pruža sintaktički šećer da eliminše potrebu za pisanje `x: x`. Možete samo napisati `x` jednom, a to će biti konvertovano u `x: x` (ili nešto slično) u pozadini. Evo iste funkcije ponovo napisane korišćenjem ove nove sintakse:

```js
const getMousePosition = (x, y) => ({ x, y });
```

# --instructions--

Koristeći kratice atributa za literalne vrednosti objekata, kreiraj i vrati objekt sa atributima `name`, `age` i `gender`.

# --hints--

`createPerson("Zodiac Hasbro", 56, "male")` trebalo bi vratiti `{name: "Zodiac Hasbro", age: 56, gender: "male"}`.

```js
assert.deepEqual(
  { name: 'Zodiac Hasbro', age: 56, gender: 'male' },
  createPerson('Zodiac Hasbro', 56, 'male')
);
```

Tvoj kod ne bi trebalo da koristi `key:value`.

```js
assert(!__helpers.removeJSComments(code).match(/:/g))
```

# --seed--

## --seed-contents--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  return {
    name: name,
    age: age,
    gender: gender
  };
  // Only change code above this line
};
```

# --solutions--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  /*return {
    name: name,
    age: age,
    gender: gender
  };*/
  return {
    name,
    age,
    gender
  };
  // Only change code above this line
};
```
