---
id: 587d7b8a367417b2b2512b4c
title: >-
  Ufumbuzi wa muundo kupitia vipengele vya REST
challengeType: 1
forumTopicId: 301218
dashedName: >-
  use-destructuring-assignment-with-the-rest-parameter-to-reassign-array-elements
---

# --description--

U nekim slučajevima vezanim za rešavanje strukture niza podataka, možemo želeti da prikupimo elemente koji ostanu u različitim pod-nizovima.

Rezultati su kao `Array.prototype.slice()`, kako je prikazano ispod:

```js
const [a, b, ...arr] = [1, 2, 3, 4, 5, 7];
console.log(a, b);
console.log(arr);
```

Konola će prikazati vrednosti za `1, 2` i `[3, 4, 5, 7]`.

Parametri `a` i `b` uzimaju prvu i drugu vrednost iz niza podataka. Nakon toga, zbog postojanja sintakse REST, `arr` dobija preostale vrednosti putem niza podataka. Element REST radi dobro samo kao poslednji parametar u listi. To jest, ne možete koristiti sintaksu REST da uhvatite podniz koji razdvaja poslednji element originalnog niza.

# --instructions--

Koristite obrazac rešavanja strukture sa sintaksom REST da emulirate ponašanje `Array.prototype.slice()`. `removeFirstTwo()` bi trebalo da vrati podredak originalnog reda `list` uklanjanjem prvih dve komponente.

# --hints--

`removeFirstTwo([1, 2, 3, 4, 5])` treba da bude `[3, 4, 5]`

```js
assert.deepEqual(removeFirstTwo([1, 2, 3, 4, 5]), [3, 4, 5]);
```

`removeFirstTwo()` ne bi trebalo da se menja `list`

```js
const _testArr = [1, 2, 3, 4, 5];
removeFirstTwo(_testArr);
assert.deepEqual(_testArr, [1, 2, 3, 4, 5])
```

`Array.slice()` Ne sme se koristiti.

```js
assert(!__helpers.removeJSComments(code).match(/\.\s*slice\s*\(/));
```

Trebalo bi da koristiš sintaksu REST.

```js
assert.match(code, /\.\.\./);
```

# --seed--

## --seed-contents--

```js
function removeFirstTwo(list) {
  return list;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```

# --solutions--

```js
function removeFirstTwo(list) {
  // comment with 'slice' to check comments are removed in tests
  const [, , ...shorterList] = list;
  return shorterList;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```
