---
id: 587d7b87367417b2b2512b40
title: Uporedite opseg ključnih reči var i let
challengeType: 1
forumTopicId: 301195
dashedName: compare-scopes-of-the-var-and-let-keywords
---

# --description--

Ako ne znate `let`, molimo pogledajte <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/explore-differences-between-the-var-and-let-keywords" target="_blank" rel="noopener noreferrer nofollow">ovaj izazov o razlici između <code>let</code> i <code>var</code></a>.

Kada deklariš parametar za ključnu reč `var`, on se deklarše globalno, ili lokalno ako je deklarisan unutar funkcije.

Ključna reč `let` ima takvo ponašanje, ali ima dodatne karakteristike. Kada deklarišete parametar za ključnu reč `let` unutar bloka, izjave ili matematičkog izraza, njen opseg je ograničen na taj blok, izjavu ili izraz.

Na primer:

```js
var numArray = [];
for (var i = 0; i < 3; i++) {
  numArray.push(i);
}
console.log(numArray);
console.log(i);
```

Ova konzola će prikazati vrednosti za `[0, 1, 2]` na `3`.

Za ključnu reč `var`, `i` je globalno objavljeno. Stoga kada se `i++` izvrši, ažurira globalni parametar. Ovaj kod je jednak ovom:

```js
var numArray = [];
var i;
for (i = 0; i < 3; i++) {
  numArray.push(i);
}
console.log(numArray);
console.log(i);
```

Na konzoli će biti prikazane vrednosti za `[0, 1, 2]` i `3`.

Ovo ponašanje će uzrokovati probleme ako kreirate funkciju i sačuvate je za kasniju upotrebu unutar opsega `for` koji koristi parametar `i`. Ovo je zato što će sačuvana funkcija i dalje referencirati vrednost globalnog parametra `i` koji je ažuriran.

```js
var printNumTwo;
for (var i = 0; i < 3; i++) {
  if (i === 2) {
    printNumTwo = function() {
      return i;
    };
  }
}
console.log(printNumTwo());
```

Na konzoli će biti prikazana vrednost `3`.

Kao što možete da vidite, ``printNumTwo()`` ispisuje 3, a ne 2. Ovo je zato što vrednost dodeljena ``i``, ažurirana sa ``printNumTwo()``, vraća globalni ``i``, a ne vrednost koja je bila u ``i`` kada je funkcija kreirana u opsegu `for`. Ključna reč ``let`` ne prati ovo ponašanje:

```js
let printNumTwo;
for (let i = 0; i < 3; i++) {
  if (i === 2) {
    printNumTwo = function() {
      return i;
    };
  }
}
console.log(printNumTwo());
console.log(i);
```

Na konzoli će se prikazati vrednost `2`, i greška da `i is not defined`.

`i` nije deklarisano jer nije deklarisano u globalnom opsegu. Deklariše se samo unutar opsega petlje `for`. `printNumTwo()` je vratio ispravnu vrednost jer su tri različita parametra `i` sa jedinstvenim vrednostima (0, 1 i 2) generisana ključnom reči `let` unutar deklaracije petlje.

# --instructions--

Ispravite kod tako da je `i` deklarisan u deklaraciji `if` drugačiji parametar od `i` deklarisanog na prvoj liniji funkcije. Osigurajte da ne koristite ključnu reč `var` nigde u vašem kodu.

Ova vežba je napravljena da pokaže razliku između toga kako ključne reči na `var` i `let` postavljaju opseg za deklarisani parametar. Kada programirate funkciju kao ona korišćena u ovoj vežbi, često je bolje koristiti različita imena parametara kako biste izbegli zbunjenost.

# --hints--

`var` ne sme biti prisutan u kodu.

```js
assert(!__helpers.removeJSComments(code).match(/var/g));
```

Kigezo `i` koje je deklarisano u izjavi `if` mora biti jednako nizu karaktera `block scope`.

```js
assert(__helpers.removeJSComments(code).match(/(i\s*=\s*).*\s*.*\s*.*\1('|")block\s*scope\2/g));
```

`checkScope()` treba da vrati sekvencu znakova `function scope`

```js
assert(checkScope() === 'function scope');
```

# --seed--

## --seed-contents--

```js
function checkScope() {
  var i = 'function scope';
  if (true) {
    i = 'block scope';
    console.log('Block scope i is: ', i);
  }
  console.log('Function scope i is: ', i);
  return i;
}
```

# --solutions--

```js
function checkScope() {
  let i = 'function scope';
  if (true) {
    let i = 'block scope';
    console.log('Block scope i is: ', i);
  }
 
  console.log('Function scope i is: ', i);
  return i;
}
```
