---
id: 587d7b8c367417b2b2512b55
title: Ponovo koristi JavaScript kod koristeći import
challengeType: 1
forumTopicId: 301208
dashedName: reuse-javascript-code-using-import
---

# --description--

`import` omogućava vam da izaberete koje delove fajla ili modula želite učitati. U prethodnoj lekciji, primeri su koristili export da proslede `add` iz fajla `math_functions.js`. Evo kako možete da ga uvezete ili koristite u drugom fajlu:

```js
import { add } from './math_functions.js';
```

Ovdje, `import` će tražiti `add` u `math_functions.js`, da dovede samo tu funkciju koju ćete koristiti, i ignorisati sve ostalo. `./` objašnjava import traženje fajla `math_functions.js` u istu mapu kao trenutni fajl. Put do fajla za vezu (`./`) i prilog fajla (`.js`) su potrebni kada koristite import na ovaj način.

Možete uvesti više od jednog elementa iz fajla tako što ćete ih postaviti u deklaraciju `import`, kao što sledi:

```js
import { add, subtract } from './math_functions.js';
```

# --instructions--

Dodajte ispravnu `import` naredbu koja će omogućiti trenutnom fajlu da koristi funkcije `uppercaseString` i `lowercaseString` koje ste exportovali u prethodnoj lekciji. Ove funkcije se nalaze u fajlu `string_functions.js`, koji je u istom direktorijumu kao i trenutni fajl.

# --hints--

Morate uvesti `uppercaseString` tačno.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

Treba da uvezete `lowercaseString` tačno.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

uppercaseString("hello");
lowercaseString("WORLD!");
```

# --solutions--

```js
import { uppercaseString, lowercaseString } from './string_functions.js';

uppercaseString("hello");
lowercaseString("WORLD!");
```
