---
id: 587d7b8d367417b2b2512b59
title: Import export podrazumevano
challengeType: 1
forumTopicId: 301205
dashedName: import-a-default-export
---

# --description--

U prethodnom izazovu, ste naučili o `export default` i njegovoj upotrebi. Da biste se import export podrazumevano, morate koristiti drugačiju sintaksu za `import`. U sledećem primeru, `add` je export podrazumevani fajl za `math_functions.js`. Evo kako da se import:

```js
import add from "./math_functions.js";
```

Sintaksa se razlikuje na jednom važnom mestu. Vrednost koja se uzima, `add`, nije okružena zaokruženim zagradama (`{}`). `add` ovde je ime parametra za bilo šta dato kao export podrazumevana vrednost fajla od `math_functions.js`. Možete koristiti proizvoljno ime ovde kada je import export podrazumevana.

# --instructions--

U sledećem kodu, import export podrazumevano iz fajla `math_functions.js`, koji se nalazi u istom direktorijumu kao i ovaj fajl. Pružite import ime `subtract`.

# --hints--

Treba da-import `subtract` pravilno od `math_functions.js`.

```js
assert(__helpers.removeJSComments(code).match(/import\s+subtract\s+from\s+('|")\.\/math_functions\.js\1/g));
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

subtract(7,4);
```

# --solutions--

```js
import subtract from "./math_functions.js";

subtract(7,4);
```
