---
id: 587d7b8b367417b2b2512b53
title: Koristi sintaksu klase objekata za deklarisanje komponente
challengeType: 1
forumTopicId: 301212
dashedName: use-class-syntax-to-define-a-constructor-function
---

# --description--

ES6 uvodi novu sintaksu za kreiranje elemenata, koristeći ključnu reč <dfn>darasa la vitu</dfn>.

U ES5, objekat se može kreirati deklarisanjem funkcije `constructor` i korišćenjem glavne ključne reči `new` za inicijalizaciju tog objekta.

U ES6, deklaracija ``class`` ima metodu ``constructor``, nazvanu glavnom ključnom rečju ``new``. Ako metoda ``constructor`` nije eksplicitno deklarisana, onda je implicitno deklarisana bez ikakvog argumenata.

```js
// Explicit constructor
class SpaceShuttle {
  constructor(targetPlanet) {
    this.targetPlanet = targetPlanet;
  }
  takeOff() {
    console.log("To " + this.targetPlanet + "!");
  }
}

// Implicit constructor 
class Rocket {
  launch() {
    console.log("To the moon!");
  }
}

const zeus = new SpaceShuttle('Jupiter');
// prints To Jupiter! in console
zeus.takeOff();

const atlas = new Rocket();
// prints To the moon! in console
atlas.launch();
```

Treba napomenuti da glavni element sa `class` najavljuje novu funkciju, kojoj se dodaje komponenta. Ova komponenta poziva kada `new` bude pozvan za kreiranje nečega novog.

**Napomena:** UpperCamelCase treba da se koristi kao konvencija za imena klasa ES6, kao što je korišćeno u `SpaceShuttle` gore.

Metoda za `constructor` je poseban način za kreiranje i inicijalizovanje nečega koje je napravljeno za klasu objekata. Vi ćete saznati više o ovome u sekciji Programiranje koja se fokusira na objekte Sertifikata algoritama za JavaScript i Strukture podataka.

# --instructions--

Koristi glavni rečnik `class` i napiši `constructor` da kreira klasu `Vegetable`.

Klasa `Vegetable` omogućava vam da kreirate povrtni objekat sa svojstvom `name` koji se prenosi na `constructor`.

# --hints--

`Vegetable` treba biti `class` sa putem `constructor` koji je objavljen.

```js
assert(
  typeof Vegetable === 'function' && typeof Vegetable.constructor === 'function'
);
```

Glavna reč za `class` mora se koristiti.

```js
assert(__helpers.removeJSComments(code).match(/class/g));
```

`Vegetable` Trebalo bi da se može inicijalizovati.

```js
assert(() => {
  const a = new Vegetable('apple');
  return typeof a === 'object';
});
```

`carrot.name` trebalo bi da vrati `carrot`.

```js
assert(carrot.name == 'carrot');
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const carrot = new Vegetable('carrot');
console.log(carrot.name); // Should display 'carrot'
```

# --solutions--

```js
class Vegetable {
  constructor(name) {
    this.name = name;
  }
}
const carrot = new Vegetable('carrot');
```
