---
id: 587d7b89367417b2b2512b49
title: Koristite model distribucije rešenja za postavljanje vrednosti parametara iz objekata
challengeType: 1
forumTopicId: 301215
dashedName: use-destructuring-assignment-to-assign-variables-from-objects
---

# --description--

Rešavanje strukture omogućava vam da dodelite novo ime parametru kada uzimate vrednost. Možete to učiniti dodeljivanjem novog imena nakon tačke prilikom postavljanja vrednosti.

Ako uzmete isti primer objekta iz prethodnog primera:

```js
const user = { name: 'John Doe', age: 34 };
```

Ovo je kako možete dodeliti nova imena parametrima u opsegu:

```js
const { name: userName, age: userAge } = user;
```

Možete pročitati kao: "dobiti vrednost od `user.name` i postaviti tu vrednost u novi parametar nazvan `userName`", itd. Vrednost `userName` će biti niz slova `John Doe`, a vrednost `userAge` će biti broj `34`.

# --instructions--

Promenite taj obrazac sa dve vrednosti i usklađeni obrazac rešavanja strukture. Još je potrebno postaviti parametre `highToday` i `highTomorrow` vrednosti za `today` do `tomorrow` iz objekta `HIGH_TEMPERATURES`.

# --hints--

Trebalo bi da uklonite sintaksu dodeljivanja vrednosti ES5.

```js
assert(
  !__helpers.removeJSComments(code).match(/highToday = HIGH_TEMPERATURES\.today/g) &&
    !__helpers.removeJSComments(code).match(/highTomorrow = HIGH_TEMPERATURES\.tomorrow/g)
);
```

Trebalo bi da koristite strukturalni rešenje za kreiranje parametra `highToday`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(today\s*:\s*highToday[^}]*|[^,]*,\s*today\s*:\s*highToday\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

Trebalo bi da koristite strukturalni rešenje za kreiranje parametra `highTomorrow`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(tomorrow\s*:\s*highTomorrow[^}]*|[^,]*,\s*tomorrow\s*:\s*highTomorrow\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

`highToday` mora biti jednako `77` i `highTomorrow` mora biti jednako `80`.

```js
assert(highToday === 77 && highTomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line
  
const highToday = HIGH_TEMPERATURES.today;
const highTomorrow = HIGH_TEMPERATURES.tomorrow; 

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today: highToday, tomorrow: highTomorrow } = HIGH_TEMPERATURES;
```
