---
id: 587d7b8b367417b2b2512b50
title: Napiši kratku funkciju za ES6
challengeType: 1
forumTopicId: 301224
dashedName: write-concise-declarative-functions-with-es6
---

# --description--

Kada definišete funkcije unutar objekata u ES5, treba da koristite ključnu reč `function` kao što sledi:

```js
const person = {
  name: "Taylor",
  sayHello: function() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

Za ES6, možete potpuno ukloniti ključnu reč `function` i dvotačku prilikom definisanja funkcija unutar objekata. Evo primera ove sintakse:

```js
const person = {
  name: "Taylor",
  sayHello() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

# --instructions--

Izmenite funkciju `setGear` unutar objekta `bicycle` da koristi kratku sintaksu pomenutu gore.

# --hints--

Standardni poziv funkcije nije korišćen.

```js
assert(!__helpers.removeJSComments(code).match(/function/));
```

`setGear` Treba da bude funkcija deklaracije.

```js
assert(
  typeof bicycle.setGear === 'function' && __helpers.removeJSComments(code).match(/setGear\s*\(.+\)\s*\{/)
);
```

`bicycle.setGear(48)` treba da promeni vrednost `gear` na 48.

```js
bicycle.setGear(48);
assert(bicycle.gear === 48);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const bicycle = {
  gear: 2,
  setGear: function(newGear) {
    this.gear = newGear;
  }
};
// Only change code above this line
bicycle.setGear(3);
console.log(bicycle.gear);
```

# --solutions--

```js
const bicycle = {
  gear: 2,
  // setGear: function(newGear) {
  setGear(newGear) {
    this.gear = newGear;
  }
};
bicycle.setGear(3);
```
