---
id: 587d7b89367417b2b2512b48
title: Koristi spread operaciju za procenu redova podataka na mestu
challengeType: 1
forumTopicId: 301222
dashedName: use-the-spread-operator-to-evaluate-arrays-in-place
---

# --description--

ES6 inaleta <dfn>raspršivački operator</dfn>, koji nam omogućava da proširimo nizove podataka i matematičke izraze na mestima gde se očekuje više parametara ili komponenti.

Kod ES5 ispod koristi `apply()` da izračuna najveću vrednost u nizu podataka:

```js
var arr = [6, 89, 3, 45];
var maximus = Math.max.apply(null, arr);
```

`maximus` ima vrednost od `89`.

Morali smo da koristimo `Math.max.apply(null, arr)` jer `Math.max(arr)` vraća `NaN`. `Math.max()` očekuje argumente razdvojene zarezima, ali ne red podataka. Spread operater čini ovu sintaksu lakšom za čitanje i održavanje.

```js
const arr = [6, 89, 3, 45];
const maximus = Math.max(...arr);
```

`maximus` biće sa vrednošću od `89`.

`...arr` Vraća otvoreni red podataka. Drugim rečima, širi red podataka. Međutim, operacija raspršivanja (spread) radi samo na mestu, kao u funkcionalskim argumentima ili u listi redova podataka. Primer:

```js
const spreaded = [...arr];
```

Međutim, sledeći kod neće raditi:

```js
const spreaded = ...arr;
```

# --instructions--

Kopiraj sav sadržaj iz `arr1` unutar drugog reda `arr2` koristeći operater širenja.

# --hints--

`arr2` treba da bude tačna kopija od `arr1`.

```js
assert(arr2.every((v, i) => v === arr1[i]) && arr2.length);
```

Operater spread `...` bi trebalo koristiti za dupliranje `arr1`.

```js
assert(__helpers.removeJSComments(code).match(/Array\(\s*\.\.\.arr1\s*\)|\[\s*\.\.\.arr1\s*\]/));
```

`arr2` mora ostati nepromijenjen kada se `arr1` promijeni.

```js
assert((arr1, arr2) => {
  arr1.push('JUN');
  return arr2.length < arr1.length;
});
```

# --seed--

## --seed-contents--

```js
const arr1 = ['JAN', 'FEB', 'MAR', 'APR', 'MAY'];
let arr2;

arr2 = [];  // Change this line

console.log(arr2);
```

# --solutions--

```js
const arr1 = ['JAN', 'FEB', 'MAR', 'APR', 'MAY'];
let arr2;

arr2 = [...arr1];
```
