---
id: 598f48a36c8c40764b4e52b3
title: Spreči promene nečega
challengeType: 1
forumTopicId: 301207
dashedName: prevent-object-mutation
---

# --description--

Kao što je viđeno u prethodnom izazovu, sama izjava `const` ne štiti vaše podatke od promena. Da bi se osiguralo da su vaši podaci nepromenljivi, JavaScript pruža funkciju `Object.freeze` radi sprečavanja promena podataka.

Bilo koji pokušaj izmene nečega će biti odbijen, i greška će biti bačena ako se skript izvršava u režimu strogog načina (strict mode).

```js
let obj = {
  name:"FreeCodeCamp",
  review:"Awesome"
};
Object.freeze(obj);
obj.review = "bad";
obj.newProp = "Test";
console.log(obj); 
```

Korišćenje `obj.review` na `obj.newProp` uzrokuje greške, jer naš editor radi u strict modu po podrazumevanju, i konzola će prikazati vrednost `{ name: "FreeCodeCamp", review: "Awesome" }`.

# --instructions--

U ovom izazovu ćete koristiti `Object.freeze` da sprečite promenu matematičkih konstanti. Morate zaključati nešto sa `MATH_CONSTANTS` kako bi niko nije mogao da promeni vrednost `PI`, doda ili obriše parametre.

# --hints--

Ne smete menjati reč `const`.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`MATH_CONSTANTS` treba da bude konstanta promenljiva (uz korišćenje `const`).

```js
assert(__helpers.removeJSComments(code).match(/const\s+MATH_CONSTANTS/g));
```

Ne možete promeniti prvobitnu deklaraciju `MATH_CONSTANTS`.

```js
assert(__helpers.removeJSComments(code).match(
   /const\s+MATH_CONSTANTS\s+=\s+{\s+PI:\s+3.14\s+};/g
));
```

`PI` mora biti jednako sa `3.14`.

```js
assert(PI === 3.14);
```

# --seed--

## --seed-contents--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  // Only change code below this line


  // Only change code above this line
  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```

# --solutions--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  Object.freeze(MATH_CONSTANTS);

  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```
