---
id: 5cdafbb0291309899753167f
title: Napravi obećanje (promise) u JavaScriptu
challengeType: 1
forumTopicId: 301197
dashedName: create-a-javascript-promise
---

# --description--

Obećanje u JavaScript je, kako se čuje - koristiš ga da bi dao obećanje da će nešto uraditi, obično ne u jednom trenutku. Kada se vežba završi, ili ispuniš svoje obećanje ili ne uspeš da uradiš to. `Promise` je komponentna funkcija, stoga moraš da koristiš glavnu reč `new` kako bi napravio jednu. Uzima funkciju, kao svoj argument, sa dve parametre - `resolve` i `reject`. Ovo su načini koji se koriste za određivanje rezultata obećanja. Sintaksa izgleda ovako:

```js
const myPromise = new Promise((resolve, reject) => {

});
```

# --instructions--

Napravite novi savez nazvan `makeServerRequest`. Popijte funkciju sa parametrima `resolve` i `reject` za komponentu.

# --hints--

Trebalo bi da postaviš obećanje na parametar nazvan `makeServerRequest`.

```js
assert(makeServerRequest instanceof Promise);
```

Vaša funkcija treba da primi funkciju sa parametrima `resolve` i `reject`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /Promise\s*\(\s*(function\s*\(\s*resolve\s*,\s*reject\s*\)\s*{|\(\s*resolve\s*,\s*reject\s*\)\s*=>\s*{)[^}]*}/g
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {

});
```
