---
id: 5cdafbd72913098997531681
title: Obradi obećanje koje je implementirano do tada
challengeType: 1
forumTopicId: 301203
dashedName: handle-a-fulfilled-promise-with-then
---

# --description--

Obećanja su najvažnija kada imate proces koji traje nepoznato vreme u vašem kodu (tj. nešto što nije trenutno), često je to zahtev servera. Kada napravite zahtev serveru, traje određeno vreme, i nakon završetka obično želite da nešto uradite sa odgovorom od servera. Ovo se može postići korišćenjem `then`.

```js
Promise.prototype.then(onFulfilled, onRejected)
```

Put za `then` priprema callback funkcije prilikom završetka obećanja – bilo da je izvršeno ili odbijeno. Jedan od handler-a za `onFulfilled` i `onRejected` će biti pokren (ili izvršen) da obradi ispunjenje ili odbijanje trenutnog obećanja. Kada se obećanje ispuni sa `resolve`, poziva se handler za `onFulfilled`.

```js
myPromise.then(result => {
  
});
```

`result` potiče iz tvrdnje napravljene putem `resolve`.

# --instructions--

Dodajte putanju za `then` za vašu obavezu. Koristite `result` kao parametar njegove funkcije odgovora i unesite `result` u konzolu.

# --hints--

Trebaće ti da pozoveš putanju `then` u obećanje.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.then\(/g)
);
```

Tvoja putanja ``then`` treba da ima funkciju povratne pozive sa ``result`` kao parametrom.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

Trebalo bi da napišete `result` u konzoli.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.then\(.*?result.*?console.log\(result\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```
