---
id: 5cdafbe72913098997531682
title: Obradi odbijeno obećanje koristeći `catch`
challengeType: 1
forumTopicId: 301204
dashedName: handle-a-rejected-promise-with-catch
---

# --description--

`catch` je metoda koja se koristi kada vaša obećanje bude aktivirano. Izvršava se odmah nakon što je metoda obećanja `reject` pozvana. Evo sintakse:

```js
myPromise.catch(error => {
  
});
```

`error` je list koji prolazi u putanji od `reject`.

# --instructions--

Dodaj putanju `catch` za vašu obavezu. Koristi `error` kao argument za njegovu funkciju odgovora i napiši `error` u konzoli.

# --hints--

Trebalo bi da pozoveš putanju `catch` u obećanje.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.catch\(/g)
);
```

Tvoja putanja za `catch` treba da ima funkciju povratnog poziva sa `error` kao njenim parametrom.

```js
assert(
  /\.catch\((function\(error\){|error|\(error\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

Trebaće ti da napišeš `error` u konzoli.

```js
assert(
  /\.catch\((function\(error\){|error|\(error\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.catch\(.*?error.*?console.log\(error\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});

makeServerRequest.catch(error => {
  console.log(error);
});
```
