---
id: 5cddbfd622f1a59093ec611d
title: Kreiraj skriptu modula
challengeType: 6
forumTopicId: 301198
dashedName: create-a-module-script
---

# --description--

JavaScript počeo sa malom ulogom u mreži koja je posebno puna HTML. Danas je veoma velika, i neke veb stranice su gotovo isključivo napravljene korišćenjem JavaScript. Da bi se učinio JavaScript modularnijim, čistijim i lakšim za održavanje; ES6 je uveo način deljenja koda između fajlova JavaScript. Ovo podrazumeva prenos dela fajla radi upotrebe u jednom ili više fajlova, i dovođenje dela koji vam je potreban, tamo gde vam treba. Da biste koristili ovu funkciju, morate kreirati skript u vašem dokumentu za HTML sa `type` od `module`. Evo primera:

```html
<script type="module" src="filename.js"></script>
```

Skripta koja koristi ovaj tip `module` sada može da koristi funkcije `import` i `export`, koje ćete naučiti u budućim izazovima.

# --instructions--

Dodaj skriptu u dokument sa HTML tipa `module` i dostavi izvorni fajl za `index.js`

# --hints--

Trebalo bi da kreirate etiketu `script`.

```js
assert(__helpers.removeJSComments(code).match(/<\s*script[^>]*>\s*<\/\s*script\s*>/g));
```

Vaša vrednost od `script` mora imati atribut `type` sa vrednošću `module`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^t]*type\s*=\s*('|")module\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

Vaša lozinka za `script` mora da ima `src` od/za `index.js`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^s]*src\s*=\s*('|")index\.js\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <!-- Only change code below this line -->

    <!-- Only change code above this line -->
  </body>
</html>
```

# --solutions--

```html
<html>
  <body>
    <script type="module" src="index.js"></script>
  </body>
</html>
```
