---
id: 5cfa550e84205a357704ccb6
title: Koristi model rešenja strukture da izvuče vrednost iz stvari
challengeType: 1
forumTopicId: 301216
dashedName: use-destructuring-assignment-to-extract-values-from-objects
---

# --description--

<dfn>Dekonstrukcija strukture</dfn> je posebna sintaksa uvedena u ES6, za dodeljivanje vrednosti na elegantan način direktno iz objekta.

Razmotri sledeći ES5 kod:

```js
const user = { name: 'John Doe', age: 34 };

const name = user.name;
const age = user.age;
```

`name` imao bi vrednost sekvence slova `John Doe`, i `age` imao bi broj `34`.

Evo ekvivalentne deklaracije niza koja koristi ES6 sintaksu za dekonstrukciju struktura:

```js
const { name, age } = user;
```

Zdravo, `name` bi imalo vrednost sekvence slova `John Doe`, a `age` bi imao broj `34`.

Ovde, parametri `name` i `age` će biti kreirani i imu dodeljene vrednosti iz njihovih odgovarajućih vrednosti objekta `user`. Vidite kako je to mnogo čistije.

Možete izvući mnogo ili malo vrednosti iz nečega, kako želite.

# --instructions--

Promeni taj maping sa dve vrednosti u odgovarajući maping strukturnog rešenja. Treba da postavi parametri `today` i `tomorrow` vrednosti od `today` do `tomorrow` iz objekta `HIGH_TEMPERATURES`.

# --hints--

Trebalo bi da uklonite sintaksu dodeljivanja vrednosti ES5.

```js
assert(
  !__helpers.removeJSComments(code).match(/today\s*=\s*HIGH_TEMPERATURES\.(today|tomorrow)/g)
);
```

Trebalo bi da koristite strukturalni rešenje za kreiranje parametra `today`.

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(today[^}]*|[^,]*,\s*today)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

Trebalo bi da koristite strukturalni rešenje za kreiranje parametra `tomorrow`.

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(tomorrow[^}]*|[^,]*,\s*tomorrow)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

`today` mora biti jednako sa `77` i `tomorrow` mora biti jednako sa `80`.

```js
assert(today === 77 && tomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line

const today = HIGH_TEMPERATURES.today;
const tomorrow = HIGH_TEMPERATURES.tomorrow;

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today, tomorrow } = HIGH_TEMPERATURES;
```
