---
id: a202eed8fc186c8434cb6d61
title: Napravi Konverter Sekvence Slova
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

U ovoj laboratoriji, napravićete jednostavan konverter za sekvence znakova koji transformiše dati niz znakova.

Na primer, `"hello"` bi trebalo da bude `"olleh"`.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratorijum.

Priče o korisniku:

1. Treba da kreirate funkciju nazvanu `reverseString` koja prima niz znakova kao argument.
2. Funkcija treba da vrati niz znakova u obrnutom redosledu.

# --hints--

Trebalo bi da imate funkciju nazvanu `reverseString`.

```js
assert.isFunction(reverseString);
```

`reverseString` Trebalo bi da primi niz znakova kao argument.

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")` Trebalo bi da vrati sekvencu slova.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` treba da vrati niz znakova `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` treba da vrati niz karaktera `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` treba da vrati sekvencu karaktera `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
