---
id: 68e529226b0defdff26c4a8a
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Sada, kreiraj instancu `Product` sa imenom `Wireless Mouse` i cenom od `50.0`. Postavi njegovu vrednost u parametar nazvan `product`. Zatim, ispiši `product` u konzolu da vidiš njegov izgled.

# --hints--

Trebalo bi da imaš parametar koji se zove `product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("product")`)
})
```

Trebalo bi da postavite vrednost primeru `Product` u parametar `product`.

```js
({
  test: () => runPython(`assert isinstance(product, Product)`)
})
```

Vaš primer za `Product` treba da ima ime `Wireless Mouse` i cenu `50.0`.

```js
({
  test: () => runPython(`
    assert product.name == 'Wireless Mouse'
    assert product.price == 50.0
  `)
})
```

Treba da ispišete `product` u konzolu.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(product)")`)
})
```

Prilikom štampanja `Product` sa imenom `'Wireless Mouse'` i cenom `50.0`, trebalo bi da prikaže `Wireless Mouse - $50.0`.

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product('Wireless Mouse', 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
```
