---
id: 68e5293bd00d2fe134f58983
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Sada dodaj drugu metodu abstracta nazvanu `apply_discount` u vašu klasu `DiscountStrategy`. Koristi dekorator `@abstractmethod` i sa dodeli mu parametre `self` i `product`.

Dodaj napomene tipa: `product` treba da bude tipa `Product`, i ta metoda bi trebalo da vrati `float`. Koristi `pass` unutar tela te metode.

# --hints--

Vaša klasa za `DiscountStrategy` mora da ima metodu nazvanu `apply_discount`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

Tvoja putanja ``apply_discount`` bi trebalo da ima dekorator ``@abstractmethod``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

Tvoja putanja ``apply_discount`` bi trebalo da ima parametre ``self`` i ``product: Product``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

Tvoj put za `apply_discount` bi trebalo da ima napomenu tipa povratka `-> float`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
