---
id: 68e529226b0defdff26c4a89
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Trenutno, kada kreirate primer `Product` i odštampate ga, videćete podrazumevanu reprezentaciju objekta koji ne pruža mnogo informacija. Na primer, može izgledati kao `<__main__.Product object at 0x7f9b8c2d1d30>`.

Da biste konvertovali podrazumevanu reprezentaciju objekta u čitljiv niz karaktera, kreirajte metod `__str__` u vašoj klasi `Product` koja vraća niz karaktera sa strukturom `name - $price` (gde `name` i `price` treba zameniti odgovarajućim vrednostima svojstava elementa).

# --hints--

Vaša klasa za `Product` mora imati metodu za `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

Tvoj put za `__str__` bi trebalo da vrati niz karaktera sa strukturom `name - $price`.

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
--fcc-editable-region--
    
--fcc-editable-region--
```
