---
id: 68e5293bd00d2fe134f58984
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Sada je vreme za kreiranje stvarnog izvršnog koda vaše apstraktne klase objekata. Kreirajte klasu objekata koja se zove `PercentageDiscount` i koja nasleđuje od `DiscountStrategy`.

Definišite metodu za ``__init__`` koja prima ``self`` i ``percent`` kao parametre. Definišite tip ``int`` za ``percent`` i postavite povratni tip te metode na ``None``. Unutar metode ``__init__``, sačuvajte ``percent`` kao svojstvo primernog elementa.

# --hints--

Treba da kreiraš klasu objekata koja se zove `PercentageDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

Tvoja klasa za `PercentageDiscount` treba da nasleduje od `DiscountStrategy`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

Tvoja klasa za `PercentageDiscount` treba da ima putanju ka `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

Tvoja putanja ``__init__`` mora imati parametre ``self`` i ``percent: int``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

Tvoja putanja ``__init__`` treba da ima atribut tipa povratak ``None``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

Trebalo bi da postaviš vrednost `percent` za `self.percent` unutar svog puta `__init__`.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
