---
id: 68e5293bd00d2fe134f58988
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Implementirajte putanju za `is_applicable` za vaš kurs/klasu `FixedAmountDiscount`. Popust na definisanim iznosom mora se koristiti samo ako je ekonomski opravdan.

Rudisha `True` ako je `product.price * 0.9` veće od `self.amount`, inače vrati `False`. Ovo osigurava da je popust manji od 90% originalne cene, pa nije preteran.

# --hints--

Vaša klasa za `FixedAmountDiscount` mora imati metod/način za `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("is_applicable")`)
})
```

Tvoja putanja ``is_applicable`` mora imati parametre ``self``, ``product: Product``, i ``user_tier: str``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_args("self, product:Product, user_tier:str")`)
})
```

Tvoja putanja ``is_applicable`` bi trebalo da ima tip vraćanja ``-> bool``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

Tvoj put za `is_applicable` treba da vrati `True` kada `product.price * 0.9` je veći od `self.amount`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(50)
    assert discount.is_applicable(test_product, "regular") == True
  `)
})
```

Tvoja putanja za `is_applicable` treba da vrati `False` kada `product.price * 0.9` je manjiše ili jednako od `self.amount`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(95)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
