---
id: 68e52994bc7ea2e3dec5309d
title: Korak 20
challengeType: 20
dashedName: step-20
---

# --description--

Za treću strategiju popusta, kreiraj klasu objekata nazvanu `PremiumUserDiscount` koja nasleđuje od `DiscountStrategy`.

Za razliku od prethodnih strategija, ovo ne zahteva `__init__` jer ne čuva nikakvu konfiguraciju. Postaviće fiksni popust od 20 procenata za sve premium korisnike.

Trenutno, kreiraj samo klasu praznih stvari koja nasleđuje od `DiscountStrategy` koristeći `pass`.

# --hints--

Treba da kreiraš klasu objekata koja se zove `PremiumUserDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PremiumUserDiscount")`)
})
```

Klasa za stvari `PremiumUserDiscount` treba da nasledi od `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").inherits_from("DiscountStrategy")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
