---
id: 68e5293bd00d2fe134f58981
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Kreiraj klasu objekata nazvanu `DiscountStrategy` koja nasleđuje od `ABC`. Ovo će biti osnovna apstraktna klasa koja opisuje interfejs svih strategija popusta.

Trenutno, koristi samo `pass` unutar tela klase da kreiraš praznu klasu.

# --hints--

Trebalo bi da kreirate klasu objekata koja se zove `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

Vaša klasa za `DiscountStrategy` mora da nasledi od `ABC`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
