---
id: 68e5293bd00d2fe134f58987
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Sada ćeš kreirati drugu strategiju popusta. Kreiraj klasu nazvanu `FixedAmountDiscount` koja nasleđuje iz `DiscountStrategy`.

Definišite metodu `__init__` koja prima parametar `self` i `amount`. Definišite parametar `amount` tipa `int` i povratni tip `None` na taj način. Unutar metode `__init__`, sačuvajte `amount` kao svojstvo primernog elementa.

# --hints--

Treba da kreiraš klasu objekata koja se zove `FixedAmountDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("FixedAmountDiscount")`)
})
```

Klasa za `FixedAmountDiscount` mora nasleđivati od `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").inherits_from("DiscountStrategy")`)
})
```

Vaša klasa za `FixedAmountDiscount` treba da ima metodu/funkciju `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("__init__")`)
})
```

Tvoja putanja ``__init__`` mora imati parametre ``self`` i ``amount: int``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_args("self, amount: int")`)
})
```

Tvoja putanja ``__init__`` bi trebalo da ima tip vraćanja ``None``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_returns("None")`)
})
```

Trebalo bi da postaviš vrednost `amount` za `self.amount` unutar svog puta `__init__`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_stmt("self.amount = amount")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
