---
id: 68e52994bc7ea2e3dec530a1
title: Korak 23
challengeType: 20
dashedName: step-23
---

# --description--

Sada vam treba način da upravljate sa mnogo strategija popusta i pronađete najbolje cene.

Kreiraj klasu koja se zove `DiscountEngine`. Ova klasa će upravljati svim strategijama popusta i izračunati najbolju cenu za proizvode.

Pružite metodu `__init__` koja prima `self` i `strategies` kao parametre, gde `strategies` ima tip `list[DiscountStrategy]`. Korišćenje `list[DiscountStrategy]` ukazuje da je parametar niz elemenata tipa `DiscountStrategy`. Tip povratne vrednosti metode bi trebalo da bude `None`. Unutar metode `__init__`, sačuvajte `strategies` kao svojstvo elementa primera.

# --hints--

Trebalo bi da napraviš klasu objekata koja se zove `DiscountEngine`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountEngine")`)
})
```

Tvoja klasa za `DiscountEngine` treba da ima metodu `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").has_function("__init__")`)
})
```

Tvoja putanja ``__init__`` mora imati parametre ``self`` i ``strategies: list[DiscountStrategy]``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_args("self, strategies: list[DiscountStrategy]")`)
})
```

Tvoja putanja ``__init__`` treba da ima tip vraćanja ``None``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_returns("None")`)
})
```

Trebalo bi da staviš `strategies` u `self.strategies` unutar svog puta `__init__`.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_stmt("self.strategies = strategies")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
