---
id: 68e524c94e2e4ab7b4bfc7e2
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovoj radionici, kreiraćete kalkulator popusta koji može da koristi različite strategije popusta za proizvode. Sistem će odrediti najbolju cenu za klijenta na osnovu mnogo pravila popusta.

Počnite definišući klasu objekata nazvanu `Product`. Dajte svojoj klasi metod `__init__` sa sledećim parametrima: `self`, `name`, i `price`. Unutar metode `__init__`, postavite vrednosti parametara `name` i `price` za atribute elemenata sa istim imenima.

# --hints--

Treba da kreiraš klasu objekata nazvanu `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

Tvoja klasa `Product` mora imati metod `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

Tvoja putanja ``__init__`` mora imati sledeće parametre: ``self``, ``name``, i ``price``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

Trebalo bi da postaviš vrednost `name` za `self.name` unutar svog puta `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

Treba da postaviš vrednost `price` u `self.price` unutar svog puta `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
