---
id: 68e5293bd00d2fe134f58982
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Kao što je diskutovano na prethodnoj lekciji, apstraktne metode su metode koje moraju biti implementirane od strane bilo koje klase koja nasleđuje od apstraktne bazne klase. One služe kao ugovor, osiguravajući da sve strategije redukcije imaju potrebne metode.

U vašoj klasi od `DiscountStrategy`, definišite metodu pod nazivom `is_applicable`. Koristite dekorator `@abstractmethod` iznad definicije te metode. Ta metoda treba da ima sledeće parametre: `self`, `product`, i `user_tier`. Koristite deklaraciju `pass` unutar tela te metode.

Dodajte specifikacije tipa kako biste učinili potpis puta jasnim: `product` mora biti tipa `Product`, `user_tier` mora biti tipa `str`.


# --hints--

Vaša klasa za `DiscountStrategy` mora da ima metodu pod nazivom `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("is_applicable")`)
})
```

Tvoja putanja ``is_applicable`` bi trebalo da ima dekorator ``@abstractmethod``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_decorators("abstractmethod")`)
})
```

Tvoja putanja ``is_applicable`` mora imati parametre ``self``, ``product: Product``, i ``user_tier: str``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

--fcc-editable-region--
class DiscountStrategy(ABC):
    pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
