---
id: 68e529226b0defdff26c4a8b
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Sada ćeš početi da gradiš interfejs sistema redukcije. Python pruža način definisanja apstraktnih klasa objekata koje se koriste kao šeme za druge klase. To se postiže korišćenjem klase `ABC` (Abstract Base Class) iz modula `abc`.

Iznad vašeg koda, dodajte sledeću izjavu za import:

```py
from abc import ABC, abstractmethod
```

Ovo se primenjuje na klasu `ABC` na dekorator `abstractmethod`, koji će se koristiti za definisanje metoda koje moraju biti implementirane od strane bilo koje klase koja nasleđuje iz vaše apstraktne klase.

# --hints--

Trebalo bi da imaš deklaraciju import navedenu na tvom kodu.

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
