---
id: 68e5293bd00d2fe134f58985
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Pošto `PercentageDiscount` nasleđuje od apstraktne klase, mora implementirati sve apstraktne metode. Počnite implementiranjem metode `is_applicable`. Metoda `is_applicable` treba da uzme `self`, `product: Product` i `user_tier: str` kao parametre. Treba da ima tip povratne vrednosti sa oznakom `bool`.

Ova metoda mora da utvrdi da li je procenat popusta validan. Vrati `True` ako je `self.percent` manji ili jednako `70`, inače vrati `False`. Ovo sprečava nevažeće popuste kao što je popust od 90 procenata.

# --hints--

Vaša klasa od `PercentageDiscount` treba da ima metodu za `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

Tvoja putanja ``is_applicable`` mora imati parametre ``self``, ``product: Product``, i ``user_tier: str``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

Tvoja putanja od `is_applicable` treba da vrati `True` kada `self.percent` je manje ili jednako `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

Tvoj put/način ``is_applicable`` treba da vrati ``False`` kada je ``self.percent`` veći od ``70``.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

Tvoja putanja ``is_applicable`` bi trebalo da ima oznaku tipa povratak ``bool``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
