---
id: 68e5293bd00d2fe134f58986
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Implementiraj metodu `apply_discount` u vašem `PercentageDiscount`.

Ovaj put treba da izračuna i vrati smanjenu cenu.

Formula je: `product.price * (1 - self.percent / 100)`

# --hints--

Tvoja klasa za `PercentageDiscount` treba da ima metodu za `apply_discount`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("apply_discount")`)
})
```

Tvoj put/način za `apply_discount` mora imati parametre `self` i `product: Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_args("self, product: Product")`)
})
```

Tvoja putanja ``apply_discount`` bi trebalo da ima tip vraćanja ``float``.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_returns("float")`)
})
```

Tvoja putanja za `apply_discount` treba da izračuna sniženu cenu koristeći formulu `product.price * (1 - self.percent / 100)`.

```js
({
  test: () => runPython(`
    test_product = Product("Test", 100.0)
    discount = PercentageDiscount(20)
    result = discount.apply_discount(test_product)
    assert result == 80.0
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
