---
id: 68e79d91daf8fb4808f93c47
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Nakon izvršavanja metode `apply_discount`, pokušajte sa svojom strategijom popusta. Kreirajte primer `PercentageDiscount` sa procentom `10` i postavite njegovu vrednost u parametar nazvan `discount`. Zatim primenite popust koristeći `discount.apply_discount(product)` i ispišite rezultate da biste videli kako radi popust.

# --hints--

Trebalo bi da kreirate parametar nazvan `discount` sa primerom `PercentageDiscount`.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

Vaš primer za `PercentageDiscount` mora biti inicijalizovan sa procentom `10`.

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

Trebalo bi da pozoveš `discount.apply_discount(product)` za štampanje rezultata.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
    
    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--
```
