---
id: 68e9207ddcba053c5cb24413
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

U Python, anotacije tipa su opcioni indikatori koji drugim programerima govore koji je tip podataka očekivan za parametar ili funkciju. Evo primera koji dodaje anotaciju tipa `str` za parametar `name`.

```py
def greet(name: str):
    return 'Hello, ' + name
```

Za parametre puta `__init__`, ažuriraj parametar `name` da ima tipsku oznaku `str` i parametar `price` da ima tipsku oznaku `float`.

# --hints--

Trebalo bi dodati podsetnik tipa `str` za parametar `name` i podsetnik tipa `float` za parametar `price` u metodi `__init__` klase objekata `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--
```
