---
id: ab6137d4e35944e21037b769
title: Kreiraj Konvertor Zaglavlja Reči
challengeType: 26
dashedName: build-a-title-case-converter
---

# --description--

U ovoj laboratoriji ćeš kreirati funkciju koja menja niz slova u format naslova (Title Case). Naslovni slučaj znači da prvo slovo svake reči postaje veliko, a ostatak reči je malim slovima.

`"Web Development Is Awesome"` je primer sekvence slova sa velikim početnim slovima.

**Cilj**: Ispunite korisničke priče navedene ispod i prođite sve testove kako biste završili laboratoriju.

Priče o korisniku:

1. Mora imati funkciju `titleCase` koja prima niz slova kao argument.
1. Funkcija `titleCase` treba da vrati niz slova gde je prvo slovo svake reči veliko, a ostatak reči malo.
1. `titleCase("I like to code")` treba da vrati `"I Like To Code"`.
1. `titleCase("javaScript is fun")` treba da vrati `"Javascript Is Fun"`.

# --hints--

Trebalo bi da kreirate funkciju nazvanu `titleCase`.

```js
assert.isFunction(titleCase);
```

`titleCase` Trebalo bi da primi samo jedan argument.

```js
assert.lengthOf(titleCase, 1);
```

`titleCase("I'm a little tea pot")` trebalo bi da vrati sekvencu karaktera.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` treba da vrati sekvencu karaktera `I'm A Little Tea Pot`.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` trebalo bi da vrati niz znakova `Short And Stout`.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` mora da vrati niz znakova `Here Is My Handle Here Is My Spout`.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}
```
