---
id: 60f02e7361b68405e27b62a5
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Na ovoj radionici će vam biti naučeno kako da dizajnirate forme kreiranjem obrasca za registraciju.

Sve osnovne strukture HTML (`DOCTYPE`, `html`, `head`, na `body`) su vam već postavljene.

Unutar `body`, dodajte zaglavlje sekcije kako biste pružili kontekst za obrazac koristeći funkciju `h1` sa tekstom `Registration Form`.

# --hints--

Trebalo bi da ubaciš `h1` unutar `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Trebalo bi mu dati tekst `h1` `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
