---
id: 60f1922fcbd2410527b3bd89
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

U prethodnim lekcijama, naučili ste kako da radite sa jedinicama vidokruga.

Zapamtite da jedinica `vh` označava dužinu vidnog polja, što je jednako 1% od `height` vidnog polja. Ovo se odnosi na visinu vidnog polja.

Sada je vreme za stilizovanje projekta korišćenjem CSS. Počnite primenom `body` `width` od `100%`, i `height` od `100vh`.

# --hints--

Trebalo bi da koristiš ekstrakter atributa `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Trebalo bi mu dati `body` `width` od `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Trebalo bi mu dati `body` `height` od `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
