---
id: 60f1a5e2d2c23707a4f9a660
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Sada, uklonite horizontalni separator, postavljanjem standardnog `body` koji je dodan od strane nekih pregledača u `margin` da postane `0`.

# --hints--

Treba da dodate `margin` unutar selektora funkcije za `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

Trebalo bi mu dati vrednost `margin` od `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
