---
id: 60f1a9cbd23023082e149fee
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

To je bolje. Sada, učinite temu prijatnom za oči, menjanjem `background-color` od `body` u `#1b1b32`. Zatim, da biste videli tekst, promenite `color` u `#f5f6f7`.

# --hints--

Trebalo bi da dodate `background-color` unutar selektora polja za `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

Treba da mu date vrednost `background-color` od `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

Treba da dodaš `color` unutar izabira elementa `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

Trebalo bi mu dati vrednost `color` od `#f5f6f7`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
