---
id: 60f5d2776c854e069560fbe6
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

`fieldset` prvi će sadržavati polja za ime, e-mail i lozinku. Počnite dodavanjem četiri elementa `label` unutar prvog `fieldset`.

# --hints--

Može da ima tri elementa `fieldset` ukupno.

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

Elementi za `fieldset` bi trebalo da budu direktni potomci elementa `form`.

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

Četiri elementa za `label` bi trebalo da budu direktni potomci elementa `fieldset`. Osigurajte da zatvorite elemente `label`.

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

Zadva poslednja elementa `fieldset` moraju biti prazni.

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

Uverite se da zatvorite elemente za `label`.

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
