---
id: 60f5c3e399ff1a05629964e4
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Kako je navedeno u zaglavlju, kreirate formular. Stoga, nakon elementa `p`, postavite `form` sa atributom komponente `action` koja cilja `https://register-demo.freecodecamp.org`.

# --hints--

Morate dodati element `form` blizu elementa `p`.

```js
assert.exists(document.querySelector('p + form'));
```

Trebalo bi mu dati atribut elementa `form` sa svojstvom `action`.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

Trebalo bi mu dati vrednost `action` od `https://register-demo.freecodecamp.org`.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

Vaša komponenta za `form` mora da ima zatvarajući tag `</form>`.

```js
assert.match(code, /<\/form\>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
