---
id: 60f8604682407e0d017bbf7f
title: Korak 21
challengeType: 0
dashedName: step-21
---

# --description--

U svakom odgovarajućem elementu od `label`, i odmah nakon elementa za `input`, dodajte razmak i dodajte sledeći tekst:

```md
Personal
Business
```

# --hints--

Trebalo bi mu dati prvi tekst od `Personal` za `label`.

```js
assert.include(document.querySelector('fieldset:nth-child(2) > label')?.innerText, 'Personal');
```

Treba da mu daš `label` za drugi tekst od `Business`.

```js
assert.include(document.querySelector('fieldset:nth-child(2) > label:nth-child(2)')?.innerText, 'Business');
```

Prvom elementu `label` treba da dodaš jedan razmak ispred teksta.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) > label')?.innerText?.[0], ' ');
```

Trebalo bi da mu/jaj daje `label` na isto mesto, ispred.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) > label:nth-child(2)')?.innerText?.[0], ' ');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      <label><input type="radio" /></label>
      <label><input type="radio" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
