---
id: 60f803d5241e6a0433a523a1
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

U prethodnim lekcijama, naučili ste kako da radite sa dimenzijama za `rem`. Zapamtite da dimenzija od `rem` predstavlja koren `em`, i da je povezana sa veličinom fonta elementa `html`.

Pošto elementi `label` su po defaultu inline, svi se prikazuju jedan pored drugog na jednoj liniji, što čini njihov tekst teškim za čitanje.

Da bi se prikazalo na različitim redovima, dodajte `display: block` u element `label`, i dodajte atribut `margin` sa vrednošću `0.5rem 0`, kako biste ih razdvojili od jedinstvenog.

# --hints--

Treba ti da koristiš alat `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

Treba da dodaš atribut `display` sa vrednošću `block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

Treba dodati svojstvo `margin` sa vrednošću `0.5rem 0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label>Enter Your First Name:</label>
      <label>Enter Your Last Name:</label>
      <label>Enter Your Email:</label>
      <label>Create a New Password:</label>
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
