---
id: 60f84ec41116b209c280ba91
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Koristeći `type="password"`, možete koristiti svojstvo karakteristike `pattern` da definišete regularni izraz koji lozinka mora pratiti kako bi bila smatrana validnom.

Dodajte svojstvo elementa `pattern` na polje za lozinku `input` da zahteva unos koji zadovoljava: `[a-z0-5]{8,}`

Gornji deo je običan izraz koji odgovara osam ili više malih slova ili brojeva `0` do `5`. Zatim, uklonite atribut elementa `minlength` i pokušajte.

# --hints--

Treba da mu dodeliš atribut elementa za lozinku `input` sa atributom elementa `pattern`.

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

Trebalo bi da mu dodeliš atribut elementu `pattern` vrednost `[a-z0-5]{8,}`.

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

Morate ukloniti atribut elementa `minlength` iz polja za lozinku `input`.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
--fcc-editable-region--
      <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
