---
id: 60f81616cff80508badf9ad5
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Trenutno bi trebalo da budete u mogućnosti za slanje formulara. Međutim, možete primetiti da se ne dešava mnogo toga.

Kako bi se forma učinila interaktivnijom, dodajte atribut elementa ``required`` za elemente ``input`` u prvom ``fieldset``.

Sada, ako pokušate da pošaljete formular bez popunjavanja neophodnih polja, videćete poruku o grešci.

# --hints--

Treba da mu daš prvi element od `input` svojstvo elementa `required`.

```js
assert.isTrue(document.querySelector('input')?.required);
```

Trebalo bi mu dati drugi atribut za `input`, atribut `required`.

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.required);
```

Trebalo bi mu dati treći atribut `input` svojstvo atributa `required`.

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.required);
```

Trebalo bi mu dodeliti četvrti element od `input` svojstvo elementa `required`.

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.required);
```

Trebalo bi dati `input` koji je tipa `submit` svojstvo komponente `required`.

```js
assert.isFalse(document.querySelector('input[type="submit"]')?.required);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
