---
id: 62ff8b9dab5ac88e4d3d43a3
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Prateći najbolje prakse pristupa, spojite elemente od `input` i elemente od `label` koristeći svojstvo elementa `for`.

Koristite `first-name`, `last-name`, `email`, na `new-password` kao vrednosti svojstava relevantnog elementa za `id`.

# --hints--

Prvi element `input` treba da ima `id` od `first-name`.

```js
assert.isTrue(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

Drugi element od `input` mora imati `id` za `last-name`.

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

Treći deo `input` treba da ima `id` od `email`.

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

Četvrti element od `input` treba da ima `id` od `new-password`.

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

Prvi element `label` treba da ima svojstvo elementa `for` sa vrednošću `first-name`.

```js
assert.isTrue(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

Drugi element `label` bi trebalo da ima svojstvo elementa `for` sa vrednošću `last-name`.

```js
assert.isTrue(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

Treće polje `label` treba da ima svojstvo polja `for` sa vrednošću `email`.

```js
assert.isTrue(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

Četvrti element od `label` treba da ima svojstvo elementa `for` sa vrednošću `new-password`.

```js
assert.isTrue(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label>Enter Your First Name: <input /></label>
      <label>Enter Your Last Name: <input /></label>
      <label>Enter Your Email: <input /></label>
      <label>Create a New Password: <input /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
