---
id: 63541ef4f96cd82e8e6c788a
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Atribut elementa `method` opisuje kako poslati podatke forme za URL, definisane u atributu elementa `action`. Podaci forme mogu se poslati putem zahteva `GET` kao parametri za URL (za `method="get"`) ili putem zahteva `POST` kao podaci u telu zahteva (za `method="post"`).

Postavite atribut elementa `method` ili pošalj podatke sa obrasca putem zahteva `POST`.

**NAPOMENA**: Naučićete više o metodama `GET` i `POST` na nadolazećim predavanjima i radionicama.

# --hints--

Niste u mogućnosti da dodate novu funkciju za `form`.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

Vaše polje za `form` mora da ima svojstvo polja `method`.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

Atribut vašeg elementa za `method` treba biti postavljen na `post`.

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
