---
id: 6880995521770716269832cf
title: Korak 23
challengeType: 0
dashedName: step-23
---

# --description--

Radijska dugmad još uvek nisu dobila vrednost. Stoga, POST zahtev poslat putem obrasca neće prikazati koji je izbor odabran.

Da biste napravili potpuno stabilan zahtev, dodajte atribut elementa `value` na sve radio dugmad sa vrednostima `personal` i `business`, odgovarajuće.

# --hints--

Trebalo bi da mu date prvom ulazu za radio atribut elementa `value` sa vrednošću `personal`.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(1) input[type="radio"]')?.value, 'personal');
```

Treba da mu date drugi unos za radio atribut `value` sa vrednošću `business`.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(2) input[type="radio"]')?.value, 'business');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      <label><input type="radio" name="account-type" /> Personal</label>
      <label><input type="radio" name="account-type" /> Business</label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
