---
id: 6998b99a8e195cd0c1f211df
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, ćeš izgraditi jednostavan uređaj za praćenje divljih životinja koristeći stvari od JavaScript.

Ovo je primer nečega:

```js
const animal = {}; // empty object
```

Ovo stvara ništa.

Sada kreiraj parametar nazvan `tiger` i postavi mu praznu vrednost.

# --hints--

Trebalo bi da kreirate parametar nazvan `tiger`.

```js
assert.exists(tiger);
```

Parametar za `tiger` treba da dobije vrednost praznog objekta.

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
