---
id: 6999b0198d18d402c36cc5fe
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

U ovom koraku, kreiraćete funkciju koja izvlači atribute iz objekta.

Ovo je primer za uklanjanje atributa korišćenjem `delete`:

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

U ovom primeru, svojstvo `canFly` je uklonjeno iz objekta `bird`.

Sada kreiraj funkciju nazvanu `removeEndangeredStatus`. Ta funkcija treba da primi jedan parametar nazvan `animal`.

Unutar funkcije, ukloni svojstvo `isEndangered` sa objekta `animal` koristeći ključnu reč `delete`. Vrati ažurirani objekat `animal`.

Nakon kreiranja funkcije, koristite `console.log` da pozovete `removeEndangeredStatus(tiger)` ili ćete videti nešto ažurirano u konzoli.

# --hints--

Trebalo bi da kreirate funkciju nazvanu `removeEndangeredStatus`.

```js
assert.isFunction(removeEndangeredStatus);
```

Funkcija od `removeEndangeredStatus` treba da ima jedan parametar, `animal`.

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Funkcija `removeEndangeredStatus` treba da vrati ažurirani objekat `animal`.

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` treba da koristi glavni termin od `delete` da ukloni atribut/karakteristiku `isEndangered`.

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

Trebalo bi da ukucate `removeEndangeredStatus(tiger)` u konzolu.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

Poziv `removeEndangeredStatus(tiger)` mora da ukloni atribut `isEndangered`.

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` Trebalo bi da koristi parametre funkcije i radi sa bilo kojim objektom.

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
