---
id: 6999a0c1698b93d5d41156ed
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada dodaj još jednu osobinu objektu `tiger` nazvanom `isEndangered`.

Postavi svojstvo `isEndangered` na `true`.

# --hints--

`tiger` trebalo bi da ima atribut `isEndangered`.

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` treba biti jednako `true`.

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
