---
id: 6999b07c34c2cde21ee6e1d8
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Na ovom koraku, kreiraćete funkciju koja proverava da li objekat ima određena svojstva.

Evo primera korišćenja `hasOwnProperty`:

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

U ovom primeru, funkcija `hasColor` koristi `hasOwnProperty` da proveri da li svojstvo `color` postoji na objektu.

Sada kreiraj funkciju nazvanu `hasHabitat`. Ta funkcija mora da prihvati jedan parametar koji se zove `animal`.

Koristi putanju `hasOwnProperty` da vrati `true` ako objekat `animal` ima atribut nazvan `"habitat"` i `false` inače.

Nakon kreiranja funkcije, koristite `console.log` da pozovete `hasHabitat(tiger)` i `hasHabitat(elephant)` kako biste videli oba rezultata u konzoli.

# --hints--

Treba da kreiraš funkciju koja se zove `hasHabitat`.

```js
assert.isFunction(hasHabitat);
```

Funkcija od `hasHabitat` treba da ima jedan parametar, `animal`.

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat` Trebalo bi koristiti način `hasOwnProperty` da proveri svojstvo `habitat`.

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

Treba da napišete `hasHabitat(tiger)` u konzoli.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

Treba da napišete `hasHabitat(elephant)` u konzoli.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

Pozvati `hasHabitat(tiger)` treba da vrati `true` nakon dodavanja staništa.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

Treba da vrati `hasHabitat(elephant)` `false` ako nema naselja.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` trebao bi koristiti parametre funkcije i raditi sa bilo čim.

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
