---
id: 6999b0d79ea938edcb032237
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

Na ovom poslednjem koraku radionice za praćenje divljine, kreiraćete funkciju koja dohvaća svojstvo objekta korišćenjem notacije uglastih zagrada.

Evo primera kako dobiti atribute koristeći notaciju uglastih zagrada:

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

U ovom primeru, notacija sa zagradama omogućava vam da pristupite atributima korišćenjem parametra.

Sada kreiraj funkciju koja se zove `getProperty`.

Funkcija mora da primi dva parametra: `animal` na `propertyName`.

Vrati vrednost svojstva koristeći notaciju uglastih zagrada.

Nakon kreiranja funkcije, koristite `console.log` da pozovete `getProperty(tiger, "species")` sa `getProperty(elephant, "age")` ili vidite vrednosti koje su vraćene u konzolu.

Dakle, radionica za praćenje divljih životinja je završena!

# --hints--

Trebalo bi da kreirate funkciju nazvanu `getProperty`.

```js
assert.isFunction(getProperty);
```

Funkcija `getProperty` mora imati dva parametra: `animal` i `propertyName`.

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty` Trebalo bi da koristi notaciju uglastih zagrada za dobijanje atributa.

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

Trebalo bi da napišete `getProperty(tiger, "species")` u konzolu.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

Treba da napišete `getProperty(elephant, "age")` u konzolu.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

Poziv `getProperty(tiger, "species")` treba da vrati `"Tiger"`.

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

Funkcija `getProperty(elephant, "age")` bi trebalo da vrati `12`.

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty` Trebalo bi da koristi parametre funkcije i radi sa bilo kojim objektom.

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
