---
id: 6999adc72c02812a28c01e31
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Na ovom koraku, ﾄ㌃ﾅ｡ kreirati funkciju koja ispravlja svojstvo objekta.

Ovde je primer za validaciju atributa unutar funkcije:

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

U ovom primeru, svojstvo za `age` je aﾅｾurirano na novu vrednost.

Sada kreiraj funkciju koja se zove `updateAge`. Ta funkcija treba da primi dva parametra: `animal` i `newAge`.

Unutar funkcije, aﾅｾuriraj svojstvo `age` objekta `animal` na `newAge`. Vrati aﾅｾurirani objekat `animal`.

Nakon kreiranja funkcije, koristite `console.log` za pozivanje `updateAge(elephant, 12)` ili vidite da je neﾅ｡to od `elephant` aﾅｾurirano u konzoli.

# --hints--

Treba da kreiraﾅ｡ funkciju koja se zove `updateAge`.

```js
assert.isFunction(updateAge);
```

Funkcija za `updateAge` mora imati dva parametra: `animal` i `newAge`.

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

`updateAge` treba koristiti notaciju taﾄ耕e za modifikovanje svojstva od `age`.

```js
assert.match(code, /animal\.age\s*=\s*newAge/);
```

Funkcija `updateAge` treba da vrati objekat `animal` koji je aﾅｾuriran.

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

Treba da napiﾅ｡ete `updateAge(elephant, 12)` u konzoli.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

Kuitisha `updateAge(elephant, 12)` treba da ispravi atribut starosti.

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` Trebalo bi da koristi parametre funkcije i radi sa bilo ﾄ絞m.

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
