---
id: 69999ff7f7862fdfb885ea5d
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Molimo izmeni svoj deo `tiger` dodavanjem novog atributa pod nazivom `age`.

Postavi svojstvo `age` na vrednost `5`.

# --hints--

`tiger` mora imati svojstvo od `age`.

```js
assert.property(tiger, "age");
```

`tiger.age` mora biti jednako `5`.

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
