---
id: 6999ac171928136168654774
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Na ovom koraku, kreiraćete funkciju za dobijanje svojstva `species` objekta.

Evo malog primera za dobijanje atributa iz objekta:

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

U ovom primeru, funkcija prima nešto kao argument i vraća njegov atribut/svojstvo `species`.

Sada kreiraj funkciju koja se zove `getSpecies`.

Funkcija mora da primi jedan parametar nazvan `animal`.

Vrati atribut `species` od `animal` koristeći tačku.

Nakon kreiranja funkcije, koristite `console.log` da pozovete `getSpecies(tiger)` ili vidite rezultate u konzoli.

# --hints--

Trebalo bi da kreirate funkciju nazvanu `getSpecies`.

```js
assert.isFunction(getSpecies);
```

Funkcija za `getSpecies` mora imati jedan parametar, `animal`.

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Trebalo bi da napišete `getSpecies(tiger)` na konzoli.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies` Trebalo bi da dobije atribut vrste koristeći dot notaciju.

```js
assert.match(code, /animal\.species/);
```

Pozivanje `getSpecies(tiger)` bi trebalo da vrati `"Tiger"`.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies` Trebalo bi da koristi parametar funkcije i radi sa bilo čime.

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
