---
id: 6999faf886634b72be7a819b
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Za ovog praćenja divljine, ažuriraćete nešto sa novim funkcijama.

Evo primera kako kreirati nešto sa svojstvima:

```js
const animal = {
  name: "Lion" // property: value
};
```

Ažurirajte sada nešto za `tiger` ili ono što ima funkciju nazvanu `species`.

Postavi vrednost na `"Tiger"`.

# --hints--

`tiger` treba da bude nešto.

```js
assert.isObject(tiger);
```

`tiger` trebalo bi imati svojstvo od `species`.

```js
assert.property(tiger, "species");
```

`tiger.species` mora biti jednako `"Tiger"`.

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
