---
id: 699a00de9f564bb0effaa14d
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Na ovom koraku, kreiraćete funkciju za dobijanje svojstva `age` objekta.

Evo primera:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

U ovom primeru, funkcija prima nešto kao argument i vraća svoje svojstvo od `age`.

Sada kreiraj funkciju koja se zove `getAge`.

Funkcija mora da primi jedan parametar nazvan `animal`.

Vrati svojstvo `age` od `animal` koristeći notaciju tačke.

Nakon kreiranja funkcije, koristite `console.log` da pozovete `getAge(tiger)` ili vidite rezultate u konzoli.

# --hints--

Trebaće ti da kreiraš funkciju koja se zove `getAge`.

```js
assert.isFunction(getAge);
```

Funkcija za `getAge` treba da ima jedan parametar, `animal`.

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge` trebalo bi da dobije svojstvo `age` korišćenjem dot notacije.

```js
assert.match(code, /animal\.age/);
```

Treba da napišete `getAge(tiger)` u konzoli.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

Pozvati `getAge(tiger)` mora da vrati `5`.

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge` Trebalo bi da koristi funkcionalski parametar i radi sa bilo čime.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
