---
id: 69ba8da6b296391d913096a0
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Sada kreiraj drugi element koji se zove `elephant`.

Dodajte sledeće karakteristike:

- `species` sa vrednošću `"Elephant"`
- `age` sa vrednošću `10`
- `isEndangered` sa vrednošću `true`

# --hints--

Trebalo bi da kreirate parametar nazvan `elephant`.

```js
assert.exists(elephant);
```

`elephant` mora imati svojstvo `species` sa vrednošću jednakom `"Elephant"`.

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant` mora imati svojstvo `age` sa vrednošću jednaku `10`.

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant` treba da ima svojstvo od `isEndangered` sa vrednošću jednaku `true`.

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
