---
id: 594810f028c0303b75339ad8
title: Zigzag matrica
challengeType: 1
forumTopicId: 302348
dashedName: zig-zag-matrix
---

# --description--

Zig-zag matrica podataka je raspored kvadrata prvih prirodnih brojeva $N^2$, gde se brojevi povećavaju sekvencijalno kada idete u zig-zag pravcu duž anti-diyagonala podatkovne matrice.

Na primer, za unos `5`, sledeći rezultati bi trebalo da se generišu:

<pre>
 0  1  5  6 14
 2  4  7 13 15
 3  8 12 16 21
 9 11 17 20 22
10 18 19 23 24
</pre>

# --instructions--

Implementirajte funkciju koja prima veličinu zizzag matrice i vraća odgovarajuću matricu kao dvodimenzionalni niz podataka.

# --before-each--

```js
const zm1 = [[0]];
const zm2 = [[0, 1], [2, 3]];
const zm5 = [
  [0, 1, 5, 6, 14],
  [2, 4, 7, 13, 15],
  [3, 8, 12, 16, 21],
  [9, 11, 17, 20, 22],
  [10, 18, 19, 23, 24]
];
```

# --hints--

ZigZagMatrix treba da bude funkcija.

```js
assert.equal(typeof ZigZagMatrix, 'function');
```

ZigZagMatrix treba da vrati red podataka.

```js
assert.equal(typeof ZigZagMatrix(1), 'object');
```

ZigZagMatrix treba da vrati red podataka interni organizovane redove podataka.

```js
assert.equal(typeof ZigZagMatrix(1)[0], 'object');
```

ZigZagMatrix(1) trebalo bi da vrati \[[0]].

```js
assert.deepEqual(ZigZagMatrix(1), zm1);
```

ZigZagMatrix(2) trebalo bi da vrati \[[0, 1], [2, 3]].

```js
assert.deepEqual(ZigZagMatrix(2), zm2);
```

ZigZagMatrix(5) treba da vrati spomenutu matricu.

```js
assert.deepEqual(ZigZagMatrix(5), zm5);
```

# --seed--

## --seed-contents--

```js
function ZigZagMatrix(n) {

  return [[], []];
}
```

# --solutions--

```js
function ZigZagMatrix(n) {
  const mtx = [];
  for (let i = 0; i < n; i++) {
    mtx[i] = [];
  }

  let i = 1;
  let j = 1;
  for (let e = 0; e < n * n; e++) {
    mtx[i - 1][j - 1] = e;
    if ((i + j) % 2 === 0) {
      // Even stripes
      if (j < n) j++;
      else i += 2;
      if (i > 1) i--;
    } else {
      // Odd stripes
      if (i < n) i++;
      else j += 2;
      if (j > 1) j--;
    }
  }
  return mtx;
}
```
