---
id: 594810f028c0303b75339ad2
title: Vektorski križni proizvod
challengeType: 1
forumTopicId: 302342
dashedName: vector-cross-product
---

# --description--

Vektor se definiše da ima tri dimenzije, predstavljene kao uporani skup od tri broja: (X, Y, Z).

# --instructions--

Napišite funkciju koja prima dva vektora (redove podataka) kao ulaz i izračunava njihov kros-proizvod. Vaša funkcija bi trebalo da vrati `null` za nevažeći unos, kao što su vektori različitih dužina.

# --hints--

`crossProduct` Trebalo bi da bude funkcija.

```js
assert.equal(typeof crossProduct, 'function');
```

`crossProduct()` trebalo bi da vrati null.

```js
assert.equal(crossProduct(), null);
```

`crossProduct([1, 2, 3], [4, 5, 6])` trebalo bi da vrati `[-3, 6, -3]`.

```js
assert.deepEqual(res12, exp12);
```

# --before-each--

```js
const tv1 = [1, 2, 3];
const tv2 = [4, 5, 6];
const res12 = crossProduct(tv1, tv2);
const exp12 = [-3, 6, -3];
```

# --seed--

## --seed-contents--

```js
function crossProduct(a, b) {

}
```

# --solutions--

```js
function crossProduct(a, b) {
  if (!a || !b) {
    return null;
  }

  // Check lengths
  if (a.length !== 3 || b.length !== 3) {
    return null;
  }

  return [
    (a[1] * b[2]) - (a[2] * b[1]),
    (a[2] * b[0]) - (a[0] * b[2]),
    (a[0] * b[1]) - (a[1] * b[0])
  ];
}
```
