---
id: 594810f028c0303b75339ad4
title: Završiti reči
challengeType: 1
forumTopicId: 302344
dashedName: word-wrap
---

# --description--

Čak i danas, sa proporcionalnim fontovima i složenim podešavanjima, postoje situacije kada je potrebno obaviti prelamanje teksta u specifičnoj koloni. Osnovna vežba je prelama teksta paragrafa na jednostavan način.

# --instructions--

Napišite funkciju koja može da preformira ovaj tekst za bilo koji broj karaktera. Zapamtite da ulazni tekst već sadrži razmake linija, koje vaša funkcija mora adekvatno obraditi. Na primer, tekst preformiran za 80 karaktera treba da izgleda ovako:

<pre>
Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX
algorithm. If your language provides this, you get easy extra credit, but you
must reference documentation indicating that the algorithm is something better
than a simple minimum length algorithm.
</pre>

# --before-each--

```js
const text =
`Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX algorithm.
If your language provides this, you get easy extra credit,
but you ''must reference documentation'' indicating that the algorithm
is something better than a simple minimum length algorithm.`;

const wrapped80 = wrap(text, 80);
const wrapped42 = wrap(text, 42);

const firstRow80 =
    'Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX';

const firstRow42 = 'Wrap text using a more sophisticated';
```

# --hints--

`wrap` to bi trebalo da bude funkcija.

```js
assert.equal(typeof wrap, 'function');
```

`wrap` trebalo bi da vrati sekvencu slova.

```js
assert.equal(typeof wrap('abc', 10), 'string');
```

`wrap(text,80)` Trebalo bi da vrati 4 linije.

```js
assert(wrapped80.split('\n').length === 4);
```

Tvoja funkcija `wrap` treba vratiti očekivani tekst.

```js
assert.equal(wrapped80.split('\n')[0], firstRow80);
```

`wrap(text,42)` inapaswa kurudisha mistari 7.

```js
assert(wrapped42.split('\n').length === 7);
```

Vaša funkcija `wrap` treba da vrati očekivani tekst.

```js
assert.equal(wrapped42.split('\n')[0], firstRow42);
```

# --seed--

## --seed-contents--

```js
function wrap(text, limit) {
  return text;
}
```

# --solutions--

```js
function wrap(text, limit) {
  const noNewlines = text.replace('\n', '');
  if (noNewlines.length > limit) {
    // find the last space within limit
    const edge = noNewlines.slice(0, limit).lastIndexOf(' ');
    if (edge > 0) {
      const line = noNewlines.slice(0, edge);
      const remainder = noNewlines.slice(edge + 1);
      return line + '\n' + wrap(remainder, limit);
    }
  }
  return text;
}
```
