---
id: 594810f028c0303b75339ad5
title: Mešavina od Y
challengeType: 1
forumTopicId: 302345
dashedName: y-combinator
---

# --description--

U programu sa snažnim performansama <a href="https://www.freecodecamp.org/news/the-principles-of-functional-programming/" target="_blank" rel="noopener noreferrer nofollow"> na lambda računanju </a>, funkcije (lambda izrazi) nemaju stanje i dozvoljeno im je da se referenciraju samo argumenti koji su definisani. Ovo se razlikuje od uobičajenog definisanja rekurentne funkcije gde je funkciji vezano stanje parametra, a to stanje se koristi unutar tela funkcije.

Kompozicija Y je bezstanična funkcija koja, kada se koristi sa drugom bezstaničnom funkcijom, vraća ponovljenu verziju te funkcije. Kompozicija Y je jednostavnija od klase funkcija kao što su ove, nazvane kompozicijom fiksnih tačaka.

# --instructions--

Definišite funkciju koja nema stanje kombinacije Y i koristite je za izračunavanje faktorijala. Funkcija `factorial(N)` je već dostupna.

# --hints--

Y bi trebalo da vrati funkciju.

```js
assert.equal(typeof Y((f) => (n) => n), 'function');
```

factorial(1) treba da vrati 1.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(1), 1);
```

factorial(2) trebalo bi da vrati 2.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(2), 2);
```

factorial(3) trebalo bi da vrati 6.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(3), 6);
```

factorial(4) trebalo bi da vrati 24.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(4), 24);
```

factorial(10) treba da vrati 3628800.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(10), 3628800);
```

# --seed--

## --seed-contents--

```js
function Y(f) {
  return function() {

  };
}

var factorial = Y(function(f) {
  return function (n) {
    return n > 1 ? n * f(n - 1) : 1;
  };
});
```

# --solutions--

```js
var Y = f => (x => x(x))(y => f(x => y(y)(x)));
```
