---
id: 594810f028c0303b75339ad7
title: Algoritam za smanjenje Zhang-Suena
challengeType: 1
forumTopicId: 302347
dashedName: zhang-suen-thinning-algorithm
---

# --description--

Ovo je algoritam koji se koristi za redukciju crno-belih slika, to jest, jedan bit po pikselu. Na primer, za ulaznu sliku:

```js
const testImage1 = [
 '                               ',
 '#########       ########       ',
 '###   ####     ####  ####      ',
 '###    ###     ###    ###      ',
 '###   ####     ###             ',
 '#########      ###             ',
 '### ####       ###    ###      ',
 '###  ####  ### ####  #### ###  ',
 '###   #### ###  ########  ###  ',
 '                               '
];
```

Dostavite smančenu sliku kao što sledi:

```js
[ '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               ' ];
```

## Algoriti

Pretpostavimo da su pikseli crne boje jediničke, a pikseli bele boje nule, i da je ulazna slika pravougaoni niz N x M od jedinica i nula. Algoritam radi za sve crne piksele P1 koje imaju osam suseda. Ti susedi su raspoređeni redosledom kao što sledi:

$$\begin{array}{|c|c|c|}
  \\hline
  P9 & P2              & P3\\\\ \\hline
  P8 & \boldsymbol{P1} & P4\\\\ \\hline
  P7 & P6              & P5\\\\ \\hline
\end{array}$$

Nije očigledno da pikseli na granici slike ne mogu imati osam punih suseda.

- Izračunaj $A(P1)$ = broj prelaza sa bele na crnu boju, ($0 \to 1$), u nizu P2, P3, P4, P5, P6, P7, P8, P9, P2. (Obrati pažnju na dodatni P2 na kraju – niz je cikličan.)
- Tafsiri $B(P1)$ = idadi ya majirani wa pikseli nyeusi wa P1. ($= \\sum(P2 \ldots P9)$)

**Hatua ya 1:**

Svi pikseli se ispituju, a oni koji istovremeno ispunjavaju sve sledeće uslove označavaju se samo u ovom koraku.

1. Pikseli ni nyeusi na ina majirani nane
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. Angalau moja kati ya $P2$, $P4$ na $P6$ ni nyeupe
5. Angalau moja kati ya $P4$, $P6$ na $P8$ ni nyeupe

Nakon prolaska kroz sliku i prikupljanja svih piksela koji ispunjavaju uslove iz prvog koraka, svi ti pikseli se menjaju u bele.

**Korak 2:**

Svi pikseli se ponovo testiraju, a oni koji ispunjavaju sve sledeće uslove identifikuju se samo u ovoj fazi.

1. Pikseli su crne i imaju osam suseda
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. Najmanje jedan od P2, P4 i P8 je beli
5. Najmanje jedan od P2, P6 i P8 je beli

Nakon obrade slike i prikupljanja svih piksela koji zadovoljavaju sve uslove koraka 2, svi pikseli koji zadovoljavaju ove uslove vraćaju se u belu boju.

**Ponavljaj:**

Ako je bilo koji piksel promenjen u ovom ciklusu faze 1 ili faze 2, onda se sve faze ponavljaju dok nijedan piksel slike nije promenjen.

# --instructions--

Napišite proceduru za implementaciju redukcije Zhang-Suena na `image`, dati niz sekvenci karaktera, gde svaki niz karaktera predstavlja jedan red slike. U nizu karaktera, `#` predstavlja crnu pikselu, a prazan prostor predstavlja belu pikselu. Funkcija treba da vrati smanjenu sliku, koristeći istu reprezentaciju.

# --before-each--

```js
const _testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
const expected1 = [
  '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               '
];
const _testImage2 = [
  '                                                          ',
  ' #################                   #############        ',
  ' ##################               ################        ',
  ' ###################            ##################        ',
  ' ########     #######          ###################        ',
  '   ######     #######         #######       ######        ',
  '   ######     #######        #######                      ',
  '   #################         #######                      ',
  '   ################          #######                      ',
  '   #################         #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######         #######       ######        ',
  ' ########     #######          ###################        ',
  ' ########     ####### ######    ################## ###### ',
  ' ########     ####### ######      ################ ###### ',
  ' ########     ####### ######         ############# ###### ',
  '                                                          '
];
const expected2 = [
  '                                                          ',
  '                                                          ',
  '    # ##########                       #######            ',
  '     ##        #                   ####       #           ',
  '     #          #                 ##                      ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     ############               #                         ',
  '     #          #               #                         ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #                            ##                      ',
  '     #                             ############           ',
  '                       ###                          ###   ',
  '                                                          ',
  '                                                          '
];
```

# --hints--

`thinImage` trebalo bi da bude funkcija.

```js
assert.equal(typeof thinImage, 'function');
```

`thinImage` Trebalo bi da vrati red podataka.

```js
assert(Array.isArray(thinImage(_testImage1)));
```

`thinImage` Trebalo bi da vrati red/liniju sekvence karaktera.

```js
assert.equal(typeof thinImage(_testImage1)[0], 'string');
```

`thinImage(testImage1)` Trebalo bi da vrati smanjenu sliku kao u primeru.

```js
assert.deepEqual(thinImage(_testImage1), expected1);
```

`thinImage(testImage2)` Trebalo bi da vrati smanjenu sliku.

```js
assert.deepEqual(thinImage(_testImage2), expected2);
```

# --seed--

## --seed-contents--

```js
function thinImage(image) {

}

const testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
```

# --solutions--

```js
function Point(x, y) {
  this.x = x;
  this.y = y;
}

const ZhangSuen = (function () {
  function ZhangSuen() {
  }

  ZhangSuen.nbrs = [[0, -1], [1, -1], [1, 0], [1, 1], [0, 1], [-1, 1], [-1, 0], [-1, -1], [0, -1]];

  ZhangSuen.nbrGroups = [[[0, 2, 4], [2, 4, 6]], [[0, 2, 6], [0, 4, 6]]];

  ZhangSuen.toWhite = [];

  ZhangSuen.main = function (image) {
    ZhangSuen.grid = new Array(image);
    for (let r = 0; r < image.length; r++) {
      ZhangSuen.grid[r] = image[r].split('');
    }
    ZhangSuen.thinImage();
    return ZhangSuen.getResult();
  };

  ZhangSuen.thinImage = function () {
    let firstStep = false;
    let hasChanged;
    do {
      hasChanged = false;
      firstStep = !firstStep;
      for (let r = 1; r < ZhangSuen.grid.length - 1; r++) {
        for (let c = 1; c < ZhangSuen.grid[0].length - 1; c++) {
          if (ZhangSuen.grid[r][c] !== '#') {
            continue;
          }
          const nn = ZhangSuen.numNeighbors(r, c);
          if (nn < 2 || nn > 6) {
            continue;
          }
          if (ZhangSuen.numTransitions(r, c) !== 1) {
            continue;
          }
          if (!ZhangSuen.atLeastOneIsWhite(r, c, firstStep ? 0 : 1)) {
            continue;
          }
          ZhangSuen.toWhite.push(new Point(c, r));
          hasChanged = true;
        }
      }
      for (let i = 0; i < ZhangSuen.toWhite.length; i++) {
        const p = ZhangSuen.toWhite[i];
        ZhangSuen.grid[p.y][p.x] = ' ';
      }
      ZhangSuen.toWhite = [];
    } while ((firstStep || hasChanged));
  };

  ZhangSuen.numNeighbors = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === '#') {
        count++;
      }
    }
    return count;
  };

  ZhangSuen.numTransitions = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === ' ') {
        if (ZhangSuen.grid[r + ZhangSuen.nbrs[i + 1][1]][c + ZhangSuen.nbrs[i + 1][0]] === '#') {
          count++;
        }
      }
    }
    return count;
  };

  ZhangSuen.atLeastOneIsWhite = function (r, c, step) {
    let count = 0;
    const group = ZhangSuen.nbrGroups[step];
    for (let i = 0; i < 2; i++) {
      for (let j = 0; j < group[i].length; j++) {
        const nbr = ZhangSuen.nbrs[group[i][j]];
        if (ZhangSuen.grid[r + nbr[1]][c + nbr[0]] === ' ') {
          count++;
          break;
        }
      }
    }
    return count > 1;
  };

  ZhangSuen.getResult = function () {
    const result = [];
    for (let i = 0; i < ZhangSuen.grid.length; i++) {
      const row = ZhangSuen.grid[i].join('');
      result.push(row);
    }
    return result;
  };
  return ZhangSuen;
}());

function thinImage(image) {
  return ZhangSuen.main(image);
}
```
