---
id: 594810f028c0303b75339ace
title: Fabrika za sakupljanje
challengeType: 1
forumTopicId: 302222
dashedName: accumulator-factory
---

# --description--

Problem koji je predstavio Paul Graham je kreiranje funkcije koja prima jedan argument (broj) i vraća drugu funkciju koja je akumulator. Vraćena funkcija akumulator takođe prima jedan numerički argument i vraća sumu svih numeričkih vrednosti koje su do sada prosleđene tom akumulatoru (uključujući početnu vrednost koja je data kada je akumulator inicijalizovan).

# --instructions--

Kreiraj funkciju koja prima broj $n$ i kreira akumulatorske funkcije koje vraćaju zbir svakog broja koji je prošao kroz njih.

**Zakon:**

Ne koristite međunarodne parametre.

**Napomena:**

Ambalaža održava spoljašnje stanje.

# --hints--

`accumulator` Trebalo bi da bude funkcija.

```js
assert(typeof accumulator === 'function');
```

`accumulator(0)` trebalo bi da vrati funkciju.

```js
assert(typeof accumulator(0) === 'function');
```

`accumulator(0)(2)` trebalo bi da vrati broj.

```js
assert(typeof accumulator(0)(2) === 'number');
```

Kada se proslede vrednosti 3, -4, 1.5 i 5, trebalo bi da se vrati 5.5.

```js
assert(testFn && testFn(5) === 5.5);
```

# --before-each--

```js
const testFn = typeof accumulator === 'function' && accumulator(3);
if (testFn) {
  testFn(-4);
  testFn(1.5);
}
```

# --seed--

## --seed-contents--

```js
function accumulator(sum) {

}
```

# --solutions--

```js
function accumulator(sum) {
  return function(n) {
    return sum += n;
  };
}
```
