---
id: 594810f028c0303b75339acc
title: Problem sa ABC
challengeType: 1
forumTopicId: 302220
dashedName: abc-problem
---

# --description--

Dostupna je kolekcija ABC dugmadi (primer, pismenica sa slova za decu). Postoji 20 dugmadi, svako sa dva slova. Kompletno pismo je osigurano na svim stranama dugmadića. Primer kolekcije dugmadi:

<pre>(B O)
(X K)
(D Q)
(C P)
(N A)
(G T)
(R E)
(T G)
(Q D)
(F S)
(J W)
(H U)
(V I)
(A N)
(O B)
(E R)
(F S)
(L Y)
(P C)
(Z M)
</pre>

# --instructions--

Implementirati funkciju koja prima sekvencu slova (reč) i određuje da li se reč može napisati koristeći dati skup tastera.

Smjernice za uzimanje u obzir:

<ul>
  <li>Nakon upotrebe karaktera na dugmetu, to dugme više ne može biti korišćeno.</li>
  <li>Funkcija ne bi trebalo da zavisi od velikih ili malih slova.</li>
</ul>

# --before-each--

```js
const words = ['bark', 'BooK', 'TReAT', 'COMMON', 'squAD', 'conFUSE'];
```

# --hints--

`canMakeWord` Trebalo bi da bude funkcija.

```js
assert(typeof canMakeWord === 'function');
```

`canMakeWord` Trebalo bi da vrati vrednost istina ili laž.

```js
assert(typeof canMakeWord('hi') === 'boolean');
```

`canMakeWord("bark")` treba da vrati true.

```js
assert(canMakeWord(words[0]));
```

`canMakeWord("BooK")` trebalo bi da vrati false.

```js
assert(!canMakeWord(words[1]));
```

`canMakeWord("TReAT")` trebalo bi da vrati true.

```js
assert(canMakeWord(words[2]));
```

`canMakeWord("COMMON")` treba da vrati false.

```js
assert(!canMakeWord(words[3]));
```

`canMakeWord("squAD")` treba da vrati true.

```js
assert(canMakeWord(words[4]));
```

`canMakeWord("conFUSE")` trebalo bi vratiti true.

```js
assert(canMakeWord(words[5]));
```

# --seed--

## --seed-contents--

```js
function canMakeWord(word) {

}
```

# --solutions--

```js
function canMakeWord(word) {
  const characters = 'BO XK DQ CP NA GT RE TG QD FS JW HU VI AN OB ER FS LY PC ZM';
  const blocks = characters.split(' ').map(pair => pair.split(''));

  const letters = [...word.toUpperCase()];
  let length = letters.length;
  const copy = new Set(blocks);

  letters.forEach(letter => {
    for (let block of copy) {
      const index = block.indexOf(letter);

      if (index !== -1) {
        length--;
        copy.delete(block);
        break;
      }
    }
  });
  return !length;
}
```
