---
id: 594810f028c0303b75339acf
title: Ackermannova funkcija
challengeType: 1
forumTopicId: 302223
dashedName: ackermann-function
---

# --description--

Funkcija Ackermanna je uobičajak primer rekurzivne funkcije, poznata posebno jer nije osnovna rekurzivna funkcija. Raste veoma brzo po vrednosti, slično veličini svog stabla poziva.

Funkcija Ackermana se obično definiše kao što sledi:

$A(m, n) = \\begin{cases} n+1 & \\mbox{if } m = 0 \\\\ A(m-1, 1) & \\mbox{if } m > 0 \\mbox{ and } n = 0 \\\\ A(m-1, A(m, n-1)) & \\mbox{if } m > 0 \\mbox{ and } n > 0. \\end{cases}$

Njevi argumenti nikad nisu bili negativni i uvek završavaju.

# --instructions--

Napišite funkciju koja vraća vrednost $A(m, n)$. Preporučuje se opcionalna preciznost (jer funkcija raste veoma brzo), ali nije obavezno.

# --hints--

`ack` Trebalo bi da bude funkcija.

```js
assert(typeof ack === 'function');
```

`ack(0, 0)` treba da vrati 1.

```js
assert(ack(0, 0) === 1);
```

`ack(1, 1)` trebalo bi da vrati 3.

```js
assert(ack(1, 1) === 3);
```

`ack(2, 5)` inapaswa kurudisha 13.

```js
assert(ack(2, 5) === 13);
```

`ack(3, 3)` trebalo bi da vrati 61.

```js
assert(ack(3, 3) === 61);
```

# --seed--

## --seed-contents--

```js
function ack(m, n) {

}
```

# --solutions--

```js
function ack(m, n) {
  return m === 0 ? n + 1 : ack(m - 1, n === 0 ? 1 : ack(m, n - 1));
}
```
