---
id: 594810f028c0303b75339ad6
title: Reprezentacija Zeckendorfovih brojeva
challengeType: 1
forumTopicId: 302346
dashedName: zeckendorf-number-representation
---

# --description--

Pošto se brojevi mogu predstaviti u pozicionom sistemu kao zbir stepeni desetice (decimalni) ili dvojke (binarni); svi pozitivni ceo brojevi mogu se predstaviti kao zbir jedinstvenih članova Fibonaccijevog niza. Napomenuti da su prvih šest različitih Fibonaccijevih brojeva: `1, 2, 3, 5, 8, 13`.

Decimalni broj jedanaest može se napisati kao `0*13 + 1*8 + 0*5 + 1*3 + 0*2 + 0*1` ili `010100` u pozicionom zapisu gde cifre predstavljaju množenje sa određenim članom sekvence. Predvodne nule se uklanjaju da bi 11 u decimalnom sistemu bilo `10100`. Međutim, 10100 nije jedini način da se stvori 11 iz Fibonaccijevih brojeva, `0*13 + 1*8 + 0*5 + 0*3 + 1*2 + 1*1` ili i 010011 bi takođe predstavljalo 11 u decimalnom sistemu. Za stvarni Zeckendorfov broj postoji dodatno ograničenje da *ne mogu se koristiti dve uzastopne Fibonaccijevih broja*, što dovodi do jedinstvenog početnog rešenja.

# --instructions--

Napiši funkciju koja generiše i vraća reprezentaciju Zeckendorfovih brojeva za `n`.

# --hints--

`zeckendorf` Trebalo bi da bude funkcija.

```js
assert.equal(typeof zeckendorf, 'function');
```

`zeckendorf(0)` trebalo bi da vrati `0`.

```js
assert.equal(zeckendorf(0), 0);

```

`zeckendorf(1)` treba da vrati `1`.

```js
assert.equal(zeckendorf(1), 1);
```

`zeckendorf(2)` treba da vrati `10`.

```js
assert.equal(zeckendorf(2), 10);
```

`zeckendorf(3)` trebalo bi da vrati `100`.

```js
assert.equal(zeckendorf(3), 100);
```

`zeckendorf(4)` trebalo bi da vrati `101`.

```js
assert.equal(zeckendorf(4), 101);
```

`zeckendorf(5)` trebalo bi da vrati `1000`.

```js
assert.equal(zeckendorf(5), 1000);
```

`zeckendorf(6)` trebalo bi da vrati `1001`.

```js
assert.equal(zeckendorf(6), 1001);
```

`zeckendorf(7)` inapaswa kurudisha `1010`.

```js
assert.equal(zeckendorf(7), 1010);
```

`zeckendorf(8)` treba da vrati `10000`.

```js
assert.equal(zeckendorf(8), 10000);
```

`zeckendorf(9)` trebalo bi da vrati `10001`.

```js
assert.equal(zeckendorf(9), 10001);
```

`zeckendorf(10)` mora da vrati `10010`.

```js
assert.equal(zeckendorf(10), 10010);
```

`zeckendorf(11)` trebalo bi da vrati `10100`.

```js
assert.equal(zeckendorf(11), 10100);
```

`zeckendorf(12)` treba da vrati `10101`.

```js
assert.equal(zeckendorf(12), 10101);
```

`zeckendorf(13)` trebalo bi da vrati `100000`.

```js
assert.equal(zeckendorf(13), 100000);
```

`zeckendorf(14)` treba da vrati `100001`.

```js
assert.equal(zeckendorf(14), 100001);
```

`zeckendorf(15)` treba da vrati `100010`.

```js
assert.equal(zeckendorf(15), 100010);
```

`zeckendorf(16)` treba da vrati `100100`.

```js
assert.equal(zeckendorf(16), 100100);
```

`zeckendorf(17)` bi trebalo da vrati `100101`.

```js
assert.equal(zeckendorf(17), 100101);
```

`zeckendorf(18)` treba da vrati `101000`.

```js
assert.equal(zeckendorf(18), 101000);
```

`zeckendorf(19)` treba da vrati `101001`.

```js
assert.equal(zeckendorf(19), 101001);
```

`zeckendorf(20)` treba da vrati `101010`.

```js
assert.equal(zeckendorf(20), 101010);
```

# --seed--

## --seed-contents--

```js
function zeckendorf(n) {

}
```

# --solutions--

```js
// zeckendorf :: Int -> Int
function zeckendorf(n) {
  const f = (m, x) => (m < x ? [m, 0] : [m - x, 1]);
  return parseInt((n === 0 ? ([0]) :
    mapAccumL(f, n, reverse(
      tail(fibUntil(n))
    ))[1]).join(''));
}

// fibUntil :: Int -> [Int]
let fibUntil = n => {
  const xs = [];
  until(
      ([a]) => a > n,
      ([a, b]) => (xs.push(a), [b, a + b]), [1, 1]
  );
  return xs;
};

let mapAccumL = (f, acc, xs) => (
  xs.reduce((a, x) => {
    const pair = f(a[0], x);

    return [pair[0], a[1].concat(pair[1])];
  }, [acc, []])
);

let until = (p, f, x) => {
  let v = x;
  while (!p(v)) v = f(v);
  return v;
};

const tail = xs => (
   xs.length ? xs.slice(1) : undefined
);

const reverse = xs => xs.slice(0).reverse();
```
