---
id: 594810f028c0303b75339acb
title: Milango 100
challengeType: 1
forumTopicId: 302217
dashedName: 100-doors
---

# --description--

Postoje 100 vrata u redu, a sva su prvobitno zatvorena. Prolaze pored tih vrata 100 puta. Prvi put prolazite pored svakog vrata i 'menjate' stanje vrata (ako je vrata zatvorena, otvorite ih; ako su otvorena, zatvorite ih). Drugi put, posetite samo sva vrata broja 2 (tj. vrata broj 2, 4, 6, ...) i promenite njihovo stanje. Treći put, posetite sva vrata broja 3 (tj. vrata broj 3, 6, 9, ...), itd., sve dok ne posetite samo vrata broj 100.

# --instructions--

Izvršite funkciju da utvrdite status vrata nakon poslednjeg prolaska. Vratite najnovije rezultate u red podataka, postavljajući samo broj vrata ako su ta vrata otvorena.

# --hints--

`getFinalOpenedDoors` Trebalo bi da bude funkcija.

```js
assert(typeof getFinalOpenedDoors === 'function');
```

`getFinalOpenedDoors` trebalo bi da vrati red podataka.

```js
assert(Array.isArray(getFinalOpenedDoors(100)));
```

`getFinalOpenedDoors` Trebalo bi da proizvede tačne rezultate.

```js
const solution = [1, 4, 9, 16, 25, 36, 49, 64, 81, 100];
assert.deepEqual(getFinalOpenedDoors(100), solution);
```

# --seed--

## --seed-contents--

```js
function getFinalOpenedDoors(numDoors) {

}
```

# --solutions--

```js
function getFinalOpenedDoors(numDoors) {
  // this is the final pattern (always squares).
  // thus, the most efficient solution simply returns an array of squares up to numDoors).
  const finalState = [];
  let i = 1;
  while (Math.pow(i, 2) <= numDoors) {
    finalState.push(Math.pow(i, 2));
    i++;
  }
  return finalState;
}
```
