---
id: 5949b579404977fbaefcd736
title: broj milijardi 9 za Boga broj kompletan
challengeType: 1
forumTopicId: 302219
dashedName: 9-billion-names-of-god-the-integer
---

# --description--

Ovo je verzija kratke priče Arthura C. Klarka.

(Izvršitelji treba da budu svesni rezultata završetka ove vežbe.)

Detaljno, definišući značenje "ime":

<ul>
  <li>Puni broj 1 ima ime 1 "1".</li>
  <li>Puni broj 2 ima imena 2 "1+1" i "2".</li>
  <li>Puni broj 3 ima imena 3 "1+1+1", "2+1" i "3".</li>
  <li>Puni broj 4 ima imena 5 "1+1+1+1", "2+1+1", "2+2", "3+1", "4".</li>
  <li>Puni broj 5 ima imena 7 "1+1+1+1+1", "2+1+1+1", "2+2+1", "3+1+1", "3+2", "4+1", "5".</li>
</ul>

Ovo se može prikazati sledećim formatom:

<pre>          1
        1   1
      1   1   1
    1   2   1   1
  1   2   2   1   1
1   3   3   2   1   1
</pre>

Gde red $n$ odgovara celobrojnom broju $n$, i za svaku poziciju (red $m$, kolona $C$) od levo ka desno, ona odgovara broju imena koja počinju sa $C$.

Treba znati da zbir serije $n$-th $P(n)$ je funkcija celobrojnog deljenja.

# --instructions--

Funkcija koja vraća zbir $n$-tog niza.

# --hints--

`numberOfNames` Trebalo bi da bude funkcija.

```js
assert(typeof numberOfNames === 'function');
```

`numberOfNames(5)` trebalo bi da bude jednako 7.

```js
assert.equal(numberOfNames(5), 7);
```

`numberOfNames(12)` treba biti jednako 77.

```js
assert.equal(numberOfNames(12), 77);
```

`numberOfNames(18)` trebalo bi da bude jednako 385.

```js
assert.equal(numberOfNames(18), 385);
```

`numberOfNames(23)` trebalo bi da bude jednako 1255.

```js
assert.equal(numberOfNames(23), 1255);
```

`numberOfNames(42)` treba da bude jednako 53174.

```js
assert.equal(numberOfNames(42), 53174);
```

`numberOfNames(123)` treba da bude jednako 2552338241.

```js
assert.equal(numberOfNames(123), 2552338241);
```

# --seed--

## --seed-contents--

```js
function numberOfNames(num) {

  return true;
}
```

# --solutions--

```js
function numberOfNames(num) {
  const cache = [
    [1]
  ];
  for (let l = cache.length; l < num + 1; l++) {
    let Aa;
    let Mi;
    const r = [0];
    for (let x = 1; x < l + 1; x++) {
      r.push(r[r.length - 1] + (Aa = cache[l - x < 0 ? cache.length - (l - x) : l - x])[(Mi = Math.min(x, l - x)) < 0 ? Aa.length - Mi : Mi]);
    }
    cache.push(r);
  }
  return cache[num][cache[num].length - 1];
}
```
