---
id: 594dc6c729e5700999302b45
title: Ravnomerni razmaci
challengeType: 1
forumTopicId: 302230
dashedName: balanced-brackets
---

# --description--

Utvrdi da li je dati niz zagrada ispravan, odnosno da li sadrži sve odgovarajuće parove otvorenih i zatvorenih zagrada u ispravnom redosledu, bez pogrešnog ugnježđivanja.

**Primjeri:**
| Ulaz | Izlaz |
| ------------------------- | -------- |
| <code>\[]</code>          | true     |
| <code>]\[</code>          | false    |
| <code>[][]</code>         | true     |
| <code>]\[]</code>         | false    |
| <code>\[]]\[\[]</code>    | false    |
| <code>\[\[\[\[]]]]</code> | true     |

# --before-each--

```js
const testCases = [
  '[]',
  ']][[[][][][]][',
  '[][[[[][][[[]]]]]]',
  '][',
  '[[[]]]][[]',
  '][[]',
  '][[][]][[[]]',
  '[[][]]][',
  '[[[]]][[]]]][][[',
  '[]][[]]][[[[][]]',
  '][]][[][',
  '[[]][[][]]',
  '[[]]',
  ']][]][[]][[[',
  '][]][][[',
  '][][',
  '[]]]',
  ''
];
```

# --hints--

`isBalanced` Trebalo bi da bude funkcija.

```js
assert(typeof isBalanced === 'function');
```

`isBalanced("[]")` treba da vrati true.

```js
assert(isBalanced(testCases[0]));
```

`isBalanced("]][[[][][][]][")` trebalo bi da vrati false.

```js
assert(!isBalanced(testCases[1]));
```

`isBalanced("[][[[[][][[[]]]]]]")` treba da vrati true.

```js
assert(isBalanced(testCases[2]));
```

`isBalanced("][")` treba da vrati false.

```js
assert(!isBalanced(testCases[3]));
```

`isBalanced("[[[]]]][[]")` trebalo bi vratiti false.

```js
assert(!isBalanced(testCases[4]));
```

`isBalanced("][[]")` trebalo bi da vrati false.

```js
assert(!isBalanced(testCases[5]));
```

`isBalanced("][[][]][[[]]")` treba da vrati false.

```js
assert(!isBalanced(testCases[6]));
```

`isBalanced("[[][]]][")` treba da vrati false.

```js
assert(!isBalanced(testCases[7]));
```

`isBalanced("[[[]]][[]]]][][[")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[8]));
```

`isBalanced("[]][[]]][[[[][]]")` inapaswa kurudisha false.

```js
assert(!isBalanced(testCases[9]));
```

`isBalanced("][]][[][")` treba da vrati false.

```js
assert(!isBalanced(testCases[10]));
```

`isBalanced("[[]][[][]]")` trebao bi da vrati true.

```js
assert(isBalanced(testCases[11]));
```

`isBalanced("[[]]")` treba da vrati true.

```js
assert(isBalanced(testCases[12]));
```

`isBalanced("]][]][[]][[[")` trebalo bi da vrati false.

```js
assert(!isBalanced(testCases[13]));
```

`isBalanced("][]][][[")` trebalo bi da vrati false.

```js
assert(!isBalanced(testCases[14]));
```

`isBalanced("][][")` trebalo bi da vrati false.

```js
assert(!isBalanced(testCases[15]));
```

`isBalanced("[]]]")` treba da vrati false.

```js
assert(!isBalanced(testCases[16]));
```

`isBalanced("")` treba da vrati true.

```js
assert(isBalanced(testCases[17]));
```

# --seed--

## --seed-contents--

```js
function isBalanced(str) {

  return true;
}
```

# --solutions--

```js
function isBalanced(str) {
  if (str === '') return true;
  let a = str;
  let b;
  do {
    b = a;
    a = a.replace(/\[\]/g, '');
  } while (a !== b);
  return !a;
}
```
