---
id: 594d8d0ab97724821379b1e6
title: Wastani/Mode
challengeType: 1
forumTopicId: 302226
dashedName: averagesmode
---

# --description--

Napiši funkciju `mode` koja pronalazi vrednost koja se najčešće pojavljuje u nizu.

Slučaj prazne kolekcije može se zanemariti. Potrebno je voditi računa o situaciji kada mod nije jedinstven.

Ako nije prikladno ili moguće podržati opštu kolekciju, koristi vektor (niz). Ako nije prikladno ili moguće podržati neodređen tip vrednosti, koristi cele brojeve.

# --before-each--

```js
const arr1 = [1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17];
const arr2 = [1, 2, 4, 4, 1];
```

# --hints--

`mode` treba da bude funkcija.

```js
assert(typeof mode === 'function');
```

`mode([1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17])` treba da bude jednako `[6]`.

```js
assert.deepEqual(mode(arr1), [6]);
```

`mode([1, 2, 4, 4, 1])` treba da bude jednako `[1, 4]`.

```js
assert.deepEqual(mode(arr2).sort(), [1, 4]);
```

# --seed--

## --seed-contents--

```js
function mode(arr) {

  return true;
}
```

# --solutions--

```js
function mode(arr) {
  const counter = {};
  let result = [];
  let max = 0;
  // for (const i in arr) {
  arr.forEach(el => {
    if (!(el in counter)) {
      counter[el] = 0;
    }
    counter[el]++;

    if (counter[el] === max) {
      result.push(el);
    }
    else if (counter[el] > max) {
      max = counter[el];
      result = [el];
    }
  });
  return result;
}
```
