---
id: 594fa2746886f41f7d8bf225
title: Topološko uređenje
challengeType: 1
forumTopicId: 302340
dashedName: topological-sort
---

# --description--

Ako vam je data mapa između stavki, a stavke koje zavise jedna od druge, topološko sortiranje postavlja stavke u redosled koji ne dozvoljava da nešto zavisi od nečega. Postoje dva poznata algoritma za topološko sortiranje: Kahnovo topološko sortiranje (1962) i pretraga dubine prvogodišta.

# --instructions--

Napišite funkciju koja će vratiti listu sa validnim redosledom za prikupljanje biblioteke iz njenog zavisnosti.

- Prihvatiti imena biblioteka kao jednu reč.
- Predmeti navedeni kao da zavise samo od nečega i koji nemaju sopstvenu zavisnost, ali čiji je način grupisacije definisan.
- Bilo koja samostalna zavisnost se ne uzima u obzir.
- Zavisnosti bez bilo kakve grupisacije se ne uzimaju u obzir.

Koristi sledeće podatke kao primer:

<pre>
Biblioteka          Zavisiosti Biblioteke
=======          ====================
des_system_lib   std synopsys std_cell_lib des_system_lib dw02 dw01 ramlib ieee
dw01             ieee dw01 dware gtech
dw02             ieee dw02 dware
dw03             std synopsys dware dw03 dw02 dw01 ieee gtech
dw04             dw04 ieee dw01 dware gtech
dw05             dw05 ieee dware
dw06             dw06 ieee dware
dw07             ieee dware
dware            ieee dware
gtech            ieee gtech
ramlib           std ieee
std_cell_lib     ieee std_cell_lib
synopsys
</pre>

Kompilacija biblioteke u VHDL jeziku zahteva da biblioteka bude kompilovana nakon bilo koje zavisne biblioteke. Podaci iznad ne bi mogli da se obrade ako, na primer, `dw04` dodate u listu zavisnosti za `dw01`.

Ulaz funkcije će biti sekvenca sa više redova, svaki red će imati ime biblioteke, praćeno njenom zavisnošću (ako postoji).

Na primer:

```js
const libsSimple =
  `aaa bbb
  bbb`;
```

# --before-each--

```js
const libsSimple =
  `aaa bbb
  bbb`;

const libsVHDL =
  `des_system_lib   std synopsys std_cell_lib des_system_lib dw02 dw01 ramlib ieee
  dw01             ieee dw01 dware gtech
  dw02             ieee dw02 dware
  dw03             std synopsys dware dw03 dw02 dw01 ieee gtech
  dw04             dw04 ieee dw01 dware gtech
  dw05             dw05 ieee dware
  dw06             dw06 ieee dware
  dw07             ieee dware
  dware            ieee dware
  gtech            ieee gtech
  ramlib           std ieee
  std_cell_lib     ieee std_cell_lib
  synopsys`;

const libsCustom =
  `a b c d
  b c d
  d c
  c base
  base`;

const libsUnorderable =
  `TestLib Base MainLib
  MainLib TestLib
  Base`;
```

# --hints--

`topologicalSort` Trebalo bi da bude funkcija.

```js
assert(typeof topologicalSort === 'function');
```

`topologicalSort(libsSimple)` Trebalo bi da vrati red podataka.

```js
assert(Array.isArray(topologicalSort(libsSimple)));
```

`topologicalSort(libsSimple)` treba vratiti `['bbb', 'aaa']`.

```js
assert.deepEqual(topologicalSort(libsSimple), ['bbb', 'aaa']);
```

`topologicalSort(libsVHDL)` trebalo bi da vrati `['ieee', 'std_cell_lib', 'gtech', 'dware', 'dw07', 'dw06', 'dw05', 'dw02', 'dw01', 'dw04', 'std', 'ramlib', 'synopsys', 'dw03', 'des_system_lib']`.

```js
assert.deepEqual(topologicalSort(libsVHDL), ['ieee', 'std_cell_lib', 'gtech', 'dware', 'dw07', 'dw06', 'dw05', 'dw02', 'dw01', 'dw04', 'std', 'ramlib', 'synopsys', 'dw03', 'des_system_lib']);
```

`topologicalSort(libsCustom)` trebalo bi da vrati `['base', 'c', 'd', 'b', 'a']`.

```js
assert.deepEqual(topologicalSort(libsCustom), ['base', 'c', 'd', 'b', 'a']);
```

`topologicalSort` Trebalo bi da ignoriše neorganizovanu zavisnost.

```js
assert.deepEqual(topologicalSort(libsUnorderable), ['Base']);
```

# --seed--

## --seed-contents--

```js
function topologicalSort(libs) {

  return true;
}
```

# --solutions--

```js
function topologicalSort(libs) {
  // A map of the input data, with the keys as the packages, and the values as
  // and array of packages on which it depends.
  const D = libs
    .split('\n')
    .map(e => e.split(' ').filter(ep => ep !== ''))
    .reduce((p, c) =>
      p.set(c[0], c.filter((e, i) => (i > 0 && e !== c[0] ? e : null))), new Map());
  [].concat(...D.values()).forEach(e => {
    D.set(e, D.get(e) || []);
  });

  // The above map rotated so that it represents a DAG of the form
  // Map {
  //    A => [ A, B, C],
  //    B => [C],
  //    C => []
  // }
  // where each key represents a node, and the array contains the edges.
  const G = [...D.keys()].reduce((p, c) =>
    p.set(
      c,
      [...D.keys()].filter(e => D.get(e).includes(c))),
    new Map()
  );

  // An array of leaf nodes; nodes with 0 in degrees.
  const Q = [...D.keys()].filter(e => D.get(e).length === 0);

  // The result array.
  const S = [];
  while (Q.length) {
    const u = Q.pop();
    S.push(u);
    G.get(u).forEach(v => {
      D.set(v, D.get(v).filter(e => e !== u));
      if (D.get(v).length === 0) {
        Q.push(v);
      }
    });
  }

  return S;
}
```
