---
id: 594faaab4e2a8626833e9c3d
title: "Razdvoj niz znakova koristeći prenos (ili: translaciju)"
challengeType: 1
forumTopicId: 302338
dashedName: tokenize-a-string-with-escaping
---

# --description--

Napišite funkciju ili program koji može da razdvoji niz karaktera svaki put kada naiđe na netransitivni delimiter uzorak.

Trebalo bi da primi tri ulazna parametra:

<ul>
  <li><strong>sekvenca slova</strong></li>
  <li><strong>svojstvo segmentacije</strong></li>
  <li><strong>svojstvo prenosa</strong></li>
</ul>

Treba da vrati listu stringova.

Sheria za kugawanya:

<ul>
  <li>Sekcije podeljene separatorima biće stavke u listi rezultata.</li>
  <li>Prazne sekcije takođe treba sačuvati, čak i na početku i na kraju.</li>
</ul>

Pravila transfera:

<ul>
  <li>"Transfer" znači svojstvo prethodno praćeno svojstvom transfera koje nema niti transfer.</li>
  <li>Svojstvo transfera koje sledi za svojstvom bez specifičnog značenja, i dalje se računa kao transfer (ali ne radi ništa specifično).</li>
  <li>Svako svojstvo transfera korišćeno za prebacivanje nečega, ne bi trebalo da bude deo rezultata.</li>
</ul>

Pokažite da vaša funkcija zadovoljava sledeći test slučaj:

Ako je niz slova

<pre>one^|uno||three^^^^|four^^^|^cuatro|</pre>

i koristeći `|` kao delimiter i `^` kao svojstvo prenosa, vaša funkcija bi trebalo da izbaci sledeći red podataka:

<pre>  ['one|uno', '', 'three^^', 'four^|cuatro', '']
</pre>

# --before-each--

```js
const testStr1 = 'one^|uno||three^^^^|four^^^|^cuatro|';
const res1 = ['one|uno', '', 'three^^', 'four^|cuatro', ''];

const testStr2 = 'a@&bcd&ef&&@@hi';
const res2 = ['a&bcd', 'ef', '', '@hi'];

const testStr3 = 'hello^|world|how^are^you^|';
const res3 = ['hello|world', 'howareyou|'];
```

# --hints--

`tokenize` Trebalo bi da bude funkcija.

```js
assert(typeof tokenize === 'function');
```

`tokenize` Trebalo bi da vrati red podataka.

```js
assert(typeof tokenize('a', 'b', 'c') === 'object');
```

`tokenize('one^|uno||three^^^^|four^^^|^cuatro|', '|', '^')` treba da vrati `['one|uno', '', 'three^^', 'four^|cuatro', '']`

```js
assert.deepEqual(tokenize(testStr1, '|', '^'), res1);
```

`tokenize('a@&bcd&ef&&@@hi', '&', '@')` treba da vrati `['a&bcd', 'ef', '', '@hi']`

```js
assert.deepEqual(tokenize(testStr2, '&', '@'), res2);
```

`tokenize('hello^|world|how^are^you^|', '|', '^')` treba da vrati `['hello|world', 'howareyou|']`

```js
assert.deepEqual(tokenize(testStr3, '|', '^'), res3);
```

# --seed--

## --seed-contents--

```js
function tokenize(str, sep, esc) {
  return true;
}
```

# --solutions--

```js
// tokenize :: String -> Character -> Character -> [String]
function tokenize(str, charDelim, charEsc) {
  const dctParse = str.split('')
    .reduce((a, x) => {
      const blnEsc = a.esc;
      const blnBreak = !blnEsc && x === charDelim;
      const blnEscChar = !blnEsc && x === charEsc;

      return {
        esc: blnEscChar,
        token: blnBreak ? '' : (
          a.token + (blnEscChar ? '' : x)
        ),
        list: a.list.concat(blnBreak ? a.token : [])
      };
    }, {
      esc: false,
      token: '',
      list: []
    });

  return dctParse.list.concat(
    dctParse.token
  );
}
```
