---
id: 5951ed8945deab770972ae56
title: Minara ya Hanoi
challengeType: 1
forumTopicId: 302341
dashedName: towers-of-hanoi
---

# --description--

Reši problem Kula Hanoja. Broj diskova biće prosleđen kao prvi argument, a zatim slede tri niske koje predstavljaju oznake stubova. Napravi ugnježdeni niz koji sadrži listu poteza u obliku `["izvor", "odredište"]`.

Na primer, argumenti `(4, 'A', 'B', 'C')` treba da proizvedu ugnježdeni niz poteza, pri čemu su prva dva poteza `[['A', 'C'], ['B', 'A']]`, što znači da je prvi potez premeštanje diska sa stuba `A` na stub `C`, a drugi sa stuba `B` na stub `A`.

# --instructions--

Napiši funkciju koja vraća niz poteza potrebnih za rešavanje Kula Hanoja.

# --before-each--

```js
const res3Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B']];
const res7First10Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B'], ['A', 'C'], ['B', 'C'], ['B', 'A']];

function getTowerResults_() {
  return {
    res3: towerOfHanoi(3, 'A', 'B', 'C'),
    res5: towerOfHanoi(5, 'X', 'Y', 'Z'),
    res7: towerOfHanoi(7, 'A', 'B', 'C')
  };
}
```

# --hints--

`towerOfHanoi` treba da bude funkcija.

```js
assert(typeof towerOfHanoi === 'function');
```

`towerOfHanoi(3, ...)` treba da vrati 7 poteza.

```js
assert(getTowerResults_().res3.length === 7);
```

`towerOfHanoi(3, 'A', 'B', 'C')` inapaswa kurejesha `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B']]`.

```js
assert.deepEqual(getTowerResults_().res3, res3Moves);
```

Deseti potez za `towerOfHanoi(5, "X", "Y", "Z")` treba da bude Y -> X.

```js
const { res5 } = getTowerResults_();
assert.deepEqual(res5[9], ['Y', 'X']);
```

Prvih deset poteza za `towerOfHanoi(7, 'A', 'B', 'C')` treba da budu `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B'], ['A','C'], ['B','C'], ['B','A']]`

```js
const { res7 } = getTowerResults_();
assert.deepEqual(res7.slice(0, 10), res7First10Moves);
```

# --seed--

## --seed-contents--

```js
function towerOfHanoi(n, a, b, c) {

  return [[]];
}
```

# --solutions--

```js
function towerOfHanoi(n, a, b, c) {
  const res = [];
  towerOfHanoiHelper(n, a, c, b, res);
  return res;
}

function towerOfHanoiHelper(n, a, b, c, res) {
  if (n > 0) {
    towerOfHanoiHelper(n - 1, a, c, b, res);
    res.push([a, c]);
    towerOfHanoiHelper(n - 1, b, a, c, res);
  }
}
```
