---
id: 5951815dd895584b06884620
title: Krugovi zadatog poluprečnika kroz dve tačke
challengeType: 1
forumTopicId: 302231
dashedName: circles-of-given-radius-through-two-points
---

# --description--

Ako su zadate dve tačke u ravni i poluprečnik, uobičajeno je moguće nacrtati dva kruga datog poluprečnika koji prolaze kroz te tačke.

**Aspekti za razmatranje:**

<ul>
<li>Poluprečnik nula treba tretirati kao slučaj u kome krug nije definisan (osim kada se tačke poklapaju).</li>
<li>Ako se tačke poklapaju, postoji beskonačno mnogo krugova koji prolaze kroz tu tačku, osim ako je poluprečnik takođe nula, kada se krug svodi na jednu tačku.</li>
<li>Ako tačke određuju prečnik, vrati samo jedan krug.</li>
<li>Ako su tačke previše udaljene, nijedan krug se ne može konstruisati.</li>
</ul>

# --instructions--

Implementiraj funkciju koja prihvata dve tačke i poluprečnik i vraća dva kruga koja prolaze kroz te tačke. Za svaki pronađeni krug vrati koordinate njegovog centra zaokružene na četiri decimalna mesta. Svaki centar vrati kao niz, a rezultat kao niz tih nizova.

**U posebnim slučajevima, vratite sledeće:**

<ul>
<li>Ako tačke određuju prečnik, vrati jednu tačku. Međutim, ako je poluprečnik takođe nula, vrati <code>"Radius Zero"</code>.</li>
<li>Ako se tačke poklapaju, vrati <code>"Coincident point. Infinite solutions"</code>.</li>
<li>Ako je rastojanje između tačaka veće od prečnika, vrati <code>"No intersection. Points further apart than circle diameter"</code>.</li>
</ul>

**Primjeri unosa:**

<pre>      p1                p2           r
0.1234, 0.9876    0.8765, 0.2345    2.0
0.0000, 2.0000    0.0000, 0.0000    1.0
0.1234, 0.9876    0.1234, 0.9876    2.0
0.1234, 0.9876    0.8765, 0.2345    0.5
0.1234, 0.9876    0.1234, 0.9876    0.0
</pre>

# --before-each--

```js
const testCases = [
  [[0.1234, 0.9876], [0.8765, 0.2345], 2.0],
  [[0.0000, 2.0000], [0.0000, 0.0000], 1.0],
  [[0.1234, 0.9876], [0.1234, 0.9876], 2.0],
  [[0.1234, 0.9876], [0.8765, 0.2345], 0.5],
  [[0.1234, 0.9876], [0.1234, 0.9876], 0.0]
];
const answers = [
  [[1.8631, 1.9742], [-0.8632, -0.7521]],
  [0, 1],
  'Coincident point. Infinite solutions',
  'No intersection. Points further apart than circle diameter',
  'Radius Zero'
];
```

# --hints--

`getCircles` treba da bude funkcija.

```js
assert(typeof getCircles === 'function');
```

`getCircles([0.1234, 0.9876], [0.8765, 0.2345], 2.0)` inapaswa kurudisha `[[1.8631, 1.9742], [-0.8632, -0.7521]]`.

```js
assert.deepEqual(getCircles(...testCases[0]), answers[0]);
```

`getCircles([0.0000, 2.0000], [0.0000, 0.0000], 1.0)` treba da vrati `[0, 1]`

```js
assert.deepEqual(getCircles(...testCases[1]), answers[1]);
```

`getCircles([0.1234, 0.9876], [0.1234, 0.9876], 2.0)` trebalo bi da vrati `Coincident point. Infinite solutions`

```js
assert.deepEqual(getCircles(...testCases[2]), answers[2]);
```

`getCircles([0.1234, 0.9876], [0.8765, 0.2345], 0.5)` trebalo bi da vrati `No intersection. Points further apart than circle diameter`

```js
assert.deepEqual(getCircles(...testCases[3]), answers[3]);
```

`getCircles([0.1234, 0.9876], [0.1234, 0.9876], 0.0)` treba vratiti `Radius Zero`

```js
assert.deepEqual(getCircles(...testCases[4]), answers[4]);
```

# --seed--

## --seed-contents--

```js
function getCircles(...args) {

  return true;
}
```

# --solutions--

```js
const hDist = (p1, p2) => Math.hypot(...p1.map((e, i) => e - p2[i])) / 2;
const pAng = (p1, p2) => Math.atan(p1.map((e, i) => e - p2[i]).reduce((p, c) => c / p, 1));
const solveF = (p, r) => t => [parseFloat((r * Math.cos(t) + p[0]).toFixed(4)), parseFloat((r * Math.sin(t) + p[1]).toFixed(4))];
const diamPoints = (p1, p2) => p1.map((e, i) => parseFloat((e + (p2[i] - e) / 2).toFixed(4)));

function getCircles(...args) {
  const [p1, p2, s] = args;
  const solve = solveF(p1, s);
  const halfDist = hDist(p1, p2);

  let msg = [];
  switch (Math.sign(s - halfDist)) {
    case 0:
      msg = s ? diamPoints(p1, p2) :
        'Radius Zero';
      break;
    case 1:
      if (!halfDist) {
        msg = 'Coincident point. Infinite solutions';
      }
      else {
        const theta = pAng(p1, p2);
        const theta2 = Math.acos(halfDist / s);
        [1, -1].map(e => solve(theta + e * theta2)).forEach(
          e => msg.push(e));
      }
      break;
    case -1:
      msg = 'No intersection. Points further apart than circle diameter';
      break;
    default:
      msg = 'Reached the default';
  }
  return msg;
}
```
