---
id: 595668ca4cfe1af2fb9818d4
title: Niz Haršad ili Niven
challengeType: 1
forumTopicId: 302281
dashedName: harshad-or-niven-series
---

# --description--

Harshad ili Niven brojevi su pozitivni ceo brojevi ≥ 1 koji se mogu podeliti sa zbirom svojih cifara bez ostatka.

Na primer, `42` je Harshad broj jer `42` je deljiv sa `(4 + 2)` bez ostataka.

Pretpostavlja se da je sekvenca definisana kao broj u rastućem redosledu.

# --instructions--

Implementirati funkciju koja generiše članove Harshad sekvence.

Korišćenje za vraćanje reda podataka sa deset uzastopnih članova, počevši od prvog Haršad broja koji prelazi `n`.

# --hints--

`isHarshadOrNiven` Trebalo bi da bude funkcija.

```js
assert(typeof isHarshadOrNiven === 'function');
```

`isHarshadOrNiven(10)` trebalo bi da vrati `[12, 18, 20, 21, 24, 27, 30, 36, 40, 42]`

```js
assert.deepEqual(isHarshadOrNiven(10), [12, 18, 20, 21, 24, 27, 30, 36, 40, 42]);
```

`isHarshadOrNiven(400)` trebalo bi da vrati `[402, 405, 407, 408, 410, 414, 420, 423, 432, 440]`

```js
assert.deepEqual(isHarshadOrNiven(400), [402, 405, 407, 408, 410, 414, 420, 423, 432, 440]);
```

`isHarshadOrNiven(1000)` treba da vrati `[1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]`

```js
assert.deepEqual(isHarshadOrNiven(1000), [1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]);
```

# --seed--

## --seed-contents--

```js
function isHarshadOrNiven(n) {
  const res = [];

  return res;
}
```

# --solutions--

```js
function isHarshadOrNiven(n) {
  function isHarshad(num) {
    let s = 0;
    const nStr = num.toString();
    for (let i = 0; i < nStr.length; ++i) {
      s += parseInt(nStr.charAt(i), 10);
    }
    return n % s === 0;
  }

  const res = [];
  let count = 0;

  while (count < 10) {
    n++;
    if (isHarshad(n)) {
      count++;
      res.push(n);
    }
  }

  return res;
}
```
