---
id: 595608ff8bcd7a50bd490181
title: Niz gradina
challengeType: 1
forumTopicId: 302279
dashedName: hailstone-sequence
---

# --description--

Niz brojeva hailstone može se generisati počevši od početnog pozitivnog celog broja, `n` za:

- Ako je `n` jednako `1`, onda se sekvenca završava
- Ako je `n` jednako `even`, onda sledeći element u sekvenci za `n` je `= n/2`
- Ako je `n` jednako `odd`, onda sledeći element u sekvenci za `n` je `= (3 * n) + 1`

Neproverena teorija Kollatza je da niz hailstona za bilo koji početni broj uvek se završava.

Hailstone niz je takođe poznat kao hailstone brojevi (jer njegove vrednosti obično padaju i rastu često kao ledeni gradovi na nebu), ili kao Collatz niz.

# --instructions--

1. Napravi funkciju koja generiše Hailstone sekvencu za brojeve
2. Tvoja funkcija mora da vrati niz podataka sa brojevima ispod `limit`, koji ima najdužu Hailstone sekvencu i dužinu te sekvencu. (Ali ne prikazuj stvarnu sekvencu!)

# --hints--

`hailstoneSequence` Trebalo bi da bude funkcija.

```js
assert(typeof hailstoneSequence === 'function');
```

`hailstoneSequence(30)` Trebalo bi da vrati red podataka.

```js
assert(Array.isArray(hailstoneSequence(30)));
```

`hailstoneSequence(30)` treba da vrati `[27, 112]`.

```js
assert.deepEqual(hailstoneSequence(30), [27, 112]);
```

`hailstoneSequence(50000)` trebalo bi da vrati `[35655, 324]`.

```js
assert.deepEqual(hailstoneSequence(50000), [35655, 324]);
```

`hailstoneSequence(100000)` treba da vrati `[77031, 351]`.

```js
assert.deepEqual(hailstoneSequence(100000), [77031, 351]);
```

# --seed--

## --seed-contents--

```js
function hailstoneSequence(limit) {
  const res = [];


  return res;
}
```

# --solutions--

```js
function hailstoneSequence (limit) {
  function hailstone(n) {
    const seq = [n];
    while (n > 1) {
      n = n % 2 ? 3 * n + 1 : n / 2;
      seq.push(n);
    }
    return seq;
  }

  let n = 0;
  let max = 0;
  for (let i = limit; --i;) {
    const seq = hailstone(i);
    const sLen = seq.length;

    if (sLen > max) {
      n = i;
      max = sLen;
    }
  }

  return [n, max];
}
```
