---
id: 59637c4d89f6786115efd814
title: Niz Hofstadtera Q
challengeType: 1
forumTopicId: 302287
dashedName: hofstadter-q-sequence
---

# --description--

Niz Hofstadtera Q je definisan kao:

$Q(1)=Q(2)=1, \\\\ Q(n)=Q\big(n-Q(n-1)\big)+Q\big(n-Q(n-2)), \\quad n>2.$

Definisano je kao Fibonacci niz, ali umesto da je sledeći broj u Fibonacci nizu zbir dva prethodna broja, u Q nizu, dva prethodna broja ukazuju na koliko daleko nazad u Q niz treba pronaći dva broja za sabiranje kako bi se dobio sledeći broj niza.

# --instructions--

Implementiraj Hofstadter Q sekvencijsku jednačinu kao funkciju. Funkcija treba da primi broj, `n`, i vrati ceo broj.

# --before-each--

```js
const testCase = [1000, 1500, 2000, 2500];
const res = [502, 755, 1005, 1261];
```

# --hints--

`hofstadterQ` Trebalo bi da bude funkcija.

```js
assert(typeof hofstadterQ === 'function');
```

`hofstadterQ()` trebalo bi da vrati `integer`

```js
assert(Number.isInteger(hofstadterQ(1000)));
```

`hofstadterQ(1000)` treba da vrati `502`

```js
assert.equal(hofstadterQ(testCase[0]), res[0]);
```

`hofstadterQ(1500)` trebalo bi da vrati `755`

```js
assert.equal(hofstadterQ(testCase[1]), res[1]);
```

`hofstadterQ(2000)` treba vratiti `1005`

```js
assert.equal(hofstadterQ(testCase[2]), res[2]);
```

`hofstadterQ(2500)` treba da vrati `1261`

```js
assert.equal(hofstadterQ(testCase[3]), res[3]);
```

# --seed--

## --seed-contents--

```js
function hofstadterQ(n) {

  return n;
}
```

# --solutions--

```js
function hofstadterQ (n) {
  const memo = [1, 1, 1];
  const Q = function (i) {
    let result = memo[i];
    if (typeof result !== 'number') {
      result = Q(i - Q(i - 1)) + Q(i - Q(i - 2));
      memo[i] = result;
    }
    return result;
  };
  return Q(n);
}
```
