---
id: 5966f99c45e8976909a85575
title: Siku ya wiki
challengeType: 1
forumTopicId: 302245
dashedName: day-of-the-week
---

# --description--

Kompanija je odlučila da svaki put kada Božić padne u nedelju zaposlenima odobri dodatne plaćene slobodne dane, tako da zajedno sa državnim praznicima ne moraju da rade naredne sedmice (između 25. decembra i 1. januara).

# --instructions--

Napišite funkciju koja prima početnu i krajnju godinu i vraća niz podataka svih godina u kojima će datum 25. decembra biti nedelja.

# --before-each--

```js
const firstSolution = [1977, 1983, 1988, 1994, 2005, 2011, 2016];
const secondSolution = [2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118];
```

# --hints--

`findXmasSunday` treba da bude funkcija.

```js
assert(typeof findXmasSunday === 'function');
```

`findXmasSunday(2000, 2100)` Trebalo bi da vrati red podataka.

```js
assert(typeof findXmasSunday(2000, 2100) === 'object');
```

`findXmasSunday(1970, 2017)` treba da vrati `[1977, 1983, 1988, 1994, 2005, 2011, 2016]`

```js
assert.deepEqual(findXmasSunday(1970, 2017), firstSolution);
```

`findXmasSunday(2008, 2121)` treba da vrati `[2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118]`

```js
assert.deepEqual(findXmasSunday(2008, 2121), secondSolution);
```

# --seed--

## --seed-contents--

```js
function findXmasSunday(start, end) {

  return true;
}
```

# --solutions--

```js
function findXmasSunday(start, end) {
  const xmasSunday = [];
  for (let year = start; year <= end; year++) {
    const xmas = new Date(year, 11, 25);
    if (xmas.getDay() === 0) {
      xmasSunday.push(year);
    }
  }
  return xmasSunday;
}
```
