---
id: 59669d08d75b60482359409f
title: Format datuma
challengeType: 1
forumTopicId: 302243
dashedName: date-format
---

# --description--

Vratite red podataka koji sadrži datum sa dva karaktera za današnji datum, sa sledećim karakteristikama:

- Prvi niz znakova mora da sadrži godinu, mesec i dan, razdvojene crticama/znakom dash (`-`).
- Godina u prvom nizu znakova mora da ima četiri cifre.
- Mesec i dan u prvom nizu znakova ne smeju imati vodeće nule.
- Naziv dana u sedmici i mesec u drugom nizu znakova ne smeju biti skraćeni.
- Dan u drugom nizu znakova ne sme imati vodeće nule.

Primeri rezultata:

```js
['2007-11-23', 'Friday, November 23, 2007']
['2021-3-2', 'Tuesday, March 2, 2021']
```

# --hints--

`getDateFormats` Trebalo bi da bude funkcija.

```js
assert(typeof getDateFormats === 'function');
```

`getDateFormats` Trebalo bi da vrati nešto.

```js
assert(typeof getDateFormats() === 'object');
```

`getDateFormats` trebalo bi da vrati red podataka sa 2 elementa.

```js
assert(getDateFormats().length === 2);
```

`getDateFormats` Trebalo bi da vrati ispravan datum u odgovarajućem formatu.

```js
const getDateSolution = () => {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
};

const dates = getDateSolution();
const equalsMessage = `message: <code>getDataFormats()</code> should return <code>["${dates[0]}", "${dates[1]}"]</code>.`;
assert.deepEqual(getDateFormats(), dates, equalsMessage);
```

# --seed--

## --seed-contents--

```js
function getDateFormats() {

  return true;
}
```

# --solutions--

```js
function getDateFormats() {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
}
```
