---
id: 59667989bf71cf555dd5d2ff
title: Izgovor S
challengeType: 1
forumTopicId: 302303
dashedName: s-expressions
---

# --description--

<a href="https://rosettacode.org/wiki/S-expressions" target="_blank" rel="noopener noreferrer nofollow">S-izrazi</a> predstavljaju jednostavan način za parsiranje i čuvanje podataka.

# --instructions--

Implementirajte jednostavan čitač/analizator S Formata koji obrađuje sekvence znakova sa i bez navodnika, celi brojeve i decimalne brojeve.

Funkcija treba da pročita tačno jedan, potencijalno ugnježden, S-izraz iz niske i da ga vrati kao (ugnježdeni) niz.

Novi redovi i ostali prazni znakovi mogu se zanemariti, osim ako se nalaze unutar niske pod navodnicima.

"`()`" unutar sekvence karaktera sa navodnicima smatra se delom sekvence karaktera, a ne prevodi se.

Podrška za escape navodnike unutar niske je opciona; zato se "`(foo"bar)`" može tretirati kao niska "`foo"bar`" ili kao greška.

Za ovaj zadatak parser ne mora da prepoznaje `\` kao escape znak, ali treba da prepoznaje brojeve ako jezik podržava odgovarajuće tipove podataka.

Zapamtite da, osim `()"` (`\` ako trčanje pomaže) i praznih mesta, nema posebnih slova. Ništa drugo nije dozvoljeno bez navodnika.

Msomaji anapaswa kuwa na uwezo wa kusoma ingizo lifuatalo

<pre>((data "quoted data" 123 4.5)
(data (!@# (4.5) "(more" "data)")))
</pre>

na kulibadilisha kuwa muundo wa data wa asili.

# --before-each--

```js
const simpleSExpr = '(data1 data2 data3)';
const simpleSolution = ['data1', 'data2', 'data3'];

const basicSExpr = '((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))';
const basicSolution = [["data","\"quoted data\"",123,4.5],["data",["!@#",[4.5],"\"(more\"","\"data)\""]]];
```

# --hints--

`parseSexpr` Trebalo bi da bude funkcija.

```js
assert(typeof parseSexpr === 'function');
```

`parseSexpr('(data1 data2 data3)')` treba da vrati `['data1', 'data2', 'data3']`

```js
assert.deepEqual(parseSexpr(simpleSExpr), simpleSolution);
```

`parseSexpr('((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))')` treba da vrati `[['data', '"quoted data"', 123, 4.5], ['data', ['!@#', [4.5], '"(more"', '"data)"']]]`.

```js
assert.deepEqual(parseSexpr(basicSExpr), basicSolution);
```

# --seed--

## --seed-contents--

```js
function parseSexpr(str) {

  return true;
}
```

# --solutions--

```js
function parseSexpr(str) {
  const t = str.match(/\s*("[^"]*"|\(|\)|"|[^\s()"]+)/g);
  for (var o, c = 0, i = t.length - 1; i >= 0; i--) {
    var n,
      ti = t[i].trim();
    if (ti == '"') return;
    else if (ti == '(') t[i] = '[', c += 1;
    else if (ti == ')') t[i] = ']', c -= 1;
    else if ((n = +ti) == ti) t[i] = n;
    else t[i] = `'${ti.replace('\'', '\\\'')}'`;
    if (i > 0 && ti != ']' && t[i - 1].trim() != '(') t.splice(i, 0, ',');
    if (!c) if (!o) o = true; else return;
  }
  return c ? undefined : eval(t.join(''));
}
```
