---
id: 59694356a6e7011f7f1c5f4e
title: Napravi karticu za FreeCell
challengeType: 1
forumTopicId: 302246
dashedName: deal-cards-for-freecell
---

# --description--

*FreeCell* je kartaška igra pasijansa koju je napravio Pol Alfil (Paul Alfille) za sistem PLATO 1978. godine. Džim Horn (Jim Horne) iz Microsofta preimenovao ju je u FreeCell i ponovo implementirao za DOS, a zatim za Windows. Ta verzija je sadržala 32000 numerisanih deljenja karata.

Kada je igra postala popularna, Džim Horn je objavio algoritam, pa su druge implementacije FreeCell-a počele da reprodukuju Microsoftova deljenja. Ta deljenja imaju brojeve od 1 do 32000. Novije Microsoft verzije imaju milion deljenja, numerisanih od 1 do 1000000, a neke implementacije dozvoljavaju i brojeve van tog opsega.

Algoritam koristi Microsoft C linearni kongruentni generator slučajnih brojeva:

<ul>
  <li>$state_{n + 1} \equiv 214013 \times state_n + 2531011 \pmod{2^{31}}$</li>
  <li>$rand_n = state_n \div 2^{16}$</li>
<li>`rand_n` je u opsegu od 0 do 32767.</li>
</ul>

Algoriti inafuata hatua hizi:

<ol>
  <li>Weka mbegu (seed) ya RNG kwa nambari ya gawio.
</li><li>Napravi niz od 52 karte: As tref, As karo, As herc, As pik, zatim dvojke svih boja i tako redom do Kralja. Indeksi niza idu od 0 do 51, pri čemu je As tref na indeksu 0, a Kralj pik na indeksu 51.</li>
  <li>Mpaka safu ya data iwe tupu:</li>
    <ul>
<li>Izaberi nasumičnu kartu pomoću <i>indeksa</i> ≡ <i>sledeći slučajni broj</i> (mod <i>dužina niza</i>).</li>
<li>Zameni tu nasumično izabranu kartu poslednjom kartom u nizu.</li>
<li>Ukloni nasumično izabranu kartu iz niza. (Dužina niza se smanjuje za 1.)</li>
      <li>Gawa kadi hii ya nasibu.</li>
    </ul>
<li>Podeli svih 52 karte licem nagore u 8 kolona. Prvih 8 karata ide u svih 8 kolona, narednih 8 na njih i tako redom.</li>
</ol>

**Primer:**

**Mpangilio wa kugawa kadi**

<pre> 1  2  3  4  5  6  7  8
 9 10 11 12 13 14 15 16
17 18 19 20 21 22 23 24
25 26 27 28 29 30 31 32
33 34 35 36 37 38 39 40
41 42 43 44 45 46 47 48
49 50 51 52</pre>

**Igra #1**

```js
[
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
]
```

**Igra #617**

```js
[
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
]
```

# --instructions--

Napiši funkciju koja prihvata broj deljenja i raspoređuje karte na isti način kao ovaj algoritam. Funkcija treba da vrati dvodimenzionalni niz koji predstavlja FreeCell tablu.

# --before-each--

```js
const replaceThis = 3;
const game1 = [
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
];
const game617 = [
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
];
```

# --hints--

`dealFreeCell` trebalo bi da bude funkcija.

```js
assert(typeof dealFreeCell === 'function');
```

`dealFreeCell(seed)` Trebalo bi da vrati nešto.

```js
assert(typeof dealFreeCell(1) === 'object');
```

`dealFreeCell(seed)` Treba da vrati red podataka sa dužinom od 7.

```js
assert(dealFreeCell(1).length === 7);
```

`dealFreeCell(1)` trebalo bi da vrati red podataka kao primer "Mchezo #1"

```js
assert.deepEqual(dealFreeCell(1), game1);
```

`dealFreeCell(617)` Trebalo bi da vrati red podataka kao što je primer "Igra #617"

```js
assert.deepEqual(dealFreeCell(617), game617);
```

# --seed--

## --seed-contents--

```js
function dealFreeCell(seed) {

  return true;
}
```

# --solutions--

```js
// RNG
function FreeCellRNG(seed) {
  return {
    lastNum: seed,
    next() {
      this.lastNum = ((214013 * this.lastNum) + 2531011) % (Math.pow(2, 31));
      return this.lastNum >> 16;
    }
  };
}
// Get cards
function getDeck() {
  const ranks = ['A', '2', '3', '4', '5', '6', '7', '8', '9', 'T', 'J', 'Q', 'K'];
  const suits = ['C', 'D', 'H', 'S'];
  const cards = [];
  for (let i = 0; i < ranks.length; i += 1) {
    for (let j = 0; j < suits.length; j += 1) {
      cards.push(`${ranks[i]}${suits[j]}`);
    }
  }
  return cards;
}
function dealFreeCell(seed) {
  const rng = FreeCellRNG(seed);
  const deck = getDeck();

  const dealtCards = [[], [], [], [], [], [], []];
  let currentColumn = 0;
  let currentRow = 0;

  let rand;
  let temp;
  let card;
  while (deck.length > 0) {
    // Choose a random card
    rand = rng.next() % deck.length;

    // Swap this random card with the last card in the array
    temp = deck[deck.length - 1];
    deck[deck.length - 1] = deck[rand];
    deck[rand] = temp;

    // Remove this card from the array
    card = deck.pop();

    // Deal this card
    dealtCards[currentRow].push(card);
    currentColumn += 1;
    if (currentColumn === 8) {
      currentColumn = 0;
      currentRow += 1;
    }
  }

  return dealtCards;
}
```
