---
id: 596a8888ab7c01048de257d5
title: Duboka kopija
challengeType: 1
forumTopicId: 302247
dashedName: deepcopy
---

# --description--

Napišite funkciju koja vraća duboku kopiju datog objekta. Ta kopija ne sme biti isti objekat koji je dat.

Ovo vežbanje će osigurati:

<ul>
  <li>Predmeti sa kriterijumima koji su akcije/funkcije</li>
  <li>Predmeti datuma ili predmeti sa kriterijumima koji su datumi</li>
  <li>Opšti tekst ili predmeti sa kriterijumima koji su opšti tekstovi</li>
  <li>Kopija primera</li>
</ul>

# --before-each--

```js
const obj1 = { test: 'test' };
const obj2 = {
  t: 'test',
  a: ['an', 'array']
};
const obj3 = {
  t: 'try',
  o: obj2
};
```

# --hints--

`deepcopy` Trebalo bi da bude funkcija.

```js
assert(typeof deepcopy === 'function');
```

`deepcopy({test: "test"})` Trebalo bi da nešto restaurira.

```js
assert(typeof deepcopy(obj1) === 'object');
```

`deepcopy` Ne smete da vraćate identičnu stvar koja je bila izdata.

```js
assert(deepcopy(obj2) != obj2);
```

Kada primi objekat sa strukturom podataka, `deepcopy` treba da vrati duboku kopiju tog objekta.

```js
assert.deepEqual(deepcopy(obj2), obj2);
```

Kada primite objekt koji sadrži nešto drugo unutra, `deepcopy` bi trebalo da vrati duboku kopiju tog objekta.

```js
assert.deepEqual(deepcopy(obj3), obj3);
```

# --seed--

## --seed-contents--

```js
function deepcopy(obj) {

  return true;
}
```

# --solutions--

```js
function deepcopy(obj) {
  return JSON.parse(JSON.stringify(obj));
}
```
