---
id: 596e457071c35c882915b3e4
title: Uporedite listu sekvenci slova
challengeType: 1
forumTopicId: 302235
dashedName: compare-a-list-of-strings
---

# --description--

Lista je uređeni skup vrednosti koji može imati duplikate. Evo primera:

```js
const list = [['AA',  'BB', 'CC'], ['AA', 'ACB', 'AA'], [], ['AA']];
````

Ako ti je data lista stringova bilo koje dužine, implementiraj funkciju za svaki od sledećih uslova:

<ul>
  <li>Proveriti da li su sve jednake po matematičkom stilu sekvence karaktera</li>
  <li>Proveriti da li je svaka sekvenca karaktera matematički manja od sledeće (što znači ako je lista u striktno rastućem redosledu)</li>
</ul>

# --before-each--

```js
const testCases = [['AA', 'AA', 'AA', 'AA'], ['AA', 'ACB', 'BB', 'CC'], [], ['AA'], ['BB', 'AA']];
```

# --hints--

`allEqual` Trebalo bi da bude funkcija.

```js
assert(typeof allEqual === 'function');
```

`azSorted` treba da bude funkcija.

```js
assert(typeof azSorted === 'function');
```

`allEqual(["AA", "AA", "AA", "AA"])` inapaswa kurudisha true.

```js
assert(allEqual(testCases[0]));
```

`azSorted(["AA", "AA", "AA", "AA"])` inapaswa kurudisha false.

```js
assert(!azSorted(testCases[0]));
```

`allEqual(["AA", "ACB", "BB", "CC"])` trebalo bi vratiti false.

```js
assert(!allEqual(testCases[1]));
```

`azSorted(["AA", "ACB", "BB", "CC"])` trebalo bi da vrati true.

```js
assert(azSorted(testCases[1]));
```

`allEqual([])` inapaswa kurudisha true.

```js
assert(allEqual(testCases[2]));
```

`azSorted([])` trebalo bi da vrati true.

```js
assert(azSorted(testCases[2]));
```

`allEqual(["AA"])` trebalo bi da vrati true.

```js
assert(allEqual(testCases[3]));
```

`azSorted(["AA"])` treba da vrati true.

```js
assert(azSorted(testCases[3]));
```

`allEqual(["BB", "AA"])` trebalo bi da vrati false.

```js
assert(!allEqual(testCases[4]));
```

`azSorted(["BB", "AA"])` trebalo bi da vrati false.

```js
assert(!azSorted(testCases[4]));
```

# --seed--

## --seed-contents--

```js
function allEqual(arr) {

  return true;
}

function azSorted(arr) {

  return true;
}
```

# --solutions--

```js
function allEqual(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] === a[i]);
  } return out;
}

function azSorted(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] < a[i]);
  } return out;
}
```
