---
id: 596e414344c3b2872167f0fe
title: Diskusija o komi
challengeType: 1
forumTopicId: 302234
dashedName: comma-quibbling
---

# --description--

<a href="https://rosettacode.org/wiki/Comma_quibbling" target="_blank" rel="noopener noreferrer nofollow">Diskusija o zarezu</a> je vežba koju je predstavio Eric Lippert na svom blogu.

# --instructions--

Napiši funkciju za generisanje sekvence karaktera koja proističe iz spajanja ulaznih reči iz liste/niza gde:

<ol>
  <li>Ulaz bez reči proizvodi sekvencu od dva karaktera u uglastim zagradama (<code>"{}"</code>)</li>
  <li>Ulaz samo jedne reči, npr. <code>["ABC"]</code>, proizvodi sekvencu karaktera te reči unutar tih dve zagrade, npr. <code>"{ABC}"</code></li>
  <li>Ulaz sa dve reči, npr. <code>["ABC", "DEF"]</code>, proizvodi sekvencu karaktera te dve reči unutar dve zagrade, razdvojene sa nizom karaktera <code>" and "</code>, npr. <code>"{ABC and DEF}"</code></li>
  <li>Ulaz sa tri ili više reči, npr. <code>["ABC", "DEF", "G", "H"]</code>, proizvodi sekvencu karaktera svih reči osim poslednje, razdvojene sa <code>", "</code> i poslednja reč razdvojena sa <code>" and "</code> i sve unutar zagrada; npr. <code>"{ABC, DEF, G and H}"</code></li>
</ol>

Testirajte svoju funkciju sa nizom ulaza da prikažete rezultate ovde, na ovoj stranici:

<ul>
  <li>[] # (Nema ulaznih reči).</li>
  <li>["ABC"]</li>
  <li>["ABC", "DEF"]</li>
  <li>["ABC", "DEF", "G", "H"]</li>
</ul>

**Napomena:** Prihvatite reči koje su sekvenca neprvih velikih slova za ovu vežbu.

# --before-each--

```js
const testCases = [[], ["ABC"], ["ABC", "DEF"], ["ABC", "DEF", "G", "H"]];
const results = ["{}", "{ABC}", "{ABC and DEF}", "{ABC, DEF, G and H}"];
```

# --hints--

`quibble` Trebalo bi da bude funkcija.

```js
assert(typeof quibble === 'function');
```

`quibble(["ABC"])` Trebalo bi da vrati sekvencu slova.

```js
assert(typeof quibble(['ABC']) === 'string');
```

`quibble([])` treba da vrati "{}"

```js
assert.equal(quibble(testCases[0]), results[0]);
```

`quibble(["ABC"])` trebalo bi da vrati `"{ABC}"`.

```js
assert.equal(quibble(testCases[1]), results[1]);
```

`quibble(["ABC", "DEF"])` treba da vrati `"{ABC and DEF}"`.

```js
assert.equal(quibble(testCases[2]), results[2]);
```

`quibble(["ABC", "DEF", "G", "H"])` inapaswa kurudisha `"{ABC, DEF, G and H}"`.

```js
assert.equal(quibble(testCases[3]), results[3]);
```

# --seed--

## --seed-contents--

```js
function quibble(words) {

  return true;
}
```

# --solutions--

```js
function quibble(words) {
  return "{" +
    words.slice(0, words.length - 1).join(", ") +
   (words.length > 1 ? " and " : "") +
   (words[words.length - 1] || '') +
  "}";
}
```
