---
id: 596fda99c69f779975a1b67d
title: Izračunaj broj pojavljivanja sekvence malih slova
challengeType: 1
forumTopicId: 302237
dashedName: count-occurrences-of-a-substring
---

# --description--

Kreirajte funkciju, ili prikažite ugrađenu funkciju, koja broji broj nepreklapajućih pojavljivanja malog niza slova unutar niza slova.

Funkcija mora da primi dva argumenta:

<ul>
  <li>Prvo polje je sekvenca slova za pretragu, i</li>
  <li>drugo je sekvenca malih slova za pretragu.</li>
</ul>

Treba nam vratiti brojač punih brojeva.

Povezivanje bi trebalo da izvuče najveći broj nepreklapajućih mečeva.

Generalno, ovo u suštini znači poređenje od leve ka desnoj ili od desne ka levoj strani.

# --before-each--

```js
const testCases = ['the three truths', 'ababababab', 'abaabba*bbaba*bbab'];
const searchString = ['th', 'abab', 'a*b'];
const results = [3, 2, 2];
```

# --hints--

`countSubstring` Trebalo bi da bude funkcija.

```js
assert(typeof countSubstring === 'function');
```

`countSubstring("the three truths", "th")` treba da vrati `3`.

```js
assert.equal(countSubstring(testCases[0], searchString[0]), results[0]);
```

`countSubstring("ababababab", "abab")` treba da vrati `2`.

```js
assert.equal(countSubstring(testCases[1], searchString[1]), results[1]);
```

`countSubstring("abaabba*bbaba*bbab", "a*b")` treba da vrati `2`.

```js
assert.equal(countSubstring(testCases[2], searchString[2]), results[2]);
```

# --seed--

## --seed-contents--

```js
function countSubstring(str, subStr) {

  return true;
}
```

# --solutions--

```js
function countSubstring(str, subStr) {
  const escapedSubStr = subStr.replace(/[.+*?^$[\]{}()|/]/g, '\\$&');
  const matches = str.match(new RegExp(escapedSubStr, 'g'));
  return matches ? matches.length : 0;
}
```
