---
id: 596fd036dc1ab896c5db98b1
title: Promeni sekunde u kombinovani vremenski format
challengeType: 1
forumTopicId: 302236
dashedName: convert-seconds-to-compound-duration
---

# --description--

Tekeleza kitendakazi ambacho:

<ul>
<li>prihvata pozitivan ceo broj koji predstavlja trajanje u sekundama kao ulaz (na primer, <code>100</code>), i</li>
<li>vraća nisku koja prikazuje isto trajanje izraženo u nedeljama, danima, satima, minutima i sekundama, kao što je opisano u nastavku (na primer, <code>1 min, 40 sec</code>).</li>
</ul>

Pokaži da prolazi sledeće tri testa:

<div style='font-size:115%; font-weight: bold;'>Majaribio</div>

| Ulazni broj | Izlazni broj |
| ----------------- | ---------------------------------- |
| 7259              | <code>2 hr, 59 sec</code>          |
| 86400             | <code>1 d</code>                   |
| 6000000           | <code>9 wk, 6 d, 10 hr, 40 min</code> |

<div style="font-size:115%; font-weight: bold;">Maelezo</div>
<ul>
  <li>
Treba koristiti sledećih pet jedinica:

| Kifungu | Kiambatisho kinachotumika kwenye matokeo | Ubadilishaji           |
| ------- | ----------------------------------------- | --------------------- |
| wiki    | <code>wk</code>                           | 1 wiki = siku 7       |
| siku    | <code>d</code>                            | 1 siku = saa 24       |
| saa     | <code>hr</code>                           | 1 saa = dakika 60     |
| dakika  | <code>min</code>                          | 1 dakika = sekunde 60 |
| sekunde | <code>sec</code>                          | ---                   |

  </li>
  <li>
Međutim, u rezultat uključi samo jedinice čija vrednost nije nula (na primer, vrati <code>1 d</code>, a ne <code>0 wk, 1 d, 0 hr, 0 min, 0 sec</code>).
  </li>
  <li>
Daj prednost većim jedinicama u odnosu na manje kad god je to moguće (na primer, vrati <code>2 min, 10 sec</code>, a ne <code>1 min, 70 sec</code> ili <code>130 sec</code>).
  </li>
  <li>
Prati format prikazan u testovima (vrednosti su poređane od najveće ka najmanjoj jedinici i razdvojene zarezom i razmakom; vrednost i jedinica svake količine razdvojene su razmakom).
  </li>
</ul>

# --before-each--

```js
const testCases = [7259, 86400, 6000000];
const results = ['2 hr, 59 sec', '1 d', '9 wk, 6 d, 10 hr, 40 min'];
```

# --hints--

`convertSeconds` treba da bude funkcija.

```js
assert(typeof convertSeconds === 'function');
```

`convertSeconds(7259)` inapaswa kurudisha `2 hr, 59 sec`.

```js
assert.equal(convertSeconds(testCases[0]), results[0]);
```

`convertSeconds(86400)` inapaswa kurudisha `1 d`.

```js
assert.equal(convertSeconds(testCases[1]), results[1]);
```

`convertSeconds(6000000)` treba da vrati `9 wk, 6 d, 10 hr, 40 min`.

```js
assert.equal(convertSeconds(testCases[2]), results[2]);
```

# --seed--

## --seed-contents--

```js
function convertSeconds(sec) {

  return true;
}
```

# --solutions--

```js
function convertSeconds(sec) {
  const localNames = ['wk', 'd', 'hr', 'min', 'sec'];
  // compoundDuration :: [String] -> Int -> String
  const compoundDuration = (labels, intSeconds) =>
    weekParts(intSeconds)
    .map((v, i) => [v, labels[i]])
    .reduce((a, x) =>
      a.concat(x[0] ? [`${x[0]} ${x[1] || '?'}`] : []), []
    )
    .join(', ');

    // weekParts :: Int -> [Int]
  const weekParts = intSeconds => [0, 7, 24, 60, 60]
    .reduceRight((a, x) => {
      const r = a.rem;
      const mod = x !== 0 ? r % x : r;

      return {
        rem: (r - mod) / (x || 1),
        parts: [mod].concat(a.parts)
      };
    }, {
      rem: intSeconds,
      parts: []
    })
    .parts;

  return compoundDuration(localNames, sec);
}
```
