---
id: 597089c87eec450c68aa1643
title: Eleza aina ya data ya msingi
challengeType: 1
forumTopicId: 302248
dashedName: define-a-primitive-data-type
---

# --description--

Tip koji funkcioniše kao ceo broj, ali ima minimalnu validnu vrednost od 1 i maksimalnu validnu vrednost od 10.

Upravljanje greškama:

<ul>
  <li>Kada pokušate da inicijalizujete <code>Num</code> sa vrednošću izvan opsega 1 - 10, treba bi se baciti <code>TypeError</code> sa porukom greške <code>'Out of range'</code>.</li>
  <li>Kada pokušate da inicijalizujete <code>Num</code> sa vrednošću koja nije broj, treba bi se baciti <code>TypeError</code> sa porukom greške <code>'Not a Number'</code>.</li>
</ul>

# --hints--

`Num` Trebalo bi da bude funkcija.

```js
assert(typeof Num === 'function');
```

`new Num(4)` treba da vrati objekat.

```js
assert(typeof new Num(4) === 'object');
```

`new Num('test')` inapaswa kutupa TypeError yenye ujumbe 'Not a Number'.

```js
assert.throws(() => new Num('test'), TypeError);
```

`new Num(0)` trebalo bi da baci TypeError sa porukom 'Out of range'.

```js
assert.throws(() => new Num(0), TypeError);
```

`new Num(-5)` inapaswa kutupa TypeError yenye ujumbe 'Out of range'.

```js
assert.throws(() => new Num(-5), TypeError);
```

`new Num(10)` inapaswa kutupa TypeError yenye ujumbe 'Out of range'.

```js
assert.throws(() => new Num(11), TypeError);
```

`new Num(20)` inapaswa kutupa TypeError yenye ujumbe 'Out of range'.

```js
assert.throws(() => new Num(20), TypeError);
```

`new Num(3) + new Num(4)` treba biti jednako 7.

```js
assert.equal(new Num(3) + new Num(4), 7);
```

`new Num(3) - new Num(4)` Trebalo bi da bude jednako -1.

```js
assert.equal(new Num(3) - new Num(4), -1);
```

`new Num(3) * new Num(4)` Trebalo bi da bude jednako 12.

```js
assert.equal(new Num(3) * new Num(4), 12);
```

`new Num(3) / new Num(4)` treba biti jednako 0.75.

```js
assert.equal(new Num(3) / new Num(4), 0.75);
```

`new Num(3) < new Num(4)` Trebalo bi da je tačno.

```js
assert(new Num(3) < new Num(4));
```

`new Num(3) > new Num(4)` Trebalo bi da nije istina.

```js
assert(!(new Num(3) > new Num(4)));
```

`(new Num(5)).toString()` Treba da vrati '5'

```js
assert.equal(new Num(5).toString(), '5');
```

# --seed--

## --seed-contents--

```js
function Num(n) {

  return n;
}
```

# --solutions--

```js
function Num(n) {
  if (isNaN(n)) {
    throw new TypeError('Not a Number');
  }
  if (n < 1 || n > 10) {
    throw new TypeError('Out of range');
  }

  this._value = +n;
}
Num.prototype.valueOf = function() { return this._value; };
Num.prototype.toString = function () { return this._value.toString(); };
```
