---
id: 59713bd26bdeb8a594fb9413
title: Konverzija valute
challengeType: 1
forumTopicId: 302238
dashedName: count-the-coins
---

# --description--

Postoje četiri tipa valuta koje se često koriste u američkim valutama:

<ul>
  <li>kvartovi (25 centi)</li>
  <li>desetinki (10 centi)</li>
  <li>niklovi (5 centi), na</li>
  <li>peniji (1 cent)</li>
</ul>

<p>Postoji šest načina da se izdaš menjač od 15 centi:</p>

<ul>
  <li>Dime i nikl</li>
  <li>Dime i pet penija</li>
  <li>tri nikla</li>
  <li>dva nikla i pet penija</li>
  <li>jedan nikl i deset penija</li>
  <li>petnaest penija</li>
</ul>

# --instructions--

Pokrenite funkciju da utvrdite koliko je načina za vraćanje novca za određeni unos, `cents`, što predstavlja iznos u američkim centima korišćenjem ovih često korišćenih valuta.

# --hints--

`countCoins` Trebalo bi da bude funkcija.

```js
assert(typeof countCoins === 'function');
```

`countCoins(15)` treba da vrati `6`.

```js
assert.equal(countCoins(15), 6);
```

`countCoins(85)` trebalo bi da vrati `163`.

```js
assert.equal(countCoins(85), 163);
```

`countCoins(100)` trebalo bi da vrati `242`.

```js
assert.equal(countCoins(100), 242);
```

# --seed--

## --seed-contents--

```js
function countCoins(cents) {

  return true;
}
```

# --solutions--

```js
function countCoins(cents) {
  const operands = [1, 5, 10, 25];
  const targetsLength = cents + 1;
  const operandsLength = operands.length;
  const t = [1];

  for (let a = 0; a < operandsLength; a++) {
    for (let b = 1; b < targetsLength; b++) {
      // initialise undefined target
      t[b] = t[b] ? t[b] : 0;

      // accumulate target + operand ways
      t[b] += (b < operands[a]) ? 0 : t[b - operands[a]];
    }
  }

  return t[targetsLength - 1];
}
```
