---
id: 597b2b2a2702b44414742771
title: Faktoria
challengeType: 1
forumTopicId: 302263
dashedName: factorial
---

# --description--

Napiši funkciju za vraćanje fabrike broja.

Generator brojeva se dobija od:

<pre><big>n! = n * (n-1) * (n-2) * ..... * 1</big>
</pre>

Na primer:

<ul>
  <li><code>3! = 3 * 2 * 1 = 6</code></li>
  <li><code>4! = 4 * 3 * 2 * 1 = 24</code></li>
</ul>

**Napomena:** `0! = 1`

# --hints--

`factorial` Trebalo bi da bude funkcija.

```js
assert(typeof factorial === 'function');
```

`factorial(2)` Trebalo bi da vrati broj.

```js
assert(typeof factorial(2) === 'number');
```

`factorial(3)` Trebalo bi da vrati 6.

```js
assert.equal(factorial(3), 6);
```

`factorial(5)` trebalo bi da vrati 120.

```js
assert.equal(factorial(5), 120);
```

`factorial(10)` trebalo bi da vrati 3,628,800.

```js
assert.equal(factorial(10), 3628800);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

}
```

# --solutions--

```js
function factorial(n) {
  let sum = 1;
  while (n > 1) {
    sum *= n;
    n--;
  }
  return sum;
}
```
