---
id: 597f1e7fbc206f0e9ba95dc4
title: Polja za kompletan broj
challengeType: 1
forumTopicId: 302265
dashedName: factors-of-an-integer
---

# --description--

Napišite funkciju koja vraća opsege punih pozitivnih brojeva kao niz podataka.

Ovi delitelji su pozitivni ceo brojevi kojim se broj koji se deli može podeliti da bi se dobio rezultat koji je pozitivan ceo broj.

# --before-each--

```js
const ans=[[1,3,5,9,15,45],[1,53],[1,2,4,8,16,32,64]];
```

# --hints--

`factors` trebalo bi da bude funkcija.

```js
assert(typeof factors === 'function');
```

`factors(45)` treba da restaurira `[1,3,5,9,15,45]`.

```js
assert.deepEqual(factors(45), ans[0]);
```

`factors(53)` treba da restaurira `[1,53]`.

```js
assert.deepEqual(factors(53), ans[1]);
```

`factors(64)` trebalo bi da vrati `[1,2,4,8,16,32,64]`.

```js
assert.deepEqual(factors(64), ans[2]);
```

# --seed--

## --seed-contents--

```js
function factors(num) {

}
```

# --solutions--

```js
function factors(num)
{
 let n_factors = [], i, sqr=Math.floor(Math.sqrt(num));

 for (i = 1; i <=sqr ; i += 1)
  if (num % i === 0)
  {
   n_factors.push(i);
   if (num / i !== i)
    n_factors.push(num / i);
  }
 n_factors.sort(function(a, b){return a - b;});
 return n_factors;
}
```
