---
id: 597f24c1dda4e70f53c79c81
title: Fibonaccijev niz
challengeType: 1
forumTopicId: 302268
dashedName: fibonacci-sequence
---

# --description--

Napišite funkciju koja generiše N-ti broj <code>n<sup>th</sup></code> u Fibonaccijevoj sekvenci.

Broj <code>n<sup>th</sup></code> u Fibonaccijevoj sekvenci se dobija sa:

<code>F<sub>n</sub> = F<sub>n-1</sub> + F<sub>n-2</sub></code>

Prva dva elementa serije su 0 i 1.

Dakle, niz je: 0, 1, 1, 2, 3, 5, 8, 13...

# --hints--

`fibonacci` Trebalo bi da bude funkcionalno.

```js
assert(typeof fibonacci === 'function');
```

`fibonacci(2)` trebalo bi da vrati broj.

```js
assert(typeof fibonacci(2) == 'number');
```

`fibonacci(3)` Trebalo bi da vrati 2.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` Trebalo bi da vrati 5.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` trebalo bi da vrati 55.

```js
assert.equal(fibonacci(10), 55);
```

# --seed--

## --seed-contents--

```js
function fibonacci(n) {

}
```

# --solutions--

```js
function fibonacci(n) {
  let a = 0, b = 1, t;
  while (--n >= 0) {
    t = a;
    a = b;
    b += t;
  }
  return a;
}
```
