---
id: 59880443fb36441083c6c20e
title: Staza Eulera
challengeType: 1
forumTopicId: 302258
dashedName: euler-method
---

# --description--

Eulerova metoda numerički aproksimira rešenja običnih diferencijalnih jednačina prvog reda (ODEs) za zadati početni uslov. To je eksplicitna metoda za rešavanje problema sa početnim uslovima (IVPs), kao što je opisano u <a href="https://www.freecodecamp.org/news/eulers-method-explained-with-examples/" title="Euler's Method Explained with Examples" target="_blank" rel="noopener noreferrer nofollow">makala hii</a>.

ODE mora biti uklonjen/izvučen u sledećem formatu:

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt} = f(t,y(t))$</big></li>
</ul>

u zavisnosti od početne vrednosti

<ul style='list-style: none;'>
  <li><big>$y(t_0) = y_0$</big></li>
</ul>

Da bismo dobili numeričko rešenje, menjamo izraz na levoj strani (LHS) koristeći aproksimaciju konačne razlike:

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt}  \approx \frac{y(t+h)-y(t)}{h}$</big></li>
</ul>

Zatim reši za $y(t+h):$

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, \frac{dy(t)}{dt}$</big></li>
</ul>

koji je jednako sa

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, f(t,y(t))$</big></li>
</ul>

Pravilo rekurzivnog rešenja je:

<ul style='list-style: none;'>
  <li><big>$y_{n+1} = y_n + h \, f(t_n, y_n)$</big></li>
</ul>

gde je $h$ veličina koraka, najvažniji parametar za tačnost rešenja. Manja veličina koraka povećava tačnost, ali takođe povećava računski trošak, pa ga uvek treba ručno odabrati u zavisnosti od postojećeg zadatka.

**Primer: Zakon hlađenja Newtona**

Newtonov zakon hlađenja opisuje kako se telo sa početnom temperaturom $T(t_0) = T_0$ hladi u okruženju sa temperaturom $T_R$:

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, \Delta T$</big></li>
</ul>

au

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, (T(t) - T_R)$</big></li>
</ul>

Kaže se da je koeficijent hlađenja $\frac{dT(t)}{dt}$ objekta povezan sa razlikovanjem temperature $\Delta T = (T(t) - T_R)$ i okruženjem.

Matematičko rešenje, koje ćemo uporediti s procijenama brojeva, je

<ul style='list-style: none;'>
  <li><big>$T(t) = T_R + (T_0 - T_R) \; e^{-k t}$</big></li>
</ul>

# --instructions--

Izvršite postupak Eulerovog metoda, a zatim ga koristite za rešavanje datog primera Njutnovog zakona hlađenja za tri različite veličine koraka:

<ul>
  <li><code>2 s</code></li>
  <li><code>5 s</code> na</li>
  <li><code>10 s</code></li>
</ul>

Zatim uporedite sa matematičkim rešenjem.

**Početna vrednost:**

<ul>
  <li>Početna temperatura <big>$T_0$</big> će biti <code>100 °C</code></li>
  <li>Temperatura sobe <big>$T_R$</big> će biti <code>20 °C</code></li>
  <li>Konstanta hlađenja <big>$k$</big> će biti <code>0.07</code></li>
  <li>Vremenski interval za izračun će biti od <code>0 s</code> do <code>100 s</code></li>
</ul>

Prvi parametar je početno stanje, drugi parametar je početna temperatura, treći parametar je prošlo vreme i četvrti parametar je veličina koraka.

# --hints--

`eulersMethod` To bi trebalo da bude funkcija.

```js
assert(typeof eulersMethod === 'function');
```

`eulersMethod(0, 100, 100, 2)` Trebalo bi da vrati broj.

```js
assert(typeof eulersMethod(0, 100, 100, 2) === 'number');
```

`eulersMethod(0, 100, 100, 2)` trebalo bi da vrati 20.0424631833732.

```js
assert.equal(eulersMethod(0, 100, 100, 2), 20.0424631833732);
```

`eulersMethod(0, 100, 100, 5)` inapaswa kurudisha 20.01449963666907.

```js
assert.equal(eulersMethod(0, 100, 100, 5), 20.01449963666907);
```

`eulersMethod(0, 100, 100, 10)` inapaswa kurudisha 20.000472392.

```js
assert.equal(eulersMethod(0, 100, 100, 10), 20.000472392);
```

# --seed--

## --seed-contents--

```js
function eulersMethod(x1, y1, x2, h) {

}
```

# --solutions--

```js
function eulersMethod(x1, y1, x2, h) {
  let x = x1;
  let y = y1;

  while ((x < x2 && x1 < x2) || (x > x2 && x1 > x2)) {
    y += h * (-0.07 * (y - 20));
    x += h;
  }

  return y;
}
```
