---
id: 598eef80ba501f1268170e1e
title: Niz Fibonaccijevih brojeva za korak n
challengeType: 1
forumTopicId: 302267
dashedName: fibonacci-n-step-number-sequences
---

# --description--

Ovaj niz brojeva je proširenje standardnog Fibonaccijevog niza, gde:

<ol>
<li>Za $n = 2$ dobija se Fibonačijev niz, sa početnim vrednostima $[1, 1]$ i relacijom $F_k^2 = F_{k-1}^2 + F_{k-2}^2$.</li>
<li>Za $n = 3$ dobija se Tribonačijev niz, sa početnim vrednostima $[1, 1, 2]$ i relacijom $F_k^3 = F_{k-1}^3 + F_{k-2}^3 + F_{k-3}^3$.</li>
<li>Za $n = 4$ dobija se Tetrabonačijev niz, sa početnim vrednostima $[1, 1, 2, 4]$ i relacijom $F_k^4 = F_{k-1}^4 + F_{k-2}^4 + F_{k-3}^4 + F_{k-4}^4$...</li>
<li>Uopšteno, za $n>2$ dobija se Fibonačijev niz reda $n$ – $F_k^n$; početne vrednosti čine prvih $n$ članova Fibonačijevog niza reda $(n-1)$, $F_k^{n-1}$, a svaki naredni član računa se kao $F_k^n = \sum_{i=1}^{n} {F_{k-i}^{(n)}}$.</li>
</ol>

Kod malih vrednosti $n$, ponekad se koriste grčki numerički prefiksi za davanje pojedinačnih imena svakoj sekvenci.

Fibonaccijev niz:

| $n$ | Jina la mfululizo | Thamani                                                 |
| --- | ----------------- | ------------------------------------------------------- |
| 2   | fibonacci         | 1 1 2 3 5 8 13 21 34 55 89 144 233 377 610 ...          |
| 3   | tribonacci        | 1 1 2 4 7 13 24 44 81 149 274 504 927 1705 3136 ...     |
| 4   | tetranacci        | 1 1 2 4 8 15 29 56 108 208 401 773 1490 2872 5536 ...   |
| 5   | pentanacci        | 1 1 2 4 8 16 31 61 120 236 464 912 1793 3525 6930 ...   |
| 6   | hexanacci         | 1 1 2 4 8 16 32 63 125 248 492 976 1936 3840 7617 ...   |
| 7   | heptanacci        | 1 1 2 4 8 16 32 64 127 253 504 1004 2000 3984 7936 ...  |
| 8   | octonacci         | 1 1 2 4 8 16 32 64 128 255 509 1016 2028 4048 8080 ...  |
| 9   | nonanacci         | 1 1 2 4 8 16 32 64 128 256 511 1021 2040 4076 8144 ...  |
| 10  | decanacci         | 1 1 2 4 8 16 32 64 128 256 512 1023 2045 4088 8172 ...  |

Povezani niz može se generisati gde su promenjene početne vrednosti: Lukasov niz sabira dve prethodne vrednosti kao Fibonaccijev niz za $n=2$, ali koristi $[2, 1]$ kao svoje početne vrednosti.

# --instructions--

Napiši funkciju koja generiše Fibonačijeve i Lukasove nizove reda `n`. Prvi argument je $n$, drugi broj elemenata koji treba vratiti, a treći određuje tip niza. Ako je treći argument `"f"`, vrati Fibonačijev niz, a ako je `"l"`, vrati Lukasov niz. Rezultat treba da bude niz.

# --hints--

`fib_luc` Trebalo bi da bude funkcija.

```js
assert(typeof fib_luc === 'function');
```

`fib_luc(2,10,"f")` trebalo bi da vrati `[1,1,2,3,5,8,13,21,34,55]`.

```js
assert.deepEqual(fib_luc(2, 10, 'f'), globalThis.ans[0]);
```

`fib_luc(3,15,"f")` treba da vrati `[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136]`.

```js
assert.deepEqual(fib_luc(3, 15, 'f'), globalThis.ans[1]);
```

`fib_luc(4,15,"f")` treba da vrati `[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536]`.

```js
assert.deepEqual(fib_luc(4, 15, 'f'), globalThis.ans[2]);
```

`fib_luc(2,10,"l")` inapaswa kurudisha `[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76]`.

```js
assert.deepEqual(fib_luc(2, 10, 'l'), globalThis.ans[3]);
```

`fib_luc(3,15,"l")` inapaswa kurudisha `[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ]`.

```js
assert.deepEqual(fib_luc(3, 15, 'l'), globalThis.ans[4]);
```

`fib_luc(4,15,"l")` treba da vrati `[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ]`.

```js
assert.deepEqual(fib_luc(4, 15, 'l'), globalThis.ans[5]);
```

`fib_luc(5,15,"l")` treba da vrati `[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]`.

```js
assert.deepEqual(fib_luc(5, 15, 'l'), globalThis.ans[6]);
```

# --before-all--

```js
globalThis.ans = [[1,1,2,3,5,8,13,21,34,55],
[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136],
[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536],
[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76],
[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ],
[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ],
[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]];
```

# --seed--

## --seed-contents--

```js
function fib_luc(n, len, w) {

}
```

# --solutions--

```js
function fib_luc(n, len, w) {
    function nacci(a, n, len) {
        while (a.length < len) {
            let sum = 0;
            for (let i = Math.max(0, a.length - n); i < a.length; i++)
                sum += a[i];
            a.push(sum);
        }
        return a;
    }
    if(w=="f"){
        return nacci(nacci([1,1], n, n), n, len);
    }else{
        return nacci(nacci([2,1], n, n), n, len);
    }
}
```
