---
id: 598eea87e5cf4b116c3ff81a
title: Vigezo vya namba ya Mersenne
challengeType: 1
forumTopicId: 302264
dashedName: factors-of-a-mersenne-number
---

# --description--

Mersenov broj je broj oblika <code>2<sup>P</sup>-1</code>.

Ako je `P` prost broj, Mersenov broj može biti Mersenov prost broj. (Ako `P` nije prost, onda ni Mersenov broj nije prost.)

U potrazi za Mersenovim prostim brojevima korisno je unapred eliminisati kandidate pronalaženjem malog delioca pre pokretanja <a href="https://rosettacode.org/wiki/Lucas-Lehmer_test" target="_blank" rel="noopener noreferrer nofollow">Lucas-Lehmerovog testa</a>, koji može dugo da traje.

Postoje veoma efikasni algoritmi za određivanje da li broj deli <code>2<sup>P</sup>-1</code> (ili, ekvivalentno, da li je <code>2<sup>P</sup> mod broj = 1</code>).

Drugi jezici već imaju ugrađenu implementaciju ove operacije snaga-po-modulu (poznate kao modPow ili slično).

U nastavku je prikazan način kako da samostalno implementiraš `modPow`:

Kwa mfano, hebu tukokotoe <code>2<sup>23</sup> mod 47</code>.

Pretvori eksponent 23 u binarni zapis, dobija se 10111. Počevši sa <code><tt>square</tt> = 1</code>, uzastopno ga kvadriraj.

Ukloni vodeći bit eksponenta, a ako je njegova vrednost 1, pomnoži `square` osnovom stepena (2), zatim izračunaj <code><tt>square</tt> modulo 47</code>.

Koristi rezultat modula iz poslednjeg koraka kao početnu vrednost za `square` u sledećem koraku:

<pre>Ondoa   Hiari
square     kidole cha juu  zidisha kwa 2  mod 47
------------  -------  -------------  ------
1*1 = 1       1  0111  1*2 = 2           2
2*2 = 4       0   111     hapana           4
4*4 = 16      1    11  16*2 = 32        32
32*32 = 1024  1     1  1024*2 = 2048    27
27*27 = 729   1        729*2 = 1458      1
</pre>

Pošto <code>2<sup>23</sup> mod 47 = 1</code>, 47 je faktor od <code>2<sup>P</sup>-1</code>.

(Da bi videli ovo, oduzmi 1 sa obe strane: <code>2<sup>23</sup>-1 = 0 mod 47</code>.)

Pošto smo pokazali da je 47 parametar, <code>2<sup>23</sup>-1</code> nije primarni broj.

Svojstva Mersenovih brojeva omogućavaju nam da dalje optimizujemo proces.

Bilo koji kriterijum `q` iz <code>2<sup>P</sup>-1</code> mora biti u formi `2kP+1`, `k` ako je pozitivan broj ili nula. Štaviše, `q` mora biti `1` ili `7 mod 8`.

Konačno, bilo koji mogući parametar `q` mora biti <a href="https://rosettacode.org/wiki/Primality_by_trial_division" target="_blank" rel="noopener noreferrer nofollow">osnovni broj</a>.

Kao i kod drugih algoritama za testno deljenje, algoritam staje kada `2kP+1 > sqrt(N)`. Ovo posebno radi za brojeve Mersenne gde `P` je osnovni broj. Na primer, <code>M<sub>4</sub>=15</code> nema kriterijume koristeći ove metode, ali je deljiv sa 3 i sa 5, niti jedan od njih ne zadovoljava `2kP+1`.

# --instructions--

Koristeći gore navedeni metod, dođite do parametra <code>2<sup>p</sup>-1</code>.

# --hints--

`check_mersenne` to bi trebalo da bude funkcija.

```js
assert(typeof check_mersenne === 'function');
```

`check_mersenne(3)` Trebalo bi da vrati sekvencu slova.

```js
assert(typeof check_mersenne(3) == 'string');
```

`check_mersenne(3)` treba vratiti sekvencu slova `M3 = 2^3-1 is prime`.

```js
assert.equal(check_mersenne(3), 'M3 = 2^3-1 is prime');
```

`check_mersenne(23)` treba da vrati sekvencu karaktera `M23 = 2^23-1 is composite with factor 47`.

```js
assert.equal(check_mersenne(23), 'M23 = 2^23-1 is composite with factor 47');
```

`check_mersenne(929)` treba da vrati sekvencu karaktera `M929 = 2^929-1 is composite with factor 13007`.

```js
assert.equal(
  check_mersenne(929),
  'M929 = 2^929-1 is composite with factor 13007'
);
```

# --seed--

## --seed-contents--

```js
function check_mersenne(p) {

}
```

# --solutions--

```js
function check_mersenne(p){
    function isPrime(value){
      for (let i=2; i < value; i++){
        if (value % i == 0){
          return false;
        }
        if (value % i != 0){
          return true;
         }
      }
    }

    function trial_factor(base, exp, mod){
      let square, bits;
      square = 1;
      bits = exp.toString(2).split('');
      for (let i=0,ln=bits.length; i<ln; i++){
        square = Math.pow(square, 2) * (bits[i] == 1 ? base : 1) % mod;
      }
      return (square == 1);
    }

    function mersenne_factor(p){
      let limit, k, q;
      limit = Math.sqrt(Math.pow(2,p) - 1);
      k = 1;
      while ((2*k*p - 1) < limit){
        q = 2*k*p + 1;
        if (isPrime(q) && (q % 8 == 1 || q % 8 == 7) && trial_factor(2,p,q)){
          return q; // q is a factor of 2**p-1
        }
        k++;
      }
      return null;
    }
  let f, result;
  result="M"+p+" = 2^"+p+"-1 is ";
  f = mersenne_factor(p);
  result+=f == null ? "prime" : "composite with factor "+f;
  return result;
}
```
