---
id: 599d1566a02b571412643b84
title: Množi na etiopski način
challengeType: 1
forumTopicId: 302257
dashedName: ethiopian-multiplication
---

# --description--

Etiopska metoda množenja je način za množenje celih brojeva koristeći samo sabiranje, umnožiti sa dva i deljenje sa dve.

**Njia:**

<ol>
<li>Uzmi dva broja koja se množe i zapiši ih na vrh dve kolone.</li>
<li>U levoj koloni uzastopno deli poslednji broj sa 2, odbacujući ostatak, i upisuj rezultat ispod prethodnog sve dok ne dobiješ vrednost <code>1</code>.</li>
<li>U desnoj koloni uzastopno udvostručuj poslednji broj i upisuj rezultat ispod prethodnog. Zaustavi se kada u levoj koloni stigneš do <code>1</code>.</li>
<li>Pregledaj dobijenu tabelu i izbaci sve redove u kojima je broj u levoj koloni paran.</li>
<li>Saberi preostale vrednosti iz desne kolone da bi dobio rezultat množenja početna dva broja.</li>
</ol>

**Primer:** `17 × 34`

<pre>17   34
</pre>

Prepolovljavanje prve kolone:

<pre>17   34
8
4
2
1
</pre>

Udvostručavanje druge kolone:

<pre>17   34
8    68
4   136
2   272
1   544
</pre>

Navedite redove čije je prve ćelije paran broj:

<pre>17   34
8    <strike>68</strike>
4   <strike>136</strike>
2   <strike>272</strike>
1   544
</pre>

Saberi preostale brojeve iz desne kolone:

<!-- markdownlint-disable MD003 -->

<pre>17   34
8    --
4   ---
2   ---
1   544
   ====
    578
</pre>

<!-- markdownlint-enable MD003 -->

Tako `17`, povećan za `34`, putem Etiopijana je `578`.

# --instructions--

Vežba je definiranje tri imenovane funkcije/metode/procedure:

<ol>
  <li>funkcija za deljenje celog broja sa dva,</li>
  <li>funkcija za množenje celog broja sa dva, i</li>
  <li>funkcija koja proverava da li je ceo broj paran</li>
</ol>

Koristite ove funkcije za kreiranje funkcije koja množi na etiopski način.

<!-- markdownlint-disable MD046-->

# --hints--

`eth_mult` Trebalo bi da bude funkcija.

```js
assert(typeof eth_mult === 'function');
```

`eth_mult(17,34)` treba da vrati `578`.

```js
assert.equal(eth_mult(17, 34), 578);
```

`eth_mult(23,46)` treba da vrati `1058`.

```js
assert.equal(eth_mult(23, 46), 1058);
```

`eth_mult(12,27)` treba da vrati `324`.

```js
assert.equal(eth_mult(12, 27), 324);
```

`eth_mult(56,98)` treba da vrati `5488`.

```js
assert.equal(eth_mult(56, 98), 5488);
```

`eth_mult(63,74)` treba da vrati `4662`.

```js
assert.equal(eth_mult(63, 74), 4662);
```

# --seed--

## --seed-contents--

```js
function eth_mult(a, b) {

}
```

# --solutions--

```js
function eth_mult(a, b) {
  let sum = 0; a = [a]; b = [b];

  let half = a => a / 2,
    double = a => a * 2,
    is_even = a => a % 2 == 0;

  while (a[0] !== 1) {
    a.unshift(Math.floor(half(a[0])));
    b.unshift(double(b[0]));
  }

  for (let i = a.length - 1; i > 0; i -= 1) {
    if (!is_even(a[i])) {
      sum += b[i];
    }
  }
  return sum + b[0];
}
```
