---
id: 59c3ec9f15068017c96eb8a3
title: Niz Fareya
challengeType: 1
forumTopicId: 302266
dashedName: farey-sequence
---

# --description--

Farejev niz <code>F<sub>n</sub></code> reda `n` predstavlja niz potpuno skraćenih razlomaka između `0` i `1` čiji su imenitelji manji ili jednaki `n`, poređanih po rastućoj vrednosti.

*Mfululizo wa Farey* ponekad se pogrešno naziva *mfululizo wa Farey*.

Kila mfululizo wa Farey:

<ul>
  <li>počinje sa vrednošću 0, predstavljeno delom $ \frac{0}{1} $</li>
  <li>završava sa vrednošću 1, predstavljeno delom $ \frac{1}{1}$.</li>
</ul>

Nizovi Fareya stepena `1` do `5` su:

<ul>
  <li style='list-style: none;'>${\bf\it{F}}_1 = \frac{0}{1}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_2 = \frac{0}{1}, \frac{1}{2}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_3 = \frac{0}{1}, \frac{1}{3}, \frac{1}{2}, \frac{2}{3}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_4 = \frac{0}{1}, \frac{1}{4}, \frac{1}{3}, \frac{1}{2}, \frac{2}{3}, \frac{3}{4}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_5 = \frac{0}{1}, \frac{1}{5}, \frac{1}{4}, \frac{1}{3}, \frac{2}{5}, \frac{1}{2}, \frac{3}{5}, \frac{2}{3}, \frac{3}{4}, \frac{4}{5}, \frac{1}{1}$</li>
</ul>

# --instructions--

Napiši funkciju koja vraća Farejev niz reda `n`. Funkcija treba da prihvati jedan argument `n` i da vrati taj niz kao niz podataka.

# --hints--

`farey` treba da bude funkcija.

```js
assert(typeof farey === 'function');
```

`farey(3)` trebalo bi da vrati red podataka

```js
assert(Array.isArray(farey(3)));
```

`farey(3)` inapaswa kurudisha `['0/1','1/3','1/2','2/3','1/1']`

```js
assert.deepEqual(farey(3),['0/1', '1/3', '1/2', '2/3', '1/1']);
```

`farey(4)` trebalo bi da vrati `['0/1','1/4','1/3','1/2','2/3','3/4','1/1']`

```js
assert.deepEqual(farey(4), ['0/1', '1/4', '1/3', '1/2', '2/3', '3/4', '1/1']);
```

`farey(5)` inapaswa kurudisha `['0/1','1/5','1/4','1/3','2/5','1/2','3/5','2/3','3/4','4/5','1/1']`

```js
assert.deepEqual(farey(5), [
  '0/1',
  '1/5',
  '1/4',
  '1/3',
  '2/5',
  '1/2',
  '3/5',
  '2/3',
  '3/4',
  '4/5',
  '1/1'
]);
```

# --seed--

## --seed-contents--

```js
function farey(n) {

}
```

# --solutions--

```js
function farey(n) {
  const sequence = [{ string: "0/1", float: 0.0 }];
  for (let i = 1; i < n; i++) {
    for (let j = n; j >= i; j--) {
      if (i === 1 || j % i > 0) {
        sequence.push({ string: `${i}/${j}`, float: i / j });
      }
    }
  }
  return sequence
    .sort((a, b) => a.float - b.float)
    .map(e => e.string)
}
```
