---
id: 59d9c6bc214c613ba73ff012
title: SEDOLs
challengeType: 1
forumTopicId: 302305
dashedName: sedols
---

# --description--

<abbr title="Stock Exchange Daily Official List">SEDOL</abbr> je lista identifikacionih brojeva obezbeđenja koje izdaje Londonova berza.

Za svaki broj SEDOL od 6 cifara, izračunaj i dodaj kontrolni broj (checksum). To jest, ako primiš niz znakova na unosu sa leve strane, tvoja funkcija treba da vrati odgovarajući niz sa desne strane:

<pre>
710889 => 7108899
B0YBKJ => B0YBKJ7
406566 => 4065663
B0YBLH => B0YBLH2
228276 => 2282765
B0YBKL => B0YBKL9
557910 => 5579107
B0YBKR => B0YBKR5
585284 => 5852842
B0YBKT => B0YBKT7
B00030 => B000300
</pre>

Proverite da je svaki unos kreiran ispravno, posebno u vezi sa validnim slovima dozvoljenim u SEDOL sekvenci. Vaša funkcija bi trebalo da vrati `null` za nevažeći unos.

# --hints--

`sedol` To treba da bude funkcija.

```js
assert(typeof sedol === 'function');
```

`sedol('a')` trebalo bi da vrati null.

```js
assert(sedol('a') === null);
```

`sedol('710889')` trebalo bi da vrati '7108899'.

```js
assert(sedol('710889') === '7108899');
```

`sedol('BOATER')` trebalo bi da vrati null.

```js
assert(sedol('BOATER') === null);
```

`sedol('228276')` trebalo bi da vrati '2282765'.

```js
assert(sedol('228276') === '2282765');
```

# --seed--

## --seed-contents--

```js
function sedol(input) {
  const checkSum = 0
  
  return checkSum;
}
```

# --solutions--

```js
function sedol(input) {
  const checkDigit = sedolCheckDigit(input);
  if (checkDigit !== null) {
    return input + checkDigit;
  }
  return null;
}

const weight = [1, 3, 1, 7, 3, 9, 1];
function sedolCheckDigit(char6) {
  if (char6.search(/^[0-9BCDFGHJKLMNPQRSTVWXYZ]{6}$/) === -1) {
    return null;
  }

  let sum = 0;
  for (let i = 0; i < char6.length; i++) {
    sum += weight[i] * parseInt(char6.charAt(i), 36);
  }
  const checkSum = (10 - (sum % 10)) % 10;
  return checkSum.toString();
}
```
