---
id: 59da22823d04c95919d46269
title: Ribarji, nazi i izazov majmuna
challengeType: 1
forumTopicId: 302304
dashedName: sailors-coconuts-and-a-monkey-problem
---

# --description--

Pet ribara je doživelo brodolom, završilo na ostrvu i tokom dana sakupilo veliku gomilu kokosa.

Te noći prvi ribar se budi i odlučuje da uzme svoj deo ranije, pa pokušava da podeli gomilu kokosa ravnomerno na pet grupa, ali otkriva da je ostao jedan kokos. Taj kokos baca majmunu, zatim sakriva jednu od pet jednakih grupa kokosa, preostale četiri grupe ponovo spaja u jednu prividnu gomilu i vraća se na spavanje.

Da bih je skratio priču, svaki ribar po redu se budi jednom noću i vrši iste radnje dijeljenja tereta kokosovora na pet grupa, otkrivajući da je ostalo jedno kokosovo jaje i dajući to preostalo majmunu, a zatim spajajući četiri grupe kako bi stvorio jedan teret.

Ujutru (nakon što je svaki od pet ribara tokom noći tajno uradio isto), preostali kokosovi orasi se ravnomerno dele na pet jednakih delova za svakog ribara. Tada se ispostavlja da se gomila može podeliti bez ostatka. (Ujutru nema kokosa za majmuna.)

# --instructions--

Napiši funkciju koja vraća najmanji mogući broj kokosa u početnoj gomili koju je tokom dana sakupilo `N` ribara. **Napomena:** Priča se odvija u svetu u kome se sakupljanje bilo kog broja kokosa tokom dana i višestruka podela gomile mogu obaviti u vremenu koje odgovara priči, tako da to ne utiče na proračun.

# --hints--

`splitCoconuts` treba da bude funkcija.

```js
assert(typeof splitCoconuts === 'function');
```

`splitCoconuts(5)` inapaswa kurudisha 3121.

```js
assert(splitCoconuts(5) === 3121);
```

`splitCoconuts(6)` inapaswa kurudisha 233275.

```js
assert(splitCoconuts(6) === 233275);
```

`splitCoconuts(7)` trebalo bi da vrati 823537.

```js
assert(splitCoconuts(7) === 823537);
```

# --seed--

## --seed-contents--

```js
function splitCoconuts(intSailors) {

  return true;
}
```

# --solutions--

```js
function splitCoconuts(intSailors) {
  let intNuts = intSailors;
  let result = splitCoconutsHelper(intNuts, intSailors);
  while (!result) {
    intNuts += 1;
    result = splitCoconutsHelper(intNuts, intSailors);
  }

  return intNuts;
}

function splitCoconutsHelper(intNuts, intSailors, intDepth) {
  const nDepth = intDepth !== undefined ? intDepth : intSailors;
  const portion = Math.floor(intNuts / intSailors);
  const remain = intNuts % intSailors;

  if (portion <= 0 || remain !== (nDepth ? 1 : 0)) {
    return null;
  }

  if (nDepth) {
    return splitCoconutsHelper(
      intNuts - portion - remain, intSailors, nDepth - 1
    );
  }

  return intNuts;
}
```
