---
id: 59e0a8df964e4540d5abe599
title: Tekeleza Brain****
challengeType: 1
forumTopicId: 302261
dashedName: execute-brain
---

# --description--

Napisati funkciju koja implementira interpetator za Brain\*\*\*.\nTa funkcija će primiti sekvencu karaktera kao argument i mora vratiti sekvencu karaktera kao rezultat.\nViše detalja je ispod:

RCBF je skup <a href="https://rosettacode.org/wiki/Brainf***" target="_blank" rel="noopener noreferrer nofollow">prevodi i kolekcija Brainf\*\*\*</a>napisani na raznim jezicima za Rosetta Kod.

Niže su linkovi za sve verzije RCBF.

Implementacija zahteva samo tačno izvršenje sledećih instrukcija:

| Amri | Maelezo |
| :--- | :--- |
| <code>></code> | Premesti kursor desno |
| <code>&lt;</code> | Premesti kursor levo |
| <code>+</code> | Povećaj vrednost memorijske ćelije ispod kursora |
| <code>-</code> | Smanji vrednost memorijske ćelije ispod kursora |
| <code>.</code> | Izvuci karakter predstavljen ćelijom koja se nalazi ispod kursora |
| <code>,</code> | Ubaci karakter i sačuvaj ga u ćeliju koja se nalazi ispod kursora |
| <code>\[</code> | Preskoci napred do odgovarajućeg <code>]</code> ako je ćelija ispod kursora 0 |
| <code>]</code> | Preskoci nazad do odgovarajućeg <code>\[</code> ako ćelija ispod kursora nije 0 |

Dozvoljena je ćelija bilo koje veličine, podrška za EOF (*E*nd-*O*-*F*ile) je opciona, kao što su ograničena ili neograničena memorija.

# --hints--

`brain(globalThis.hello)` Trebalo bi da vrati sekvencu znakova

```js
assert(typeof brain(globalThis.hello) === 'string');
```

`brain("++++++[>++++++++++<-]>+++++.")` Trebalo bi da vrati "A"

```js
assert.equal(brain('++++++[>++++++++++<-]>+++++.'), 'A');
```

`brain(globalThis.hello)` trebalo bi da vrati `Hello World!\n`

```js
assert.equal(brain(globalThis.hello), 'Hello World!\n');
```

`brain(globalThis.fib)` treba da vrati `1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89`

```js
assert.equal(brain(globalThis.fib), '1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89');
```

# --before-all--

```js
globalThis.fib=`+

++

+++

++++

+>+>>

>>++++

+++++++

++++++++

+++++++++

++++++++++

++++++>++++

++++++++++++

+++++++++++++

+++<<<<<<[>[>>

>>>>+>+<<<<<<<-

]>>>>>>>[<<<<<<<

+>>>>>>>-]<[>++++

++++++[-<-[>>+>+<<

<-]>>>[<<<+>>>-]+<[

>[-]<[-]]>[<<[>>>+<<

<-]>>[-]]<<]>>>[>>+>+

<<<-]>>>[<<<+>>>-]+<[>

[-]<[-]]>[<<+>>[-]]<<<<

<<<]>>>>>[++++++++++++++

+++++++++++++++++++++++++

+++++++++.[-]]++++++++++<[

->-<]>+++++++++++++++++++++

+++++++++++++++++++++++++++.

[-]<<<<<<<<<<<<[>>>+>+<<<<-]>

>>>[<<<<+>>>>-]<-[>>.>.<<<[-]]

<<[>>+>+<<<-]>>>[<<<+>>>-]<<[<+

>-]>[<+>-]<<<-]`;
globalThis.hello='++++++++[>++++[>++>+++>+++>+<<<<-]>+>+>->>+[<]<-]>>.>---.+++++++..+++.>>.<-.<.+++.------.--------.>>+.>++.'
```

# --seed--

## --seed-contents--

```js
function brain(prog) {

}
```

# --solutions--

```js
function brain(prog){
  var output="";
    var code; // formatted code
  var ip = 0; // current instruction within code
  var nest = 0; // current bracket nesting (for Out button)
  var ahead = []; // locations of matching brackets

  var data = [0]; // data array (mod by +, -)
  var dp = 0; // index into data (mod by <, >)

  var inp = 0; // current input character (fetch with ,)
  var quit = 0;
    var commands = {
    '>':function() { if (++dp >= data.length) data[dp]=0 },
    '<':function() { if (--dp < 0) quit++ },
    '+':function() { ++data[dp] },
    '-':function() { --data[dp] },
    '[':function() { if (!data[dp]) ip = ahead[ip]; else ++nest },
    ']':function() { if ( data[dp]) ip = ahead[ip]; else --nest },
    ',':function() {
        var c = document.getElementById("input").value.charCodeAt(inp++);
        data[dp] = isNaN(c) ? 0 : c; // EOF: other options are -1 or no change
    },
    '.':function() {
            output+=String.fromCharCode(data[dp]);
            /*var s = document.getElementById("output").innerHTML)
             + String.fromCharCode(data[dp]);
            s = s.replace(/\n/g,"<br>").replace(/ /g,"&amp;nbsp;");
            document.getElementById("output").innerHTML = s;*/
        },
    };

    let ar=prog.split('');
    var st = [], back, error = -1;
    for (ip=0; ip<ar.length; ip++) {
        switch(ar[ip]) {
        case '[':
            st.push(ip);
            break;
        case ']':
            if (st.length == 0) error = ip;
            back = st.pop();
            ahead[ip] = back;
            ahead[back] = ip;
            break;
        }
    }

    for(ip=0;ip<ar.length;ip++){
    if(commands.hasOwnProperty(ar[ip]))
          commands[ar[ip]]();
    }

    return output;
}
```
