---
id: 59f40b17e79dbf1ab720ed7a
title: Brojevi za odeljenja
challengeType: 1
forumTopicId: 302249
dashedName: department-numbers
---

# --description--

Postoji grad dizajniran sa odličnom organizacijom, koji je odlučio da dodeli brojeve svim svojim odeljenjima:

<ul>
  <li>Idara ya polisi</li>
  <li>Idara ya usafi</li>
  <li>Idara ya zima moto</li>
</ul>

Svaka odeljenost može imati broj između 1 i 7 (uključujući).

Tri broja za odeljenje moraju biti jedinstveni (različiti međusobno) i moraju da se poveća do broja 12.

Šef policije ne voli brojeve koji nisu za grešku i želi da njegov odeljak ima standardizovan broj.

# --instructions--

Napišite program koji vraća sve validne kombinacije kao niz podataka.

```js
[2, 3, 7] [2, 4, 6] [2, 6, 4]
[2, 7, 3] [4, 1, 7] [4, 2, 6]
[4, 3, 5] [4, 5, 3] [4, 6, 2]
[4, 7, 1] [6, 1, 5] [6, 2, 4]
[6, 4, 2] [6, 5, 1]
```

# --before-each--

```js
const nums = [1, 2, 3, 4, 5, 6, 7];
const total = 12;
const len = 14;
const result = [
  [2, 3, 7],
  [2, 4, 6],
  [2, 6, 4],
  [2, 7, 3],
  [4, 1, 7],
  [4, 2, 6],
  [4, 3, 5],
  [4, 5, 3],
  [4, 6, 2],
  [4, 7, 1],
  [6, 1, 5],
  [6, 2, 4],
  [6, 4, 2],
  [6, 5, 1]
];
```

# --hints--

`combinations` Trebalo bi da bude funkcija.

```js
assert(typeof combinations === 'function');
```

`combinations([1, 2, 3], 6)` treba da vrati red podataka.

```js
assert(Array.isArray(combinations([1, 2, 3], 6)));
```

`combinations([1, 2, 3, 4, 5, 6, 7], 12)` treba da vrati niz dužine 14.

```js
assert(combinations(nums, total).length === len);
```

`combinations([1, 2, 3, 4, 5, 6, 7], 12)` treba da vrati sve važeće kombinacije.

```js
assert.deepEqual(combinations(nums, total), result);
```

# --seed--

## --seed-contents--

```js
function combinations(possibleNumbers, total) {

  return true;
}
```

# --solutions--

```js
function combinations(possibleNumbers, total) {
  let firstNumber;
  let secondNumber;
  let thirdNumber;
  const allCombinations = [];

  for (let i = 0; i < possibleNumbers.length; i += 1) {
    firstNumber = possibleNumbers[i];

    if (firstNumber % 2 === 0) {
      for (let j = 0; j < possibleNumbers.length; j += 1) {
        secondNumber = possibleNumbers[j];

        if (j !== i && firstNumber + secondNumber <= total) {
          thirdNumber = total - firstNumber - secondNumber;

          if (thirdNumber !== firstNumber && thirdNumber !== secondNumber && possibleNumbers.includes(thirdNumber)) {
            allCombinations.push([firstNumber, secondNumber, thirdNumber]);
          }
        }
      }
    }
  }
  return allCombinations;
}
```
