---
id: 5a23c84252665b21eecc8010
title: Algoritmi za sortiranje / shell sort
challengeType: 1
forumTopicId: 302317
dashedName: sorting-algorithmsshell-sort
---

# --description--

Napiši funkciju koja sortira niz elemenata koristeći Shell sort, algoritam sortiranja sa opadajućim razmacima. Funkcija treba da vrati sortirani niz.

Shell sort (poznat i kao Shellsort ili Shellova metoda) dobio je ime po svom autoru, Donaldu Shellu, koji je objavio ovaj algoritam 1959. godine.

Shell sort je niz prolaza insertion sorta nad podnizovima definisanim razmakom. Veličina razmaka se smanjuje nakon svakog prolaza sve dok ne postane 1.

Za korak 1, sortiranje je uobičajen tip umetanja elemenata, ali ovoga puta podaci su osigurani da su skoro sortirani, što predstavlja "najbolje stanje" za umetanje elemenata.

Svaki niz razmaka će uspešno sortirati podatke sve dok se završava sa 1, ali neki daju bolje rezultate od drugih.

Eksperimentalna istraživanja pokazala su da geometrijski niz razmaka sa odnosom približno 2,2 daje dobre rezultate u praksi.

# --hints--

`shellSort` trebalo bi da bude funkcija.

```js
assert(typeof shellSort == 'function');
```

`shellSort([25, 32, 12, 7, 20])` treba da vrati niz.

```js
assert(Array.isArray(shellSort([25, 32, 12, 7, 20])));
```

`shellSort([25, 32, 12, 7, 20])` inapaswa kurudisha `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(shellSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`shellSort([38, 45, 35, 8, 13])` inapaswa kurudisha `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(shellSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`shellSort([43, 36, 20, 34, 24])` inapaswa kurudisha `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(shellSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`shellSort([12, 33, 26, 18, 1, 16, 38])` inapaswa kurudisha `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(shellSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`shellSort([3, 39, 48, 16, 1, 4, 29])` treba da vrati `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(shellSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function shellSort(a) {

}
```

# --solutions--

```js
function shellSort(a) {
  for (var h = a.length; h > 0; h = parseInt(h / 2)) {
    for (var i = h; i < a.length; i++) {
      var k = a[i];
      for (var j = i; j >= h && k < a[j - h]; j -= h) a[j] = a[j - h];
      a[j] = k;
    }
  }
  return a;
}
```
