---
id: 5a23c84252665b21eecc8046
title: Razlika u ravnoteži
challengeType: 1
forumTopicId: 16086
dashedName: symmetric-difference
---

# --description--

Razmatrajući dva skupa $A$ i $B$, izračunajte $(A \setminus B) \cup (B \setminus A)$. To jest, navedite elemente koji su u $A$ ili $B$, ali ne i u oba. Ovaj skup se naziva simetrična razlika $A$ i $B$. Drugim rečima: $(A \cup B) \setminus (A \cap B)$ (skup elemenata koji su u najmanje jednom od $A$ ili $B$, izuzimajući skup elemenata koji su u oba, $A$ i $B$).

Primer:

Za set `A = [1, 2, 3]`, na `B = [1, 3, 4]`, razlika između *A* i *B* je `[2, 4]`.

# --instructions--

Napiši funkciju koja prima dva niza podataka kao parametre i vraća razliku. Sortiraj niz rezultata pre povratka.

# --hints--

`symmetricDifference` Trebalo bi da bude funkcija.

```js
assert(typeof symmetricDifference == 'function');
```

`symmetricDifference(["John", "Bob", "Mary", "Serena"], ["Jim", "Mary", "John", "Bob"])` Trebalo bi da vrati red podataka.

```js
assert(
  Array.isArray(
    symmetricDifference(
      ['John', 'Bob', 'Mary', 'Serena'],
      ['Jim', 'Mary', 'John', 'Bob']
    )
  )
);
```

`symmetricDifference(["John", "Bob", "Mary", "Serena"], ["Jim", "Mary", "John", "Bob"])` treba da vrati `["Jim", "Serena"]`.

```js
assert.deepEqual(
  symmetricDifference(
    ['John', 'Bob', 'Mary', 'Serena'],
    ['Jim', 'Mary', 'John', 'Bob']
  ),
  ['Jim', 'Serena']
);
```

`symmetricDifference([1, 2, 3], [3, 4])` treba da vrati `[1, 2, 4]`.

```js
assert.deepEqual(symmetricDifference([1, 2, 3], [3, 4]), [1, 2, 4]);
```

`symmetricDifference([1, 2, 3, 4, 5], [3, 4, 8, 7])` treba da vrati `[1, 2, 5, 7, 8]`.

```js
assert.deepEqual(symmetricDifference([1, 2, 3, 4, 5], [3, 4, 8, 7]), [
  1,
  2,
  5,
  7,
  8
]);
```

`symmetricDifference([1, 2, 3, 4, 5, 6, 7, 8], [1, 3, 5, 6, 7, 8, 9])` trebalo bi da vrati `[2, 4, 9]`.

```js
assert.deepEqual(
  symmetricDifference([1, 2, 3, 4, 5, 6, 7, 8], [1, 3, 5, 6, 7, 8, 9]),
  [2, 4, 9]
);
```

`symmetricDifference([1, 2, 4, 7, 9], [2, 3, 7, 8, 9])` treba da vrati `[1, 3, 4, 8]`.

```js
assert.deepEqual(symmetricDifference([1, 2, 4, 7, 9], [2, 3, 7, 8, 9]), [
  1,
  3,
  4,
  8
]);
```

# --seed--

## --seed-contents--

```js
function symmetricDifference(A, B) {

}
```

# --solutions--

```js
function symmetricDifference(A, B) {
  function relative_complement(A, B) {
    return A.filter(function(elem) {
      return B.indexOf(elem) == -1;
    });
  }

  function unique(ary) {
    var u = ary.concat().sort();
    for (var i = 1; i < u.length; ) {
      if (u[i - 1] === u[i]) u.splice(i, 1);
      else i++;
    }
    return u;
  }

  return unique(
    relative_complement(A, B).concat(relative_complement(B, A))
  ).sort();
}
```
