---
id: 5a23c84252665b21eecc8036
title: Uklonite kontrolni kod i proširene slova iz niza znakova
challengeType: 1
forumTopicId: 302327
dashedName: strip-control-codes-and-extended-characters-from-a-string
---

# --description--

Vežba je uklanjanje kontrolnih kodova i proširenih znakova iz niza karaktera. Rešenje treba da pokaže kako postići svaki od sledećih rezultata: Niz karaktera iz kojeg su uklonjeni kontrolni kodovi i prošireni znakovi. U ASCII-ju, kontrolni kodovi imaju decimalne brojeve od 0 do 31 i 127. U sistemu koji zavisi od ASCII-ja, ako se kontrolni kodovi uklone, niz će sadržati samo karaktere unutar decimalnog opsega od 32 do 126 u ASCII tabeli. U sistemu koji ne zavisi od ASCII-ja, smatramo da su karakteri koji nemaju odgovarajući simbol u ASCII tabeli (unutar decimalnog opsega 32 do 126) prošireni znakovi radi svrhe ove vežbe.

# --hints--

`strip` Trebalo bi da bude funkcija.

```js
assert(typeof strip == 'function');
```

`strip("abc")` trebalo bi da vrati sekvencu znakova.

```js
assert(typeof strip('abc') == 'string');
```

`strip("\ba\x00b\n\rc\fd\xc3")` treba da vrati `"abcd"`.

```js
assert.equal(strip('\ba\x00b\n\rc\fd\xc3'), 'abcd');
```

`strip("\u0000\n abc\u00E9def\u007F")` trebalo bi da vrati `" abcdef"`.

```js
assert.equal(strip('\u0000\n abc\u00E9def\u007F'), ' abcdef');
```

`strip("a\n\tb\u2102d\u2147f")` trebalo bi vratiti `"abdf"`.

```js
assert.equal(strip('a\n\tb\u2102d\u2147f'), 'abdf');
```

`strip("Français.")` treba da vrati `"Franais."`.

```js
assert.equal(strip('Français.'), 'Franais.');
```

`strip("123\tabc\u0007DEF\u007F+-*/€æŧðłþ")` treba da vrati `"123abcDEF+-*/"`.

```js
assert.equal(strip('123\tabc\u0007DEF\u007F+-*/€æŧðłþ'), '123abcDEF+-*/');
```

# --seed--

## --seed-contents--

```js
function strip(s) {

}
```

# --solutions--

```js
function strip(s) {
  return s
    .split('')
    .filter(function(x) {
      var n = x.charCodeAt(0);

      return 31 < n && 127 > n;
    })
    .join('');
}
```
