---
id: 5a23c84252665b21eecc7ee0
title: Fabrika sa leve strane
challengeType: 1
forumTopicId: 302302
dashedName: left-factorials
---

# --description--

**Faktorijal sa leve strane**, $ !n $, može značiti *subfaktorijal* ili *zbir faktorijala*. Taj isti simbol može se koristiti na zbunjujući način u dva različita značenja. Ponekad, *subfaktorijali* (takođe poznati kao *derangementi*) mogu koristiti jedan od ovih simbola:

<ul>
  <li>$!n`$</li>
  <li>$!n$</li>
  <li>$n¡$</li>
</ul>

(Možda nije očigledno golim okom, ali poslednji primer koristi usmeren znak uzvika.) Ova vežba će koristiti ovu formulu za **faktoria ya kushoto**:

$ !n = \\sum\_{k=0}^{n-1} k! $

gde je $!0 = 0$

# --instructions--

Napišite funkciju koja računa faktorijel datevenog broja.

# --hints--

`leftFactorial` Trebalo bi da bude funkcija.

```js
assert(typeof leftFactorial == 'function');
```

`leftFactorial(0)` trebalo bi da vrati broj.

```js
assert(typeof leftFactorial(0) == 'number');
```

`leftFactorial(0)` treba da vrati `0`.

```js
assert.equal(leftFactorial(0), 0);
```

`leftFactorial(1)` trebalo bi da vrati `1`.

```js
assert.equal(leftFactorial(1), 1);
```

`leftFactorial(2)` treba da vrati `2`.

```js
assert.equal(leftFactorial(2), 2);
```

`leftFactorial(3)` trebalo bi da vrati `4`.

```js
assert.equal(leftFactorial(3), 4);
```

`leftFactorial(10)` inapaswa kurudisha `409114`.

```js
assert.equal(leftFactorial(10), 409114);
```

`leftFactorial(17)` inapaswa kurudisha `22324392524314`.

```js
assert.equal(leftFactorial(17), 22324392524314);
```

`leftFactorial(19)` treba da vrati `6780385526348314`.

```js
assert.equal(leftFactorial(19), 6780385526348314);
```

# --seed--

## --seed-contents--

```js
function leftFactorial(n) {

}
```

# --solutions--

```js
function leftFactorial(n) {
  if (n == 0) return 0;
  if (n == 1) return 1;

  // Note: for n>=20, the result may not be correct.
  // This is because JavaScript uses 53 bit integers and
  // for n>=20 result becomes too large.

  let res = 2,
    fact = 2;
  for (var i = 2; i < n; i++) {
    res += fact;
    fact *= i + 1;
  }

  return res;
}
```
