---
id: 5a23c84252665b21eecc7edc
title: Poslednji petak svakog meseca
challengeType: 1
forumTopicId: 302299
dashedName: last-friday-of-each-month
---

# --description--

Napiši funkciju koja vraća datum poslednjeg petka u određenom mesecu određene godine.

# --hints--

`lastFriday` Trebalo bi da bude funkcija.

```js
assert(typeof lastFriday == 'function');
```

`lastFriday(2018, 1)` Trebalo bi da vrati broj.

```js
assert(typeof lastFriday(2018, 1) == 'number');
```

`lastFriday(2018, 1)` treba da restaurira `26`.

```js
assert.equal(lastFriday(2018, 1), 26);
```

`lastFriday(2017, 2)` inapaswa kurejesha `24`.

```js
assert.equal(lastFriday(2017, 2), 24);
```

`lastFriday(2012, 3)` treba da vrati `30`.

```js
assert.equal(lastFriday(2012, 3), 30);
```

`lastFriday(1900, 4)` inapaswa kurejesha `27`.

```js
assert.equal(lastFriday(1900, 4), 27);
```

`lastFriday(2000, 5)` treba da restaurira `26`.

```js
assert.equal(lastFriday(2000, 5), 26);
```

`lastFriday(2006, 6)` treba da obnovi `30`.

```js
assert.equal(lastFriday(2006, 6), 30);
```

`lastFriday(2010, 7)` treba da se ponovo postavi `30`.

```js
assert.equal(lastFriday(2010, 7), 30);
```

`lastFriday(2005, 8)` trebalo bi da restaurira `26`.

```js
assert.equal(lastFriday(2005, 8), 26);
```

# --seed--

## --seed-contents--

```js
function lastFriday(year, month) {

}
```

# --solutions--

```js
function lastFriday(year, month) {
  var i, last_day;
  i = 0;
  while (true) {
    last_day = new Date(year, month, i);
    if (last_day.getDay() === 5) {
      return last_day.getDate();
    }
    i -= 1;
  }
}
```
