---
id: 5a23c84252665b21eecc7edf
title: Najmanši broj koji je deljiv sa svima
challengeType: 1
forumTopicId: 302301
dashedName: least-common-multiple
---

# --description--

Najmanji broj koji je deljiv sa 12 i 18 je 36, jer je 12 delilatelj (12 × 3 = 36), a 18 je delilatelj (18 × 2 = 36), i ne postoji manji pozitivan ceo broj od 36 koji ima oba delitelja. Kao poseban slučaj, ako su m ili n nula, onda je najmanji broj koji je deljiv sa oba nula. Jedan način za izračunavanje najmanjeg broja koji je deljiv sa oba je da se ponavljano množi m dok ne dobijete broj koji je takođe višestrukost n. Ako već imate gcd za <a href="https://rosettacode.org/wiki/Greatest_common_divisor" target="_blank" rel="noopener noreferrer nofollow">mgawanyiko mkubwa kabisa</a>, onda ova formula izračunava lcm.

$$
\\operatorname{lcm}(m, n) = \\frac{|m \\times n|}{\\operatorname{gcd}(m, n)}
$$

# --instructions--

Izračunaj najmanji zajednički sadržalac svih brojeva u nizu celih brojeva. Za brojeve *m* i *n*, najmanji zajednički sadržalac je najmanji pozitivan ceo broj koji je deljiv i sa *m* i sa *n*.

# --hints--

`LCM` Trebalo bi da bude funkcija.

```js
assert(typeof LCM == 'function');
```

`LCM([2, 4, 8])` treba da vrati broj.

```js
assert(typeof LCM([2, 4, 8]) == 'number');
```

`LCM([2, 4, 8])` inapaswa kurudisha `8`.

```js
assert.equal(LCM([2, 4, 8]), 8);
```

`LCM([4, 8, 12])` treba da vrati `24`.

```js
assert.equal(LCM([4, 8, 12]), 24);
```

`LCM([3, 4, 5, 12, 40])` inapaswa kurudisha `120`.

```js
assert.equal(LCM([3, 4, 5, 12, 40]), 120);
```

`LCM([11, 33, 90])` inapaswa kurudisha `990`.

```js
assert.equal(LCM([11, 33, 90]), 990);
```

`LCM([-50, 25, -45, -18, 90, 447])` trebalo bi vratiti `67050`.

```js
assert.equal(LCM([-50, 25, -45, -18, 90, 447]), 67050);
```

# --seed--

## --seed-contents--

```js
function LCM(A) {

}
```

# --solutions--

```js
function LCM(A) {
  var n = A.length,
    a = Math.abs(A[0]);
  for (var i = 1; i < n; i++) {
    var b = Math.abs(A[i]),
      c = a;
    while (a && b) {
      a > b ? (a %= b) : (b %= a);
    }
    a = Math.abs(c * A[i]) / (a + b);
  }
  return a;
}
```
