---
id: 5a23c84252665b21eecc8002
title: Algoritmi za sortiranje/sorting Bogosort
challengeType: 1
forumTopicId: 302311
dashedName: sorting-algorithmsbogosort
---

# --description--

Sortiraj koristeći Bogosort na niz brojeva.

Bogosort je metoda za nasumično mešanje kolekcije sve dok se ne sortira.

"Bogosort" je potpuno beskoristan algoritam koji se koristi samo kao šala.

Njegovo prosečno vreme izvršavanja je O(n!) jer je verovatnoća da bilo koja kombinacija skupa završi u ispravnom redosledu približno jedna od faktoriela *n*, a najgori slučaj je beskonačan jer nema garancije da će nasumična kombinacija proizvesti poretan niz.

Najbolji slučaj je O(n) jer jedan prolaz kroz elemente može biti dovoljan da ih poređane (ili: sortiraju).

Lažni kod:

<pre><b>kada nije</b> InOrder(list) <b>uraditi</b>
  Shuffle(list)
<b>završiti</b>
</pre>

# --hints--

`bogosort` Trebalo bi da bude funkcija.

```js
assert(typeof bogosort == 'function');
```

`bogosort([25, 32, 12, 7, 20])` trebalo bi da vrati red podataka.

```js
assert(Array.isArray(bogosort([25, 32, 12, 7, 20])));
```

`bogosort([25, 32, 12, 7, 20])` trebalo bi da vrati `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(bogosort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`bogosort([38, 45, 35, 8, 13])` treba da vrati `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(bogosort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`bogosort([43, 36, 20, 34, 24])` trebalo bi da vrati `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(bogosort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`bogosort([12, 33, 26, 18, 1, 16, 38])` trebalo bi da vrati `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(bogosort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`bogosort([3, 39, 48, 16, 1, 4, 29])` treba da vrati `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(bogosort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function bogosort(v) {

}
```

# --solutions--

```js
function bogosort(v) {
  function shuffle(v) {
    for (
      var j, x, i = v.length;
      i;
      j = Math.floor(Math.random() * i), x = v[--i], v[i] = v[j], v[j] = x
    );
    return v;
  }

  function isSorted(v) {
    for (var i = 1; i < v.length; i++) {
      if (v[i - 1] > v[i]) {
        return false;
      }
    }
    return true;
  }
  var sorted = false;
  while (sorted == false) {
    v = shuffle(v);
    sorted = isSorted(v);
  }
  return v;
}
```
