---
id: 5a23c84252665b21eecc7ec4
title: JortSort
challengeType: 1
forumTopicId: 302293
dashedName: jortsort
---

# --description--

jortSort je alat za sortiranje koji omogućava korisniku da radi i osigurava efikasnost jer više nikada neće morati ponovo da sortiraju. Originalno je predstavljen od strane Jen "Moneydollars" Schiffer na poznatom kongresu JSConf2014.

jortSort bi trebalo da bude funkcija koja prima niz uporedivih elemenata kao svoj argument. Zatim sortira taj niz u uzlaznom redosledu i upoređuje sortirani niz sa prvobitnim nizom. Ako su nizovi identični (što znači da je originalni niz već bio sortiran), funkcija vraća true. Ako nisu identični (što znači da originalni niz nije bio sortiran), funkcija vraća false.

# --hints--

`jortsort` Trebalo bi da bude funkcija.

```js
assert(typeof jortsort == 'function');
```

`jortsort([1,2,3,4,5])` treba da vrati vrednost tačno ili netačno.

```js
assert(typeof jortsort([1, 2, 3, 4, 5]) == 'boolean');
```

`jortsort([1,2,3,4,5])` treba da restaurira `true`.

```js
assert.equal(jortsort([1, 2, 3, 4, 5]), true);
```

`jortsort([1,2,13,4,5])` treba da restaurira `false`.

```js
assert.equal(jortsort([1, 2, 13, 4, 5]), false);
```

`jortsort([12,4,51,2,4])` treba da vrati `false`.

```js
assert.equal(jortsort([12, 4, 51, 2, 4]), false);
```

`jortsort([1,2])` treba da vrati `true`.

```js
assert.equal(jortsort([1, 2]), true);
```

`jortsort([5,4,3,2,1])` treba da resta/vrati `false`.

```js
assert.equal(jortsort([5, 4, 3, 2, 1]), false);
```

`jortsort([1,1,1,1,1])` treba da vrati `true`.

```js
assert.equal(jortsort([1, 1, 1, 1, 1]), true);
```

# --seed--

## --seed-contents--

```js
function jortsort(array) {

}
```

# --solutions--

```js
function jortsort(array) {
  // sort the array
  var originalArray = array.slice(0);
  array.sort( function(a,b){return a - b} );

  // compare to see if it was originally sorted
  for (var i = 0; i < originalArray.length; ++i) {
    if (originalArray[i] !== array[i]) return false;
  }

  return true;
};
```
