---
id: 5a23c84252665b21eecc7e7b
title: izvori/brojevi visine
challengeType: 1
forumTopicId: 302275
dashedName: generatorexponential
---

# --description--

Generator je izvršivi objekat (kao funkcija ili procedura) sa kodom koji proizvodi sekvencu vrednosti, tako da svaki put kada pozovete generator, vraća se sledeća vrednost u sekvenci.

Generatori se često grade na coroutine-ima ili objektima kako bi je unutrašnje stanje tog objekta održavano "po defaultu".

Generatori se često koriste u situacijama gde sekvenca može biti beskonačna, i gde je moguće izračunati sledeću vrednost sekvence koristeći samo ograničeno stanje.

# --instructions--

Napiši funkciju koja koristi generatore za generisanje kvadrata i brojeva visine. Kreiraj novi generator koji filtrira sve brojeve visine iz generatora kvadrata.

Funkcija mora da vrati vrednost od \\( n^{th} \\) iz filtrisanog izvora.

Na primer, za \\(n=7\\), funkcija treba da vrati 81, jer bi niz bio 4, 9, 16, 25, 36, 49, 81. Broj 64 je izostavljen zato što je kubni broj.

# --hints--

`exponentialGenerator` Trebalo bi da bude funkcija.

```js
assert(typeof exponentialGenerator == 'function');
```

`exponentialGenerator()` Trebalo bi da vrati broj.

```js
assert(typeof exponentialGenerator(10) == 'number');
```

`exponentialGenerator(10)` trebalo bi da vrati `144`.

```js
assert.equal(exponentialGenerator(10), 144);
```

`exponentialGenerator(12)` inapaswa kurudisha `196`.

```js
assert.equal(exponentialGenerator(12), 196);
```

`exponentialGenerator(14)` trebalo bi da vrati `256`.

```js
assert.equal(exponentialGenerator(14), 256);
```

`exponentialGenerator(20)` treba vratiti `484`.

```js
assert.equal(exponentialGenerator(20), 484);
```

`exponentialGenerator(25)` bi trebalo da vrati `784`.

```js
assert.equal(exponentialGenerator(25), 784);
```

# --seed--

## --seed-contents--

```js
function exponentialGenerator(n) {

}
```

# --solutions--

```js
function exponentialGenerator(n){
  function* PowersGenerator(m) {
    var n=0;
    while(1) {
        yield Math.pow(n, m);
        n += 1;
    }
  }

  function* FilteredGenerator(g, f){
    var value = g.next().value;
    var filter = f.next().value;
    while(1) {
        if( value < filter ) {
            yield value;
            value = g.next().value;
        } else if ( value > filter ) {
            filter = f.next().value;
        } else {
            value = g.next().value;
            filter = f.next().value;
        }
    }
  }

  var squares = PowersGenerator(2);
  var cubes = PowersGenerator(3);

  var filtered = FilteredGenerator(squares, cubes);

  var curr=0;
  for(var i=0;i<n;i++) curr=filtered.next();

  return curr.value;
}
```
