---
id: 5a23c84252665b21eecc7ec2
title: umbali wa Jaro
challengeType: 1
forumTopicId: 302292
dashedName: jaro-distance
---

# --description--

Jaro rastojanje je mera sličnosti između dve niske. Što je Jaro rastojanje veće, niske su sličnije. Rezultat je normalizovan tako da `0` označava potpunu različitost, a `1` potpuno poklapanje.

**Ufafanuzi**

Jaro rastojanje \\( d_j \\) između dve niske \\(s_1\\) i \\(s_2\\) iznosi

\\begin{align}d_j = \\begin{cases}0& & \\text{ikiwa }m=0 \\\\\\\\{\\frac {1}{3}}\\left({\\frac {m}{|s\_{1}|}}+{\\frac {m}{|s\_{2}|}}+{\\frac {m-t}{m}}\\right)& & \\text{vinginevyo}\\end{cases}\\end{align}

Gde:

<ul>
  <li>\(m\) ni idadi ya <i>herufi zinazofanana</i>;</li>
<li>\(t\) predstavlja polovinu broja <i>transpozicija</i>.</li>
</ul>

Dva znaka iz \\(s_1\\) i \\(s_2\\) smatraju se *podudarnim* samo ako su jednaki i nisu udaljeni više od \\(\\left\\lfloor\\frac{\\max(|s_1|,|s_2|)}{2}\\right\\rfloor-1\\).

Svaki znak iz \\(s_1\\) poredi se sa svim odgovarajućim znakovima u \\(s_2\\). Broj podudarnih znakova koji se nalaze na različitim pozicijama, podeljen sa 2, predstavlja broj *transpozicija*.

**Mfano**

Ako su niske \\(s_1\\) *DWAYNE* i \\(s_2\\) *DUANE*, dobijamo:

<ul>
  <li>\(m = 4\)</li>
  <li>\(|s_1| = 6\)</li>
  <li>\(|s_2| = 5\)</li>
  <li>\(t = 0\)</li>
</ul>

Tunapata alama ya Jaro: \\(d_j = \\frac{1}{3}\\left(\\frac{4}{6} + \\frac{4}{5} + \\frac{4-0}{4}\\right) = 0.822\\).

# --instructions--

Napiši funkciju koja prihvata dve niske kao argumente i vraća odgovarajuće Jaro rastojanje.

# --hints--

`jaro` Trebalo bi da bude funkcija.

```js
assert(typeof jaro == 'function');
```

`jaro("MARTHA", "MARHTA")` to treba da vrati broj.

```js
assert(typeof jaro('MARTHA', 'MARHTA') == 'number');
```

`jaro("MARTHA", "MARHTA")` inapaswa kurudisha `0.9444444444444445`.

```js
assert.equal(jaro('MARTHA', 'MARHTA'), 0.9444444444444445);
```

`jaro("DIXON", "DICKSONX")` inapaswa kurudisha `0.7666666666666666`.

```js
assert.equal(jaro('DIXON', 'DICKSONX'), 0.7666666666666666);
```

`jaro("JELLYFISH", "SMELLYFISH")` inapaswa kurudisha `0.8962962962962964`.

```js
assert.equal(jaro('JELLYFISH', 'SMELLYFISH'), 0.8962962962962964);
```

`jaro("HELLOS", "CHELLO")` inapaswa kurudisha `0.888888888888889`.

```js
assert.equal(jaro('HELLOS', 'CHELLO'), 0.888888888888889);
```

`jaro("ABCD", "BCDA")` inapaswa kurudisha `0.8333333333333334`.

```js
assert.equal(jaro('ABCD', 'BCDA'), 0.8333333333333334);
```

# --seed--

## --seed-contents--

```js
function jaro(s, t) {

}
```

# --solutions--

```js
function jaro(s, t) {
  var s_len = s.length;
  var t_len = t.length;

  if (s_len == 0 && t_len == 0) return 1;

  var match_distance = Math.max(s_len, t_len) / 2 - 1;

  var s_matches = new Array(s_len);
  var t_matches = new Array(t_len);

  var matches = 0;
  var transpositions = 0;

  for (var i = 0; i < s_len; i++) {
    var start = Math.max(0, i - match_distance);
    var end = Math.min(i + match_distance + 1, t_len);

    for (var j = start; j < end; j++) {
      if (t_matches[j]) continue;
      if (s.charAt(i) != t.charAt(j)) continue;
      s_matches[i] = true;
      t_matches[j] = true;
      matches++;
      break;
    }
  }

  if (matches == 0) return 0;

  var k = 0;
  for (var i = 0; i < s_len; i++) {
    if (!s_matches[i]) continue;
    while (!t_matches[k]) k++;
    if (s.charAt(i) != t.charAt(k)) transpositions++;
    k++;
  }

  return ((matches / s_len) +
    (matches / t_len) +
    ((matches - transpositions / 2.0) / matches)) / 3.0;
}
```
