---
id: 5a23c84252665b21eecc7eca
title: Brojevi Kaprekara
challengeType: 1
forumTopicId: 302296
dashedName: kaprekar-numbers
---

# --description--

Puni pozitivan broj je Kaprekarov broj ako:

<ul>
  <li>Jedan je, ili,</li>
  <li>Decimalna reprezentacija njegovog kvadrata može se podeliti na dva dela koja sadrže pozitivne cele brojeve čija je suma prirodan broj. </li>
</ul>

Zapamtite da podjelu koja rezultira frakcijom sa samo nultom nije prihvatljiva, jer se nula ne smatra pozitivnom.

Primer brojeva Kaprekara:

<ul>
  <li><code>2223</code> je Kaprekarov broj, jer <code>2223 * 2223 = 4941729</code>, <code>4941729</code> može se podeliti na <code>494</code> i <code>1729</code>, a <code>494 + 1729 = 2223</code></li>
  <li>Niz Kaprekarovih brojeva poznat je kao A006886 i počinje sa <code>1, 9, 45, 55, ...</code></li>
</ul>

# --instructions--

Napišite funkciju koja prima broj $n$, bazu $bs$, i vraća `true` ako je broj Kaprekarov broj za zadatu bazu. Inače, funkcija treba da vrati `false`.

# --hints--

`isKaprekar` trebalo bi da bude funkcija.

```js
assert(typeof isKaprekar == 'function');
```

`isKaprekar(1, 10)` trebalo bi da vrati vrednost tačno ili netačno.

```js
assert(typeof isKaprekar(1, 10) == 'boolean');
```

`isKaprekar(1, 10)` treba vratiti `true`.

```js
assert.equal(isKaprekar(1, 10), true);
```

`isKaprekar(9, 10)` treba da vrati `true`.

```js
assert.equal(isKaprekar(9, 10), true);
```

`isKaprekar(2223, 10)` treba da vrati `true`.

```js
assert.equal(isKaprekar(2223, 10), true);
```

`isKaprekar(22823, 10)` trebalo bi da vrati `false`.

```js
assert.equal(isKaprekar(22823, 10), false);
```

`isKaprekar(9, 17)` inapaswa kurudisha `false`.

```js
assert.equal(isKaprekar(9, 17), false);
```

`isKaprekar(225, 17)` trebalo bi da vrati `true`.

```js
assert.equal(isKaprekar(225, 17), true);
```

`isKaprekar(999, 17)` trebalo bi da vrati `false`.

```js
assert.equal(isKaprekar(999, 17), false);
```

# --seed--

## --seed-contents--

```js
function isKaprekar(n, bs) {

}
```

# --solutions--

```js
function isKaprekar(n, bs) {
  if (n < 1) return false;
  if (n == 1) return true;
  for (var a = n * n, b = 0, s = 1; a; s *= bs) {
    b += (a % bs) * s;
    a = Math.floor(a / bs);
    if (b && a + b == n) return true;
  }
  return false;
}
```
