---
id: 5a23c84252665b21eecc8017
title: Soundex
challengeType: 1
forumTopicId: 302320
dashedName: soundex
---

# --description--

**Soundeks algoritam** se odnosi na *zvuk* reči. Kreira reprezentaciju reči koje zvuče slično.

Koristi se za pretraživanje <em>imena</em> i <em>adrese</em>. To znači da osoba koja popunjava <em>ime</em> može obratiti pažnju na kako zvuči umesto ispravljanja pravopisa za <em>imena</em>.

Na primer:

Ako čujete ime `Quenci` prvi put, i ako ga pogrešno izgovorite, dobićete kod **Soundex** `Q520`.

Kada pravilno ispišete ime `Quincy` još jednom, i dalje ćete dobiti isti kod `Q520`, što znači da može povezati različite izgovore jednog imena za jednu osobu bez potrebe dodavanja svakog pravopisa.

Ovo su pravila:

<ul>
<li>Ako samoglasnik (A, E, I, O, U) razdvaja dva suglasnika sa istim Soundex kodom, kodira se suglasnik sa desne strane samoglasnika. Tymczak dobija kod T-522 (T, 5 za M, 2 za C, Z se zanemaruje (pogledaj pravilo „Jedan pored drugog“ iznad), 2 za K). Pošto samoglasnik „A“ razdvaja Z i K, K dobija kod.</li>
<li>Ako slova „H“ ili „W“ razdvajaju dva suglasnika sa istim Soundex kodom, suglasnik sa desne strane se ne kodira. Na primer: Ashcraft dobija kod A-261 (A, 2 za S, C se zanemaruje, 6 za R, 1 za F), a ne A-226.</li>
</ul>

# --instructions--

Napišite funkciju koja prima niz znakova kao argument i vraća taj niz znakova konvertovan u kod/šifru.

# --hints--

`soundex` Trebalo bi da bude funkcija.

```js
assert(typeof soundex == 'function');
```

`soundex("Soundex")` treba da vrati string.

```js
assert(typeof soundex('Soundex') == 'string');
```

`soundex("Soundex")` inapaswa kurudisha `"S532"`.

```js
assert.equal(soundex('Soundex'), 'S532');
```

`soundex("Example")` treba da vrati `"E251"`.

```js
assert.equal(soundex('Example'), 'E251');
```

`soundex("Sownteks")` trebalo bi da vrati `"S532"`.

```js
assert.equal(soundex('Sownteks'), 'S532');
```

`soundex("Ekzampul")` trebalo bi vratiti `"E251"`.

```js
assert.equal(soundex('Ekzampul'), 'E251');
```

`soundex("Euler")` treba da vrati `"E460"`.

```js
assert.equal(soundex('Euler'), 'E460');
```

`soundex("Gauss")` treba da vrati `"G200"`.

```js
assert.equal(soundex('Gauss'), 'G200');
```

`soundex("Hilbert")` treba da vrati `"H416"`.

```js
assert.equal(soundex('Hilbert'), 'H416');
```

`soundex("Knuth")` treba da vrati `"K530"`.

```js
assert.equal(soundex('Knuth'), 'K530');
```

`soundex("Lloyd")` treba da vrati `"L300"`.

```js
assert.equal(soundex('Lloyd'), 'L300');
```

`soundex("Lukasiewicz")` treba da vrati `"L222"`.

```js
assert.equal(soundex('Lukasiewicz'), 'L222');
```

# --seed--

## --seed-contents--

```js
function soundex(s) {

}
```

# --solutions--

```js
function soundex(s) {
  var a = s.toLowerCase().split('');
  var f = a.shift(),
    r = '',
    codes = {
      a: '',
      e: '',
      i: '',
      o: '',
      u: '',
      b: 1,
      f: 1,
      p: 1,
      v: 1,
      c: 2,
      g: 2,
      j: 2,
      k: 2,
      q: 2,
      s: 2,
      x: 2,
      z: 2,
      d: 3,
      t: 3,
      l: 4,
      m: 5,
      n: 5,
      r: 6
    };
  r =
    f +
    a
      .map(function(v, i, a) {
        return codes[v];
      })
      .filter(function(v, i, a) {
        return i === 0 ? v !== codes[f] : v !== a[i - 1];
      })
      .join('');

  return (r + '000').slice(0, 4).toUpperCase();
}
```
