---
id: 5a23c84252665b21eecc803f
title: Izračunaj cifre kompletnog broja
challengeType: 1
forumTopicId: 302331
dashedName: sum-digits-of-an-integer
---

# --description--

Napiši funkciju koja prima niz slova kao parametar. Ovaj niz slova predstavlja broj koji može biti u bilo kojoj bazi (manje od 37) i vrati zbir njegovih cifara.

<ul>
  <li><b>1</b><sub>10</sub> Ukupoluja kuwa <b>1</b></li>
  <li><b>1234</b><sub>10</sub> Ukupoluja kuwa <b>10</b></li>
  <li><b>fe</b><sub>16</sub> Ukupoluja kuwa <b>29</b></li>
  <li><b>f0e</b><sub>16</sub> Ukupoluja kuwa <b>29</b></li>
</ul>

# --hints--

`sumDigits` Trebalo bi da bude funkcija.

```js
assert(typeof sumDigits == 'function');
```

`sumDigits("1")` Trebalo bi da vrati broj.

```js
assert(typeof sumDigits('1') == 'number');
```

`sumDigits("1")` treba da vrati `1`.

```js
assert.equal(sumDigits('1'), 1);
```

`sumDigits("12345")` mora vratiti `15`.

```js
assert.equal(sumDigits('12345'), 15);
```

`sumDigits("254")` trebalo bi da vrati `11`.

```js
assert.equal(sumDigits('254'), 11);
```

`sumDigits("fe")` trebalo bi da vrati `29`.

```js
assert.equal(sumDigits('fe'), 29);
```

`sumDigits("f0e")` treba da vrati `29`.

```js
assert.equal(sumDigits('f0e'), 29);
```

`sumDigits("999ABCXYZ")` treba da vrati `162`.

```js
assert.equal(sumDigits('999ABCXYZ'), 162);
```

# --seed--

## --seed-contents--

```js
function sumDigits(n) {

}
```

# --solutions--

```js
function sumDigits(n) {
  n += '';
  for (var s = 0, i = 0, e = n.length; i < e; i += 1)
    s += parseInt(n.charAt(i), 36);
  return s;
}
```
