---
id: 5a23c84252665b21eecc7eb0
title: I kabl za E, osim nakon C
challengeType: 1
forumTopicId: 302288
dashedName: i-before-e-except-after-c
---

# --description--

**<a href="http://www.rosettacode.org/wiki/I_before_E_except_after_C" target="_blank" rel="noopener noreferrer nofollow">„I pre E, osim posle C“</a>** predstavlja opšte pravilo engleskog pravopisa. Ako nisi siguran da li se reč piše sa `ei` ili `ie`, pravilo kaže da je ispravno `ie`, osim kada se ispred nalazi slovo `c`, kada može biti `ei`.

Koristeći date reči, proverite da li oba manja dela fraze mogu biti istinita sama po sebi:

<ol>
  <li>
    <i>"I kabla ya E wakati haikutanguliwi na C".</i>
  </li>
  <li>
    <i>"E kabla ya I wakati ikitanguliwa na C".</i>
  </li>
</ol>

Ako su oba manja predikata tačna, onda se može reći da je početna teza istinita.

# --instructions--

Napiši funkciju koja prihvata reč i proverava da li ona prati ovo pravilo. Funkcija treba da vrati `true` ako ga poštuje, odnosno `false` ako ga ne poštuje.

# --hints--

`IBeforeExceptC` treba da bude funkcija.

```js
assert(typeof IBeforeExceptC == 'function');
```

`IBeforeExceptC("receive")` treba da vrati `true` ili `false`.

```js
assert(typeof IBeforeExceptC('receive') == 'boolean');
```

`IBeforeExceptC("receive")` inapaswa kurudisha `true`.

```js
assert.equal(IBeforeExceptC('receive'), true);
```

`IBeforeExceptC("science")` inapaswa kurudisha `false`.

```js
assert.equal(IBeforeExceptC('science'), false);
```

`IBeforeExceptC("imperceivable")` inapaswa kurudisha `true`.

```js
assert.equal(IBeforeExceptC('imperceivable'), true);
```

`IBeforeExceptC("inconceivable")` inapaswa kurudisha `true`.

```js
assert.equal(IBeforeExceptC('inconceivable'), true);
```

`IBeforeExceptC("insufficient")` treba vratiti `false`.

```js
assert.equal(IBeforeExceptC('insufficient'), false);
```

`IBeforeExceptC("omniscient")` trebalo bi da vrati `false`.

```js
assert.equal(IBeforeExceptC('omniscient'), false);
```

# --seed--

## --seed-contents--

```js
function IBeforeExceptC(word) {

}
```

# --solutions--

```js
function IBeforeExceptC(word)
{
    if(word.indexOf("c")==-1 && word.indexOf("ie")!=-1)
        return true;
    else if(word.indexOf("cei")!=-1)
        return true;
    return false;
}
```
