---
id: 5a23c84252665b21eecc7e82
title: Najveći zajednički delilac
challengeType: 1
forumTopicId: 302277
dashedName: greatest-common-divisor
---

# --description--

Napiši funkciju koja vraća najveći zajednički delilac za dva celobrojka.

# --hints--

`gcd` Treba da bude funkcija.

```js
assert(typeof gcd == 'function');
```

`gcd(24,36)` Trebalo bi da vrati broj.

```js
assert(typeof gcd(24, 36) == 'number');
```

`gcd(24,36)` treba vratiti `12`.

```js
assert.equal(gcd(24, 36), 12);
```

`gcd(30,48)` treba vratiti `6`.

```js
assert.equal(gcd(30, 48), 6);
```

`gcd(10,15)` trebalo bi da vrati `5`.

```js
assert.equal(gcd(10, 15), 5);
```

`gcd(100,25)` treba da vrati `25`.

```js
assert.equal(gcd(100, 25), 25);
```

`gcd(13,250)` trebalo bi vratiti `1`.

```js
assert.equal(gcd(13, 250), 1);
```

`gcd(1300,250)` trebalo bi da vrati `50`.

```js
assert.equal(gcd(1300, 250), 50);
```

# --seed--

## --seed-contents--

```js
function gcd(a, b) {

}
```

# --solutions--

```js
function gcd(a, b) {
  return b==0 ? Math.abs(a):gcd(b, a % b);
}
```
