---
id: 5a23c84252665b21eecc8013
title: Algoritmi za sortiranje niza
challengeType: 1
forumTopicId: 302319
dashedName: sorting-algorithmsstrand-sort
---

# --description--

**Strand sort** formira sortirane podnizove koji se zatim spajaju u konačan rezultat.

Zamislite rednu grupu podataka `unsortedArray = [3, 1, 4, 2]`. Izaberite prvi element `3` i kopirajte ga u drugu rednu grupu podataka. Pronađite bilo koji veći element koji sledi ovaj element. Kada pronađete veći element, u ovom slučaju `4`, kopirajte ga u drugu rednu grupu podataka, `[3, 4]`, i uporedite sledeće elemente sa ovom novom vrednošću, `4`.

Kada dođeš do kraja niza, izdvoji elemente koje si prikupio, `[3, 4]`, i počni ponovo od prvog preostalog elementa u `unsortedArray`, u ovom slučaju `1`.

Prateći ovaj proces, rezultira dve strukturirane linije podataka, `[3, 4]` i `[1, 2]`. Povežite ove dve linije da biste stvorili `strandSortedArray`.

```js
const unsortedArray = [3, 1, 4, 2];
const strandsortedArray = [1, 2, 3, 4];
```

Napišite funkciju za sortiranje niza podataka koristeći **Strand sort**. Funkcija treba da vrati sortirani niz podataka.


# --hints--

`strandSort` trebalo bi da bude funkcija.

```js
assert(typeof strandSort == 'function');
```

`strandSort([25, 32, 12, 7, 20])` Trebalo bi da vrati red podataka.

```js
assert(Array.isArray(strandSort([25, 32, 12, 7, 20])));
```

`strandSort([25, 32, 12, 7, 20])` trebalo bi da vrati `[7, 12, 20, 25, 32]`.

```js
assert.deepEqual(strandSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`strandSort([38, 45, 35, 8, 13])` trebalo bi da vrati `[8, 13, 35, 38, 45]`.

```js
assert.deepEqual(strandSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`strandSort([43, 36, 20, 34, 24])` inapaswa kurudisha `[20, 24, 34, 36, 43]`.

```js
assert.deepEqual(strandSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`strandSort([12, 33, 26, 18, 1, 16, 38])` trebalo bi da vrati `[1, 12, 16, 18, 26, 33, 38]`.

```js
assert.deepEqual(strandSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`strandSort([3, 39, 48, 16, 1, 4, 29])` treba da vrati `[1, 3, 4, 16, 29, 39, 48]`.

```js
assert.deepEqual(strandSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function strandSort(list) {

}
```

# --solutions--

```js
function strandSort(list) {
  function merge(left, right) {
    var result = [];
    while (left.length != 0 && right.length != 0) {
      if (left[0] <= right[0]) result.push(left.shift());
      else result.push(right.shift());
    }
    result.push.apply(result, left);
    result.push.apply(result, right);
    return result;
  }

  if (list.length <= 1) return list;
  var result = [];
  while (list.length > 0) {
    var sorted = [];
    sorted.push(list.shift());
    var len = list.length;
    for (var i = 1; i < len; i++) {
      var elem = list[i];
      if (sorted[i - 1] <= elem) {
        sorted.push(elem);
        sorted.splice(i, 1);
      }
    }

    result = merge(result, sorted);
  }
  return result;
}
```
