---
id: 5a23c84252665b21eecc8014
title: Utulivu wa kupanga
challengeType: 1
forumTopicId: 302308
dashedName: sort-stability
---

# --description--

Kada sortiraš zapise u tabeli po određenoj koloni ili polju, stabilno sortiranje (stable sort) uvek čuva redosled zapisa koji imaju isti ključ.

Na primer, u ovoj tabeli sa državama i gradovima, stabilno sortiranje po drugoj koloni, gradu, će zadržati Birmingham iz Amerike iznad Birminghama iz Engleske. (Iako nasumično sortiranje *može*, u ovom slučaju, postavljanje Birminghama iz Amerike iznad Birminghama iz Engleske, stabilan algoritam sortiranja to garantuje).

<pre>UK  London
US  New York
US  Birmingham
UK  Birmingham
</pre>

Isto tako, stabilno sortiranje samo po prvoj koloni daće "UK London" kao prvi element i "US Birmingham" kao poslednji (zato što će redosled elemenata sa istom prvom rečju – "UK" ili "US" – ostati nepromenjen).

# --instructions--

Napišite funkciju koja prima niz podataka sa dve dimenzije kao parametar. Svaki element ima dva sastava, kao što je prikazano u gornjem primeru. Funkcija treba da sortira niz podataka kako je prethodno opisano i vrati sortirani niz.

# --hints--

`stableSort` treba da bude funkcija.

```js
assert(typeof stableSort == 'function');
```

`stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])` Trebalo bi da vrati red podataka.

```js
assert(
  Array.isArray(
    stableSort([
      ['UK', 'London'],
      ['US', 'New York'],
      ['US', 'Birmingham'],
      ['UK', 'Birmingham']
    ])
  )
);
```

`stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])` inapaswa kurudisha `[["US", "Birmingham"], ["UK", "Birmingham"], ["UK", "London"], ["US", "New York"]]`.

```js
assert.deepEqual(
  stableSort([
    ['UK', 'London'],
    ['US', 'New York'],
    ['US', 'Birmingham'],
    ['UK', 'Birmingham']
  ]),
  [
    ['US', 'Birmingham'],
    ['UK', 'Birmingham'],
    ['UK', 'London'],
    ['US', 'New York']
  ]
);
```

`stableSort([[2, 2], [1, 2], [1, 4], [1, 5]])` treba da vrati `[[2, 2], [1, 2], [1, 4], [1, 5]]`.

```js
assert.deepEqual(
  stableSort([
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]),
  [
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]
);
```

`stableSort([[11, 55], [12, 45], [11, 45], [32, 45]])` inapaswa kurudisha `[[12, 45], [11, 45], [32, 45], [11, 55]]`.

```js
assert.deepEqual(
  stableSort([
    [11, 55],
    [12, 45],
    [11, 45],
    [32, 45]
  ]),
  [
    [12, 45],
    [11, 45],
    [32, 45],
    [11, 55]
  ]
);
```

`stableSort([[10, 22], [1, 2], [1, 4], [1, 5], [10, 9]])` inapaswa kurudisha `[[1, 2], [1, 4], [1, 5], [10, 9], [10, 22]]`.

```js
assert.deepEqual(
  stableSort([
    [10, 22],
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9]
  ]),
  [
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9],
    [10, 22]
  ]
);
```

`stableSort([[55, 54], [12, 22], [31, 43], [31, 54], [10, 49]])` treba da vrati `[[12, 22], [31, 43], [10, 49], [55, 54], [31, 54]]`.

```js
assert.deepEqual(
  stableSort([
    [55, 54],
    [12, 22],
    [31, 43],
    [31, 54],
    [10, 49]
  ]),
  [
    [12, 22],
    [31, 43],
    [10, 49],
    [55, 54],
    [31, 54]
  ]
);
```

# --seed--

## --seed-contents--

```js
function stableSort(arr) {

}
```

# --solutions--

```js
function stableSort(arr) {
  arr.sort(function(a, b) {
    return a[1] < b[1] ? -1 : a[1] > b[1] ? 1 : 0;
  });
  return arr;
}
```
