---
id: 5a23c84252665b21eecc7ec5
title: Problem Josephusa
challengeType: 1
forumTopicId: 302294
dashedName: josephus-problem
---

# --description--

Josephusov problem je matematička zagonetka koja glasi: `n` zatvorenika stoji u krugu, numerisani od `0` do `n-1`.

Izvršilac trči oko tog kruga, počevši od zatvorenika broja $0$, uklanjajući svakog zatvorenika sa $k$ i ubijajući ga.

Ovaj proces se nastavlja, krug postaje sve manji i manji, dok ne ostane samo jedan zatvorenik, koji zatim biva oslobođen.

Na primer, ako ima zatvorenika $n=5$ i $k=2$, sekvenca ubijenih zatvorenika (nazovimo je „sekvenca ubijanja“) biće 1, 3, 0, i 4, a preživeli će biti #2.

Ako je za bilo koje $n, k > 0$, pronađi koji element će na kraju biti jedan od njih.

U jednom slučaju bilo je 41 zatvorenik, a svaki treći zatvorenik je bio pogubljen ($k=3$).

Među njima je bio i domišljati čovek po imenu Josephus, koji je rešio problem, zauzeo mesto koje mu je omogućilo da preživi i ostao živ da ispriča priču.

Koja je tačka bila njega?

# --instructions--

Napiši funkciju koja uzima početni broj zatvorenika i `k` kao parametre i vraća broj zatvorenika koji je najverovatnije preživeli.

# --hints--

`josephus` To bi trebalo da bude funkcija.

```js
assert(typeof josephus == 'function');
```

`josephus(30,3)` treba da vrati broj.

```js
assert(typeof josephus(30, 3) == 'number');
```

`josephus(30,3)` treba da vrati `28`.

```js
assert.equal(josephus(30, 3), 28);
```

`josephus(30,5)` treba da vrati `2`.

```js
assert.equal(josephus(30, 5), 2);
```

`josephus(20,2)` treba da vrati `8`.

```js
assert.equal(josephus(20, 2), 8);
```

`josephus(17,6)` treba da vrati `1`.

```js
assert.equal(josephus(17, 6), 1);
```

`josephus(29,4)` treba da vrati `1`.

```js
assert.equal(josephus(29, 4), 1);
```

# --seed--

## --seed-contents--

```js
function josephus(init, kill) {

}
```

# --solutions--

```js
function josephus(init, kill) {
  const arr = Array.from(Array(init).keys());
  let curr = -1
  while (arr.length > 1) {
    curr = (curr + kill) % arr.length;
    arr.splice(curr, 1);
    curr--;
  }
  return arr[0];
}
```
