---
id: 5a23c84252665b21eecc7ede
title: Godina letenja
challengeType: 1
forumTopicId: 302300
dashedName: leap-year
---

# --description--

Odredi da li je zadata godina proletnog skoka (leap year) u kalendaru Gregoria.

# --hints--

`isLeapYear` Trebalo bi da bude funkcija.

```js
assert(typeof isLeapYear == 'function');
```

`isLeapYear()` Trebalo bi da vrati vrednost tačno ili netačno.

```js
assert(typeof isLeapYear(2018) == 'boolean');
```

`isLeapYear(2018)` treba da vrati `false`.

```js
assert.equal(isLeapYear(2018), false);
```

`isLeapYear(2016)` treba da vrati `true`.

```js
assert.equal(isLeapYear(2016), true);
```

`isLeapYear(2000)` treba vratiti `true`.

```js
assert.equal(isLeapYear(2000), true);
```

`isLeapYear(1900)` treba da vrati `false`.

```js
assert.equal(isLeapYear(1900), false);
```

`isLeapYear(1996)` treba da vrati `true`.

```js
assert.equal(isLeapYear(1996), true);
```

`isLeapYear(1800)` treba da vrati `false`.

```js
assert.equal(isLeapYear(1800), false);
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

}
```

# --solutions--

```js
function isLeapYear(year) {
  return year % 100 === 0 ? year % 400 === 0 : year % 4 === 0;
}
```
