---
id: 5a23c84252665b21eecc8040
title: Saberi dvostruko od 3 i dvostruko od 5
challengeType: 1
forumTopicId: 302332
dashedName: sum-multiples-of-3-and-5
---

# --description--

Cilj je napisati funkciju koja pronalazi zbir svih pozitivnih višekratnika brojeva 3 ili 5 koji su manji od *n*.

# --hints--

`sumMults` treba da bude funkcija.

```js
assert(typeof sumMults == 'function');
```

`sumMults(10)` Trebalo bi da vrati broj.

```js
assert(typeof sumMults(10) == 'number');
```

`sumMults(10)` inapaswa kurudisha `23`.

```js
assert.equal(sumMults(10), 23);
```

`sumMults(100)` trebalo bi da vrati `2318`.

```js
assert.equal(sumMults(100), 2318);
```

`sumMults(1000)` treba da vrati `233168`.

```js
assert.equal(sumMults(1000), 233168);
```

`sumMults(10000)` inapaswa kurudisha `23331668`.

```js
assert.equal(sumMults(10000), 23331668);
```

`sumMults(100000)` inapaswa kurudisha `2333316668`.

```js
assert.equal(sumMults(100000), 2333316668);
```

# --seed--

## --seed-contents--

```js
function sumMults(n) {

}
```

# --solutions--

```js
function sumMults(n) {
  var sum = 0;
  for (var i = 1; i < n; i++) {
    if (i % 3 == 0 || i % 5 == 0) sum += i;
  }
  return sum;
}
```
