---
id: 5a23c84252665b21eecc7e78
title: FizzBuzz uopšteno
challengeType: 1
forumTopicId: 302273
dashedName: general-fizzbuzz
---

# --description--

Napiši opštu verziju od <a href="https://rosettacode.org/wiki/FizzBuzz" target="_blank" rel="noopener noreferrer nofollow">FizzBuzz</a> koja radi za bilo koju listu parametara, zajedno sa njihovim vrednostima.

Ovo je implementacija "fizzbuzz"-a gde su pravila igre data korisniku. Napravi funkciju koja ovo izvršava. Funkcija treba da primi dva parametra.

Prvi test će biti red podataka sa pravilima FizzBuzz-a. Na primer: `[ [3, "Fizz"] , [5, "Buzz"] ]`.

Ovo pokazuje da `Fizz` treba da se ispruša ako je broj često 3 i `Buzz` ako je često 5. Ako je često oba, onda se niz slova treba kombinovati u redosledu definisanom u redu podataka. U ovom slučaju, `FizzBuzz` ako je broj često 3 i 5.

Drugi kriterijum je broj koji funkcija mora vratiti sekvencu karaktera, kao što je opisano iznad.

# --hints--

`genFizzBuzz` trebalo bi da bude funkcija.

```js
assert(typeof genFizzBuzz == 'function');
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 6)` Trebalo bi da vrati sekvencu slova.

```js
assert(
  typeof genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    6
  ) == 'string'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 6)` treba da vrati `"Fizz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    6
  ),
  'Fizz'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 10)` treba da vrati `"Buzz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    10
  ),
  'Buzz'
);
```

`genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 12)` trebalo bi da vrati `"Buzz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    12
  ),
  'Buzz'
);
```

`genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 13)` treba da vrati `"13"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    13
  ),
  '13'
);
```

`genFizzBuzz([[3, "Buzz"],[5, "Fizz"]], 15)` trebalo bi da vrati `"BuzzFizz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Buzz'],
      [5, 'Fizz']
    ],
    15
  ),
  'BuzzFizz'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"]], 15)` trebalo bi da vrati `"FizzBuzz"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz']
    ],
    15
  ),
  'FizzBuzz'
);
```

`genFizzBuzz([[3, "Fizz"],[5, "Buzz"],[7, "Baxx"]], 105)` bi trebalo da vrati `"FizzBuzzBaxx"`.

```js
assert.equal(
  genFizzBuzz(
    [
      [3, 'Fizz'],
      [5, 'Buzz'],
      [7, 'Baxx']
    ],
    105
  ),
  'FizzBuzzBaxx'
);
```

# --seed--

## --seed-contents--

```js
function genFizzBuzz(rules, num) {

}
```

# --solutions--

```js
function genFizzBuzz(rules, num) {
  let res='';
  rules.forEach(function (e) {
    if(num % e[0] == 0)
      res+=e[1];
  })

  if(res==''){
    res=num.toString();
  }

  return res;
}
```
