---
id: 5a23c84252665b21eecc7eb1
title: Matriksi ya utambulisho
challengeType: 1
forumTopicId: 302290
dashedName: identity-matrix
---

# --description--

*Identitetna matrica* je kvadratna matrica dimenzija \\( n \\times n \\), gde su svi dijagonalni elementi `1`s (jedinica), a svi ostali elementi su `0`s (nula).

<ul>
  <li style='list-style: none;'>\(\displaystyle I_{n}=\begin{bmatrix} 1 &#x26; 0 &#x26; 0 \cr 0 &#x26; 1 &#x26; 0 \cr 0 &#x26; 0 &#x26; 1 \cr \end{bmatrix}\)</li>
</ul>

# --instructions--

Napišite funkciju koja prima broj `n` kao argument i vraća matricu identiteta dimenzija \\( n \\times n \\).

# --before-each--

```js
let results=[[ [ 1 ] ],
[ [ 1, 0 ], [ 0, 1 ] ],
[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ],
[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]]
```

# --hints--

`idMatrix` Trebalo bi da bude funkcija.

```js
assert(typeof idMatrix == 'function');
```

`idMatrix(1)` trebalo bi da vrati red podataka.

```js
assert(Array.isArray(idMatrix(1)));
```

`idMatrix(1)` trebalo bi da vrati `[ [ 1 ] ]`.

```js
assert.deepEqual(idMatrix(1), results[0]);
```

`idMatrix(2)` inapaswa kurudisha `[ [ 1, 0 ], [ 0, 1 ] ]`.

```js
assert.deepEqual(idMatrix(2), results[1]);
```

`idMatrix(3)` inapaswa kurudisha `[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ]`.

```js
assert.deepEqual(idMatrix(3), results[2]);
```

`idMatrix(4)` inapaswa kurudisha `[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]`.

```js
assert.deepEqual(idMatrix(4), results[3]);
```

# --seed--

## --seed-contents--

```js
function idMatrix(n) {

}
```

# --solutions--

```js
function idMatrix(n) {
    return Array.apply(null, new Array(n)).map(function (x, i, xs) {
        return xs.map(function (_, k) {
            return i === k ? 1 : 0;
        })
    });
}
```
