---
id: 5a23c84252665b21eecc8028
title: mfululizo wa stern-brocot
challengeType: 1
forumTopicId: 302324
dashedName: stern-brocot-sequence
---

# --description--

Za ovu vežbu, Stern-Brocot niz treba ga generisati algoritam sličan onom koji se koristi za generisanje <a href="https://rosettacode.org/wiki/Fibonacci_sequence" target="_blank" rel="noopener noreferrer nofollow">Fibonaccijev niz</a>.

<ol>
  <li>Wanachama wa kwanza na wa pili wa mfululizo wote ni 1:</li>
    <ul><li>1, 1</li></ul>
  <li>Anza kwa kuzingatia mwanachama wa pili wa mfululizo</li>
<li>Saberi posmatrani član niza i njegov prethodni član, (1 + 1) = 2, a zatim dodaj rezultat na kraj niza.</li>
    mfululizo:</li>
    <ul><li>1, 1, 2</li></ul>
  <li>Ongeza mwanachama unaozingatiwa wa mfululizo mwisho wa mfululizo:</li>
    <ul><li>1, 1, 2, 1</li></ul>
  <li>Zingatia mwanachama anayefuata wa mfululizo, (mwanachama wa tatu yaani 2)</li>
  <li>ENDELEA NA HATUA YA 3 </li>
    <ul>
      <li></li>
      <li> ─── Kupanua mzunguko mwingine tunapata: ───</li>
      <li></li>
    </ul>
<li>Saberi posmatrani član niza i njegov prethodni član, (2 + 1) = 3, a zatim dodaj rezultat na kraj niza.</li>
    mfululizo:</li>
    <ul><li>1, 1, 2, 1, 3</li></ul>
  <li>Ongeza mwanachama unaozingatiwa wa mfululizo mwisho wa mfululizo:</li>
    <ul><li>1, 1, 2, 1, 3, 2</li></ul>
  <li>Zingatia mwanachama anayefuata wa mfululizo, (mwanachama wa nne yaani 1)</li>
</ol>

# --instructions--

Napravite funkciju koja vraća poziciju u Stern-Brocot sekvenci gde se $ n $ susreću prvi put, a gde je sekvenca generisana metodom navedenim gore. Napomenite da ova sekvenca koristi indeks koji počinje od 1.

# --hints--

`sternBrocot` Trebalo bi da bude funkcija.

```js
assert(typeof sternBrocot == 'function');
```

`sternBrocot(2)` treba da vrati broj.

```js
assert(typeof sternBrocot(2) == 'number');
```

`sternBrocot(2)` treba da vrati `3`.

```js
assert.equal(sternBrocot(2), 3);
```

`sternBrocot(3)` inapaswa kurudisha `5`.

```js
assert.equal(sternBrocot(3), 5);
```

`sternBrocot(5)` inapaswa kurudisha `11`.

```js
assert.equal(sternBrocot(5), 11);
```

`sternBrocot(7)` trebalo bi da vrati `19`.

```js
assert.equal(sternBrocot(7), 19);
```

`sternBrocot(10)` inapaswa kurudisha `39`.

```js
assert.equal(sternBrocot(10), 39);
```

# --seed--

## --seed-contents--

```js
function sternBrocot(num) {

}
```

# --solutions--

```js
function sternBrocot(num) {
  function f(n) {
    return n < 2
      ? n
      : n & 1
      ? f(Math.floor(n / 2)) + f(Math.floor(n / 2 + 1))
      : f(Math.floor(n / 2));
  }

  function gcd(a, b) {
    return a ? (a < b ? gcd(b % a, a) : gcd(a % b, b)) : b;
  }
  var n;
  for (n = 1; f(n) != num; n++);
  return n;
}
```
