---
id: 5a7dad05be01840e1778a0d1
title: Fractran
challengeType: 1
forumTopicId: 302270
dashedName: fractran
---

# --description--

FRACTRAN je jedinstven programski jezik sa Turingovom sposobnošću, dizajniran od strane matematičara Johna Hortona Conwaya.

Program FRACTRAN je uređeni spisak pozitivnih komponenti $P = (f_1, f_2, \ldots, f_m)$, zajedno sa unosom početnog pozitivnog celog broja $n$.

Program se izvršava ažuriranjem celog broja $n$, kao što sledi:

<ul>
  <li>za prvi član, $f_i$, u nizu gde je $nf_i$ ceo broj, zameni $n$ sa $nf_i$;</li>
  <li>ponavljaj ovo pravilo dok ne postoji ni jedan član u nizu koji generiše ceo broj kada se pomnoži sa $n$, zatim stani.</li>
</ul>

Conway izdaje program osnovnih brojeva u FRACTRAN:

$\\dfrac{17}{91}$, $\\dfrac{78}{85}$, $\\dfrac{19}{51}$, $\\dfrac{23}{38}$, $\\dfrac{29}{33}$, $\\dfrac{77}{29}$, $\\dfrac{95}{23}$, $\\dfrac{77}{19}$, $\\dfrac{1}{17}$, $\\dfrac{11}{13}$, $\\dfrac{13}{11}$, $\\dfrac{15}{14}$, $\\dfrac{15}{2}$, $\\dfrac{55}{1}$

Počevši od $n=2$, ovaj FRACTRAN program će promeniti $n$ u $15=2\\times (\\frac{15}{2})$, zatim $825=15\\times (\\frac{55}{1})$, generišući sledeću sekvencu celih brojeva:

$2$, $15$, $825$, $725$, $1925$, $2275$, $425$, $390$, $330$, $290$, $770$, $\\ldots$

Nakon 2, ova sekvenca ima sledeće potencije od 2:

$2^2=4$, $2^3=8$, $2^5=32$, $2^7=128$, $2^{11}=2048$, $2^{13}=8192$, $2^{17}=131072$, $2^{19}=524288$, $\\ldots$

koji su stepeni brojača 2.

# --instructions--

Napišite funkciju koja prima string 'fractran' kao parametar sekvence karaktera i vraća prvih 10 karaktera stringa kao niz podataka. Ako rezultat nema 10 karaktera, vratite te karaktere kako jesu.

# --hints--

`fractran` Trebalo bi da bude funkcija.

```js
assert(typeof fractran == 'function');
```

`fractran("3/2, 1/3")` Trebalo bi da vrati red podataka.

```js
assert(Array.isArray(fractran('3/2, 1/3')));
```

`fractran("3/2, 1/3")` trebalo bi da vrati `[ 2, 3, 1 ]`.

```js
assert.deepEqual(fractran('3/2, 1/3'), [2, 3, 1]);
```

`fractran("3/2, 5/3, 1/5")` treba da vrati `[ 2, 3, 5, 1 ]`.

```js
assert.deepEqual(fractran('3/2, 5/3, 1/5'), [2, 3, 5, 1]);
```

`fractran("3/2, 6/3")` treba da vrati `[ 2, 3, 6, 9, 18, 27, 54, 81, 162, 243 ]`.

```js
assert.deepEqual(fractran('3/2, 6/3'), [2, 3, 6, 9, 18, 27, 54, 81, 162, 243]);
```

`fractran("2/7, 7/2")` trebalo bi vratiti `[ 2, 7, 2, 7, 2, 7, 2, 7, 2, 7 ]`.

```js
assert.deepEqual(fractran('2/7, 7/2'), [2, 7, 2, 7, 2, 7, 2, 7, 2, 7]);
```

`fractran("17/91, 78/85, 19/51, 23/38, 29/33, 77/29, 95/23, 77/19, 1/17, 11/13, 13/11, 15/14, 15/2, 55/1")` treba da vrati `[ 2, 15, 825, 725, 1925, 2275, 425, 390, 330, 290 ]`.

```js
assert.deepEqual(
  fractran(
    '17/91, 78/85, 19/51, 23/38, 29/33, 77/29, 95/23, 77/19, 1/17, 11/13, 13/11, 15/14, 15/2, 55/1'
  ),
  [2, 15, 825, 725, 1925, 2275, 425, 390, 330, 290]
);
```

# --seed--

## --seed-contents--

```js
function fractran(progStr) {

}
```

# --solutions--

```js
function fractran(progStr){
  var num = new Array();
  var den = new Array();
  var val ;
  var out="";
  function compile(prog){
    var regex = /\s*(\d*)\s*\/\s*(\d*)\s*(.*)/m;
    while(regex.test(prog)){
      num.push(regex.exec(prog)[1]);
      den.push(regex.exec(prog)[2]);
      prog = regex.exec(prog)[3];
    }
  }

  function step(val){
    var i=0;
    while(i<den.length && val%den[i] != 0) i++;
    return num[i]*val/den[i];
  }

  var seq=[]

  function exec(val){
    var i = 0;
    while(val && i<limit){
      seq.push(val)
      val = step(val);
      i ++;
    }
  }

  // Main
  compile(progStr);
  var limit = 10;
  exec(2);
  return seq;
}
```
