---
id: 5e4ce2f5ac708cc68c1df261
title: linearno kongruentni generator sekvenci
challengeType: 1
forumTopicId: 385266
dashedName: linear-congruential-generator
---

# --description--

Generator sekvenca linearne kongruentne funkcije (LCG) je <em>algoritam</em> koji generiše niz pseudo-nasumičnih brojeva korišćenjem izraza neopadajućeg linearnog segmenta. Svi generatori linearne kongruentne funkcije koriste ovu formulu:

$$r_{n + 1} = (a \times r_n + c) \bmod m$$

Gde:

<ul>
<li>$ r_0 $ je seme.</li>
<li>$r_1$, $r_2$, $r_3$, ..., su nasumični brojevi sreće.</li>
<li>$a$, $c$, $m$ su konstante.</li>
</ul>

Ako neko pažljivo odabere vrednosti $a$, $c$ i $m$, onda će generator proizvesti ravnomernu distribuciju celih brojeva od $0 do $m - 1$.

Brojevi <abbr title="linear congruential generator">LCG</abbr> imaju lošu kvalitetu. $r_n$ i $r\_{n + 1}$ nisu proizvoljni, kao što bi to bili pravi nasumi brojevi. Svako ko zna $r_n$ može predvideti $r\_{n + 1}$, stoga <abbr title="linear congruential generator">LCG</abbr> nije siguran za enkripciju. Međutim, <abbr title="linear congruential generator">LCG</abbr> je i dalje dovoljno dobar za jednostavne zadatke kao što su test primijalnosti Millera-Rabina ili ciklusi u FreeCellu. Među prednostima <abbr title="linear congruential generator">LCG</abbr>, osoba može lako regenerisati sekvencu brojeva, od iste $r_0$. Takođe, osoba može regenerisati tu sekvencu u različitom programskom jeziku, jer je formula veoma jednostavna.

# --instructions--

Napiši funkciju koja prima $r_0,a,c,m,n$ kao parametre i vraća $r_n$.

# --hints--

`linearCongGenerator` trebalo bi da bude funkcija.

```js
assert(typeof linearCongGenerator == 'function');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` Trebalo bi da vrati broj.

```js
assert(typeof linearCongGenerator(324, 1145, 177, 2148, 3) == 'number');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` treba da vrati `855`.

```js
assert.equal(linearCongGenerator(324, 1145, 177, 2148, 3), 855);
```

`linearCongGenerator(234, 11245, 145, 83648, 4)` trebalo bi da vrati `1110`.

```js
assert.equal(linearCongGenerator(234, 11245, 145, 83648, 4), 1110);
```

`linearCongGenerator(85, 11, 1234, 214748, 5)` treba da vrati `62217`.

```js
assert.equal(linearCongGenerator(85, 11, 1234, 214748, 5), 62217);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 1)` treba da vrati `12345`.

```js
assert.equal(linearCongGenerator(0, 1103515245, 12345, 2147483648, 1), 12345);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 2)` trebalo bi da vrati `1406932606`.

```js
assert.equal(
  linearCongGenerator(0, 1103515245, 12345, 2147483648, 2),
  1406932606
);
```

# --seed--

## --seed-contents--

```js
function linearCongGenerator(r0, a, c, m, n) {

}
```

# --solutions--

```js
function linearCongGenerator(r0, a, c, m, n) {
    for (let i = 0; i < n; i++) {
        r0 = (a * r0 + c) % m;
    }
    return r0;
}
```
