---
id: 5e4ce2eaac708cc68c1df260
title: Levenshtein distanca
challengeType: 1
forumTopicId: 385264
dashedName: levenshtein-distance
---

# --description--

U teoriji informacija i računarstvu, **Levenštajnovo rastojanje** je mera razlike između dva niza (odnosno rastojanje uređivanja). Levenštajnovo rastojanje između dva niza znakova definiše se kao najmanji broj izmena potrebnih da se jedan niz pretvori u drugi, pri čemu su dozvoljene operacije umetanje, brisanje ili zamena jednog znaka.

Primer:

Levenštajnovo rastojanje između "**kitten**" i "**sitting**" iznosi 3, jer su potrebne tri izmene da bi se jedan niz pretvorio u drugi, a to nije moguće uraditi sa manje od tri izmene.

<ul>
  <li><strong>k</strong>itten   <strong>s</strong>itten    (kubadilisha 'k' na 's')</li>
  <li>sitt<strong>e</strong>n   sitt<strong>i</strong>n    (kubadilisha 'e' na 'i')</li>
  <li>sittin   sittin<strong>g</strong>    (kuongeza 'g' mwishoni).</li>
</ul>

*Umbali wa Levenshtein kati ya "**rosettacode**", "**raisethysword**" ni **8**.*

*Rastojanje između dva niza ostaje isto čak i kada se oba niza obrnu.*

# --instructions--

Napiši funkciju koja vraća Levenštajnovo rastojanje između dva niza znakova prosleđena kao argumenti.

# --hints--

`levenshtein` treba da bude funkcija.

```js
assert(typeof levenshtein == 'function');
```

`levenshtein("mist", "dist")` Trebalo bi da se vrati broj.

```js
assert(typeof levenshtein('mist', 'dist') == 'number');
```

`levenshtein("mist", "dist")` inapaswa kurejesha `1`.

```js
assert.equal(levenshtein('mist', 'dist'), 1);
```

`levenshtein("tier", "tor")` inapaswa kurejesha `2`.

```js
assert.equal(levenshtein('tier', 'tor'), 2);
```

`levenshtein("kitten", "sitting")` inapaswa kurejesha `3`.

```js
assert.equal(levenshtein('kitten', 'sitting'), 3);
```

`levenshtein("stop", "tops")` trebalo bi da restoreuje `2`.

```js
assert.equal(levenshtein('stop', 'tops'), 2);
```

`levenshtein("rosettacode", "raisethysword")` bi trebalo da vrati `8`.

```js
assert.equal(levenshtein('rosettacode', 'raisethysword'), 8);
```

`levenshtein("mississippi", "swiss miss")` treba da restaurira `8`.

```js
assert.equal(levenshtein('mississippi', 'swiss miss'), 8);
```

# --seed--

## --seed-contents--

```js
function levenshtein(a, b) {

}
```

# --solutions--

```js
function levenshtein(a, b) {
  var t = [], u, i, j, m = a.length, n = b.length;
  if (!m) { return n; }
  if (!n) { return m; }
  for (j = 0; j <= n; j++) { t[j] = j; }
  for (i = 1; i <= m; i++) {
    for (u = [i], j = 1; j <= n; j++) {
      u[j] = a[i - 1] === b[j - 1] ? t[j - 1] : Math.min(t[j - 1], t[j], u[j - 1]) + 1;
    } t = u;
  } return u[n];
}
```
