---
id: 5e6dd14797f5ce267c2f19d0
title: Mfululizo wa tazama-na-sema
challengeType: 1
forumTopicId: 385277
dashedName: look-and-say-sequence
---

# --description--

Mfululizo wa tazama na sema ni mfululizo wa nambari uliofafanuliwa kwa njia ya kurudia.

Ufafanuzi wa mfululizo

<ul><li>Chukua nambari ya desimali</li>
<li><span>Pogledaj</span> broj i vizuelno grupiši uzastopne iste cifre.</li>
<li><span>Izgovori</span> broj sleva nadesno, grupu po grupu: koliko se puta cifra pojavljuje, a zatim samu cifru.</li></ul><span> Tako se dobija sledeći broj u nizu.</span>

Primer:

<ul><li>Kuanzia na nambari 1, una <span>moja</span> 1 ambayo hutengeneza 11</li>
<li>Kuanzia na 11, una <span>mbili</span> 1. Yaani: 21</li>
<li>Počevši od 21, imaš <span>jednu</span> dvojku, pa <span>jednu</span> jedinicu. To jest: (12)(11), što daje 1211.</li>
<li>Počevši od 1211, imaš <span>jednu</span> jedinicu, <span>jednu</span> dvojku, pa <span>dve</span> jedinice. To jest: (11)(12)(21), što daje 111221.</li></ul>

# --instructions--

Napiši funkciju koja prihvata nisku znakova kao argument, analizira je i vraća dobijenu nisku.

# --hints--

`lookAndSay` treba da bude funkcija.

```js
assert(typeof lookAndSay == 'function');
```

`lookAndSay("1")` treba da vrati nisku.

```js
assert(typeof lookAndSay('1') == 'string');
```

`lookAndSay("1")` treba da vrati `"11"`.

```js
assert.equal(lookAndSay('1'), '11');
```

`lookAndSay("11")` inapaswa kurudisha `"21"`.

```js
assert.equal(lookAndSay('11'), '21');
```

`lookAndSay("21")` inapaswa kurudisha `"1211"`.

```js
assert.equal(lookAndSay('21'), '1211');
```

`lookAndSay("1211")` inapaswa kurudisha `"111221"`.

```js
assert.equal(lookAndSay('1211'), '111221');
```

`lookAndSay("3542")` trebalo bi da vrati `"13151412"`.

```js
assert.equal(lookAndSay('3542'), '13151412');
```

# --seed--

## --seed-contents--

```js
function lookAndSay(str) {

}
```

# --solutions--

```js
function lookAndSay(str) {
    return str.replace(/(.)\1*/g, function(seq, p1) {
      return seq.length.toString() + p1;
    });
}
```
