---
id: 5e6dee7749a0b85a3f1fc7d5
title: Lucas-Lehmer test
challengeType: 1
forumTopicId: 385281
dashedName: lucas-lehmer-test
---

# --description--

Test Lucas-Lehmera: za nepar primbroj $p$, broj Mersena $2^p-1$ je prim brojište ako i samo ako $2^p-1$ deli $S(p-1)$ gde $S(n+1)=(S(n))^2-2$, i $S(1)=4$.

# --instructions--

Napišite funkciju koja vraća da li je dat Mersenov broj premij ili nije.

# --hints--

`lucasLehmer` Trebalo bi da bude funkcija.

```js
assert(typeof lucasLehmer == 'function');
```

`lucasLehmer(11)` Trebalo bi da vrati vrednost istina ili laž.

```js
assert(typeof lucasLehmer(11) == 'boolean');
```

`lucasLehmer(11)` treba da restaurira `false`.

```js
assert.equal(lucasLehmer(11), false);
```

`lucasLehmer(15)` trebalo bi da vrati `false`.

```js
assert.equal(lucasLehmer(15), false);
```

`lucasLehmer(13)` trebalo bi da vrati `true`.

```js
assert.equal(lucasLehmer(13), true);
```

`lucasLehmer(17)` treba da restaura `true`.

```js
assert.equal(lucasLehmer(17), true);
```

`lucasLehmer(19)` inapaswa kurejesha `true`.

```js
assert.equal(lucasLehmer(19), true);
```

`lucasLehmer(21)` trebalo bi da vrati `false`.

```js
assert.equal(lucasLehmer(21), false);
```

# --seed--

## --seed-contents--

```js
function lucasLehmer(p) {

}
```

# --solutions--

```js
function lucasLehmer(p) {
    function isPrime(p) {
        if (p == 2)
            return true;
        else if (p <= 1 || p % 2 == 0)
            return false;
        else {
            var to = Math.sqrt(p);
            for (var i = 3; i <= to; i += 2)
                if (p % i == 0)
                    return false;
            return true;
        }
    }

    function isMersennePrime(p) {
        if (p == 2)
            return true;
        else {
            var m_p = Math.pow(2, p) - 1
            var s = 4;
            for (var i = 3; i <= p; i++)
                s = (s * s - 2) % m_p
            return s == 0;
        }
    }

    return isPrime(p) && isMersennePrime(p)
}
```
