---
id: 5e6dd1278e6ca105cde40ea9
title: Kraći niz koji je sličan
challengeType: 1
forumTopicId: 385271
dashedName: longest-common-subsequence
---

# --description--

**Najduža zajednička podsekvenca** (ili **LCS**) skupova A i B je najduži skup elemenata iz A i B koji su zajednički između ta dva skupa i koji su u istoj sekvenci u svakom skupu. Na primer, sekvenca `1234` i `1224533324` ima LCS od `1234`:

<u>1234</u>

<u>12</u>245<u>3</u>332<u>4</u>

Na primer sekvence slova, razmotrite niz `thisisatest` i `testing123testing`. LCS će biti `tsitest`:

<u>t</u>hi<u>si</u>sa<u>test</u>

<u>t</u>e<u>s</u>t<u>i</u>ng123<u>test</u>ing.

Tvoj kod mora da obrađuje samo niz karaktera.

# --instructions--

Napišite funkciju koja detektuje razliku između velikih i malih slova i vraća LCS za sekvencu od dva slova. Nije potrebno prikazivati više LCS-ova.

# --hints--

`lcs` Trebalo bi da bude funkcija.

```js
assert(typeof lcs == 'function');
```

`lcs("thisisatest", "testing123testing")` Trebalo bi da ponovo kreira sekvencu slova.

```js
assert(typeof lcs('thisisatest', 'testing123testing') == 'string');
```

`lcs("thisisatest", "testing123testing")` treba da obnovi `"tsitest"`.

```js
assert.equal(lcs('thisisatest', 'testing123testing'), 'tsitest');
```

`lcs("ABCDGH", "AEDFHR")` treba da vrati `"ADH"`.

```js
assert.equal(lcs('ABCDGH', 'AEDFHR'), 'ADH');
```

`lcs("AGGTAB", "GXTXAYB")` trebalo bi da vrati `"GTAB"`.

```js
assert.equal(lcs('AGGTAB', 'GXTXAYB'), 'GTAB');
```

`lcs("BDACDB", "BDCB")` treba da ponovo učita `"BDCB"`.

```js
assert.equal(lcs('BDACDB', 'BDCB'), 'BDCB');
```

`lcs("ABAZDC", "BACBAD")` trebalo bi da restaurira `"ABAD"`.

```js
assert.equal(lcs('ABAZDC', 'BACBAD'), 'ABAD');
```

# --seed--

## --seed-contents--

```js
function lcs(a, b) {

}
```

# --solutions--

```js
function lcs(a, b) {
  var aSub = a.substring(0, a.length - 1);
  var bSub = b.substring(0, b.length - 1);

  if (a.length === 0 || b.length === 0) {
    return '';
  } else if (a.charAt(a.length - 1) === b.charAt(b.length - 1)) {
    return lcs(aSub, bSub) + a.charAt(a.length - 1);
  } else {
    var x = lcs(a, bSub);
    var y = lcs(aSub, b);
    return (x.length > y.length) ? x : y;
  }
}
```
