---
id: 5e94a54cc7b022105bf0fd2c
title: Ciklus reči
challengeType: 1
forumTopicId: 393913
dashedName: word-frequency
---

# --description--

Uzimajući u obzir kompletnu sekvencu slova i brojeva, vratite najčešće reč u fajlu (uz broj pojavljivanja) u décrošujućem redosledu.

# --instructions--

Napišite funkciju koja broji pojavljivanje svake reči i vraća najčešće reči zajedno sa njihovom frekvencijom u opadajućem redosledu.

Funkcija treba da vrati niz podataka sa dva elementa za svaki element u sledećem formatu: `[word, freq]`. `word` treba da bude mala slova reči, a `freq` je broj koji predstavlja brojač/količinu.

Funkcija treba da vrati prazan niz podataka, ukoliko nije dostavljena sekvenca karaktera.

Funkcija ne bi trebalo da razlikuje razlike između velikih i malih slova; na primer, niz karaktera "Hello" i "hello" treba da se tretiraju isto.

Možete prepoznati reči sa specijalnim karakterima kao što su donje linije, diakritici, navodnici, zarezi, itd., kao zasebne reči.

Na primer, uzimajući u obzir niz slova "Hello hello goodbye", vaša funkcija bi trebalo da vrati `[['hello', 2], ['goodbye', 1]]`.

# --hints--

`wordFrequency` Trebalo bi da bude funkcija.

```js
assert.isFunction(wordFrequency);
```

`wordFrequency` trebalo bi da vrati red podataka.

```js
assert.isArray(wordFrequency('test'));
```

`wordFrequency("Hello hello world", 2)` treba da vrati `[['hello', 2], ['world', 1]]`

```js
const example_1 = 'Hello hello world';
const example_1_solution = [['hello', 2], ['world', 1]];

assert.deepEqual(wordFrequency(example_1, 2), example_1_solution);
```

`wordFrequency("The quick brown fox jumped over the lazy dog", 1)` treba vratiti `[['the', 2]]`

```js
const example_2 = 'The quick brown fox jumped over the lazy dog';
const example_2_solution = [['the', 2]];

assert.deepEqual(wordFrequency(example_2, 1), example_2_solution);
```

`wordFrequency("Opensource opensource open-source open source", 1)` treba da vrati `[['opensource', 2]]`

```js
const example_3 = 'Opensource opensource open-source open source';
const example_3_solution = [['opensource', 2]];

assert.deepEqual(wordFrequency(example_3, 1), example_3_solution);
```

`wordFrequency("Apple App apply aPP aPPlE", 3)` treba da vrati `[['app', 2], ['apple', 2], ['apply', 1]]` ili `[['apple', 2], ['app', 2], ['apply', 1]]`

```js
const example_4 = 'Apple App apply aPP aPPlE';
const example_4_solution = [["app", 2],["apple", 2],["apply", 1]];

const arr = wordFrequency(example_4, 3);
assert.includeDeepMembers(arr, example_4_solution);
```

`wordFrequency("c d a d c a b d d c", 4)` treba da vrati `[['d', 4], ['c', 3], ['a', 2], ['b', 1]]`

```js
const example_5 = 'c d a d c a b d d c';
const example_5_solution = [['d', 4], ['c', 3], ['a', 2], ['b', 1]];

assert.deepEqual(wordFrequency(example_5, 4), example_5_solution);
```

`wordFrequency("", 5)` treba da vrati `[]`

```js
const example_6 = '';
const example_6_solution = [];

assert.deepEqual(wordFrequency(example_6, 5), example_6_solution);
```

# --seed--

## --seed-contents--

```js
function wordFrequency(txt, n) {

}
```

# --solutions--

```js
function wordFrequency(txt, n) {
  var words = txt.split(/\s+/);
  var wordCount = {};
  words.forEach(word => {
    if (word == '') {
      return;
    }
    const lowerWord = word.toLowerCase();
    wordCount[lowerWord] =
      lowerWord in wordCount ? wordCount[lowerWord] + 1 : 1;
  });

  var wordsArray = [];
  for (let [word, count] of Object.entries(wordCount)) {
    wordsArray.push([word, count]);
  }

  wordsArray.sort((a, b) => {
    if (a[1] !== b[1]) {
      return b[1] - a[1];
    } else if (a[0] !== b[0]) {
      return a[0] < b[0] ? -1 : 1;
    }
    return 0;
  });
  return wordsArray.slice(0, n);
}
```
