---
id: 5e9ddb06ec35240f39657419
title: FizzBuzz
challengeType: 1
forumTopicId: 385370
dashedName: fizzbuzz
---

# --description--

Napišite program koji generiše niz podataka punih brojeva od 1 do 100 (uključujući). Ali:

<ul>
<li>Za brojeve deljive sa 3 dodaj <code>"Fizz"</code> u niz umesto broja.</li>
<li>Za brojeve deljive sa 5 dodaj <code>"Buzz"</code> u niz umesto broja.</li>
<li>Za brojeve deljive i sa 3 i sa 5 dodaj <code>"FizzBuzz"</code> u niz umesto broja.</li>
</ul>

# --instructions--

Tvoj program treba da vrati niz sa rezultatima prema gore navedenim pravilima.

# --hints--

`fizzBuzz` treba da bude funkcija.

```js
assert(typeof fizzBuzz == 'function');
```

`fizzBuzz()` treba da vrati niz.

```js
assert(Array.isArray(fizzBuzz()) == true);
```

Nambari zinazogawika kwa 3 pekee zinapaswa kurudisha `"Fizz"`.

```js
assert.equal(fizzBuzz()[2], 'Fizz');
```

Nambari zinazogawika kwa 5 pekee zinapaswa kurudisha `"Buzz"`.

```js
assert.equal(fizzBuzz()[99], 'Buzz');
```

Brojevi koji su deljivi sa 3 i 5 treba da vrate `"FizzBuzz"`.

```js
assert.equal(fizzBuzz()[89], 'FizzBuzz');
```

Brojevi koji nisu deljivi ni sa 3 ni sa 5 treba da ostanu nepromenjeni.

```js
assert.equal(fizzBuzz()[12], 13);
```

# --seed--

## --seed-contents--

```js
function fizzBuzz() {

}
```

# --solutions--

```js
function fizzBuzz() {
    let res=[];
    for (let i =1; i < 101; i++) {
        if (i % 3 === 0  && i % 5 === 0) {
            res.push("FizzBuzz");
        }
        else if (i % 3 === 0) {
            res.push("Fizz");
        }
        else if (i % 5 === 0) {
            res.push("Buzz");
        } 
        else {
            res.push(i);
        }
    }
    return res;
}
```
