---
id: 5ea28156e79528a9ab248f27
title: Provera Luhna za brojeve kreditnih kartica
challengeType: 1
forumTopicId: 385284
dashedName: luhn-test-of-credit-card-numbers
---

# --description--

Luhn algoritam ga koriste neke kompanije za kreditne kartice da bi razlikovale važeće brojeve kreditnih kartica od onih koji mogu biti nasumičan izbor cifara.

Kompanije koje koriste brojeve kreditnih kartica koji su verifikovani Luhnovim testom, poseduju brojeve koji prolaze sledeći test:

<ol>
  <li> Pretvorite sekvencu cifara u tu broj.</li>
  <li> Uzmite prvu, treću, ... i svaku drugu cifru koja nije na težini u pretvorenim ciframa i saberi ih kako biste dobili delimični zbir s1</li>
  <li> Uzimanjem druge, četvrte... i svake druge cifre na težini u pretvorenim ciframa:</li>
    <ol>
      <li> Udvojte svaku cifru i saberi cifre ako je rezultat veći od devetke kako biste dobili delimične zbirove za ponderisane cifre.</li>
      <li> Saberi delimične zbirove ponderisanih cifara kako biste dobili s2.</li>
    </ol>
  <li> Ako s1 + s2 završava nulom, onda je originalni broj u formatu validnog broja kreditne kartice, kao što ga potvrđuje Luhnov test.</li>
</ol>

Na primer, ako je testni broj 49927398716:

```md
Reverse the digits:
  61789372994
Sum the odd digits:
  6 + 7 + 9 + 7 + 9 + 4 = 42 = s1
The even digits:
    1, 8, 3, 2, 9
  Two times each even digit:
    2, 16, 6, 4, 18
  Sum the digits of each multiplication:
    2, 7, 6, 4, 9
  Sum the last:
    2 + 7 + 6 + 4 + 9 = 28 = s2

s1 + s2 = 70 which ends in zero which means that 49927398716 passes the Luhn test.
```

# --instructions--

Napiši funkciju koja validira broj po Lhunovom testu. Vrati true ako je validan broj. Inače, vrati false.

# --hints--

`luhnTest` Trebalo bi da bude funkcija.

```js
assert(typeof luhnTest === 'function');
```

`luhnTest("4111111111111111")` Trebalo bi da vrati vrednost tačno ili netačno.

```js
assert(typeof luhnTest('4111111111111111') === 'boolean');
```

`luhnTest("4111111111111111")` treba da vrati `true`.

```js
assert.equal(luhnTest('4111111111111111'), true);
```

`luhnTest("4111111111111112")` trebalo bi da vrati `false`.

```js
assert.equal(luhnTest('4111111111111112'), false);
```

`luhnTest("49927398716")` treba da vrati `true`.

```js
assert.equal(luhnTest('49927398716'), true);
```

`luhnTest("49927398717")` treba da vrati `false`.

```js
assert.equal(luhnTest('49927398717'), false);
```

`luhnTest("1234567812345678")` treba vratiti `false`.

```js
assert.equal(luhnTest('1234567812345678'), false);
```

`luhnTest("1234567812345670")` inapaswa kurudisha `true`.

```js
assert.equal(luhnTest('1234567812345670'), true);
```

# --seed--

## --seed-contents--

```js
function luhnTest(str) {

}
```

# --solutions--

```js
function luhnTest(str) {
  var luhnArr = [0, 2, 4, 6, 8, 1, 3, 5, 7, 9];
  var counter = 0;
  var incNum;
  var odd = false;
  var temp = String(str).replace(/[^\d]/g, '');
  if (temp.length == 0) return false;
  for (var i = temp.length - 1; i >= 0; --i) {
    incNum = parseInt(temp.charAt(i), 10);
    counter += (odd = !odd) ? incNum : luhnArr[incNum];
  }
  return counter % 10 == 0;
}
```
