---
id: 5eaf48389ee512d4d103684b
title: Samoodveštajući brojevi
challengeType: 1
forumTopicId: 385289
dashedName: self-describing-numbers
---

# --description--

Postoji nekoliko kompletnih brojeva poznatih kao „samoodređujući brojevi“ ili „autokataloški brojevi“.

Potpuni broj se kaže da je „samoodopisivan“ ako poseduje svojstvo da, kada se pozicije cifara indeksiraju od 0 do N-1, cifra na svakoj poziciji je jednaka broju puta koliko se ta cifra pojavljuje u tom broju.

Na primer, **2020** je broj sa četiri cifre koji se opisuje sam:

<ul>
    <li> pozicija 0 ima vrednost 2 i postoji 0 dve u broju; </li>
    <li> pozicija 1 ima vrednost 0 i nema 1 u broju; </li>
    <li> pozicija 2 ima vrednost 2 i postoji 2 dve; </li>
    <li> pozicija 3 ima vrednost 0 i nema 3 u broju; </li>
</ul>

Samodeskriptivni brojevi &lt; 100,000,000 su: 1210, 2020, 21200, 3211000, 42101000.

# --instructions--

Napisi funkciju koja prima ceo pozitivan broj kao parametar. Ako je to broj koji se sam opisuje, vrati true. Inače, vrati false.

# --hints--

`isSelfDescribing` trebalo bi da bude funkcija.

```js
assert(typeof isSelfDescribing == 'function');
```

`isSelfDescribing()` Treba vratiti vrednost tačno ili netačno.

```js
assert(typeof isSelfDescribing(2020) == 'boolean');
```

`isSelfDescribing(2020)` treba da vrati `true`.

```js
assert.equal(isSelfDescribing(2020), true);
```

`isSelfDescribing(3021)` trebalo bi da vrati `false`.

```js
assert.equal(isSelfDescribing(3021), false);
```

`isSelfDescribing(3211000)` treba da vrati `true`.

```js
assert.equal(isSelfDescribing(3211000), true);
```

# --seed--

## --seed-contents--

```js
function isSelfDescribing(n) {

}
```

# --solutions--

```js
function isSelfDescribing(n) {
    let digits = String(n).split("");
    digits = digits.map(function(e) {return parseInt(e)});
    let count = digits.map((x) => {return 0})
    digits.forEach((d) =>{
        if (d >= count.length) {
            return false
        }
        count[d] += 1;
    });

     if (digits === count) {
        return true;
    }
    if (digits.length != count.length) {
        return false;
    }
    
    for (let i=0; i< digits.length; i++){
      if (digits[i] !== count[i]) {
        return false;
      }
    }
    return true;
}
```
