---
id: 5eb3e4b20aa93c437f9e9717
title: Skup stvarnih brojeva
challengeType: 1
forumTopicId: 385322
dashedName: set-of-real-numbers
---

# --description--

Svi realni brojevi čine beskonačni skup ℝ. Među njegovim podskupovima, jednostavan skup je konveksni skup, svaki predstavljen kao interval između dva realna broja *a* i *b* gde je *a* ≤ *b*. Zapravo postoji četiri slučaja značenja „između“, u zavisnosti od otvorenih ili zatvorenih granica:

<ul>
  <li>[<i>a</i>, <i>b</i>]: {<i>x</i> | <i>a</i> ≤ <i>x</i> na <i>x</i> ≤ <i>b</i> }</li>
  <li>(<i>a</i>, <i>b</i>): {<i>x</i> | <i>a</i> < <i>x</i> na <i>x</i> < <i>b</i> }</li>
  <li>[<i>a</i>, <i>b</i>): {<i>x</i> | <i>a</i> ≤ <i>x</i> na <i>x</i> < <i>b</i> }</li>
  <li>(<i>a</i>, <i>b</i>]: {<i>x</i> | <i>a</i> < <i>x</i> na <i>x</i> ≤ <i>b</i> }</li>
</ul>

Zapamti da ako je *a* = *b*, samo od ovih četiri \[*a*, *a*] će imati sadržaj.

**Zoezi**

<ul>
<li>Osmisli način predstavljanja proizvoljnog skupa realnih brojeva, pri čemu „proizvoljan“ ima značenje opisano u napomenama o implementaciji.</li>
<li>Obezbedi način izvođenja sledećih standardnih operacija nad skupovima (`x` je realan broj; `A` i `B` su skupovi):</li>
  <ul>
    <li>
      <i>x</i> ∈ <i>A</i>: amua kama <i>x</i> ni kipengele cha <i>A</i><br>
Primer: 1 pripada intervalu [1, 2), dok 2, 3, ... ne pripadaju.
    </li>
    <li>
      <i>A</i> ∪ <i>B</i>: muungano wa <i>A</i> na <i>B</i>, yaani {<i>x</i> | <i>x</i> ∈ <i>A</i> au <i>x</i> ∈ <i>B</i>}<br>
      mfano: [0, 2) ∪ (1, 3) = [0, 3); [0, 1) ∪ (2, 3] = sawa, [0, 1) ∪ (2, 3]
    </li>
    <li>
      <i>A</i> ∩ <i>B</i>: mkusanyiko wa <i>A</i> na <i>B</i>, yaani {<i>x</i> | <i>x</i> ∈ <i>A</i> na <i>x</i> ∈ <i>B</i>}<br>
      mfano: [0, 2) ∩ (1, 3) = (1, 2); [0, 1) ∩ (2, 3] = seti tupu
    </li>
    <li>
<i>A</i> - <i>B</i>: razlika skupova <i>A</i> i <i>B</i>, zapisuje se i kao <i>A</i> \ <i>B</i>, odnosno {<i>x</i> | <i>x</i> ∈ <i>A</i> i <i>x</i> ∉ <i>B</i>}<br>
      mfano: [0, 2) − (1, 3) = [0, 1]
    </li>
  </ul>
</ul>

# --instructions--

Napiši funkciju koja prihvata dva objekta, nisku i niz kao argumente. Objekti predstavljaju skupove i imaju svojstva: `low`, `high` i `rangeType`.

`rangeType` može imati vrednosti 0, 1, 2 i 3 za `CLOSED`, `BOTH_OPEN`, `LEFT_OPEN` na `RIGHT_OPEN`, respektivno. Funkcija bi trebalo da implementira skup koristeći ove podatke.

Niska predstavlja operaciju koja treba da se izvrši nad skupovima. Može biti: `"union"`, `"intersect"` ili `"subtract"` (razlika).

Nakon izvršavanja operacije, funkcija bi trebalo da proveri da li vrednosti u nizu podataka postoje u dobijenom skupu i da sačuva tačno ili netačno (boolean) vrednost u niz podataka. Funkcija bi trebalo da vrati ovaj niz tačnih ili netačnih vrednosti.

# --hints--

`realSet` Trebalo bi da bude funkcija.

```js
assert(typeof realSet == 'function');
```

`realSet({"low":0, "high":1, "rangeType":2}, {"low":0, "high":2, "rangeType":3}, "union", [1, 2, 3])` treba da vrati niz.

```js
assert(
  Array.isArray(
    realSet(
      { low: 0, high: 1, rangeType: 2 },
      { low: 0, high: 2, rangeType: 3 },
      'union',
      [1, 2, 3]
    )
  )
);
```

`realSet({"low":0, "high":1, "rangeType":2}, {"low":0, "high":2, "rangeType":3}, "union", [1, 2, 3])` treba da vrati `[true, false, false]`.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 1, rangeType: 2 },
    { low: 0, high: 2, rangeType: 3 },
    'union',
    [1, 2, 3]
  ),
  [true, false, false]
);
```

`realSet({"low":0, "high":2, "rangeType":3}, {"low":1, "high":2, "rangeType":2}, "intersect", [0, 1, 2])` trebalo bi da vrati `[false, false, false]`.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 2, rangeType: 3 },
    { low: 1, high: 2, rangeType: 2 },
    'intersect',
    [0, 1, 2]
  ),
  [false, false, false]
);
```

`realSet({"low":0, "high":3, "rangeType":3}, {"low":0, "high":1, "rangeType":1}, "subtract", [0, 1, 2])` treba vratiti `[true, true, true]`.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 3, rangeType: 3 },
    { low: 0, high: 1, rangeType: 1 },
    'subtract',
    [0, 1, 2]
  ),
  [true, true, true]
);
```

`realSet({"low":0, "high":3, "rangeType":3}, {"low":0, "high":1, "rangeType":0}, "subtract", [0, 1, 2])` trebalo bi da vrati `[false, false, true]`.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 3, rangeType: 3 },
    { low: 0, high: 1, rangeType: 0 },
    'subtract',
    [0, 1, 2]
  ),
  [false, false, true]
);
```

`realSet({"low":0, "high":33, "rangeType":1}, {"low":30, "high":31, "rangeType":0}, "intersect", [30, 31, 32])` treba vratiti `[true, true, false]`.

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 33, rangeType: 1 },
    { low: 30, high: 31, rangeType: 0 },
    'intersect',
    [30, 31, 32]
  ),
  [true, true, false]
);
```

# --seed--

## --seed-contents--

```js
function realSet(set1, set2, operation, values) {

}
```

# --solutions--

```js
function realSet(set1, set2, operation, values) {
  const RangeType = {
    CLOSED: 0,
    BOTH_OPEN: 1,
    LEFT_OPEN: 2,
    RIGHT_OPEN: 3
  };

  function Predicate(test) {
    this.test = test;
    this.or = function(other) {
      return new Predicate(t => this.test(t) || other.test(t));
    };
    this.and = function(other) {
      return new Predicate(t => this.test(t) && other.test(t));
    };
    this.negate = function() {
      return new Predicate(t => !this.test(t));
    };
  }

  function RealSet(start, end, rangeType, predF) {
    this.low = start;
    this.high = end;

    if (predF) {
      this.predicate = new Predicate(predF);
    } else {
      this.predicate = new Predicate(d => {
        switch (rangeType) {
          case RangeType.CLOSED:
            return start <= d && d <= end;
          case RangeType.BOTH_OPEN:
            return start < d && d < end;
          case RangeType.LEFT_OPEN:
            return start < d && d <= end;
          case RangeType.RIGHT_OPEN:
            return start <= d && d < end;
        }
      });
    }

    this.contains = function(d) {
      return this.predicate.test(d);
    };

    this.union = function(other) {
      var low2 = Math.min(this.low, other.low);
      var high2 = Math.max(this.high, other.high);
      return new RealSet(low2, high2, null, d =>
        this.predicate.or(other.predicate).test(d)
      );
    };

    this.intersect = function(other) {
      var low2 = Math.min(this.low, other.low);
      var high2 = Math.max(this.high, other.high);
      return new RealSet(low2, high2, null, d =>
        this.predicate.and(other.predicate).test(d)
      );
    };

    this.subtract = function(other) {
      return new RealSet(this.low, this.high, null, d =>
        this.predicate.and(other.predicate.negate()).test(d)
      );
    };
  }
  set1 = new RealSet(set1.low, set1.high, set1.rangeType);
  set2 = new RealSet(set2.low, set2.high, set2.rangeType);
  var result = [];
  values.forEach(function(value) {
    result.push(set1[operation](set2).contains(value));
  });
  return result;
}
```
