---
id: 5eb3e4a21f462f409d656c73
title: Samoreferencijalni niz
challengeType: 1
forumTopicId: 385317
dashedName: self-referential-sequence
---

# --description--

Postoje nekoliko načina za generisanje rekurentne sekvence. Jedan od najčešće korišćenih metoda (mfululizo wa <a href="https://rosettacode.org/wiki/Look-and-say_sequence" target="_blank" rel="noopener noreferrer nofollow">Look-and-say</a>) je početi sa pozitivnim celobrojnim brojem, a zatim generisati sledeći element kombinovanjem grupi prebrojenih sličnih susednih cifara:

<pre>0, 10, 1110, 3110, 132110, 1113122110, 311311222110 ...</pre>

Generisani elementi rastu dužinom po geometrijskom skalovanju i nikada se nisu spojili.

Drug način za generisanje rekurentne sekvence je skraćivanje/koristiti prethodni element.

Izračunajte koliko je cifara svake vrste prisutno, a zatim spojite ukupan broj i cifre za svaku cifru poređane po numeričkom redosledu. Zapamtite da su prvih pet elemenata isti kao prethodna sekvenca.

<pre>0, 10, 1110, 3110, 132110, 13123110, 23124110 ...</pre>

Sortiraj brojeve od najvećeg do najmanjeg. Ne uključujte brojeve koji nisu vidljivi u prethodnom elementu.

Na osnovu vrednosti seme, niz generisan ovim putem uvek konvergira na konstantnu vrednost ili kratak ciklus. (Za naše svrhe, konvergencija znači da je element sličan prethodno viđenom elementu.) Prikazani niz, sa vrednošću seme 0, konvergira ka konstantnoj vrednosti 1433223110 nakon 11 ciklusa. Najbrža vrednost seme koja konvergirira je 22. Može da se stabilizuje nakon prvog elementa. (Sledeći element je 22, koji je već viđen.)

# --instructions--

Napišite funkciju koja prima početnu vrednost kao parametar, generiše rekurzivnu sekvencu dok ne konvergira, i vraća je kao niz podataka.

# --hints--

`selfReferential` Trebalo bi da bude funkcija.

```js
assert(typeof selfReferential === 'function');
```

`selfReferential(40)` Trebalo bi da vrati red podataka.

```js
assert(Array.isArray(selfReferential(40)));
```

`selfReferential(40)` treba da vrati `["40", "1410", "142110", "14123110", "1413124110", "2413125110", "151413224110", "152413225110", "251413324110", "152423224110", "152413423110"]`.

```js
assert.deepEqual(selfReferential(40), [
  '40',
  '1410',
  '142110',
  '14123110',
  '1413124110',
  '2413125110',
  '151413224110',
  '152413225110',
  '251413324110',
  '152423224110',
  '152413423110'
]);
```

`selfReferential(132110)` bi trebalo da vrati `["132110", "13123110", "23124110", "1413223110", "1423224110", "2413323110", "1433223110"]`.

```js
assert.deepEqual(selfReferential(132110), [
  '132110',
  '13123110',
  '23124110',
  '1413223110',
  '1423224110',
  '2413323110',
  '1433223110'
]);
```

`selfReferential(132211)` treba vratiti `["132211", "132231", "232221", "134211", "14131231", "14231241", "24132231", "14233221"]`.

```js
assert.deepEqual(selfReferential(132211), [
  '132211',
  '132231',
  '232221',
  '134211',
  '14131231',
  '14231241',
  '24132231',
  '14233221'
]);
```

`selfReferential(1413223110)` treba da vrati `["1413223110", "1423224110", "2413323110", "1433223110"]`.

```js
assert.deepEqual(selfReferential(1413223110), [
  '1413223110',
  '1423224110',
  '2413323110',
  '1433223110'
]);
```

`selfReferential(251413126110)` treba da vrati `["251413126110", "16151413225110", "16251413226110", "26151413325110", "16251423225110", "16251413424110", "16153413225110"]`.

```js
assert.deepEqual(selfReferential(251413126110), [
  '251413126110',
  '16151413225110',
  '16251413226110',
  '26151413325110',
  '16251423225110',
  '16251413424110',
  '16153413225110'
]);
```

# --seed--

## --seed-contents--

```js
function selfReferential(n) {

}
```

# --solutions--

```js
function selfReferential(n) {
  var descending,
    i,
    incr,
    j,
    max_i,
    max_len,
    max_seq,
    seq,
    sequence,
    indexOf =
      [].indexOf ||
      function(item) {
        for (var i = 0, l = this.length; i < l; i++) {
          if (i in this && this[i] === item) return i;
        }
        return -1;
      };

  sequence = function(n) {
    var c, cnt, cnts, d, digit, i, j, l, len, new_cnts, ref, s, seq;
    cnts = {};
    ref = n.toString();
    for (j = 0, len = ref.length; j < len; j++) {
      c = ref[j];
      d = parseInt(c);
      incr(cnts, d);
    }
    seq = [ref];
    while (true) {
      s = '';
      for (i = l = 9; l >= 0; i = --l) {
        if (cnts[i]) {
          s += '' + cnts[i] + i;
        }
      }
      if (indexOf.call(seq, s) >= 0) {
        break;
      }
      seq.push(s);
      new_cnts = {};
      for (digit in cnts) {
        cnt = cnts[digit];
        incr(new_cnts, cnt);
        incr(new_cnts, digit);
      }
      cnts = new_cnts;
    }
    return seq;
  };

  incr = function(h, k) {
    if (h[k] == null) {
      h[k] = 0;
    }
    return (h[k] += 1);
  };

  descending = function(n) {
    var tens;
    if (n < 10) {
      return true;
    }
    tens = n / 10;
    if (n % 10 > tens % 10) {
      return false;
    }
    return descending(tens);
  };

  return sequence(n);
}
```
