---
id: 5eb3e497b8d6d7f63c5517ea
title: Pretraga liste zapisa
challengeType: 1
forumTopicId: 385315
dashedName: search-a-list-of-records
---

# --description--

Rekordi imaju atribute koji opisuju entitet. Svaki atribut ima ime i vrednost. Na primer, osoba može imati atribut `age` sa vrednošću 25. Ključna operacija na listi rekorda je dohvaćanje rekorda sa određenom vrednošću atributa.

# --instructions--

Napiši funkciju koja prihvata nisku znakova kao argument. Funkcija treba da vrati indeks elementa u `list` čija se vrednost svojstva `name` poklapa sa prosleđenom niskom.

# --hints--

`searchCity` Trebalo bi da bude funkcija.

```js
assert(typeof searchCity === 'function');
```

`searchCity("Dar Es Salaam")` Trebalo bi da vrati broj.

```js
assert(typeof searchCity('Dar Es Salaam') === 'number');
```

`searchCity("Dar Es Salaam")` treba vratiti `6`.

```js
assert.equal(searchCity('Dar Es Salaam'), 6);
```

`searchCity("Casablanca")` inapaswa kurudisha `9`.

```js
assert.equal(searchCity('Casablanca'), 9);
```

`searchCity("Cairo")` bi trebalo da vrati `1`.

```js
assert.equal(searchCity('Cairo'), 1);
```

`searchCity("Mogadishu")` inapaswa kurudisha `4`.

```js
assert.equal(searchCity('Mogadishu'), 4);
```

`searchCity("Lagos")` inapaswa kurudisha `0`.

```js
assert.equal(searchCity('Lagos'), 0);
```

# --seed--

## --seed-contents--

```js
function searchCity(name) {

}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```

# --solutions--

```js
function searchCity(name) {
  return list.findIndex(item => item.name === name);
}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```
