---
id: 68507cadbf36aa089a84ce2d
title: Implementiraj Luhnov algoritam
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

Algoritam Luhna, poznat takođe kao algoritmija „modulo 10“ ili „mod 10“, jednostavna je formula provere sume koja se koristi za validaciju različitih identifikacionih brojeva, kao što su brojevi kreditnih kartica. Ovo su koraci za validaciju broja korišćenjem Luhnovog algoritma:

- Počevši od desne strane i uklanjanjem poslednje cifre (proveravajuća cifra), pomnožite vrednost svake ostale cifre sa dva.
- Ako je rezultat množenja cifara veći od `9`, saberi te cifre da dobiješ samo jednu cifru. Inače, možete oduzeti `9` od rezultata.
- Uzmite zbir svih cifara, uključujući proveravajuću cifru.
- Ako je zbir svih cifara deljiv sa `10`, onda je broj validan; inače nije validan.

Na primer, uzmimo broj `453914881`. Koraci verifikacije pomoću Luhnovog algoritma su:

 ```md 
 Nambari ya akaunti  4   5   3   9   1   4   8   8   1 
 Zidisha kila nyingine 4  10   3  18   1   8   8  16   1 
 Jumlisha tarakimu 4   1   3   9   1   8   8   7   1 
 ``` 

- Zatim saberi sve brojeve, `4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`.
- Pošto `42` nije višestruki od `10`, broj nije validan.

U ovoj laboratoriji, ćeš napraviti validator kreditnih kartica koristeći Luhn-ov algoritam.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Morate definisati funkciju nazvanu `verify_card_number` koja prima niz cifara (koji predstavlja broj kartice) i proverava da li je validan prema Luhnovom algoritmu.

2. Unutar funkcije za `verify_card_number`:

- Treba da obradiš sve razmake ili crtice koje mogu postojati u broju kartice koji se šalje.
- Vrati `VALID!` ako je broj kartice validan; inače, vrati `INVALID!`.

Ako utvrdite da je kompletirano, trebalo bi da vidite sledeću poruku u zavisnosti od unosa:

| Broj kartice | Poruka |
| :--- | :--- |
| 453914889 | HALAL! |
| 4111-1111-1111-1111 | HALAL! |
| 1234 5678 9012 3456 | NIJE HALAL! |


# --hints--

Trebalo bi da imaš funkciju nazvanu `verify_card_number`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')` trebao bi da vrati `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')` treba da vrati `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')` treba da vrati `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')` trebalo bi vratiti `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

`verify_card_number` treba da vrati `VALID!` kada je dato sa validnim brojem kreditne kartice.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

`verify_card_number` trebalo bi da vrati `INVALID!` kada mu se prosledi nevažeći broj kreditne kartice.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
