---
id: 61fd70336ebb3e4f62ee81ba
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Tabele za HTML koriste element `caption` da bi objasnile o čemu se tabela radi. Element `caption` bi trebalo da bude prvo dete od `table`, ali se može postaviti korišćenjem atributa CSS za `caption-side`.

Dodaj karakteristiku `caption` u svoj prvi `table`, i dodaj tekst `Assets`.

# --hints--

Vaš prvi element od `table` treba da ima element `caption`.

```js
assert.strictEqual(
  document.querySelector('table')?.children?.[0]?.localName,
  'caption'
);
```

Vaš element za `caption` treba da ima tekst sa `Assets`.

```js
assert.strictEqual(document.querySelector('caption')?.textContent.trim(), 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
