---
id: 61fd6ab779390f49148773bb
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Ispod vašeg elementa `h1`, kreirajte element `div`. Dodajte atribut elementu `id` postavljen na `years`. Ako želite da je ovaj određeni element sakriven od čitača ekrana, dodajte mu atribut elementa `aria-hidden` postavljen na `true`.

# --hints--

Trebalo bi da napraviš novi element `div` nakon tvog elementa `h1`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling?.localName,
  'div'
);
```

Tvoj novi element od `div` treba da ima svojstvo atributa `id` koje je postavljeno na `years`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('id'), 'years');
```

Vaš novi element `div` mora imati svojstvo elementa `aria-hidden` postavljeno na `true`.

```js
assert.strictEqual(
  document.querySelector('div')?.getAttribute('aria-hidden'),
  'true'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
