---
id: 61fd66c687e610436494c6f1
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Unutar tvog elementa `section`, dodaj komponentu `h1` koja sadrži konfigurisan element unutar `span`.

# --hints--

Karakteristika tvoga `section` treba da ima karakteristiku `h1`.

```js
assert.strictEqual(
  document.querySelector('section')?.children?.[0]?.localName,
  'h1'
);
```

Vaša funkcija `h1` trebalo bi da ima funkciju/element `span`.

```js
assert.strictEqual(
  document.querySelector('h1')?.children?.[0]?.localName,
  'span'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
