---
id: 61fd5a93fd62bb35968adeab
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Počnite sa svojim projektom dodeljivanjem vaše karakteristike `head`, karakteristike `link` za vaš stilski dokument.

# --hints--

Tvoj kod mora imati komponentu od `link`.

```js
assert.match(code, /<link/);
```

Tvoj element `link` treba da bude unutar tvog elementa `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Element vašeg `link` treba da ima svojstvo elementa `rel` sa vrednošću `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

Element vašeg `link` mora imati svojstvo elementa `href` sa vrednošću `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
