---
id: 61fd67a656743144844941cb
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Čitači ekrana najavljuju elemente HTML prema protoku dokumenta. Na kraju želimo da izračun stanja ima sekcijski naslov "Balance Sheet" i podnaslov "AcmeWidgetCorp". Međutim, ovaj raspored se ne može dobro razumeti ako ga čitač ekrana najavi.

Dodaj `class="flex"` atribut postojećem `span` elementu, zatim unutar njega dodaj dva `span` elementa. Prvom dodaj tekst `AcmeWidgetCorp`, a drugom tekst `Balance Sheet`. CSS će promeniti redosled prikaza, ali će HTML redosled biti važniji za čitače ekrana.

# --hints--

Vaš element `span` koji postoji mora imati svojstvo `class` sa vrednošću `flex`.

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

Tvoja komponenta sa `span` koja postoji mora imati dve nove komponente od `span` unutar sebe.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

Vaši novi elementi `span` ne bi trebalo da imaju svojstvo `class`.

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

Tvoja nova prva komponenta za `span` mora imati tekst `AcmeWidgetCorp`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

Tvoja nova druga sekcija za `span` treba da ima tekst `Balance Sheet`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
