---
id: 61fd78621573aa5e8b512f5e
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

U vašem trećem elementu sa `tr`, dodajte element/polje `th` sa tekstom `Savings Funds set aside for emergencies.`. Umotajte taj tekst, osim `Savings `, u element/polje `span` sa atributom `class` postavljeno na `description`.

Nakon toga, dodaj tri elementa od `td` sa sledećim tekstom (u redosledu): `$500`, `$650`, `$728`. Ukloni treći element od `td` atributa/svojstva koje je postavljeno na `class`, a to je `current`.

# --hints--

Vaš treći element `tr` mora imati atribut/element `th`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isNotNull(tableRow?.querySelector('th'));
```

Vaše polje `th` treba da ima tekst `Savings Funds set aside for emergencies.`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Savings Funds set aside for emergencies.'
);
```

Treba da obradiš tekst iz `Funds set aside for emergencies.` u polju `span`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'Funds set aside for emergencies.'
);
```

Tvoj element za `span` treba da ima svojstvo `class` postavljeno na `description`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('description')
);
```

Trebalo bi da ima tri elementa za `td`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

Vaš prvi element `td` bi trebalo da ima tekst od `$500`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$500');
```

Drugi vaš element za `td` mora da sadrži tekst od `$650`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$650');
```

Treći vaš komponent za `td` mora imati tekst od `$728`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$728');
```

Treći element vašeg `td` mora imati `class` postavljen na `current`.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="total"></tr>
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
