---
id: 61fd6b7c83dbf54a08cf0498
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Unutar vašeg elementa od `div`, dodajte tri elementa od `span`. Svaki treba da ima svojstvo elementa od `class` koje je postavljeno na `year`, i dodajte sledeći tekst (u redosledu): `2019`, `2020`, i `2021`.

# --hints--

Vaš element za `div` mora imati tri komponente za `span`.

```js
assert.lengthOf(document.querySelector('div')?.children, 3);
```

Svaki element `span` mora imati svojstvo elementa `class` postavljeno na `year`.

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert.isTrue(span?.classList?.contains('year')));
```

Tvoj prvi element `span` treba da sadrži tekst `2019`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[0]?.textContent.trim(),
  '2019'
);
```

Tvoj drugi element `span` treba da sadrži tekst `2020`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[1]?.textContent.trim(),
  '2020'
);
```

Tvoj treći element `span` treba da sadrži tekst `2021`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[2]?.textContent.trim(),
  '2021'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
