---
id: 61fd6cc9475a784b7776233e
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Ispod vašeg elementa `div`, dodajte novi element `div` koji ima `class` postavljen na `table-wrap`. Ovo će biti kontejner za vaše tabele.

Unutar toga, dodajte tri elementa od `table`. Koristićete CSS da ih napravite u jednu tabelu, ali korišćenje različitih tabela pomoći će čitaču ekrana da razume tok dokumenta.

# --hints--

Trebalo bi da kreirate novu komponentu za `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Vaša nova funkcija za `div` treba da ima `class` postavljeno na `table-wrap`.

```js
assert.strictEqual(document.querySelector('.table-wrap')?.localName, 'div');
```

Vaš element za `.table-wrap` treba da dođe nakon vašeg postojećeg `div`.

```js
assert.isTrue(
  document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap')
);
```

Vaša komponenta za `.table-wrap` treba da ima tri elementa od `table`.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert.lengthOf(children, 3);
children.forEach(child => assert.strictEqual(child?.localName, 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
