---
id: 61fd75ea7f663457612dba02
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Unutar svakog vašeg novog elementa sa `th`, postavite u element/komponentu `span` koji ima svojstvo elementa `class` postavljeno na `sr-only year`. Unesite sledeći tekst (u ovom redosledu): `2019`, `2020`, i `2021`.

Treći vaš element za `th` atributa elementa `class` koji je postavljen na `current`.

Ostavite polje za `td` prazno. Ovaj element služi samo za osiguranje da vaša tabela ima raspored od četiri reda i povezuje odgovarajuće zaglavlja i redove.

# --hints--

Svaki element vašeg `th` mora imati element `span`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.lengthOf(th?.children, 1);
  assert.strictEqual(th?.children?.[0]?.localName, 'span');
});
```

Svaki tvoj novi element `span` mora imati svojstvo elementa `class` postavljeno na `sr-only year`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert.isTrue(th?.children?.[0]?.classList?.contains('sr-only'));
  assert.isTrue(th?.children?.[0]?.classList?.contains('year'));
});
```

Prvi deo tvog `span` treba da sadrži tekst `2019`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent.trim(),
  '2019'
);
```

Drugi vaš element za `span` treba da ima tekst `2020`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent.trim(),
  '2020'
);
```

Treći element vašeg `span` mora imati tekst `2021`.

```js
assert.strictEqual(
  document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent.trim(),
  '2021'
);
```

Treći atribut vašeg `th` treba da ima `class` postavljen na `current`.

```js
assert.isTrue(
  document
    .querySelector('table')
    ?.querySelectorAll('th')?.[2]
    ?.classList?.contains('current')
);
```

Vaše polje `td` treba da bude prazno.

```js
const table = document.querySelector('table');
const tableData = table.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent.trim(), '');
assert.isEmpty(tableData?.children);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
