---
id: 61fd71d596e8f253b9408b39
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Komponenta `tr` koristi se za prikazivanje reda tabele. Dodajte komponentu `tr` unutar vaše komponente `thead`. U vašoj novoj komponenti `tr`, dodajte komponentu `td`, praćenu sa tri komponente `th`.

Element `td` prikazuje podatnu ćeliju, dok element `th` prikazuje zaglavnu ćeliju.

# --hints--

Funkcija vašeg `thead` bi trebalo da ima funkciju `tr`.

```js
assert.strictEqual(
  document.querySelector('thead')?.children?.[0]?.localName,
  'tr'
);
```

Element vašeg `tr` bi trebalo da ima element `td` kao prvi potomak.

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[0]?.localName,
  'td'
);
```

Vaš element za `tr` treba da ima tri komponente od `th`, nakon elementa `td`.

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[1]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[2]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[3]?.localName,
  'th'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
