---
id: a39963a4c10bc8b4d4f06d7e
title: Izvršiti funkciju za uklanjanje vrednosti
challengeType: 26
dashedName: implement-a-value-remover-function
---

# --description--

U ovom laboratoriju, kreiraćeš funkciju koja prima inicijalni niz podataka kao svoj prvi argument, praćen jedan ili više argumenata koji predstavljaju vrednosti za uklanjanje.

**Cilj**: Popunite priče korisnika ispod i prođite sve testove kako biste završili ovaj rad.

**Priče korisnika:**

1. Mora da ima funkciju `destroyer` koja prima niz podataka i jedan ili više dodatnih argumenata.
1. Funkcija `destroyer` bi trebalo da vrati novi niz koji ne sadrži sve elemente iz prvog argumenta koji odgovaraju bilo kojim naknadnim argumentima.
1. Funkcija mora da prihvati nepoznat broj argumenata.

# --hints--

Trebalo bi da imaš funkciju `destroyer`.

```js
assert.isFunction(destroyer);
```

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)` treba da vrati `[1, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)` treba vratiti `[1, 5, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)` treba vratiti `[1]`.

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)` trebalo bi vratiti `[]`.

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)` trebalo bi da vrati `["hamburger"]`.

```js
assert.deepEqual(destroyer(["tree", "hamburger", 53], "tree", 53), ["hamburger"]);
```

`destroyer( ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan" )` bi trebalo da vrati `[12, 92, 65]`.

```js
assert.deepEqual(
  destroyer(
    ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"],
    "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan"
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function destroyer(arr, ...valsToRemove) {
  return arr.filter(item => !valsToRemove.includes(item));
}
```
